Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (64783)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 762
% Passive clauses: 8045
% Generated clauses: 20204
% Final active clauses: 752
% Final passive clauses: 6840
% Input formulas: 44
% Initial clauses: 176
% 
% Duplicate literals: 619
% Trivial inequalities: 163
% Fw subsumption resolutions: 2032
% Fw demodulations: 458
% Bw demodulations: 9
% 
% Simple tautologies: 237
% Equational tautologies: 61
% Forward subsumptions: 7930
% 
% Binary resolution: 9718
% Factoring: 1405
% Forward superposition: 3466
% Backward superposition: 1917
% Self superposition: 25
% Equality factoring: 28
% Equality resolution: 38
% 
% Split clauses: 602
% Split components: 1219
% SAT solver clauses: 1080
% SAT solver unit clauses: 32
% SAT solver binary clauses: 227
% 
% Sat splits: 602
% 
% TWLsolver clauses: 1074
% TWLsolver calls for satisfiability: 467
% 
% Memory used [KB]: 8827
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (64807)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 419
% Passive clauses: 2217
% Generated clauses: 12868
% Final active clauses: 402
% Final passive clauses: 1712
% Input formulas: 44
% Initial clauses: 176
% 
% Duplicate literals: 264
% Trivial inequalities: 116
% Fw subsumption resolutions: 1056
% Fw demodulations: 735
% Bw demodulations: 96
% Condensations: 315
% Global subsumptions: 18
% 
% Simple tautologies: 330
% Equational tautologies: 24
% Forward subsumptions: 4928
% 
% Binary resolution: 2773
% Factoring: 72
% Forward superposition: 2928
% Backward superposition: 2019
% Self superposition: 55
% Equality factoring: 30
% Equality resolution: 27
% 
% Split clauses: 159
% Split components: 339
% SAT solver clauses: 3042
% SAT solver unit clauses: 190
% SAT solver binary clauses: 90
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 10
% 
% Sat splits: 159
% Sat splitting refutations: 12
% 
% TWLsolver clauses: 2970
% TWLsolver calls for satisfiability: 1737
% 
% Memory used [KB]: 12665
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f25973,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f243,f428,f242,f430,f432,f426,f241,f433,f240,f429,f425,f2252,f2254,f2614,f3918,f3922,f3973,f3977,f4075,f4076,f4170,f4873,f3920,f4171,f4874,f3975,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4343,f4344,f4972,f4887,f4974,f4891,f4976,f4894,f4978,f4897,f4980,f4900,f4982,f4903,f4984,f4906,f4986,f4909,f4988,f4912,f4990,f4915,f4992,f4918,f4994,f4921,f4996,f4924,f4998,f4927,f5000,f4930,f5002,f4933,f5004,f4936,f4515,f4960,f4958,f4516,f4963,f4961,f5078,f5079,f5081,f5083,f5084,f5085,f5089,f5091,f5092,f5093,f5929,f5930,f5936,f5934,f5931,f5932,f5939,f5933,f5941,f5938,f6012,f6013,f6019,f6017,f6014,f6015,f6022,f6016,f6024,f6021,f6113,f6114,f6131,f6130,f6371,f6372,f6379,f6376,f6373,f6374,f6382,f6375,f6384,f6377,f6381,f6466,f6467,f6479,f6477,f7096,f7097,f7098,f7099,f7102,f7105,f7106,f7320,f7321,f7322,f7323,f7326,f7329,f7330,f7334,f7335,f10112,f10114,f10117,f10120,f10122,f10124,f10127,f10130,f10133,f10136,f10237,f10238,f10240,f10243,f10248,f10250,f10253,f10264,f10266,f10268,f10082,f3974,f10269,f10265,f10272,f10263,f10275,f10278,f10281,f10284,f10285,f10290,f10291,f10296,f10297,f10254,f10280,f9915,f10300,f9908,f10301,f9907,f10304,f9904,f10305,f9903,f10306,f9902,f10307,f9901,f10303,f10308,f10339,f10340,f10341,f10342,f10343,f10344,f10346,f10347,f10348,f10349,f10453,f10455,f10457,f10461,f10463,f10465,f10468,f10469,f13539,f13540,f13534,f13541,f13543,f13544,f13522,f13548,f13521,f13549,f13520,f13550,f13519,f13551,f13518,f13552,f13570,f13571,f13575,f13576,f13583,f13584,f13585,f10273,f13586,f13587,f13589,f13591,f13595,f13597,f13599,f13601,f13603,f13605,f13607,f13609,f13611,f13613,f13615,f13617,f13621,f13623,f13625,f13627,f13629,f13631,f13633,f13635,f13637,f13639,f13641,f13643,f13645,f13647,f13671,f13673,f13697,f13699,f13711,f13723,f13753,f13808,f13841,f13896,f13997,f13999,f14001,f14004,f14006,f14010,f14012,f14019,f14021,f14026,f14028,f14046,f14048,f14050,f14053,f14055,f14063,f14065,f14070,f14072,f14086,f14088,f14090,f14099,f14101,f14106,f14108,f14146,f14148,f14150,f14152,f14154,f14159,f14161,f14163,f14165,f14170,f14172,f14174,f14176,f14198,f14200,f14203,f14208,f14210,f14212,f14229,f14233,f14237,f14239,f14349,f14363,f14381,f14395,f14397,f14414,f14427,f14444,f14457,f14459,f14461,f14463,f14465,f15845,f15846,f15847,f14514,f14516,f14518,f14520,f14522,f14524,f14526,f14528,f14532,f14534,f14536,f14538,f14540,f14542,f14544,f14546,f14603,f14605,f14607,f14609,f14611,f14613,f14615,f14617,f14619,f14621,f14623,f14625,f14644,f14663,f14685,f14690,f14692,f14701,f14703,f14711,f14713,f14715,f14717,f14739,f14744,f14746,f14755,f14757,f14765,f14767,f14769,f14771,f14782,f14791,f14802,f14811,f14813,f14835,f14837,f14851,f14853,f14875,f14877,f14891,f14893,f14919,f14921,f14935,f14937,f14952,f14971,f14986,f15005,f15027,f15049,f15078,f15080,f15082,f15117,f15119,f15121,f15123,f15136,f15138,f15140,f15151,f15153,f15164,f15166,f15177,f15179,f15190,f15192,f15216,f15240,f15260,f15262,f15264,f15276,f15278,f15280,f15282,f15322,f15324,f15326,f15328,f15330,f15332,f15357,f15360,f15385,f15388,f15407,f15409,f15427,f15429,f15431,f15444,f15456,f15468,f15470,f15482,f15484,f15505,f15529,f15555,f15579,f15591,f15593,f15595,f15597,f15599,f15601,f15613,f15615,f15617,f15619,f15621,f15623,f15625,f15629,f15640,f15642,f15648,f15650,f15656,f15658,f15661,f15663,f15669,f15671,f15677,f15679,f15681,f15684,f15686,f15706,f15716,f15725,f15732,f15743,f15745,f15747,f15754,f15756,f15763,f16009,f16011,f16012,f16013,f16014,f16015,f16150,f16151,f16557,f16558,f16560,f16561,f16562,f16564,f16565,f16566,f16567,f16568,f16569,f16570,f16571,f16572,f16573,f16574,f16575,f16576,f16582,f16583,f16584,f16585,f16588,f16589,f16590,f16591,f16593,f16594,f16596,f16597,f16602,f16603,f16605,f16606,f16607,f16609,f16611,f16613,f16615,f16616,f16617,f16618,f16619,f16620,f16621,f16624,f16625,f16626,f16631,f16632,f16633,f16639,f16640,f16641,f16642,f16645,f16646,f16647,f16648,f16650,f16651,f16653,f16657,f16658,f16660,f16661,f16663,f16664,f16675,f16677,f16681,f16683,f16685,f16689,f16692,f16694,f16697,f16703,f16709,f16711,f16714,f16716,f16727,f16731,f16734,f16738,f16740,f16755,f16757,f16772,f16774,f16778,f16786,f16790,f16792,f16798,f16801,f16822,f16824,f16826,f16828,f16834,f16839,f16842,f16844,f16846,f16848,f16852,f16855,f16856,f16857,f16858,f16859,f16860,f16861,f16862,f16863,f16864,f16865,f16866,f16867,f16868,f16869,f16870,f16871,f16872,f16873,f16874,f16875,f16876,f16877,f16878,f16879,f16880,f16881,f16882,f16883,f16884,f16885,f16886,f16887,f16888,f16889,f16890,f16891,f16896,f16897,f16898,f16913,f16920,f16929,f16946,f16953,f16954,f16960,f16971,f16974,f16978,f16980,f16982,f16984,f16991,f16993,f16994,f16995,f16996,f16997,f16998,f16999,f17003,f17005,f17007,f17008,f17009,f17010,f17011,f17012,f17013,f17014,f17015,f17016,f17017,f17018,f17019,f17020,f17021,f17022,f17023,f17024,f17025,f17026,f17027,f17028,f17029,f17030,f17031,f17032,f17033,f17034,f17035,f17036,f17037,f17038,f17039,f17040,f17041,f17042,f17043,f17044,f17045,f17046,f17047,f17048,f17049,f17050,f17051,f17052,f17053,f17054,f17055,f17056,f17057,f17058,f17059,f17060,f17061,f17062,f17063,f17064,f17065,f17066,f17067,f17068,f17069,f17070,f17071,f17072,f17073,f17074,f17075,f17076,f17077,f17078,f17079,f17080,f17081,f17082,f17083,f17084,f17085,f17086,f17087,f17088,f17089,f17090,f17091,f17092,f17093,f17094,f17095,f17096,f17097,f17098,f17099,f17100,f17101,f17102,f17103,f17104,f17105,f17106,f17107,f17108,f17109,f17110,f17111,f17112,f17113,f17114,f17115,f17116,f17117,f17118,f17119,f17120,f17121,f17122,f17123,f17124,f17125,f17126,f17127,f17128,f17129,f17130,f17131,f17132,f17133,f17134,f17135,f17136,f17137,f17138,f17139,f17140,f17141,f17142,f17143,f17144,f17145,f17146,f17147,f17148,f17149,f17150,f17151,f17152,f17153,f17154,f17155,f17156,f17157,f17158,f17159,f17160,f17161,f17162,f17163,f17164,f17165,f17166,f17167,f17168,f17169,f17170,f17171,f17172,f17173,f17174,f17175,f17176,f17177,f17178,f17179,f17180,f17181,f17182,f17183,f17184,f17185,f17186,f17187,f17188,f17189,f17190,f17191,f17192,f17193,f17194,f17195,f17196,f17197,f17198,f17199,f17200,f17201,f17202,f17203,f17204,f17205,f17206,f17207,f17208,f17209,f17210,f17211,f17212,f17213,f17214,f17215,f17216,f17217,f17218,f17219,f17220,f17221,f17222,f17223,f17224,f17225,f17226,f17227,f17228,f17229,f17230,f17231,f17232,f17233,f17234,f17235,f17236,f17237,f17238,f17239,f17240,f17241,f17242,f17243,f17244,f17245,f17246,f17247,f17248,f17249,f17250,f17251,f17252,f17253,f17254,f17255,f17256,f17257,f17258,f17259,f17260,f17261,f17262,f17263,f17264,f17265,f17266,f17267,f17268,f17269,f17270,f17271,f17272,f17273,f17274,f17275,f17276,f17277,f17278,f17279,f17280,f17281,f17282,f17283,f17284,f17285,f17286,f17287,f17288,f17289,f17290,f17291,f17292,f17293,f17294,f17295,f17296,f17297,f17298,f17299,f17300,f17301,f17302,f17303,f17304,f17305,f17306,f17307,f17308,f17309,f17310,f17311,f17312,f17313,f17314,f17315,f17316,f17317,f17318,f17319,f17320,f17321,f17322,f17323,f17324,f17325,f17326,f17327,f17328,f17329,f17330,f17331,f17332,f17333,f17334,f17335,f17336,f17337,f17338,f17339,f17340,f17341,f17342,f17343,f17344,f17345,f17346,f17347,f17348,f17349,f17350,f17351,f17389,f17391,f17398,f17401,f17403,f17405,f17412,f17413,f17415,f17416,f17417,f17423,f17424,f17425,f17426,f17427,f17430,f17431,f17432,f17433,f17438,f17439,f17440,f17441,f17442,f17443,f17444,f17445,f17446,f17447,f17448,f17449,f17450,f17451,f17452,f17453,f17454,f17455,f17456,f17457,f17458,f17459,f17460,f17461,f17462,f17463,f17464,f17465,f17466,f17467,f17468,f17469,f17470,f17471,f17472,f17473,f17474,f17475,f17480,f17489,f17497,f17499,f17501,f17503,f17504,f17505,f17506,f17511,f17512,f17520,f17525,f17526,f17528,f17530,f17532,f17533,f17534,f17535,f17536,f17537,f17543,f17545,f17547,f17549,f17552,f17555,f17556,f17557,f17558,f17559,f17560,f17561,f17562,f17577,f17585,f17586,f17587,f17592,f17601,f17602,f17609,f17610,f17611,f17612,f17613,f17614,f17615,f17616,f17617,f17620,f17621,f17622,f17623,f17624,f17625,f17626,f17627,f17628,f17629,f17630,f17631,f17632,f17633,f17634,f17635,f17636,f17637,f17638,f17639,f17644,f17645,f17646,f17647,f17648,f17649,f17650,f17651,f17652,f17653,f17654,f17655,f17656,f17657,f17658,f17659,f17660,f17661,f17662,f17663,f17664,f17665,f17666,f17667,f17668,f17669,f17670,f17671,f17672,f17673,f17677,f17681,f17690,f17691,f17693,f17694,f17695,f17696,f17697,f17698,f17699,f17700,f17704,f17708,f17709,f17717,f17718,f17719,f17721,f17722,f17723,f17724,f17725,f17726,f17727,f17728,f17729,f17730,f17731,f17732,f17733,f17734,f17735,f17736,f17737,f17739,f17741,f17742,f17743,f17744,f17745,f17746,f17747,f17752,f17753,f17754,f17755,f17756,f17757,f17758,f17759,f17760,f17761,f17762,f17763,f17764,f17765,f17766,f17767,f17768,f17769,f17770,f17771,f17772,f17773,f17774,f17775,f17776,f17777,f17778,f17779,f17780,f17781,f17782,f17783,f17784,f17785,f17786,f17787,f17788,f17789,f17790,f17791,f17792,f17793,f17794,f17795,f17796,f17797,f17798,f17799,f17800,f17801,f17802,f17803,f17804,f17805,f17806,f17807,f17808,f17809,f17810,f17811,f17812,f17813,f17814,f17815,f17816,f17817,f17818,f17819,f17820,f17821,f17822,f17823,f17824,f17825,f17826,f17827,f17828,f17829,f17830,f17831,f17832,f17833,f17834,f17835,f17836,f17837,f17838,f17839,f17840,f17841,f17842,f17843,f17844,f17845,f17846,f17847,f17848,f17849,f17850,f17851,f17852,f17853,f17854,f17855,f17856,f17857,f17858,f17859,f17860,f17861,f17862,f17863,f17864,f17865,f17866,f17867,f17868,f17869,f17870,f17871,f17872,f17873,f17874,f17875,f17876,f17877,f17878,f17879,f17880,f17881,f17882,f17883,f17884,f17885,f17886,f17887,f17888,f17889,f17890,f17891,f17892,f17893,f17894,f17895,f17896,f17897,f17898,f17899,f17900,f17901,f17902,f17903,f17904,f17905,f17906,f17907,f17908,f17909,f17910,f17911,f17912,f17913,f17914,f23929,f23931,f23932,f23933,f23934,f23935,f23936,f23937,f23938,f23939,f23940,f23941,f23942,f23943,f23944,f23945,f23946,f23947,f23948,f23949,f23950,f24087,f24096,f24098,f24100,f24102,f24104,f24105,f24107,f24109,f24111,f24113,f24114,f24115,f24261,f24262,f24264,f24267,f24276,f24278,f24297,f24299,f24300,f25836,f25813,f4888,f25919,f25956,f25972])).
fof(f25972,plain,(
  $false | ($spl5 | $spl10 | $spl60)),
  inference(subsumption_resolution,[],[f4959,f25096])).
fof(f25096,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | ($spl5 | $spl10)),
  inference(forward_demodulation,[],[f24578,f3920])).
fof(f24578,plain,(
  s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) != s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl5),
  inference(unit_resulting_resolution,[],[f3917,f434,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f104,f146,f145,f144,f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f144,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f145,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f146,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f104,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X14,X15,X16] : ((p(s(bool,X16)) <=> s(bool,X14) = s(bool,X15)) <=> ((p(s(bool,X16)) | p(s(bool,X15)) | p(s(bool,X14))) & (p(s(bool,X16)) | ~p(s(bool,X14)) | ~p(s(bool,X15))) & (p(s(bool,X15)) | ~p(s(bool,X14)) | ~p(s(bool,X16))) & (p(s(bool,X14)) | ~p(s(bool,X15)) | ~p(s(bool,X16)))))),
  file('thm.pred_set.UNIQUE_MEMBER_SING','thm.sat.sat_2|49|_prove_ltb')).
fof(f434,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f247,f244])).
fof(f244,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.UNIQUE_MEMBER_SING','thm.bool.TRUTH')).
fof(f247,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f21])).
fof(f21,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.pred_set.UNIQUE_MEMBER_SING','thm.bool.NOT_CLAUSES|split|1')).
fof(f3917,plain,(
  sP5('F_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl5),
  inference(unit_resulting_resolution,[],[f434,f2603,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f216])).
fof(f2603,plain,(
  s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) != s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl5),
  inference(unit_resulting_resolution,[],[f430,f341])).
fof(f341,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK17(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK17(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK17(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK17(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.UNIQUE_MEMBER_SING',aHL_EXT)).
fof(f4959,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl60),
  inference(cnf_transformation,[],[f4959_D])).
fof(f4959_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f25956,plain,(
  $false | ($spl8 | $spl68 | $spl71)),
  inference(subsumption_resolution,[],[f25955,f2031])).
fof(f2031,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f244,f1427,f311])).
fof(f311,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.pred_set.UNIQUE_MEMBER_SING','thm.bool.IMP_ANTISYM_AX')).
fof(f1427,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))))))) )),
  inference(equality_resolution,[],[f306])).
fof(f306,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))) )),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X7,X9,X8] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X9),s(fun(X6,bool),X8)))))) <=> (s(X6,X7) = s(X6,X9) | p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.UNIQUE_MEMBER_SING','thm.pred_set.IN_INSERT')).
fof(f25955,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | ($spl8 | $spl68 | $spl71)),
  inference(forward_demodulation,[],[f5937,f25920])).
fof(f25920,plain,(
  s(sK12,sK13) = s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)) | ($spl8 | $spl68)),
  inference(subsumption_resolution,[],[f25713,f244])).
fof(f25713,plain,(
  ~p(s(bool,'T_0')) | s(sK12,sK13) = s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)) | ($spl8 | $spl68)),
  inference(superposition,[],[f433,f5935])).
fof(f5935,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | $spl68),
  inference(cnf_transformation,[],[f5935_D])).
fof(f5935_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f5937,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | $spl71),
  inference(cnf_transformation,[],[f5937_D])).
fof(f5937_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f25919,plain,(
  s(sK12,sK13) = s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)) | $spl174),
  inference(cnf_transformation,[],[f25919_D])).
fof(f25919_D,plain,(
  s(sK12,sK13) = s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f4888,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl24),
  inference(cnf_transformation,[],[f4888_D])).
fof(f4888_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f25813,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(sK12,sK13) = s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12))) ) | ($spl8 | $spl68)),
  inference(duplicate_literal_removal,[],[f25714])).
fof(f25714,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(sK12,sK13) = s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl68)),
  inference(superposition,[],[f4966,f5935])).
fof(f4966,plain,(
  ( ! [X130,X129] : (s(bool,X130) = s(bool,'const.bool.IN_2'(s(sK12,X129),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X129) | s(bool,'T_0') = s(bool,X130)) ) | $spl8),
  inference(subsumption_resolution,[],[f4351,f244])).
fof(f4351,plain,(
  ( ! [X130,X129] : (~p(s(bool,'T_0')) | s(sK12,sK13) = s(sK12,X129) | s(bool,X130) = s(bool,'const.bool.IN_2'(s(sK12,X129),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X130)) ) | $spl8),
  inference(superposition,[],[f433,f635])).
fof(f635,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f251,f251])).
fof(f251,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.UNIQUE_MEMBER_SING',aHL_BOOL_CASES)).
fof(f25836,plain,(
  $false | ($spl8 | $spl61 | $spl68 | $spl71)),
  inference(subsumption_resolution,[],[f25830,f2031])).
fof(f25830,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | ($spl8 | $spl61 | $spl68 | $spl71)),
  inference(backward_demodulation,[],[f25697,f5937])).
fof(f25697,plain,(
  s(sK12,sK13) = s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)) | ($spl8 | $spl61 | $spl68)),
  inference(unit_resulting_resolution,[],[f5935,f13555])).
fof(f13555,plain,(
  ( ! [X215] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X215),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X215)) ) | ($spl8 | $spl61)),
  inference(subsumption_resolution,[],[f13300,f4960])).
fof(f13300,plain,(
  ( ! [X215] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X215),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X215) | s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) ) | ($spl8 | $spl61)),
  inference(superposition,[],[f4960,f4966])).
fof(f24300,plain,(
  $false | ($spl69 | $spl72)),
  inference(subsumption_resolution,[],[f24243,f244])).
fof(f24243,plain,(
  ~p(s(bool,'T_0')) | ($spl69 | $spl72)),
  inference(backward_demodulation,[],[f24116,f24070])).
fof(f24070,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)))) | $spl69),
  inference(unit_resulting_resolution,[],[f244,f5936,f311])).
fof(f24116,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1)) ) | ($spl69 | $spl72)),
  inference(subsumption_resolution,[],[f24085,f5939])).
fof(f24085,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X1)) ) | $spl69),
  inference(trivial_inequality_removal,[],[f24078])).
fof(f24078,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X1)) ) | $spl69),
  inference(superposition,[],[f5936,f635])).
fof(f24299,plain,(
  $false | ($spl8 | $spl69 | $spl72)),
  inference(global_subsumption,[],[f24298,f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2032,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f6113,f6114,f6371,f6372,f6373,f6374,f6375,f6466,f6467,f6737,f6740,f6767,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2031,f4402,f4263,f6665,f4352,f4265,f4135,f4264,f4353,f6189,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f307,f306,f305,f4964,f4965,f2026,f6481,f4134,f4401,f6759,f6757,f6755,f6729,f6727,f6725,f6710,f6709,f6705,f6670,f6604,f6603,f6597,f6560,f6559,f6555,f6508,f6506,f6504,f6503,f6496,f6494,f6492,f6491,f6455,f6424,f6175,f6129,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f5936,f5939,f24072,f24071,f24070,f24068,f24067,f24116,f24260,f24263,f24266,f24296])).
fof(f24296,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl69 | $spl72)),
  inference(subsumption_resolution,[],[f24203,f24116])).
fof(f24203,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK16(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl69 | $spl72)),
  inference(backward_demodulation,[],[f24116,f304])).
fof(f24266,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl69 | $spl72)),
  inference(subsumption_resolution,[],[f24136,f24116])).
fof(f24136,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK16('const.pred_set.EMPTY_0',X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl69 | $spl72)),
  inference(backward_demodulation,[],[f24116,f1979])).
fof(f24263,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl69 | $spl72)),
  inference(subsumption_resolution,[],[f24133,f24116])).
fof(f24133,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK16(X1,'const.pred_set.EMPTY_0',X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl69 | $spl72)),
  inference(backward_demodulation,[],[f24116,f1973])).
fof(f24260,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl69 | $spl72)),
  inference(subsumption_resolution,[],[f24117,f24116])).
fof(f24117,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl69 | $spl72)),
  inference(backward_demodulation,[],[f24116,f263])).
fof(f24067,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'T_0') | $spl69),
  inference(unit_resulting_resolution,[],[f434,f5936,f361])).
fof(f24068,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'T_0') | $spl69),
  inference(unit_resulting_resolution,[],[f244,f5936,f363])).
fof(f24071,plain,(
  sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | $spl69),
  inference(unit_resulting_resolution,[],[f434,f5936,f361])).
fof(f24072,plain,(
  ~sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | $spl69),
  inference(unit_resulting_resolution,[],[f244,f5936,f363])).
fof(f263,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.pred_set.UNIQUE_MEMBER_SING','thm.bool.EQ_CLAUSES|split|3')).
fof(f264,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f165])).
fof(f505,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f302,f264])).
fof(f1420,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,X9) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X9),s(fun(bool,bool),X10)))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f306,f251])).
fof(f1421,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X12),s(fun(bool,bool),X13)))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f306,f251])).
fof(f1425,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,X9) | p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X10)))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f306,f251])).
fof(f1426,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X11),s(fun(bool,bool),X13)))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f306,f251])).
fof(f1458,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X9)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X10),s(fun(bool,bool),X9)))))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f307,f251])).
fof(f1459,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X12)))) | p(s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X13),s(fun(bool,bool),X12)))))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f307,f251])).
fof(f1668,plain,(
  ( ! [X4,X5,X3] : (s(X5,f(s(bool,'F_0'))) != s(X5,g(s(bool,'F_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,'T_0') = s(bool,sK17(X3,X4,X5,bool))) )),
  inference(superposition,[],[f341,f251])).
fof(f1670,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(X3,sK17(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'T_0') = s(bool,g(s(X3,sK17(X4,X5,bool,X3))))) )),
  inference(superposition,[],[f341,f251])).
fof(f1732,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X12),s(fun(bool,bool),X13)))))) | s(bool,'F_0') = s(bool,X12) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f305,f251])).
fof(f1734,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X18),s(fun(bool,bool),X19)))))) | s(bool,X17) = s(bool,X18) | p(s(bool,'const.bool.IN_2'(s(bool,X17),s(fun(bool,bool),X19)))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f305,f251])).
fof(f1739,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X11),s(fun(bool,bool),X13)))))) | s(bool,'F_0') = s(bool,X12) | p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),X13)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f305,f251])).
fof(f1741,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'const.bool.IN_2'(s(bool,X18),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X19)))))) | s(bool,X17) = s(bool,X18) | p(s(bool,'const.bool.IN_2'(s(bool,X18),s(fun(bool,bool),X19)))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f305,f251])).
fof(f1973,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,sK16(X1,'const.pred_set.EMPTY_0',X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) )),
  inference(superposition,[],[f304,f505])).
fof(f1975,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X5,sK16(X6,X7,X5)),s(fun(X5,bool),X6))) | s(fun(X5,bool),X6) = s(fun(X5,bool),X7) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X5,sK16(X6,X7,X5)),s(fun(X5,bool),X7)))) )),
  inference(superposition,[],[f304,f251])).
fof(f1977,plain,(
  ( ! [X2,X3] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2))) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X3))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK16(X2,X3,bool))) )),
  inference(superposition,[],[f304,f251])).
fof(f1979,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,sK16('const.pred_set.EMPTY_0',X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) )),
  inference(superposition,[],[f304,f505])).
fof(f1981,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X5,sK16(X6,X7,X5)),s(fun(X5,bool),X7))) | s(fun(X5,bool),X6) = s(fun(X5,bool),X7) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X5,sK16(X6,X7,X5)),s(fun(X5,bool),X6)))) )),
  inference(superposition,[],[f304,f251])).
fof(f6129,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1))))))) )),
  inference(condensation,[],[f6128])).
fof(f6128,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(condensation,[],[f6127])).
fof(f6127,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X8),s(fun(bool,bool),X11)))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f6110,f635])).
fof(f6110,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X9) != s(bool,X10) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X8),s(fun(bool,bool),X11)))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1420,f635])).
fof(f6175,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1))))))) )),
  inference(condensation,[],[f6174])).
fof(f6174,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X23)))))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X21)) )),
  inference(trivial_inequality_removal,[],[f6169])).
fof(f6169,plain,(
  ( ! [X23,X21,X22] : (s(bool,X21) != s(bool,X21) | p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X23)))))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f1421,f251])).
fof(f6424,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X4),s(fun(bool,bool),X3)))))) | s(bool,'T_0') = s(bool,X2) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6414,f244])).
fof(f6414,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X4),s(fun(bool,bool),X3)))))) | s(bool,'T_0') = s(bool,X2) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1458,f635])).
fof(f6455,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X4),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X3) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6450,f244])).
fof(f6450,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X4),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X3) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1459,f635])).
fof(f6491,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(fun(bool,bool),'const.pred_set.EMPTY_0') = s(fun(bool,bool),X0) | s(bool,'T_0') = s(bool,sK16('const.pred_set.EMPTY_0',X0,bool))) )),
  inference(superposition,[],[f1977,f505])).
fof(f6492,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2))) | s(fun(bool,bool),X1) = s(fun(bool,bool),X2) | s(bool,'T_0') = s(bool,sK16(X1,X2,bool)) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1977,f635])).
fof(f6494,plain,(
  ( ! [X8,X7,X9] : (s(bool,X8) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X9))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X9) | s(bool,'T_0') = s(bool,sK16(X7,X9,bool)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1977,f635])).
fof(f6496,plain,(
  ( ! [X14,X13] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X14))) | s(fun(bool,bool),X13) = s(fun(bool,bool),X14) | s(bool,'T_0') = s(bool,sK16(X13,X14,bool)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X13)))) )),
  inference(superposition,[],[f1977,f251])).
fof(f6503,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(fun(bool,bool),'const.pred_set.EMPTY_0') = s(fun(bool,bool),X0) | s(bool,'T_0') = s(bool,sK16(X0,'const.pred_set.EMPTY_0',bool))) )),
  inference(superposition,[],[f1977,f505])).
fof(f6504,plain,(
  ( ! [X2,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2))) | s(fun(bool,bool),X1) = s(fun(bool,bool),X2) | s(bool,'T_0') = s(bool,sK16(X2,X1,bool)) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1977,f635])).
fof(f6506,plain,(
  ( ! [X8,X7,X9] : (s(bool,X8) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X9))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X9) | s(bool,'T_0') = s(bool,sK16(X9,X7,bool)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1977,f635])).
fof(f6508,plain,(
  ( ! [X14,X13] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X14))) | s(fun(bool,bool),X13) = s(fun(bool,bool),X14) | s(bool,'T_0') = s(bool,sK16(X14,X13,bool)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X13)))) )),
  inference(superposition,[],[f1977,f251])).
fof(f6555,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1))))) )),
  inference(duplicate_literal_removal,[],[f6554])).
fof(f6554,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f6553])).
fof(f6553,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f6533,f4937])).
fof(f6533,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))))) | s(bool,'F_0') = s(bool,X0) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1732,f635])).
fof(f6559,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,X20),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X19),s(fun(bool,bool),X21)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X21)))) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f6538,f465])).
fof(f6538,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,X20),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X19),s(fun(bool,bool),X21)))))) | s(bool,'T_0') = s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X21)))) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f1732,f635])).
fof(f6560,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X0) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X1),s(fun(bool,bool),X2))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f6543,f244])).
fof(f6543,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X0) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X1),s(fun(bool,bool),X2))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1732,f635])).
fof(f6597,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X19),s(fun(bool,bool),X20)))))) | s(bool,'T_0') = s(bool,X21) | p(s(bool,'const.bool.IN_2'(s(bool,X21),s(fun(bool,bool),X20)))) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(duplicate_literal_removal,[],[f6586])).
fof(f6586,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X19),s(fun(bool,bool),X20)))))) | s(bool,'T_0') = s(bool,X21) | p(s(bool,'const.bool.IN_2'(s(bool,X21),s(fun(bool,bool),X20)))) | s(bool,'T_0') = s(bool,X21) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f1734,f635])).
fof(f6603,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1))))) )),
  inference(condensation,[],[f6602])).
fof(f6602,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f6601])).
fof(f6601,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(condensation,[],[f6581])).
fof(f6581,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X1)))))) | s(bool,X0) = s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1734,f635])).
fof(f6604,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f6591,f244])).
fof(f6591,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1734,f635])).
fof(f6670,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X1) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X1),s(fun(bool,bool),X2))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f6655,f244])).
fof(f6655,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X0) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X1) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X1),s(fun(bool,bool),X2))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1739,f635])).
fof(f6705,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X2)))) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f6704])).
fof(f6704,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X20)))))) | s(bool,'T_0') = s(bool,X21) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X20)))) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(duplicate_literal_removal,[],[f6686])).
fof(f6686,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X20)))))) | s(bool,'T_0') = s(bool,X21) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X20)))) | s(bool,'T_0') = s(bool,X21) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f1741,f635])).
fof(f6709,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f6708])).
fof(f6708,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f6681])).
fof(f6681,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))) | s(bool,X0) = s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1741,f635])).
fof(f6710,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f6697,f244])).
fof(f6697,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1741,f635])).
fof(f6725,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X0))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK16(X0,X1,bool)),s(fun(bool,bool),X1))) | s(bool,X2) = s(bool,sK16(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1975,f635])).
fof(f6727,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X6))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK16(X6,X7,bool)),s(fun(bool,bool),X7))) | s(bool,'T_0') = s(bool,sK16(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1975,f635])).
fof(f6729,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X12))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK16(X12,X13,bool)),s(fun(bool,bool),X13))) | s(bool,'T_0') = s(bool,sK16(X12,X13,bool))) )),
  inference(superposition,[],[f1975,f251])).
fof(f6755,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK16(X0,X1,bool)),s(fun(bool,bool),X0))) | s(bool,X2) = s(bool,sK16(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1981,f635])).
fof(f6757,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X7))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK16(X6,X7,bool)),s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,sK16(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f1981,f635])).
fof(f6759,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X13))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK16(X12,X13,bool)),s(fun(bool,bool),X12))) | s(bool,'T_0') = s(bool,sK16(X12,X13,bool))) )),
  inference(superposition,[],[f1981,f251])).
fof(f4401,plain,(
  ( ! [X273,X275,X272,X274] : (s(bool,'T_0') != s(bool,f(s(X272,sK17(X273,X274,bool,X272)))) | s(fun(X272,bool),X273) = s(fun(X272,bool),X274) | s(bool,X275) = s(bool,g(s(X272,sK17(X273,X274,bool,X272)))) | s(bool,'T_0') = s(bool,X275)) )),
  inference(superposition,[],[f341,f635])).
fof(f4134,plain,(
  ( ! [X280,X279,X281,X278] : (s(bool,X281) != s(bool,f(s(X278,sK17(X279,X280,bool,X278)))) | s(fun(X278,bool),X279) = s(fun(X278,bool),X280) | s(bool,'T_0') = s(bool,X281) | s(bool,'T_0') = s(bool,g(s(X278,sK17(X279,X280,bool,X278))))) )),
  inference(superposition,[],[f341,f635])).
fof(f6481,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK17(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(X0,sK17(X1,X2,bool,X0))))) )),
  inference(condensation,[],[f6480])).
fof(f6480,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK17(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK17(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f6473,f251])).
fof(f6473,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK17(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK17(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f1670,f635])).
fof(f2026,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))))))))) )),
  inference(unit_resulting_resolution,[],[f1427,f307])).
fof(f4965,plain,(
  ( ! [X118,X116,X120,X119,X117] : (s(X116,X117) = s(X116,X118) | p(s(bool,'const.bool.IN_2'(s(X116,X117),s(fun(X116,bool),X119)))) | s(bool,X120) = s(bool,'const.bool.IN_2'(s(X116,X117),s(fun(X116,bool),'const.pred_set.INSERT_2'(s(X116,X118),s(fun(X116,bool),X119))))) | s(bool,'T_0') = s(bool,X120)) )),
  inference(subsumption_resolution,[],[f4348,f244])).
fof(f4348,plain,(
  ( ! [X118,X116,X120,X119,X117] : (~p(s(bool,'T_0')) | s(X116,X117) = s(X116,X118) | p(s(bool,'const.bool.IN_2'(s(X116,X117),s(fun(X116,bool),X119)))) | s(bool,X120) = s(bool,'const.bool.IN_2'(s(X116,X117),s(fun(X116,bool),'const.pred_set.INSERT_2'(s(X116,X118),s(fun(X116,bool),X119))))) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f305,f635])).
fof(f4964,plain,(
  ( ! [X111,X114,X112,X115,X113] : (p(s(bool,'const.bool.IN_2'(s(X111,X112),s(fun(X111,bool),'const.pred_set.INSERT_2'(s(X111,X114),s(fun(X111,bool),X113)))))) | s(bool,X115) = s(bool,'const.bool.IN_2'(s(X111,X112),s(fun(X111,bool),X113))) | s(bool,'T_0') = s(bool,X115)) )),
  inference(subsumption_resolution,[],[f4347,f244])).
fof(f4347,plain,(
  ( ! [X111,X114,X112,X115,X113] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X111,X112),s(fun(X111,bool),'const.pred_set.INSERT_2'(s(X111,X114),s(fun(X111,bool),X113)))))) | s(bool,X115) = s(bool,'const.bool.IN_2'(s(X111,X112),s(fun(X111,bool),X113))) | s(bool,'T_0') = s(bool,X115)) )),
  inference(superposition,[],[f307,f635])).
fof(f305,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f193])).
fof(f307,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))) )),
  inference(cnf_transformation,[],[f193])).
fof(f303,plain,(
  ( ! [X4,X2,X0,X1] : (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK16(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK16(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X8,X12] : (s(fun(X6,bool),X8) = s(fun(X6,bool),X12) <=> ! [X7] : s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8))) = s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X12))))),
  file('thm.pred_set.UNIQUE_MEMBER_SING','thm.pred_set.EXTENSION')).
fof(f304,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK16(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK16(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f191])).
fof(f4085,plain,(
  ( ! [X138,X136,X139,X137] : (~p(s(bool,'const.bool.IN_2'(s(bool,X137),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X138),s(fun(bool,bool),X139)))))) | s(bool,X136) = s(bool,X138) | p(s(bool,'const.bool.IN_2'(s(bool,X136),s(fun(bool,bool),X139)))) | s(bool,'T_0') = s(bool,X137) | s(bool,'T_0') = s(bool,X136)) )),
  inference(superposition,[],[f305,f635])).
fof(f4086,plain,(
  ( ! [X142,X140,X143,X141] : (~p(s(bool,'const.bool.IN_2'(s(bool,X142),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X141),s(fun(bool,bool),X143)))))) | s(bool,X140) = s(bool,X142) | p(s(bool,'const.bool.IN_2'(s(bool,X142),s(fun(bool,bool),X143)))) | s(bool,'T_0') = s(bool,X141) | s(bool,'T_0') = s(bool,X140)) )),
  inference(superposition,[],[f305,f635])).
fof(f4087,plain,(
  ( ! [X146,X144,X147,X145] : (s(bool,X145) != s(bool,X146) | p(s(bool,'const.bool.IN_2'(s(bool,X144),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X146),s(fun(bool,bool),X147)))))) | s(bool,'T_0') = s(bool,X145) | s(bool,'T_0') = s(bool,X144)) )),
  inference(superposition,[],[f306,f635])).
fof(f4088,plain,(
  ( ! [X151,X149,X150,X148] : (s(bool,X149) != s(bool,X150) | p(s(bool,'const.bool.IN_2'(s(bool,X150),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X148),s(fun(bool,bool),X151)))))) | s(bool,'T_0') = s(bool,X149) | s(bool,'T_0') = s(bool,X148)) )),
  inference(superposition,[],[f306,f635])).
fof(f4089,plain,(
  ( ! [X154,X152,X155,X153] : (~p(s(bool,'const.bool.IN_2'(s(bool,X153),s(fun(bool,bool),X154)))) | p(s(bool,'const.bool.IN_2'(s(bool,X152),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X155),s(fun(bool,bool),X154)))))) | s(bool,'T_0') = s(bool,X153) | s(bool,'T_0') = s(bool,X152)) )),
  inference(superposition,[],[f307,f635])).
fof(f4261,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,'const.bool.IN_2'(s(bool,X101),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X102),s(fun(bool,bool),X103)))))) | s(bool,'T_0') = s(bool,X102) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X103)))) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f305,f635])).
fof(f4262,plain,(
  ( ! [X107,X105,X108,X106] : (~p(s(bool,'const.bool.IN_2'(s(bool,X106),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X105),s(fun(bool,bool),X107)))))) | s(bool,'T_0') = s(bool,X106) | p(s(bool,'const.bool.IN_2'(s(bool,X106),s(fun(bool,bool),X107)))) | s(bool,X105) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f305,f635])).
fof(f4354,plain,(
  ( ! [X142,X140,X141,X139] : (s(bool,'T_0') != s(bool,X140) | p(s(bool,'const.bool.IN_2'(s(bool,X139),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X140),s(fun(bool,bool),X141)))))) | s(bool,X139) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) )),
  inference(superposition,[],[f306,f635])).
fof(f4355,plain,(
  ( ! [X146,X144,X145,X143] : (s(bool,'T_0') != s(bool,X144) | p(s(bool,'const.bool.IN_2'(s(bool,X144),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X143),s(fun(bool,bool),X145)))))) | s(bool,X143) = s(bool,X146) | s(bool,'T_0') = s(bool,X146)) )),
  inference(superposition,[],[f306,f635])).
fof(f4356,plain,(
  ( ! [X149,X147,X150,X148] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X148)))) | p(s(bool,'const.bool.IN_2'(s(bool,X147),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X149),s(fun(bool,bool),X148)))))) | s(bool,X147) = s(bool,X150) | s(bool,'T_0') = s(bool,X150)) )),
  inference(superposition,[],[f307,f635])).
fof(f6189,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X1),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(condensation,[],[f6188])).
fof(f6188,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'const.bool.IN_2'(s(bool,X10),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X8),s(fun(bool,bool),X11)))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f6166,f251])).
fof(f6166,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.bool.IN_2'(s(bool,X10),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X8),s(fun(bool,bool),X11)))))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1421,f635])).
fof(f4353,plain,(
  ( ! [X138,X136,X135,X137] : (~p(s(bool,'const.bool.IN_2'(s(bool,X136),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X137)))))) | s(bool,X135) = s(bool,X136) | p(s(bool,'const.bool.IN_2'(s(bool,X136),s(fun(bool,bool),X137)))) | s(bool,X135) = s(bool,X138) | s(bool,'T_0') = s(bool,X138)) )),
  inference(superposition,[],[f305,f635])).
fof(f4264,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,X113) != s(bool,X114) | p(s(bool,'const.bool.IN_2'(s(bool,X114),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X115)))))) | s(bool,X113) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f306,f635])).
fof(f4135,plain,(
  ( ! [X282,X283,X284] : (s(bool,'const.bool.IN_2'(s(bool,X284),s(fun(bool,bool),X282))) != s(bool,'const.bool.IN_2'(s(bool,X284),s(fun(bool,bool),X283))) | s(fun(bool,bool),X282) = s(fun(bool,bool),X283) | s(bool,'T_0') = s(bool,X284) | s(bool,'T_0') = s(bool,sK16(X282,X283,bool))) )),
  inference(superposition,[],[f304,f635])).
fof(f4265,plain,(
  ( ! [X118,X120,X119,X117] : (~p(s(bool,'const.bool.IN_2'(s(bool,X117),s(fun(bool,bool),X118)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X119),s(fun(bool,bool),X118)))))) | s(bool,X117) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f307,f635])).
fof(f4352,plain,(
  ( ! [X134,X132,X133,X131] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X132),s(fun(bool,bool),X133)))))) | s(bool,X131) = s(bool,X132) | p(s(bool,'const.bool.IN_2'(s(bool,X131),s(fun(bool,bool),X133)))) | s(bool,X131) = s(bool,X134) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f305,f635])).
fof(f6665,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X1),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f6635,f4937])).
fof(f6635,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X1),s(fun(bool,bool),X2)))))) | s(bool,'F_0') = s(bool,X0) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1739,f635])).
fof(f4263,plain,(
  ( ! [X111,X109,X112,X110] : (s(bool,X109) != s(bool,X110) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X110),s(fun(bool,bool),X111)))))) | s(bool,X109) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f306,f635])).
fof(f4402,plain,(
  ( ! [X277,X276,X278] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X276))) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X277))) | s(fun(bool,bool),X276) = s(fun(bool,bool),X277) | s(bool,X278) = s(bool,sK16(X276,X277,bool)) | s(bool,'T_0') = s(bool,X278)) )),
  inference(superposition,[],[f304,f635])).
fof(f2033,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f1427,f317])).
fof(f2034,plain,(
  ( ! [X2,X0,X3,X1] : (sP0('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f1427,f318])).
fof(f2035,plain,(
  ( ! [X2,X0,X1] : (~sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f244,f1427,f332])).
fof(f2038,plain,(
  ( ! [X2,X0,X1] : (~sP1('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f244,f1427,f332])).
fof(f2039,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f1427,f1427,f332])).
fof(f2041,plain,(
  ( ! [X2,X0,X1] : (sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f434,f1427,f333])).
fof(f2043,plain,(
  ( ! [X2,X0,X1] : (sP1('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f434,f1427,f336])).
fof(f2047,plain,(
  ( ! [X2,X0,X1] : (~sP5('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f244,f434,f1427,f345])).
fof(f2048,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1427,f434,f1427,f345])).
fof(f2050,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f443,f438,f449,f1427,f346])).
fof(f2051,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f244,f434,f1427,f349])).
fof(f2057,plain,(
  ( ! [X2,X0,X1] : (~sP4('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f244,f434,f1427,f349])).
fof(f2058,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP4('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1427,f434,f1427,f349])).
fof(f2063,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f1427,f350])).
fof(f2064,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f244,f434,f1427,f353])).
fof(f2070,plain,(
  ( ! [X2,X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f244,f434,f1427,f353])).
fof(f2071,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP3('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1427,f434,f1427,f353])).
fof(f2076,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f1427,f354])).
fof(f2077,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f1427,f358])).
fof(f2078,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f1427,f359])).
fof(f2079,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f1427,f360])).
fof(f2081,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f505,f1427,f362])).
fof(f2083,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f505,f1427,f362])).
fof(f2084,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f465,f1427,f363])).
fof(f2086,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f504,f1427,f363])).
fof(f2087,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f465,f1427,f363])).
fof(f2089,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f504,f1427,f363])).
fof(f2090,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'F_0')) )),
  inference(unit_resulting_resolution,[],[f434,f1427,f366])).
fof(f2091,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0',X3)) )),
  inference(unit_resulting_resolution,[],[f434,f1427,f367])).
fof(f2098,plain,(
  ( ! [X2,X0,X1] : (~sP6('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f244,f434,f1427,f372])).
fof(f2104,plain,(
  ( ! [X2,X0,X1] : (~sP6('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f244,f434,f1427,f372])).
fof(f2105,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP6('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1427,f434,f1427,f372])).
fof(f2110,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f1427,f373])).
fof(f2111,plain,(
  ( ! [X2,X0,X1] : (sP7('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f244,f244,f1427,f378])).
fof(f2116,plain,(
  ( ! [X2,X0,X1] : (sP7('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f244,f244,f1427,f378])).
fof(f2118,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f244,f1427,f1427,f378])).
fof(f2121,plain,(
  ( ! [X2,X0,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f244,f244,f1427,f378])).
fof(f2122,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1427,f244,f1427,f378])).
fof(f2123,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f244,f1427,f1427,f378])).
fof(f2124,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),X8)))))) )),
  inference(unit_resulting_resolution,[],[f1427,f1427,f1427,f378])).
fof(f2125,plain,(
  ( ! [X2,X0,X1] : (~sP7('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f244,f434,f1427,f381])).
fof(f2130,plain,(
  ( ! [X2,X0,X1] : (~sP7('F_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f244,f434,f1427,f381])).
fof(f2131,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP7('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f1427,f434,f1427,f381])).
fof(f2135,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('F_0',X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f434,f1427,f382])).
fof(f2137,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f434,f1427,f383])).
fof(f2141,plain,(
  ( ! [X2,X0,X1] : (~sP8('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f434,f434,f1427,f389])).
fof(f2144,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f1427,f390])).
fof(f2145,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f1427,f391])).
fof(f2146,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('T_0',X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f244,f1427,f394])).
fof(f2150,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.INSERT_2'(s(X4,X5),s(fun(X4,bool),X6)))))) )),
  inference(unit_resulting_resolution,[],[f1427,f1427,f394])).
fof(f2151,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f244,f1427,f395])).
fof(f2155,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),X6)) )),
  inference(unit_resulting_resolution,[],[f1427,f1427,f395])).
fof(f2156,plain,(
  ( ! [X2,X0,X1] : (~sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f434,f434,f1427,f396])).
fof(f2157,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f434,f1427,f400])).
fof(f2174,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f434,f520,f1427,f405])).
fof(f2175,plain,(
  ( ! [X2,X0,X1] : (~sP10('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f434,f244,f1427,f406])).
fof(f2181,plain,(
  ( ! [X2,X0,X1] : (~sP10('T_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f434,f244,f1427,f406])).
fof(f2182,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP10('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f434,f1427,f1427,f406])).
fof(f2187,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f1427,f408])).
fof(f2188,plain,(
  ( ! [X2,X0,X1] : (sP11('F_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f434,f434,f1427,f410])).
fof(f2191,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f244,f1427,f412])).
fof(f2195,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),X6)) )),
  inference(unit_resulting_resolution,[],[f1427,f1427,f412])).
fof(f2196,plain,(
  ( ! [X2,X0,X1] : (~sP11('T_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f244,f434,f1427,f413])).
fof(f2199,plain,(
  ( ! [X2,X0,X1] : (~sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f244,f434,f1427,f413])).
fof(f2200,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f1427,f434,f1427,f413])).
fof(f2201,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'F_0')) )),
  inference(unit_resulting_resolution,[],[f460,f1427,f416])).
fof(f2203,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f519,f1427,f416])).
fof(f2204,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f462,f1427,f416])).
fof(f2205,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'T_0')) )),
  inference(unit_resulting_resolution,[],[f456,f1427,f417])).
fof(f2206,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f457,f1427,f417])).
fof(f6767,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK16(X1,'const.pred_set.EMPTY_0',X0)),s(fun(X0,bool),X1)))) )),
  inference(trivial_inequality_removal,[],[f6760])).
fof(f6760,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK16(X1,'const.pred_set.EMPTY_0',X0)),s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f1981,f505])).
fof(f6740,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK16(X1,X2,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK16(X1,X2,X0)),s(fun(X0,bool),X1)))) )),
  inference(condensation,[],[f6739])).
fof(f6739,plain,(
  ( ! [X12,X10,X13,X11] : (s(fun(X10,bool),X11) = s(fun(X10,bool),X12) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X10,sK16(X11,X12,X10)),s(fun(X10,bool),X12))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X10,sK16(X11,X12,X10)),s(fun(X10,bool),X11))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f6733,f251])).
fof(f6733,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'F_0') != s(bool,X13) | s(fun(X10,bool),X11) = s(fun(X10,bool),X12) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X10,sK16(X11,X12,X10)),s(fun(X10,bool),X12))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X10,sK16(X11,X12,X10)),s(fun(X10,bool),X11))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f1975,f635])).
fof(f6737,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK16('const.pred_set.EMPTY_0',X1,X0)),s(fun(X0,bool),X1)))) )),
  inference(trivial_inequality_removal,[],[f6730])).
fof(f6730,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK16('const.pred_set.EMPTY_0',X1,X0)),s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f1975,f505])).
fof(f4403,plain,(
  ( ! [X280,X282,X279,X281] : (s(X281,f(s(bool,'T_0'))) != s(X281,g(s(bool,'T_0'))) | s(fun(bool,X281),X279) = s(fun(bool,X281),X280) | s(bool,X282) = s(bool,sK17(X279,X280,X281,bool)) | s(bool,'T_0') = s(bool,X282)) )),
  inference(superposition,[],[f341,f635])).
fof(f4350,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X125,sK16(X126,X127,X125)),s(fun(X125,bool),X127))) | s(fun(X125,bool),X126) = s(fun(X125,bool),X127) | s(bool,X128) = s(bool,'const.bool.IN_2'(s(X125,sK16(X126,X127,X125)),s(fun(X125,bool),X126))) | s(bool,'T_0') = s(bool,X128)) )),
  inference(superposition,[],[f304,f635])).
fof(f4349,plain,(
  ( ! [X123,X121,X124,X122] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X121,sK16(X122,X123,X121)),s(fun(X121,bool),X122))) | s(fun(X121,bool),X122) = s(fun(X121,bool),X123) | s(bool,X124) = s(bool,'const.bool.IN_2'(s(X121,sK16(X122,X123,X121)),s(fun(X121,bool),X123))) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f304,f635])).
fof(f4957,plain,(
  ( ! [X90,X88,X87,X89] : (~sP8(X87,X88,X89) | sP9(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f4336,f244])).
fof(f4336,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP8(X87,X88,X89) | sP9(X87,X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f417,f635])).
fof(f4956,plain,(
  ( ! [X85,X83,X86,X84] : (~sP10(X83,X84,X85) | sP11(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f4335,f244])).
fof(f4335,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP10(X83,X84,X85) | sP11(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f416,f635])).
fof(f4955,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f4334,f244])).
fof(f4334,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f413,f635])).
fof(f4954,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | sP11(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f4333,f244])).
fof(f4333,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | sP11(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f412,f635])).
fof(f4953,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | sP11(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f4332,f244])).
fof(f4332,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X72)) | sP11(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f411,f635])).
fof(f4952,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f4331,f244])).
fof(f4331,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f406,f635])).
fof(f4951,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f4330,f244])).
fof(f4330,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f396,f635])).
fof(f4950,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f4329,f244])).
fof(f4329,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f395,f635])).
fof(f4949,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | sP9(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f4328,f244])).
fof(f4328,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP9(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f394,f635])).
fof(f4948,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f4327,f244])).
fof(f4327,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f389,f635])).
fof(f4947,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f4326,f244])).
fof(f4326,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f381,f635])).
fof(f4946,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f4325,f244])).
fof(f4325,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f372,f635])).
fof(f4945,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f4324,f244])).
fof(f4324,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f353,f635])).
fof(f4944,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f4323,f244])).
fof(f4323,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f349,f635])).
fof(f4943,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4322,f244])).
fof(f4322,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f345,f635])).
fof(f4942,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f4321,f244])).
fof(f4321,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f336,f635])).
fof(f4941,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f4320,f244])).
fof(f4320,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f333,f635])).
fof(f4940,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f4319,f244])).
fof(f4319,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f318,f635])).
fof(f4939,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f4318,f244])).
fof(f4318,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f317,f635])).
fof(f4938,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f4317,f244])).
fof(f4317,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f311,f635])).
fof(f4937,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4313,f244])).
fof(f4313,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f263,f635])).
fof(f4266,plain,(
  ( ! [X121,X122] : (s(bool,X121) != s(bool,'const.bool.IN_2'(s(bool,X121),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,X121) = s(bool,X122) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f504,f635])).
fof(f4232,plain,(
  ( ! [X6,X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X5)) | s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f259,f635])).
fof(f4136,plain,(
  ( ! [X288,X285,X287,X286] : (s(X287,f(s(bool,X288))) != s(X287,g(s(bool,X288))) | s(fun(bool,X287),X285) = s(fun(bool,X287),X286) | s(bool,'T_0') = s(bool,X288) | s(bool,'T_0') = s(bool,sK17(X285,X286,X287,bool))) )),
  inference(superposition,[],[f341,f635])).
fof(f4082,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,X132) != s(bool,'const.bool.IN_2'(s(X129,sK16(X130,X131,X129)),s(fun(X129,bool),X131))) | s(fun(X129,bool),X130) = s(fun(X129,bool),X131) | s(bool,'T_0') = s(bool,X132) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X129,sK16(X130,X131,X129)),s(fun(X129,bool),X130)))) )),
  inference(superposition,[],[f304,f635])).
fof(f4081,plain,(
  ( ! [X127,X125,X128,X126] : (s(bool,X128) != s(bool,'const.bool.IN_2'(s(X125,sK16(X126,X127,X125)),s(fun(X125,bool),X126))) | s(fun(X125,bool),X126) = s(fun(X125,bool),X127) | s(bool,'T_0') = s(bool,X128) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X125,sK16(X126,X127,X125)),s(fun(X125,bool),X127)))) )),
  inference(superposition,[],[f304,f635])).
fof(f4020,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f635,f635])).
fof(f2032,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))))) = s(bool,'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f1427,f1427,f311])).
fof(f1428,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2))))))) )),
  inference(forward_demodulation,[],[f1416,f505])).
fof(f1416,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2))))))) )),
  inference(unit_resulting_resolution,[],[f505,f306])).
fof(f1380,plain,(
  ( ! [X0,X1] : (~sP2('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1035,f515,f439,f434,f348])).
fof(f1322,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1037,f514,f441,f434,f347])).
fof(f1325,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f446,f439,f448,f434,f347])).
fof(f1252,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f446,f513,f450,f244,f346])).
fof(f1254,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f515,f513,f450,f244,f346])).
fof(f1260,plain,(
  ( ! [X0,X1] : (~sP4('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f938,f515,f448,f244,f346])).
fof(f1263,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f443,f438,f449,f244,f346])).
fof(f1211,plain,(
  ( ! [X2,X3,X1] : (s(bool,X3) = s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1207,f434])).
fof(f1207,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'F_0')) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f704,f505])).
fof(f1149,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1137,f401])).
fof(f1137,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f244,f434,f244,f406])).
fof(f1123,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f463,f244,f434,f405])).
fof(f1125,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f520,f244,f434,f405])).
fof(f1042,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f505,f251])).
fof(f1035,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f434,f505,f364])).
fof(f1036,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f244,f505,f362])).
fof(f1037,plain,(
  ( ! [X0,X1] : (~sP5('F_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f434,f505,f364])).
fof(f1038,plain,(
  ( ! [X0,X1] : (sP5('T_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f244,f505,f362])).
fof(f975,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f966,f384])).
fof(f966,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f434,f434,f244,f389])).
fof(f954,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f458,f434,f434,f388])).
fof(f956,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f518,f434,f434,f388])).
fof(f938,plain,(
  ( ! [X0,X1] : (~sP5('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f244,f504,f363])).
fof(f939,plain,(
  ( ! [X0,X1] : (sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f434,f504,f361])).
fof(f941,plain,(
  ( ! [X0,X1] : (~sP5('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f244,f504,f363])).
fof(f942,plain,(
  ( ! [X0,X1] : (sP5('F_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f434,f504,f361])).
fof(f900,plain,(
  ( ! [X2,X0,X1] : (sP11('T_0',X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f244,f519,f416])).
fof(f878,plain,(
  ( ! [X2,X0,X1] : (sP7('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f434,f517,f418])).
fof(f868,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f858,f365])).
fof(f858,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f434,f244,f244,f372])).
fof(f846,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f451,f244,f244,f368])).
fof(f842,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f837,f342])).
fof(f837,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f434,f434,f434,f357])).
fof(f833,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f822,f343])).
fof(f822,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f434,f244,f244,f353])).
fof(f801,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f434,f244,f244,f349])).
fof(f756,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f244,f465,f363])).
fof(f760,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f244,f465,f363])).
fof(f737,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f434,f465,f361])).
fof(f739,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f434,f465,f361])).
fof(f731,plain,(
  ( ! [X2,X0,X1] : (sP7('F_0',X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f434,f516,f418])).
fof(f704,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f698,f434])).
fof(f698,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f308,f251])).
fof(f599,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f244,f244,f412])).
fof(f592,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f434,f244,f411])).
fof(f589,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f244,f434,f400])).
fof(f588,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f434,f434,f399])).
fof(f583,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f244,f244,f395])).
fof(f575,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f244,f244,f394])).
fof(f573,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f244,f434,f383])).
fof(f572,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f244,f434,f382])).
fof(f563,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f434,f434,f377])).
fof(f552,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f434,f434,f376])).
fof(f545,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f244,f434,f367])).
fof(f540,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f244,f434,f366])).
fof(f504,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f302,f259])).
fof(f506,plain,(
  ( ! [X0,X1] : (~sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f434,f302,f331])).
fof(f508,plain,(
  ( ! [X0,X1] : (~sP1('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f434,f302,f331])).
fof(f509,plain,(
  ( ! [X2,X0,X3,X1] : (~sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f302,f302,f331])).
fof(f510,plain,(
  ( ! [X0,X1] : (sP1('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f244,f302,f333])).
fof(f511,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f244,f302,f336])).
fof(f512,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),X3)) )),
  inference(unit_resulting_resolution,[],[f302,f351])).
fof(f513,plain,(
  ( ! [X2,X0,X3,X1] : (sP4('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f302,f352])).
fof(f514,plain,(
  ( ! [X2,X0,X3,X1] : (sP3(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),X3)) )),
  inference(unit_resulting_resolution,[],[f302,f355])).
fof(f515,plain,(
  ( ! [X2,X0,X3,X1] : (sP3('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f302,f356])).
fof(f516,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),X3)) )),
  inference(unit_resulting_resolution,[],[f302,f374])).
fof(f517,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f302,f375])).
fof(f518,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f302,f392])).
fof(f519,plain,(
  ( ! [X2,X0,X3,X1] : (sP10(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f302,f407])).
fof(f520,plain,(
  ( ! [X2,X0,X3,X1] : (sP10('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),X2,X3)) )),
  inference(unit_resulting_resolution,[],[f302,f409])).
fof(f521,plain,(
  ( ! [X0,X1] : (~sP0('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f434,f302,f435])).
fof(f524,plain,(
  ( ! [X0,X1] : (~sP0('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f434,f302,f435])).
fof(f525,plain,(
  ( ! [X2,X0,X3,X1] : (~sP0('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f302,f302,f435])).
fof(f499,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f434,f434,f435])).
fof(f495,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f434,f244,f336])).
fof(f489,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f434,f244,f333])).
fof(f486,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f244,f244,f332])).
fof(f485,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f434,f434,f331])).
fof(f465,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f434,f259])).
fof(f463,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f434,f409])).
fof(f462,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f244,f408])).
fof(f460,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f434,f407])).
fof(f458,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f434,f392])).
fof(f457,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f244,f391])).
fof(f456,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f244,f390])).
fof(f454,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f434,f375])).
fof(f452,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f434,f374])).
fof(f451,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f244,f373])).
fof(f450,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f244,f360])).
fof(f449,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f244,f359])).
fof(f448,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f244,f358])).
fof(f446,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f434,f356])).
fof(f444,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f434,f355])).
fof(f443,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f244,f354])).
fof(f441,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f434,f352])).
fof(f439,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f434,f351])).
fof(f438,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f244,f350])).
fof(f437,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f244,f318])).
fof(f436,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f244,f317])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f138,f155,f154])).
fof(f154,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f155,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f138,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X14,X15,X16] : ((p(s(bool,X16)) <=> (p(s(bool,X15)) => p(s(bool,X14)))) <=> ((p(s(bool,X16)) | p(s(bool,X15))) & (p(s(bool,X16)) | ~p(s(bool,X14))) & (~p(s(bool,X15)) | p(s(bool,X14)) | ~p(s(bool,X16)))))),
  file('thm.pred_set.UNIQUE_MEMBER_SING','thm.sat.sat_27_prove_ltb')).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f239])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f239])).
fof(f413,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f239])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f236])).
fof(f236,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f235])).
fof(f235,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f154])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f237])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f237])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f237])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f233])).
fof(f233,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f155])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f234])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f234])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f402])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f234])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f234])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f234])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f234])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f108,f152,f151])).
fof(f151,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f152,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f108,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X14,X15,X16] : ((p(s(bool,X16)) <=> (p(s(bool,X15)) | p(s(bool,X14)))) <=> ((p(s(bool,X16)) | ~p(s(bool,X15))) & (p(s(bool,X16)) | ~p(s(bool,X14))) & (p(s(bool,X15)) | p(s(bool,X14)) | ~p(s(bool,X16)))))),
  file('thm.pred_set.UNIQUE_MEMBER_SING','thm.sat.sat_25_prove_ltb')).
fof(f394,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f395,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f389,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f228])).
fof(f228,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f227])).
fof(f227,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f151])).
fof(f390,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f229])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f229])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f382,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f225])).
fof(f225,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f152])).
fof(f383,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f226])).
fof(f384,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f226])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f385])).
fof(f385,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f226])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f226])).
fof(f387,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f226])).
fof(f388,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f226])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f106,f149,f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f149,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f106,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X14,X15,X16] : ((p(s(bool,X16)) <=> (p(s(bool,X15)) & p(s(bool,X14)))) <=> ((p(s(bool,X16)) | ~p(s(bool,X15)) | ~p(s(bool,X14))) & (p(s(bool,X15)) | ~p(s(bool,X16))) & (p(s(bool,X14)) | ~p(s(bool,X16)))))),
  file('thm.pred_set.UNIQUE_MEMBER_SING','thm.sat.sat_23_prove_ltb')).
fof(f377,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f223])).
fof(f378,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f223])).
fof(f381,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f223])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f148])).
fof(f373,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f221])).
fof(f374,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f221])).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f221])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f218])).
fof(f218,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f217])).
fof(f217,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f149])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f219])).
fof(f367,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f219])).
fof(f368,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f219])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f371])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f219])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f216])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f216])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f214])).
fof(f214,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f143])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f215])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f215])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f215])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f212])).
fof(f212,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f144])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f213])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f213])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f213])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f209])).
fof(f209,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f145])).
fof(f350,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f210])).
fof(f351,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f210])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f210])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f206])).
fof(f206,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f146])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f207])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f207])).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f207])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f207])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f207])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f207])).
fof(f331,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f201])).
fof(f201,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f200])).
fof(f200,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f332,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f202])).
fof(f333,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f202])).
fof(f336,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f202])).
fof(f435,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f316,f434])).
fof(f316,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f317,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f195])).
fof(f318,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f195])).
fof(f308,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f124])).
fof(f302,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X7] : ~p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.UNIQUE_MEMBER_SING','thm.pred_set.NOT_IN_EMPTY')).
fof(f259,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.pred_set.UNIQUE_MEMBER_SING','thm.bool.EQ_CLAUSES|split|0')).
fof(f260,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f163])).
fof(f24298,plain,(
  ( ! [X4] : (s(sK12,sK13) = s(sK12,X4)) ) | ($spl8 | $spl69 | $spl72)),
  inference(subsumption_resolution,[],[f24238,f244])).
fof(f24238,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(sK12,sK13) = s(sK12,X4)) ) | ($spl8 | $spl69 | $spl72)),
  inference(backward_demodulation,[],[f24116,f433])).
fof(f24297,plain,(
  $false | ($spl69 | $spl72)),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2032,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f6113,f6114,f6371,f6372,f6373,f6374,f6375,f6466,f6467,f6737,f6740,f6767,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2031,f4402,f4263,f6665,f4352,f4265,f4135,f4264,f4353,f6189,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f307,f306,f305,f4964,f4965,f2026,f6481,f4134,f4401,f6759,f6757,f6755,f6729,f6727,f6725,f6710,f6709,f6705,f6670,f6604,f6603,f6597,f6560,f6559,f6555,f6508,f6506,f6504,f6503,f6496,f6494,f6492,f6491,f6455,f6424,f6175,f6129,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f5936,f5939,f24072,f24071,f24070,f24068,f24067,f24116,f24260,f24263,f24266,f24296])).
fof(f24278,plain,(
  $false | ($spl8 | $spl61 | $spl69 | $spl72)),
  inference(global_subsumption,[],[f24277,f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2032,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f6113,f6114,f6371,f6372,f6373,f6374,f6375,f6466,f6467,f6737,f6740,f6767,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2031,f4402,f4263,f6665,f4352,f4265,f4135,f4264,f4353,f6189,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f307,f306,f305,f4964,f4965,f2026,f6481,f4134,f4401,f6759,f6757,f6755,f6729,f6727,f6725,f6710,f6709,f6705,f6670,f6604,f6603,f6597,f6560,f6559,f6555,f6508,f6506,f6504,f6503,f6496,f6494,f6492,f6491,f6455,f6424,f6175,f6129,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f5936,f5939,f24072,f24071,f24070,f24068,f24067,f24116,f24260,f24263,f24266])).
fof(f24277,plain,(
  ( ! [X7] : (s(fun(sK12,bool),sK14) = s(fun(sK12,bool),X7) | s(sK12,sK13) = s(sK12,sK16(sK14,X7,sK12))) ) | ($spl8 | $spl61 | $spl69 | $spl72)),
  inference(subsumption_resolution,[],[f24184,f24116])).
fof(f24184,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,X7,sK12)),s(fun(sK12,bool),X7))) | s(fun(sK12,bool),sK14) = s(fun(sK12,bool),X7) | s(sK12,sK13) = s(sK12,sK16(sK14,X7,sK12))) ) | ($spl8 | $spl61 | $spl69 | $spl72)),
  inference(backward_demodulation,[],[f24116,f21490])).
fof(f21490,plain,(
  ( ! [X7] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,X7,sK12)),s(fun(sK12,bool),X7))) | s(fun(sK12,bool),sK14) = s(fun(sK12,bool),X7) | s(sK12,sK13) = s(sK12,sK16(sK14,X7,sK12))) ) | ($spl8 | $spl61)),
  inference(superposition,[],[f304,f13542])).
fof(f13542,plain,(
  ( ! [X24] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X24),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X24)) ) | ($spl8 | $spl61)),
  inference(subsumption_resolution,[],[f13231,f4960])).
fof(f13231,plain,(
  ( ! [X24] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X24),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X24) | s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) ) | ($spl8 | $spl61)),
  inference(superposition,[],[f4966,f7425])).
fof(f7425,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl61),
  inference(unit_resulting_resolution,[],[f434,f4960,f704])).
fof(f24276,plain,(
  $false | ($spl8 | $spl61 | $spl69 | $spl72)),
  inference(global_subsumption,[],[f24275,f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2032,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f6113,f6114,f6371,f6372,f6373,f6374,f6375,f6466,f6467,f6737,f6740,f6767,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2031,f4402,f4263,f6665,f4352,f4265,f4135,f4264,f4353,f6189,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f307,f306,f305,f4964,f4965,f2026,f6481,f4134,f4401,f6759,f6757,f6755,f6729,f6727,f6725,f6710,f6709,f6705,f6670,f6604,f6603,f6597,f6560,f6559,f6555,f6508,f6506,f6504,f6503,f6496,f6494,f6492,f6491,f6455,f6424,f6175,f6129,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f5936,f5939,f24072,f24071,f24070,f24068,f24067,f24116,f24260,f24263,f24266])).
fof(f24275,plain,(
  ( ! [X5] : (s(fun(sK12,bool),sK14) = s(fun(sK12,bool),X5) | s(sK12,sK13) = s(sK12,sK16(X5,sK14,sK12))) ) | ($spl8 | $spl61 | $spl69 | $spl72)),
  inference(subsumption_resolution,[],[f24183,f24116])).
fof(f24183,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(X5,sK14,sK12)),s(fun(sK12,bool),X5))) | s(fun(sK12,bool),sK14) = s(fun(sK12,bool),X5) | s(sK12,sK13) = s(sK12,sK16(X5,sK14,sK12))) ) | ($spl8 | $spl61 | $spl69 | $spl72)),
  inference(backward_demodulation,[],[f24116,f21488])).
fof(f21488,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(X5,sK14,sK12)),s(fun(sK12,bool),X5))) | s(fun(sK12,bool),sK14) = s(fun(sK12,bool),X5) | s(sK12,sK13) = s(sK12,sK16(X5,sK14,sK12))) ) | ($spl8 | $spl61)),
  inference(superposition,[],[f304,f13542])).
fof(f24267,plain,(
  $false | ($spl69 | $spl72)),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2032,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f6113,f6114,f6371,f6372,f6373,f6374,f6375,f6466,f6467,f6737,f6740,f6767,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2031,f4402,f4263,f6665,f4352,f4265,f4135,f4264,f4353,f6189,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f307,f306,f305,f4964,f4965,f2026,f6481,f4134,f4401,f6759,f6757,f6755,f6729,f6727,f6725,f6710,f6709,f6705,f6670,f6604,f6603,f6597,f6560,f6559,f6555,f6508,f6506,f6504,f6503,f6496,f6494,f6492,f6491,f6455,f6424,f6175,f6129,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f5936,f5939,f24072,f24071,f24070,f24068,f24067,f24116,f24260,f24263,f24266])).
fof(f24264,plain,(
  $false | ($spl69 | $spl72)),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2032,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f6113,f6114,f6371,f6372,f6373,f6374,f6375,f6466,f6467,f6737,f6740,f6767,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2031,f4402,f4263,f6665,f4352,f4265,f4135,f4264,f4353,f6189,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f307,f306,f305,f4964,f4965,f2026,f6481,f4134,f4401,f6759,f6757,f6755,f6729,f6727,f6725,f6710,f6709,f6705,f6670,f6604,f6603,f6597,f6560,f6559,f6555,f6508,f6506,f6504,f6503,f6496,f6494,f6492,f6491,f6455,f6424,f6175,f6129,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f5936,f5939,f24072,f24071,f24070,f24068,f24067,f24116,f24260,f24263])).
fof(f24262,plain,(
  $false | ($spl69 | $spl72)),
  inference(subsumption_resolution,[],[f24119,f244])).
fof(f24119,plain,(
  ~p(s(bool,'T_0')) | ($spl69 | $spl72)),
  inference(backward_demodulation,[],[f24116,f434])).
fof(f24261,plain,(
  $false | ($spl69 | $spl72)),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2032,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f6113,f6114,f6371,f6372,f6373,f6374,f6375,f6466,f6467,f6737,f6740,f6767,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2031,f4402,f4263,f6665,f4352,f4265,f4135,f4264,f4353,f6189,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f307,f306,f305,f4964,f4965,f2026,f6481,f4134,f4401,f6759,f6757,f6755,f6729,f6727,f6725,f6710,f6709,f6705,f6670,f6604,f6603,f6597,f6560,f6559,f6555,f6508,f6506,f6504,f6503,f6496,f6494,f6492,f6491,f6455,f6424,f6175,f6129,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f5936,f5939,f24072,f24071,f24070,f24068,f24067,f24116,f24260])).
fof(f24115,plain,(
  $false | ($spl69 | $spl75)),
  inference(subsumption_resolution,[],[f24037,f5941])).
fof(f24037,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | $spl69),
  inference(unit_resulting_resolution,[],[f5936,f251])).
fof(f24114,plain,(
  $false | ($spl69 | $spl75)),
  inference(subsumption_resolution,[],[f24039,f5941])).
fof(f24039,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | $spl69),
  inference(unit_resulting_resolution,[],[f465,f5936,f635])).
fof(f24113,plain,(
  $false | ($spl69 | $spl75)),
  inference(subsumption_resolution,[],[f24112,f5941])).
fof(f24112,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | $spl69),
  inference(forward_demodulation,[],[f24040,f505])).
fof(f24040,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl69),
  inference(unit_resulting_resolution,[],[f504,f5936,f635])).
fof(f24111,plain,(
  $false | ($spl5 | $spl69 | $spl75 | $spl80)),
  inference(subsumption_resolution,[],[f24110,f5941])).
fof(f24110,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl69 | $spl80)),
  inference(forward_demodulation,[],[f24041,f6050])).
fof(f6050,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl80)),
  inference(unit_resulting_resolution,[],[f6025,f251])).
fof(f6025,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl80)),
  inference(backward_demodulation,[],[f6021,f2606])).
fof(f2606,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | $spl5),
  inference(unit_resulting_resolution,[],[f430,f304])).
fof(f24041,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl69 | $spl80)),
  inference(unit_resulting_resolution,[],[f6025,f5936,f635])).
fof(f24109,plain,(
  $false | ($spl61 | $spl69 | $spl75)),
  inference(subsumption_resolution,[],[f24108,f5941])).
fof(f24108,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl61 | $spl69)),
  inference(forward_demodulation,[],[f24043,f7425])).
fof(f24043,plain,(
  s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl61 | $spl69)),
  inference(unit_resulting_resolution,[],[f4960,f5936,f635])).
fof(f24107,plain,(
  $false | ($spl65 | $spl69 | $spl75)),
  inference(subsumption_resolution,[],[f24106,f5941])).
fof(f24106,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl65 | $spl69)),
  inference(forward_demodulation,[],[f24044,f11587])).
fof(f11587,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl65),
  inference(unit_resulting_resolution,[],[f434,f4963,f704])).
fof(f24044,plain,(
  s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl65 | $spl69)),
  inference(unit_resulting_resolution,[],[f4963,f5936,f635])).
fof(f24105,plain,(
  $false | ($spl69 | $spl75)),
  inference(subsumption_resolution,[],[f24046,f5941])).
fof(f24046,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | $spl69),
  inference(unit_resulting_resolution,[],[f465,f5936,f635])).
fof(f24104,plain,(
  $false | ($spl69 | $spl75)),
  inference(subsumption_resolution,[],[f24103,f5941])).
fof(f24103,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | $spl69),
  inference(forward_demodulation,[],[f24047,f505])).
fof(f24047,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)))) ) | $spl69),
  inference(unit_resulting_resolution,[],[f504,f5936,f635])).
fof(f24102,plain,(
  $false | ($spl5 | $spl69 | $spl75 | $spl80)),
  inference(subsumption_resolution,[],[f24101,f5941])).
fof(f24101,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl69 | $spl80)),
  inference(forward_demodulation,[],[f24048,f6050])).
fof(f24048,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl69 | $spl80)),
  inference(unit_resulting_resolution,[],[f6025,f5936,f635])).
fof(f24100,plain,(
  $false | ($spl61 | $spl69 | $spl75)),
  inference(subsumption_resolution,[],[f24099,f5941])).
fof(f24099,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl61 | $spl69)),
  inference(forward_demodulation,[],[f24050,f7425])).
fof(f24050,plain,(
  s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl61 | $spl69)),
  inference(unit_resulting_resolution,[],[f4960,f5936,f635])).
fof(f24098,plain,(
  $false | ($spl65 | $spl69 | $spl75)),
  inference(subsumption_resolution,[],[f24097,f5941])).
fof(f24097,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl65 | $spl69)),
  inference(forward_demodulation,[],[f24051,f11587])).
fof(f24051,plain,(
  s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl65 | $spl69)),
  inference(unit_resulting_resolution,[],[f4963,f5936,f635])).
fof(f24096,plain,(
  $false | ($spl69 | $spl75)),
  inference(subsumption_resolution,[],[f24053,f5941])).
fof(f24053,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | $spl69),
  inference(unit_resulting_resolution,[],[f434,f5936,f704])).
fof(f24087,plain,(
  $false | ($spl5 | $spl69 | $spl75 | $spl80)),
  inference(subsumption_resolution,[],[f24086,f5941])).
fof(f24086,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl69 | $spl80)),
  inference(forward_demodulation,[],[f24065,f6050])).
fof(f24065,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl69 | $spl80)),
  inference(unit_resulting_resolution,[],[f5936,f6067])).
fof(f6067,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl5 | $spl80)),
  inference(trivial_inequality_removal,[],[f6061])).
fof(f6061,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl5 | $spl80)),
  inference(superposition,[],[f6025,f635])).
fof(f23950,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23924,f2281])).
fof(f2281,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f425,f349])).
fof(f23924,plain,(
  sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f22125])).
fof(f22125,plain,(
  sP4('T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'F_0') | ($spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f8212,f344])).
fof(f8212,plain,(
  sP5('T_0','F_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f244,f5967,f362])).
fof(f5967,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f5942,f251])).
fof(f5942,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f5938,f2604])).
fof(f2604,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | $spl5),
  inference(unit_resulting_resolution,[],[f430,f304])).
fof(f23896,plain,(
  s(sK12,sK13) = s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)) | $spl70),
  inference(subsumption_resolution,[],[f23895,f434])).
fof(f23895,plain,(
  p(s(bool,'F_0')) | s(sK12,sK13) = s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)) | $spl70),
  inference(forward_demodulation,[],[f23894,f505])).
fof(f23894,plain,(
  s(sK12,sK13) = s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) | $spl70),
  inference(subsumption_resolution,[],[f23770,f244])).
fof(f23770,plain,(
  ~p(s(bool,'T_0')) | s(sK12,sK13) = s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)) | p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) | $spl70),
  inference(superposition,[],[f305,f5938])).
fof(f23949,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23923,f2300])).
fof(f2300,plain,(
  ( ! [X0,X1] : (sP2('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f425,f360])).
fof(f23923,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','F_0') | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f22120])).
fof(f22120,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'F_0','F_0') | ($spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f434,f446,f441,f8211,f347])).
fof(f8211,plain,(
  ~sP5('F_0','F_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f434,f5967,f364])).
fof(f23948,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23922,f2299])).
fof(f2299,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f425,f359])).
fof(f23922,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f22103])).
fof(f22103,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'F_0') | ($spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f434,f444,f441,f8206,f348])).
fof(f8206,plain,(
  ~sP5('F_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'F_0') | ($spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f434,f5967,f364])).
fof(f23947,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23921,f8843])).
fof(f8843,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2260,f3531,f362])).
fof(f3531,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14))))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2255,f307])).
fof(f2255,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f425,f307])).
fof(f2260,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f425,f311])).
fof(f23921,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0')) ) | ($spl0 | $spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f8865])).
fof(f8865,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'T_0')) ) | ($spl0 | $spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f5942,f3531,f363])).
fof(f23946,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23920,f8852])).
fof(f8852,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2260,f3531,f362])).
fof(f23920,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f8859])).
fof(f8859,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f5942,f3531,f363])).
fof(f23945,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23918,f9049])).
fof(f9049,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f3531,f3187])).
fof(f3187,plain,(
  ( ! [X8,X9] : (~sP5(X8,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X9) | p(s(bool,X9)) | ~p(s(bool,X8))) ) | $spl0),
  inference(subsumption_resolution,[],[f3115,f244])).
fof(f3115,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | p(s(bool,X9)) | ~sP5(X8,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X9)) ) | $spl0),
  inference(superposition,[],[f345,f2260])).
fof(f23918,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0')) ) | ($spl0 | $spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f8846])).
fof(f8846,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'F_0')) ) | ($spl0 | $spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f5967,f3531,f362])).
fof(f23944,plain,(
  $false | ($spl0 | $spl5 | $spl26 | $spl70)),
  inference(subsumption_resolution,[],[f23917,f2264])).
fof(f2264,plain,(
  ~sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f244,f425,f332])).
fof(f23917,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl26 | $spl70)),
  inference(backward_demodulation,[],[f23896,f8377])).
fof(f8377,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl26 | $spl70)),
  inference(subsumption_resolution,[],[f8267,f434])).
fof(f8267,plain,(
  p(s(bool,'F_0')) | sP1('const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl26 | $spl70)),
  inference(superposition,[],[f4891,f5967])).
fof(f23943,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23916,f2645])).
fof(f2645,plain,(
  ~sP5('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f2281,f344])).
fof(f23916,plain,(
  sP5('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f8212])).
fof(f23942,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23915,f2712])).
fof(f2712,plain,(
  sP5('F_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f446,f441,f2300,f347])).
fof(f23915,plain,(
  ~sP5('F_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f8211])).
fof(f23941,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23914,f2272])).
fof(f2272,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f425,f345])).
fof(f23914,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f8207])).
fof(f8207,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'F_0') | ($spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f244,f5967,f362])).
fof(f23940,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23913,f2692])).
fof(f2692,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f434,f444,f441,f2299,f348])).
fof(f23913,plain,(
  ~sP5('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f8206])).
fof(f23939,plain,(
  $false | ($spl0 | $spl5 | $spl22 | $spl70)),
  inference(subsumption_resolution,[],[f23912,f2266])).
fof(f2266,plain,(
  ~sP1('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f425,f332])).
fof(f23912,plain,(
  sP1('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl5 | $spl22 | $spl70)),
  inference(backward_demodulation,[],[f23896,f8184])).
fof(f8184,plain,(
  sP1('T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl22 | $spl70)),
  inference(unit_resulting_resolution,[],[f5942,f5496])).
fof(f5496,plain,(
  ( ! [X4] : (sP1('T_0',X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl22),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4263,f4402,f4352,f4265,f4135,f4264,f4353,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f4887,f5495])).
fof(f5495,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | sP1('T_0',X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl22),
  inference(subsumption_resolution,[],[f5486,f259])).
fof(f5486,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | sP1('T_0',X4) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X5)) ) | $spl22),
  inference(superposition,[],[f4887,f635])).
fof(f23938,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23911,f2292])).
fof(f2292,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f425,f353])).
fof(f23911,plain,(
  sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f7881])).
fof(f7881,plain,(
  sP3('T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'F_0') | ($spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f5976,f343])).
fof(f5976,plain,(
  sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f434,f5942,f361])).
fof(f23937,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23910,f2287])).
fof(f2287,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f425,f353])).
fof(f23910,plain,(
  sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f7851])).
fof(f7851,plain,(
  sP3('const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'T_0','F_0') | ($spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f5971,f343])).
fof(f5971,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f434,f5942,f361])).
fof(f23936,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23909,f2286])).
fof(f2286,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f425,f350])).
fof(f23909,plain,(
  ~sP4('T_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f7836])).
fof(f7836,plain,(
  ~sP4('T_0','T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f244,f450,f443,f5970,f346])).
fof(f5970,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f244,f5942,f363])).
fof(f23935,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23907,f2667])).
fof(f2667,plain,(
  ~sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f2292,f343])).
fof(f23907,plain,(
  sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f5976])).
fof(f23934,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23906,f2275])).
fof(f2275,plain,(
  sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f443,f438,f449,f425,f346])).
fof(f23906,plain,(
  ~sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f5975])).
fof(f5975,plain,(
  ~sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70)),
  inference(unit_resulting_resolution,[],[f244,f5942,f363])).
fof(f23933,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23905,f2662])).
fof(f2662,plain,(
  ~sP5('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f2287,f343])).
fof(f23905,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f5971])).
fof(f23932,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23904,f2654])).
fof(f2654,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f244,f443,f450,f2286,f346])).
fof(f23904,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f5970])).
fof(f23931,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23930,f465])).
fof(f23930,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl5 | $spl70)),
  inference(forward_demodulation,[],[f23903,f2260])).
fof(f23903,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f5967])).
fof(f23929,plain,(
  $false | ($spl0 | $spl5 | $spl70)),
  inference(subsumption_resolution,[],[f23902,f2260])).
fof(f23902,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70)),
  inference(backward_demodulation,[],[f23896,f5942])).
fof(f17914,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13590,f13619])).
fof(f13619,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f10273,f10273,f311])).
fof(f13590,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f10273,f263])).
fof(f17913,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13594,f244])).
fof(f13594,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f465,f10273,f311])).
fof(f17912,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13596,f244])).
fof(f13596,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f504,f10273,f311])).
fof(f17911,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f13598,f244])).
fof(f13598,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5942,f10273,f311])).
fof(f17910,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f13600,f244])).
fof(f13600,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6025,f10273,f311])).
fof(f17909,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13602,f244])).
fof(f13602,plain,(
  ~p(s(bool,'T_0')) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f4960,f10273,f311])).
fof(f17908,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13604,f244])).
fof(f13604,plain,(
  ~p(s(bool,'T_0')) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f4963,f10273,f311])).
fof(f17907,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13606,f10273])).
fof(f13606,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f465,f10273,f311])).
fof(f17906,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13608,f10273])).
fof(f13608,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f504,f10273,f311])).
fof(f17905,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f13610,f10273])).
fof(f13610,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)))) | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5942,f10273,f311])).
fof(f17904,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f13612,f10273])).
fof(f13612,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)))) | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6025,f10273,f311])).
fof(f17903,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13614,f10273])).
fof(f13614,plain,(
  ~p(s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f4960,f10273,f311])).
fof(f17902,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13616,f10273])).
fof(f13616,plain,(
  ~p(s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f4963,f10273,f311])).
fof(f17901,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13620,f244])).
fof(f13620,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f465,f10273,f311])).
fof(f17900,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13622,f244])).
fof(f13622,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f504,f10273,f311])).
fof(f17899,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f13624,f244])).
fof(f13624,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5942,f10273,f311])).
fof(f17898,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f13626,f244])).
fof(f13626,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6025,f10273,f311])).
fof(f17897,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13628,f244])).
fof(f13628,plain,(
  ~p(s(bool,'T_0')) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f4960,f10273,f311])).
fof(f17896,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13630,f244])).
fof(f13630,plain,(
  ~p(s(bool,'T_0')) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f4963,f10273,f311])).
fof(f17895,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13632,f10273])).
fof(f13632,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f465,f10273,f311])).
fof(f17894,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13634,f10273])).
fof(f13634,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f504,f10273,f311])).
fof(f17893,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f13636,f10273])).
fof(f13636,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)))) | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5942,f10273,f311])).
fof(f17892,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f13638,f10273])).
fof(f13638,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)))) | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6025,f10273,f311])).
fof(f17891,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13640,f10273])).
fof(f13640,plain,(
  ~p(s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f4960,f10273,f311])).
fof(f17890,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13642,f10273])).
fof(f13642,plain,(
  ~p(s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f4963,f10273,f311])).
fof(f17889,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13648,f10273])).
fof(f13648,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f489,f10273,f332])).
fof(f17888,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13649,f10273])).
fof(f13649,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f510,f10273,f332])).
fof(f17887,plain,(
  $false | ($spl5 | $spl22 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f13650,f10273])).
fof(f13650,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)))) | ($spl5 | $spl22 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f8184,f10273,f332])).
fof(f17886,plain,(
  $false | ($spl5 | $spl22 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f13651,f10273])).
fof(f13651,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)))) | ($spl5 | $spl22 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f8185,f10273,f332])).
fof(f8185,plain,(
  sP1('T_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl22 | $spl80)),
  inference(unit_resulting_resolution,[],[f6025,f5496])).
fof(f17885,plain,(
  $false | ($spl22 | $spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13652,f10273])).
fof(f13652,plain,(
  ~p(s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) | ($spl22 | $spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f8187,f10273,f332])).
fof(f8187,plain,(
  sP1('T_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | ($spl22 | $spl61)),
  inference(unit_resulting_resolution,[],[f4960,f5496])).
fof(f17884,plain,(
  $false | ($spl22 | $spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13653,f10273])).
fof(f13653,plain,(
  ~p(s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | ($spl22 | $spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11589,f10273,f332])).
fof(f11589,plain,(
  sP1('T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl22 | $spl65)),
  inference(unit_resulting_resolution,[],[f4963,f5496])).
fof(f17883,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13654,f244])).
fof(f13654,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f495,f10273,f332])).
fof(f17882,plain,(
  $false | ($spl105 | $spl110)),
  inference(subsumption_resolution,[],[f13655,f10273])).
fof(f13655,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f9383,f10273,f332])).
fof(f9383,plain,(
  sP1('F_0',f(s(bool,'T_0'))) | $spl105),
  inference(unit_resulting_resolution,[],[f434,f8728])).
fof(f8728,plain,(
  ( ! [X6] : (sP1(X6,f(s(bool,'T_0'))) | p(s(bool,X6))) ) | $spl105),
  inference(subsumption_resolution,[],[f8615,f244])).
fof(f8615,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1(X6,f(s(bool,'T_0')))) ) | $spl105),
  inference(superposition,[],[f336,f6844])).
fof(f6844,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl105),
  inference(unit_resulting_resolution,[],[f434,f6479,f704])).
fof(f17881,plain,(
  $false | ($spl10 | $spl110)),
  inference(subsumption_resolution,[],[f13656,f10273])).
fof(f13656,plain,(
  ~p(s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) | ($spl10 | $spl110)),
  inference(unit_resulting_resolution,[],[f9472,f10273,f332])).
fof(f9472,plain,(
  sP1('F_0',g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl10),
  inference(unit_resulting_resolution,[],[f434,f7027])).
fof(f7027,plain,(
  ( ! [X6] : (sP1(X6,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | p(s(bool,X6))) ) | $spl10),
  inference(subsumption_resolution,[],[f6918,f244])).
fof(f6918,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1(X6,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) ) | $spl10),
  inference(superposition,[],[f336,f3920])).
fof(f17880,plain,(
  $false | ($spl14 | $spl110)),
  inference(subsumption_resolution,[],[f13657,f10273])).
fof(f13657,plain,(
  ~p(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | ($spl14 | $spl110)),
  inference(unit_resulting_resolution,[],[f12239,f10273,f332])).
fof(f12239,plain,(
  sP1('F_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl14),
  inference(unit_resulting_resolution,[],[f434,f11331])).
fof(f11331,plain,(
  ( ! [X6] : (sP1(X6,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | p(s(bool,X6))) ) | $spl14),
  inference(subsumption_resolution,[],[f11238,f244])).
fof(f11238,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1(X6,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl14),
  inference(superposition,[],[f336,f3975])).
fof(f17879,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13658,f10273])).
fof(f13658,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2269,f10273,f332])).
fof(f2269,plain,(
  sP1('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f425,f336])).
fof(f17878,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13659,f10273])).
fof(f13659,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3545,f10273,f332])).
fof(f3545,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f2255,f336])).
fof(f17877,plain,(
  $false | ($spl26 | $spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13660,f244])).
fof(f13660,plain,(
  ~p(s(bool,'T_0')) | ($spl26 | $spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f8392,f10273,f332])).
fof(f8392,plain,(
  sP1(f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'T_0') | ($spl26 | $spl61)),
  inference(unit_resulting_resolution,[],[f4960,f5521])).
fof(f5521,plain,(
  ( ! [X4] : (sP1(X4,'T_0') | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4263,f4402,f4352,f4265,f4135,f4264,f4353,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f4891,f5520])).
fof(f5520,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | sP1(X4,'T_0') | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(subsumption_resolution,[],[f5511,f259])).
fof(f5511,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | sP1(X4,'T_0') | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X5)) ) | $spl26),
  inference(superposition,[],[f4891,f635])).
fof(f17876,plain,(
  $false | ($spl26 | $spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13661,f244])).
fof(f13661,plain,(
  ~p(s(bool,'T_0')) | ($spl26 | $spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11590,f10273,f332])).
fof(f11590,plain,(
  sP1(f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0') | ($spl26 | $spl65)),
  inference(unit_resulting_resolution,[],[f4963,f5521])).
fof(f17875,plain,(
  $false | ($spl105 | $spl110)),
  inference(subsumption_resolution,[],[f13662,f10273])).
fof(f13662,plain,(
  ~p(s(bool,'F_0')) | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f9336,f10273,f332])).
fof(f9336,plain,(
  sP1(f(s(bool,'T_0')),'F_0') | $spl105),
  inference(unit_resulting_resolution,[],[f434,f8727])).
fof(f8727,plain,(
  ( ! [X5] : (sP1(f(s(bool,'T_0')),X5) | p(s(bool,X5))) ) | $spl105),
  inference(subsumption_resolution,[],[f8614,f244])).
fof(f8614,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | sP1(f(s(bool,'T_0')),X5)) ) | $spl105),
  inference(superposition,[],[f333,f6844])).
fof(f17874,plain,(
  $false | ($spl10 | $spl110)),
  inference(subsumption_resolution,[],[f13663,f10273])).
fof(f13663,plain,(
  ~p(s(bool,'F_0')) | ($spl10 | $spl110)),
  inference(unit_resulting_resolution,[],[f9408,f10273,f332])).
fof(f9408,plain,(
  sP1(g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f434,f7026])).
fof(f7026,plain,(
  ( ! [X5] : (sP1(g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),X5) | p(s(bool,X5))) ) | $spl10),
  inference(subsumption_resolution,[],[f6917,f244])).
fof(f6917,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | sP1(g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),X5)) ) | $spl10),
  inference(superposition,[],[f333,f3920])).
fof(f17873,plain,(
  $false | ($spl14 | $spl110)),
  inference(subsumption_resolution,[],[f13664,f10273])).
fof(f13664,plain,(
  ~p(s(bool,'F_0')) | ($spl14 | $spl110)),
  inference(unit_resulting_resolution,[],[f12210,f10273,f332])).
fof(f12210,plain,(
  sP1(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f434,f11330])).
fof(f11330,plain,(
  ( ! [X5] : (sP1(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X5) | p(s(bool,X5))) ) | $spl14),
  inference(subsumption_resolution,[],[f11237,f244])).
fof(f11237,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | sP1(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X5)) ) | $spl14),
  inference(superposition,[],[f333,f3975])).
fof(f17872,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13665,f244])).
fof(f13665,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f511,f10273,f332])).
fof(f17871,plain,(
  $false | ($spl5 | $spl26 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f13666,f244])).
fof(f13666,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl26 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f8377,f10273,f332])).
fof(f17870,plain,(
  $false | ($spl5 | $spl26 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f13667,f244])).
fof(f13667,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl26 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f8390,f10273,f332])).
fof(f8390,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl26 | $spl80)),
  inference(unit_resulting_resolution,[],[f6025,f5521])).
fof(f17869,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13668,f10273])).
fof(f13668,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2268,f10273,f332])).
fof(f2268,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f434,f425,f333])).
fof(f17868,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13669,f10273])).
fof(f13669,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3544,f10273,f332])).
fof(f3544,plain,(
  ( ! [X0] : (sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f2255,f333])).
fof(f17867,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13674,f244])).
fof(f13674,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f489,f10273,f332])).
fof(f17866,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13675,f244])).
fof(f13675,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f510,f10273,f332])).
fof(f17865,plain,(
  $false | ($spl5 | $spl22 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f13676,f244])).
fof(f13676,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl22 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f8184,f10273,f332])).
fof(f17864,plain,(
  $false | ($spl5 | $spl22 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f13677,f244])).
fof(f13677,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl22 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f8185,f10273,f332])).
fof(f17863,plain,(
  $false | ($spl22 | $spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13678,f244])).
fof(f13678,plain,(
  ~p(s(bool,'T_0')) | ($spl22 | $spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f8187,f10273,f332])).
fof(f17862,plain,(
  $false | ($spl22 | $spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13679,f244])).
fof(f13679,plain,(
  ~p(s(bool,'T_0')) | ($spl22 | $spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11589,f10273,f332])).
fof(f17861,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13680,f10273])).
fof(f13680,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f495,f10273,f332])).
fof(f17860,plain,(
  $false | ($spl105 | $spl110)),
  inference(subsumption_resolution,[],[f13681,f10273])).
fof(f13681,plain,(
  ~p(s(bool,'F_0')) | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f9383,f10273,f332])).
fof(f17859,plain,(
  $false | ($spl10 | $spl110)),
  inference(subsumption_resolution,[],[f13682,f10273])).
fof(f13682,plain,(
  ~p(s(bool,'F_0')) | ($spl10 | $spl110)),
  inference(unit_resulting_resolution,[],[f9472,f10273,f332])).
fof(f17858,plain,(
  $false | ($spl14 | $spl110)),
  inference(subsumption_resolution,[],[f13683,f10273])).
fof(f13683,plain,(
  ~p(s(bool,'F_0')) | ($spl14 | $spl110)),
  inference(unit_resulting_resolution,[],[f12239,f10273,f332])).
fof(f17857,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13684,f10273])).
fof(f13684,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2269,f10273,f332])).
fof(f17856,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13685,f10273])).
fof(f13685,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3545,f10273,f332])).
fof(f17855,plain,(
  $false | ($spl26 | $spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13686,f10273])).
fof(f13686,plain,(
  ~p(s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) | ($spl26 | $spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f8392,f10273,f332])).
fof(f17854,plain,(
  $false | ($spl26 | $spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13687,f10273])).
fof(f13687,plain,(
  ~p(s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | ($spl26 | $spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11590,f10273,f332])).
fof(f17853,plain,(
  $false | ($spl105 | $spl110)),
  inference(subsumption_resolution,[],[f13688,f10273])).
fof(f13688,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f9336,f10273,f332])).
fof(f17852,plain,(
  $false | ($spl10 | $spl110)),
  inference(subsumption_resolution,[],[f13689,f10273])).
fof(f13689,plain,(
  ~p(s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) | ($spl10 | $spl110)),
  inference(unit_resulting_resolution,[],[f9408,f10273,f332])).
fof(f17851,plain,(
  $false | ($spl14 | $spl110)),
  inference(subsumption_resolution,[],[f13690,f10273])).
fof(f13690,plain,(
  ~p(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | ($spl14 | $spl110)),
  inference(unit_resulting_resolution,[],[f12210,f10273,f332])).
fof(f17850,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13691,f10273])).
fof(f13691,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f511,f10273,f332])).
fof(f17849,plain,(
  $false | ($spl5 | $spl26 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f13692,f10273])).
fof(f13692,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)))) | ($spl5 | $spl26 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f8377,f10273,f332])).
fof(f17848,plain,(
  $false | ($spl5 | $spl26 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f13693,f10273])).
fof(f13693,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)))) | ($spl5 | $spl26 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f8390,f10273,f332])).
fof(f17847,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13694,f10273])).
fof(f13694,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2268,f10273,f332])).
fof(f17846,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13695,f10273])).
fof(f13695,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3544,f10273,f332])).
fof(f17845,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13724,f244])).
fof(f13724,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f739,f434,f10273,f345])).
fof(f17844,plain,(
  $false | ($spl105 | $spl110)),
  inference(subsumption_resolution,[],[f13725,f10273])).
fof(f13725,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f6836,f434,f10273,f345])).
fof(f6836,plain,(
  sP5('F_0',f(s(bool,'T_0')),'F_0') | $spl105),
  inference(unit_resulting_resolution,[],[f434,f6479,f361])).
fof(f17843,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13726,f10273])).
fof(f13726,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2692,f434,f10273,f345])).
fof(f17842,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13727,f10273])).
fof(f13727,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6211,f434,f10273,f345])).
fof(f6211,plain,(
  ( ! [X0] : (sP5('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f444,f441,f3585,f348])).
fof(f3585,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))),X2)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2255,f359])).
fof(f17841,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13728,f10273])).
fof(f13728,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f1325,f434,f10273,f345])).
fof(f17840,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13729,f10273])).
fof(f13729,plain,(
  ~p(s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f9531,f434,f10273,f345])).
fof(f9531,plain,(
  sP5('T_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f244,f7425,f362])).
fof(f17839,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13730,f10273])).
fof(f13730,plain,(
  ~p(s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f12270,f434,f10273,f345])).
fof(f12270,plain,(
  sP5('T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | $spl65),
  inference(unit_resulting_resolution,[],[f244,f11587,f362])).
fof(f17838,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13731,f10273])).
fof(f13731,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f1036,f434,f10273,f345])).
fof(f17837,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13732,f10273])).
fof(f13732,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2676,f434,f10273,f345])).
fof(f2676,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f434,f446,f439,f2298,f347])).
fof(f2298,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f425,f358])).
fof(f17836,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13733,f10273])).
fof(f13733,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2301,f434,f10273,f345])).
fof(f2301,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f505,f425,f362])).
fof(f17835,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13734,f10273])).
fof(f13734,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6193,f434,f10273,f345])).
fof(f6193,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f446,f439,f3584,f347])).
fof(f3584,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2255,f358])).
fof(f17834,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13735,f10273])).
fof(f13735,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3587,f434,f10273,f345])).
fof(f3587,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f505,f2255,f362])).
fof(f17833,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13809,f10273])).
fof(f13809,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f739,f434,f10273,f345])).
fof(f17832,plain,(
  $false | ($spl105 | $spl110)),
  inference(subsumption_resolution,[],[f13810,f10273])).
fof(f13810,plain,(
  ~p(s(bool,'F_0')) | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f6836,f434,f10273,f345])).
fof(f17831,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13811,f10273])).
fof(f13811,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2692,f434,f10273,f345])).
fof(f17830,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13812,f10273])).
fof(f13812,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6211,f434,f10273,f345])).
fof(f17829,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13813,f244])).
fof(f13813,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f1325,f434,f10273,f345])).
fof(f17828,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13814,f244])).
fof(f13814,plain,(
  ~p(s(bool,'T_0')) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f9531,f434,f10273,f345])).
fof(f17827,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13815,f244])).
fof(f13815,plain,(
  ~p(s(bool,'T_0')) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f12270,f434,f10273,f345])).
fof(f17826,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13816,f244])).
fof(f13816,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f1036,f434,f10273,f345])).
fof(f17825,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13817,f10273])).
fof(f13817,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2676,f434,f10273,f345])).
fof(f17824,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13818,f10273])).
fof(f13818,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2301,f434,f10273,f345])).
fof(f17823,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13819,f10273])).
fof(f13819,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6193,f434,f10273,f345])).
fof(f17822,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13820,f10273])).
fof(f13820,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3587,f434,f10273,f345])).
fof(f17821,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13897,f14464])).
fof(f14464,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f10273,f360])).
fof(f13897,plain,(
  ~sP2('F_0','T_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f760,f443,f438,f10273,f346])).
fof(f17820,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13898,f14464])).
fof(f13898,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f941,f443,f438,f10273,f346])).
fof(f17819,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f13899,f14464])).
fof(f13899,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'T_0','T_0') | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5975,f443,f438,f10273,f346])).
fof(f17818,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f13900,f14464])).
fof(f13900,plain,(
  ~sP2('const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)),'T_0','T_0') | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6058,f443,f438,f10273,f346])).
fof(f6058,plain,(
  ~sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl80)),
  inference(unit_resulting_resolution,[],[f244,f6025,f363])).
fof(f17817,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13901,f14464])).
fof(f13901,plain,(
  ~sP2(f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'T_0','T_0') | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7435,f443,f438,f10273,f346])).
fof(f7435,plain,(
  ~sP5('T_0','T_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl61),
  inference(unit_resulting_resolution,[],[f244,f4960,f363])).
fof(f17816,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13902,f14464])).
fof(f13902,plain,(
  ~sP2(f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0','T_0') | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11597,f443,f438,f10273,f346])).
fof(f11597,plain,(
  ~sP5('T_0','T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl65),
  inference(unit_resulting_resolution,[],[f244,f4963,f363])).
fof(f17815,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13903,f14464])).
fof(f13903,plain,(
  ~sP2('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f2297,f438,f10273,f346])).
fof(f2297,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f425,f354])).
fof(f2303,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f465,f425,f363])).
fof(f17814,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13904,f14464])).
fof(f13904,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f2297,f438,f10273,f346])).
fof(f2304,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f504,f425,f363])).
fof(f17813,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13905,f14464])).
fof(f13905,plain,(
  ( ! [X0] : (~sP2('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f3583,f438,f10273,f346])).
fof(f3583,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2255,f354])).
fof(f3591,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f465,f2255,f363])).
fof(f17812,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13906,f14464])).
fof(f13906,plain,(
  ( ! [X2,X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f3583,f438,f10273,f346])).
fof(f3592,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f504,f2255,f363])).
fof(f17811,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13907,f14464])).
fof(f13907,plain,(
  ~sP2('F_0','T_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f760,f444,f438,f10273,f346])).
fof(f17810,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13908,f14464])).
fof(f13908,plain,(
  ~sP2('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f444,f438,f10273,f346])).
fof(f17809,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13909,f14464])).
fof(f13909,plain,(
  ( ! [X0] : (~sP2('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f444,f438,f10273,f346])).
fof(f17808,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13910,f14464])).
fof(f13910,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f941,f514,f438,f10273,f346])).
fof(f17807,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13911,f14464])).
fof(f13911,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f514,f438,f10273,f346])).
fof(f17806,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13912,f14464])).
fof(f13912,plain,(
  ( ! [X2,X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f514,f438,f10273,f346])).
fof(f17805,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13913,f14464])).
fof(f13913,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f443,f2286,f10273,f346])).
fof(f17804,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13914,f14464])).
fof(f13914,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f2297,f2286,f10273,f346])).
fof(f2273,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f425,f434,f425,f345])).
fof(f17803,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13915,f14464])).
fof(f13915,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f444,f2286,f10273,f346])).
fof(f17802,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13916,f14464])).
fof(f13916,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f444,f2286,f10273,f346])).
fof(f17801,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13917,f14464])).
fof(f13917,plain,(
  ( ! [X0] : (~sP2('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f443,f3570,f10273,f346])).
fof(f3570,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2255,f350])).
fof(f3548,plain,(
  ( ! [X0] : (~sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f2255,f345])).
fof(f17800,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13918,f14464])).
fof(f13918,plain,(
  ( ! [X0] : (~sP2('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f444,f3570,f10273,f346])).
fof(f17799,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13919,f14464])).
fof(f13919,plain,(
  ~sP2('F_0','T_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f760,f444,f439,f10273,f346])).
fof(f17798,plain,(
  $false | ($spl105 | $spl110)),
  inference(subsumption_resolution,[],[f13920,f14464])).
fof(f13920,plain,(
  ~sP2('F_0',f(s(bool,'T_0')),'T_0') | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f6837,f444,f439,f10273,f346])).
fof(f6837,plain,(
  ~sP5('T_0',f(s(bool,'T_0')),'F_0') | $spl105),
  inference(unit_resulting_resolution,[],[f244,f6479,f363])).
fof(f17797,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13921,f14464])).
fof(f13921,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f444,f439,f10273,f346])).
fof(f17796,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13922,f14464])).
fof(f13922,plain,(
  ( ! [X0] : (~sP2('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f444,f439,f10273,f346])).
fof(f17795,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13924,f14464])).
fof(f13924,plain,(
  ~sP2('F_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'F_0') | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f9530,f444,f439,f10273,f346])).
fof(f9530,plain,(
  ~sP5('F_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f434,f7425,f364])).
fof(f17794,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13925,f14464])).
fof(f13925,plain,(
  ~sP2('F_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f12269,f444,f439,f10273,f346])).
fof(f12269,plain,(
  ~sP5('F_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | $spl65),
  inference(unit_resulting_resolution,[],[f434,f11587,f364])).
fof(f17793,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13927,f14464])).
fof(f13927,plain,(
  ~sP2('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f444,f439,f10273,f346])).
fof(f17792,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13928,f14464])).
fof(f13928,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f444,f439,f10273,f346])).
fof(f17791,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13929,f14464])).
fof(f13929,plain,(
  ( ! [X0] : (~sP2('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f444,f439,f10273,f346])).
fof(f17790,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13930,f14464])).
fof(f13930,plain,(
  ~sP2('F_0','T_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f760,f443,f439,f10273,f346])).
fof(f17789,plain,(
  $false | ($spl105 | $spl110)),
  inference(subsumption_resolution,[],[f13931,f14464])).
fof(f13931,plain,(
  ~sP2('F_0',f(s(bool,'T_0')),'T_0') | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f6837,f443,f439,f10273,f346])).
fof(f17788,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13932,f14464])).
fof(f13932,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f443,f439,f10273,f346])).
fof(f17787,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13933,f14464])).
fof(f13933,plain,(
  ( ! [X0] : (~sP2('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f443,f439,f10273,f346])).
fof(f17786,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13934,f14464])).
fof(f13934,plain,(
  ~sP2('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f2297,f439,f10273,f346])).
fof(f17785,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13935,f14464])).
fof(f13935,plain,(
  ~sP2('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f2297,f439,f10273,f346])).
fof(f17784,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13936,f14464])).
fof(f13936,plain,(
  ( ! [X0] : (~sP2('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f3583,f439,f10273,f346])).
fof(f17783,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13939,f14464])).
fof(f13939,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f941,f514,f512,f10273,f346])).
fof(f17782,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13941,f14464])).
fof(f13941,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f514,f512,f10273,f346])).
fof(f17781,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13942,f14464])).
fof(f13942,plain,(
  ( ! [X2,X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f514,f512,f10273,f346])).
fof(f17780,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13943,f14464])).
fof(f13943,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f941,f443,f512,f10273,f346])).
fof(f17779,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13944,f14464])).
fof(f13944,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f2297,f512,f10273,f346])).
fof(f17778,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13945,f14464])).
fof(f13945,plain,(
  ( ! [X2,X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f3583,f512,f10273,f346])).
fof(f17777,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13948,f14464])).
fof(f13948,plain,(
  ~sP2('F_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'F_0') | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f9530,f444,f441,f10273,f346])).
fof(f17776,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13949,f14464])).
fof(f13949,plain,(
  ~sP2('F_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f12269,f444,f441,f10273,f346])).
fof(f17775,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13954,f14464])).
fof(f13954,plain,(
  ~sP2(f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'F_0','F_0') | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f9535,f446,f441,f10273,f346])).
fof(f9535,plain,(
  ~sP5('F_0','F_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl61),
  inference(unit_resulting_resolution,[],[f434,f7425,f364])).
fof(f17774,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13955,f14464])).
fof(f13955,plain,(
  ~sP2(f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0','F_0') | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f12274,f446,f441,f10273,f346])).
fof(f12274,plain,(
  ~sP5('F_0','F_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl65),
  inference(unit_resulting_resolution,[],[f434,f11587,f364])).
fof(f17773,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13957,f14396])).
fof(f14396,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f10273,f350])).
fof(f13957,plain,(
  ~sP4('T_0','F_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f760,f443,f448,f10273,f346])).
fof(f17772,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13958,f14396])).
fof(f13958,plain,(
  ( ! [X0,X1] : (~sP4('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f941,f443,f448,f10273,f346])).
fof(f17771,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f13959,f14396])).
fof(f13959,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5975,f443,f448,f10273,f346])).
fof(f17770,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f13960,f14396])).
fof(f13960,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6058,f443,f448,f10273,f346])).
fof(f17769,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13961,f14396])).
fof(f13961,plain,(
  ~sP4('T_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'T_0') | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7435,f443,f448,f10273,f346])).
fof(f17768,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13962,f14396])).
fof(f13962,plain,(
  ~sP4('T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0') | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11597,f443,f448,f10273,f346])).
fof(f17767,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13964,f14396])).
fof(f13964,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2645,f443,f448,f10273,f346])).
fof(f17766,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13965,f14396])).
fof(f13965,plain,(
  ( ! [X0] : (~sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6035,f443,f448,f10273,f346])).
fof(f6035,plain,(
  ( ! [X0] : (~sP5('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3564,f344])).
fof(f3564,plain,(
  ( ! [X0] : (~sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f2255,f349])).
fof(f17765,plain,(
  $false | ($spl105 | $spl110)),
  inference(subsumption_resolution,[],[f13967,f14396])).
fof(f13967,plain,(
  ~sP4('T_0','F_0',f(s(bool,'T_0'))) | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f6837,f443,f448,f10273,f346])).
fof(f17764,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13968,f14396])).
fof(f13968,plain,(
  ~sP4('T_0','T_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7431,f443,f448,f10273,f346])).
fof(f7431,plain,(
  ~sP5('T_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'T_0') | $spl61),
  inference(unit_resulting_resolution,[],[f244,f4960,f363])).
fof(f17763,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13969,f14396])).
fof(f13969,plain,(
  ~sP4('T_0','T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11593,f443,f448,f10273,f346])).
fof(f11593,plain,(
  ~sP5('T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0') | $spl65),
  inference(unit_resulting_resolution,[],[f244,f4963,f363])).
fof(f17762,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f13971,f14396])).
fof(f13971,plain,(
  ~sP4('T_0','T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5970,f443,f448,f10273,f346])).
fof(f17761,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f13972,f14396])).
fof(f13972,plain,(
  ~sP4('T_0','T_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6053,f443,f448,f10273,f346])).
fof(f6053,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl80)),
  inference(unit_resulting_resolution,[],[f244,f6025,f363])).
fof(f17760,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13973,f14396])).
fof(f13973,plain,(
  ~sP4('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f443,f448,f10273,f346])).
fof(f17759,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13974,f14396])).
fof(f13974,plain,(
  ( ! [X0] : (~sP4('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f443,f448,f10273,f346])).
fof(f17758,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13975,f14396])).
fof(f13975,plain,(
  ~sP4('T_0','F_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f760,f444,f448,f10273,f346])).
fof(f17757,plain,(
  $false | ($spl105 | $spl110)),
  inference(subsumption_resolution,[],[f13976,f14396])).
fof(f13976,plain,(
  ~sP4('T_0','F_0',f(s(bool,'T_0'))) | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f6837,f444,f448,f10273,f346])).
fof(f17756,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13977,f14396])).
fof(f13977,plain,(
  ~sP4('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f444,f448,f10273,f346])).
fof(f17755,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13978,f14396])).
fof(f13978,plain,(
  ( ! [X0] : (~sP4('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f444,f448,f10273,f346])).
fof(f17754,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13979,f14396])).
fof(f13979,plain,(
  ( ! [X0,X1] : (~sP4('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f941,f514,f448,f10273,f346])).
fof(f17753,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13981,f14396])).
fof(f13981,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2645,f446,f448,f10273,f346])).
fof(f17752,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f13982,f14396])).
fof(f13982,plain,(
  ( ! [X0] : (~sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6035,f446,f448,f10273,f346])).
fof(f17747,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13990,f14458])).
fof(f14458,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f10273,f354])).
fof(f13990,plain,(
  ~sP3('T_0','F_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f760,f438,f448,f10273,f346])).
fof(f17746,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f13991,f14458])).
fof(f13991,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f941,f438,f448,f10273,f346])).
fof(f17745,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f13992,f14458])).
fof(f13992,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5975,f438,f448,f10273,f346])).
fof(f17744,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f13993,f14458])).
fof(f13993,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6058,f438,f448,f10273,f346])).
fof(f17743,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f13994,f14458])).
fof(f13994,plain,(
  ~sP3('T_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'T_0') | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7435,f438,f448,f10273,f346])).
fof(f17742,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f13995,f14458])).
fof(f13995,plain,(
  ~sP3('T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0') | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11597,f438,f448,f10273,f346])).
fof(f17741,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14002,f14458])).
fof(f14002,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f2286,f448,f10273,f346])).
fof(f17739,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14008,f14458])).
fof(f14008,plain,(
  ( ! [X0] : (~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f3570,f448,f10273,f346])).
fof(f17737,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14014,f14458])).
fof(f14014,plain,(
  ~sP3('T_0','F_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f760,f439,f448,f10273,f346])).
fof(f17736,plain,(
  $false | ($spl105 | $spl110)),
  inference(subsumption_resolution,[],[f14015,f14458])).
fof(f14015,plain,(
  ~sP3(f(s(bool,'T_0')),'F_0','T_0') | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f6837,f439,f448,f10273,f346])).
fof(f17735,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14016,f14458])).
fof(f14016,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f439,f448,f10273,f346])).
fof(f17734,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14017,f14458])).
fof(f14017,plain,(
  ( ! [X0] : (~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f439,f448,f10273,f346])).
fof(f17733,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14024,f14458])).
fof(f14024,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f941,f512,f448,f10273,f346])).
fof(f17732,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14031,f14396])).
fof(f14031,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f2297,f2298,f10273,f346])).
fof(f17731,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14032,f14396])).
fof(f14032,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f2297,f2298,f10273,f346])).
fof(f17730,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14034,f14396])).
fof(f14034,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3791,f2297,f2298,f10273,f346])).
fof(f3791,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f2282,f344])).
fof(f2282,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f425,f434,f425,f349])).
fof(f17729,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14035,f14396])).
fof(f14035,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f2297,f2298,f10273,f346])).
fof(f17728,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14036,f14396])).
fof(f14036,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2306,f2297,f2298,f10273,f346])).
fof(f2306,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f504,f425,f363])).
fof(f17727,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14037,f14396])).
fof(f14037,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f444,f2298,f10273,f346])).
fof(f17726,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14038,f14396])).
fof(f14038,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f444,f2298,f10273,f346])).
fof(f17725,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14039,f14396])).
fof(f14039,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f514,f2298,f10273,f346])).
fof(f17724,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14041,f14396])).
fof(f14041,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3791,f446,f2298,f10273,f346])).
fof(f17723,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14042,f14396])).
fof(f14042,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2306,f515,f2298,f10273,f346])).
fof(f17722,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14043,f14458])).
fof(f14043,plain,(
  ~sP3('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f438,f2298,f10273,f346])).
fof(f17721,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14044,f14458])).
fof(f14044,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f438,f2298,f10273,f346])).
fof(f17719,plain,(
  $false | ($spl0 | $spl28 | $spl110)),
  inference(subsumption_resolution,[],[f14047,f15705])).
fof(f15705,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | ($spl28 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f10273,f4894])).
fof(f14047,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f444,f438,f2298,f10273,f346])).
fof(f17718,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14049,f14604])).
fof(f14604,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f504,f10273,f363])).
fof(f14049,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f514,f438,f2298,f10273,f346])).
fof(f17717,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14051,f14458])).
fof(f14051,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f2286,f2298,f10273,f346])).
fof(f17709,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14060,f14458])).
fof(f14060,plain,(
  ~sP3('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f439,f2298,f10273,f346])).
fof(f17708,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14061,f14458])).
fof(f14061,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f439,f2298,f10273,f346])).
fof(f17704,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14068,f14458])).
fof(f14068,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f512,f2298,f10273,f346])).
fof(f17700,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14075,f14396])).
fof(f14075,plain,(
  ( ! [X0] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f3583,f3584,f10273,f346])).
fof(f17699,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14076,f14396])).
fof(f14076,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f3583,f3584,f10273,f346])).
fof(f17698,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14078,f14396])).
fof(f14078,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3594,f3583,f3584,f10273,f346])).
fof(f3594,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f504,f2255,f363])).
fof(f17697,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14079,f14396])).
fof(f14079,plain,(
  ( ! [X0] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f444,f3584,f10273,f346])).
fof(f17696,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14080,f14396])).
fof(f14080,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f514,f3584,f10273,f346])).
fof(f17695,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14082,f14396])).
fof(f14082,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3594,f515,f3584,f10273,f346])).
fof(f17694,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14083,f14458])).
fof(f14083,plain,(
  ( ! [X0] : (~sP3('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f438,f3584,f10273,f346])).
fof(f17693,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14084,f14458])).
fof(f14084,plain,(
  ( ! [X2,X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f438,f3584,f10273,f346])).
fof(f17691,plain,(
  $false | ($spl0 | $spl28 | $spl110)),
  inference(subsumption_resolution,[],[f14087,f15705])).
fof(f14087,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f444,f438,f3584,f10273,f346])).
fof(f17690,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14089,f14604])).
fof(f14089,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f514,f438,f3584,f10273,f346])).
fof(f17681,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14097,f14458])).
fof(f14097,plain,(
  ( ! [X0] : (~sP3('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f439,f3584,f10273,f346])).
fof(f17677,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14104,f14458])).
fof(f14104,plain,(
  ( ! [X2,X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f512,f3584,f10273,f346])).
fof(f17673,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14111,f14396])).
fof(f14111,plain,(
  ~sP4('T_0','F_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f760,f443,f449,f10273,f346])).
fof(f17672,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14112,f14396])).
fof(f14112,plain,(
  ( ! [X0,X1] : (~sP4('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f941,f443,f449,f10273,f346])).
fof(f17671,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14113,f14396])).
fof(f14113,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5975,f443,f449,f10273,f346])).
fof(f17670,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f14114,f14396])).
fof(f14114,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6058,f443,f449,f10273,f346])).
fof(f17669,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f14115,f14396])).
fof(f14115,plain,(
  ~sP4('T_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'T_0') | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7435,f443,f449,f10273,f346])).
fof(f17668,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f14116,f14396])).
fof(f14116,plain,(
  ~sP4('T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0') | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11597,f443,f449,f10273,f346])).
fof(f17667,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14117,f14396])).
fof(f14117,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f2297,f449,f10273,f346])).
fof(f17666,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14118,f14396])).
fof(f14118,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f2297,f449,f10273,f346])).
fof(f17665,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14119,f14396])).
fof(f14119,plain,(
  ( ! [X0] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f3583,f449,f10273,f346])).
fof(f17664,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14120,f14396])).
fof(f14120,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f3583,f449,f10273,f346])).
fof(f17663,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14121,f14396])).
fof(f14121,plain,(
  ~sP4('T_0','F_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f760,f444,f449,f10273,f346])).
fof(f17662,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14122,f14396])).
fof(f14122,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f444,f449,f10273,f346])).
fof(f17661,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14123,f14396])).
fof(f14123,plain,(
  ( ! [X0] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f444,f449,f10273,f346])).
fof(f17660,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14124,f14396])).
fof(f14124,plain,(
  ( ! [X0,X1] : (~sP4('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f941,f514,f449,f10273,f346])).
fof(f17659,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14125,f14396])).
fof(f14125,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f514,f449,f10273,f346])).
fof(f17658,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14126,f14396])).
fof(f14126,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f514,f449,f10273,f346])).
fof(f17657,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14127,f14458])).
fof(f14127,plain,(
  ~sP3('T_0','F_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f760,f438,f449,f10273,f346])).
fof(f17656,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14128,f14458])).
fof(f14128,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f941,f438,f449,f10273,f346])).
fof(f17655,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14129,f14458])).
fof(f14129,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5975,f438,f449,f10273,f346])).
fof(f17654,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f14130,f14458])).
fof(f14130,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6058,f438,f449,f10273,f346])).
fof(f17653,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f14131,f14458])).
fof(f14131,plain,(
  ~sP3('T_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'T_0') | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7435,f438,f449,f10273,f346])).
fof(f17652,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f14132,f14458])).
fof(f14132,plain,(
  ~sP3('T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0') | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11597,f438,f449,f10273,f346])).
fof(f17651,plain,(
  $false | ($spl10 | $spl110)),
  inference(subsumption_resolution,[],[f14135,f14458])).
fof(f14135,plain,(
  ~sP3('T_0',g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'F_0') | ($spl10 | $spl110)),
  inference(unit_resulting_resolution,[],[f6901,f438,f449,f10273,f346])).
fof(f6901,plain,(
  ~sP5('F_0','T_0',g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl10),
  inference(unit_resulting_resolution,[],[f434,f3920,f364])).
fof(f17650,plain,(
  $false | ($spl14 | $spl110)),
  inference(subsumption_resolution,[],[f14136,f14458])).
fof(f14136,plain,(
  ~sP3('T_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | ($spl14 | $spl110)),
  inference(unit_resulting_resolution,[],[f11222,f438,f449,f10273,f346])).
fof(f11222,plain,(
  ~sP5('F_0','T_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl14),
  inference(unit_resulting_resolution,[],[f434,f3975,f364])).
fof(f17649,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14137,f14458])).
fof(f14137,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2667,f438,f449,f10273,f346])).
fof(f17648,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14138,f14458])).
fof(f14138,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6141,f438,f449,f10273,f346])).
fof(f6141,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3577,f343])).
fof(f3577,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f2255,f353])).
fof(f17647,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14139,f14458])).
fof(f14139,plain,(
  ~sP3('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f438,f449,f10273,f346])).
fof(f17646,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14140,f14458])).
fof(f14140,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f438,f449,f10273,f346])).
fof(f17645,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14141,f14458])).
fof(f14141,plain,(
  ( ! [X0] : (~sP3('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f438,f449,f10273,f346])).
fof(f17644,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14142,f14458])).
fof(f14142,plain,(
  ( ! [X2,X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f438,f449,f10273,f346])).
fof(f17639,plain,(
  $false | ($spl28 | $spl110)),
  inference(subsumption_resolution,[],[f14151,f15705])).
fof(f14151,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f444,f438,f449,f10273,f346])).
fof(f17638,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14153,f14604])).
fof(f14153,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f514,f438,f449,f10273,f346])).
fof(f17637,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14155,f14458])).
fof(f14155,plain,(
  ~sP3('T_0','F_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f760,f439,f449,f10273,f346])).
fof(f17636,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14156,f14458])).
fof(f14156,plain,(
  ~sP3('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f439,f449,f10273,f346])).
fof(f17635,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14157,f14458])).
fof(f14157,plain,(
  ( ! [X0] : (~sP3('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f439,f449,f10273,f346])).
fof(f17634,plain,(
  $false | ($spl28 | $spl110)),
  inference(subsumption_resolution,[],[f14158,f15705])).
fof(f14158,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f444,f439,f449,f10273,f346])).
fof(f17633,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14160,f760])).
fof(f14160,plain,(
  sP5('T_0','T_0','F_0') | $spl110),
  inference(unit_resulting_resolution,[],[f443,f439,f449,f10273,f346])).
fof(f17632,plain,(
  $false | ($spl0 | $spl28 | $spl110)),
  inference(subsumption_resolution,[],[f14162,f15705])).
fof(f14162,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2297,f439,f449,f10273,f346])).
fof(f17631,plain,(
  $false | ($spl0 | $spl28 | $spl110)),
  inference(subsumption_resolution,[],[f14164,f15705])).
fof(f14164,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3583,f439,f449,f10273,f346])).
fof(f17630,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14166,f14458])).
fof(f14166,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f941,f512,f449,f10273,f346])).
fof(f17629,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14167,f14458])).
fof(f14167,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f512,f449,f10273,f346])).
fof(f17628,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14168,f14458])).
fof(f14168,plain,(
  ( ! [X2,X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f512,f449,f10273,f346])).
fof(f17627,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14169,f14604])).
fof(f14169,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f514,f512,f449,f10273,f346])).
fof(f17626,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14171,f941])).
fof(f14171,plain,(
  ( ! [X0,X1] : (sP5('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f443,f512,f449,f10273,f346])).
fof(f17625,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14173,f14604])).
fof(f14173,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2297,f512,f449,f10273,f346])).
fof(f17624,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14175,f14604])).
fof(f14175,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3583,f512,f449,f10273,f346])).
fof(f17623,plain,(
  $false | ($spl10 | $spl110)),
  inference(subsumption_resolution,[],[f14179,f14458])).
fof(f14179,plain,(
  ~sP3('T_0',g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'F_0') | ($spl10 | $spl110)),
  inference(unit_resulting_resolution,[],[f6901,f441,f449,f10273,f346])).
fof(f17622,plain,(
  $false | ($spl14 | $spl110)),
  inference(subsumption_resolution,[],[f14180,f14458])).
fof(f14180,plain,(
  ~sP3('T_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | ($spl14 | $spl110)),
  inference(unit_resulting_resolution,[],[f11222,f441,f449,f10273,f346])).
fof(f17621,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14181,f14458])).
fof(f14181,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2667,f441,f449,f10273,f346])).
fof(f17620,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14182,f14458])).
fof(f14182,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6141,f441,f449,f10273,f346])).
fof(f17617,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14187,f16177])).
fof(f16177,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2084])).
fof(f16152,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16033,f16016])).
fof(f16016,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(global_subsumption,[],[f15859,f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f6113,f6114,f6371,f6372,f6373,f6374,f6375,f6466,f6467,f6737,f6740,f6767,f6189,f6481,f6759,f6729,f6508,f6496,f6175,f6129,f1981,f1977,f1975,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f434,f264,f4263,f4402,f4264,f4356,f4355,f4354,f4261,f4088,f4087,f6665,f4353,f4352,f4265,f4262,f4135,f4089,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f6757,f6755,f6727,f6725,f6710,f6709,f6705,f6670,f6604,f6603,f6597,f6560,f6559,f6555,f6506,f6504,f6503,f6494,f6492,f6491,f6455,f6424,f1979,f1973,f505,f263,f10273])).
fof(f15859,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f15769,f5984])).
fof(f5984,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl5 | $spl70)),
  inference(trivial_inequality_removal,[],[f5978])).
fof(f5978,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl5 | $spl70)),
  inference(superposition,[],[f5942,f635])).
fof(f15769,plain,(
  ( ! [X0] : (s(sK12,sK13) = s(sK12,X0)) ) | ($spl8 | $spl110)),
  inference(unit_resulting_resolution,[],[f10273,f433])).
fof(f16033,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK16(X1,'const.pred_set.EMPTY_0',X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16016,f1973])).
fof(f14187,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f444,f513,f449,f10273,f346])).
fof(f17616,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14188,f16178])).
fof(f16178,plain,(
  ( ! [X4,X0,X3,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2086])).
fof(f14188,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f514,f513,f449,f10273,f346])).
fof(f17615,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14189,f14396])).
fof(f14189,plain,(
  ~sP4('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f443,f2299,f10273,f346])).
fof(f17614,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14190,f14396])).
fof(f14190,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f2297,f2299,f10273,f346])).
fof(f17613,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14191,f14396])).
fof(f14191,plain,(
  ~sP4('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f444,f2299,f10273,f346])).
fof(f17612,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14192,f14396])).
fof(f14192,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f444,f2299,f10273,f346])).
fof(f17611,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14193,f14458])).
fof(f14193,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f2286,f2299,f10273,f346])).
fof(f17610,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14195,f14458])).
fof(f14195,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3804,f2286,f2299,f10273,f346])).
fof(f3804,plain,(
  ~sP5('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f2293,f343])).
fof(f2293,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f425,f434,f425,f353])).
fof(f17609,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14196,f14458])).
fof(f14196,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f2286,f2299,f10273,f346])).
fof(f17602,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14205,f14458])).
fof(f14205,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f439,f2299,f10273,f346])).
fof(f17601,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14206,f14458])).
fof(f14206,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f439,f2299,f10273,f346])).
fof(f17592,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14219,f14458])).
fof(f14219,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3804,f441,f2299,f10273,f346])).
fof(f17587,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14224,f14396])).
fof(f14224,plain,(
  ( ! [X0] : (~sP4('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f443,f3585,f10273,f346])).
fof(f17586,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14225,f14396])).
fof(f14225,plain,(
  ( ! [X0] : (~sP4('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f444,f3585,f10273,f346])).
fof(f17585,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14226,f14458])).
fof(f14226,plain,(
  ( ! [X0] : (~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f3570,f3585,f10273,f346])).
fof(f17577,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14235,f14458])).
fof(f14235,plain,(
  ( ! [X0] : (~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f439,f3585,f10273,f346])).
fof(f17562,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f14252,f14396])).
fof(f14252,plain,(
  ~sP4('T_0','T_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7431,f443,f450,f10273,f346])).
fof(f17561,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f14253,f14396])).
fof(f14253,plain,(
  ~sP4('T_0','T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11593,f443,f450,f10273,f346])).
fof(f17560,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14255,f14396])).
fof(f14255,plain,(
  ~sP4('T_0','T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5970,f443,f450,f10273,f346])).
fof(f17559,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f14256,f14396])).
fof(f14256,plain,(
  ~sP4('T_0','T_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6053,f443,f450,f10273,f346])).
fof(f17558,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14258,f14396])).
fof(f14258,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2306,f2297,f450,f10273,f346])).
fof(f17557,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14260,f14396])).
fof(f14260,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3594,f3583,f450,f10273,f346])).
fof(f17556,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14265,f14396])).
fof(f14265,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2306,f515,f450,f10273,f346])).
fof(f17555,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14266,f14396])).
fof(f14266,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3594,f515,f450,f10273,f346])).
fof(f17552,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f17551,f938])).
fof(f17551,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14272,f16152])).
fof(f14272,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f443,f2286,f450,f10273,f346])).
fof(f17549,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f17548,f17355])).
fof(f17355,plain,(
  ( ! [X0] : (~sP5(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(forward_demodulation,[],[f17354,f15769])).
fof(f17354,plain,(
  ( ! [X0] : (~sP5(X0,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(forward_demodulation,[],[f14620,f16152])).
fof(f14620,plain,(
  ( ! [X0] : (~sP5(X0,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)),'T_0')) ) | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6025,f10273,f363])).
fof(f17548,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14274,f16152])).
fof(f14274,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3583,f2286,f450,f10273,f346])).
fof(f17547,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f17546,f938])).
fof(f17546,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14276,f16152])).
fof(f14276,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f443,f3570,f450,f10273,f346])).
fof(f17545,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f17544,f17355])).
fof(f17544,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14277,f16152])).
fof(f14277,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2297,f3570,f450,f10273,f346])).
fof(f17543,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f17542,f17355])).
fof(f17542,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14278,f16152])).
fof(f14278,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3583,f3570,f450,f10273,f346])).
fof(f17537,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14296,f14396])).
fof(f14296,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2645,f443,f2300,f10273,f346])).
fof(f17536,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14297,f14396])).
fof(f14297,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3791,f2297,f2300,f10273,f346])).
fof(f17535,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14298,f14396])).
fof(f14298,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2645,f446,f2300,f10273,f346])).
fof(f17534,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14299,f14396])).
fof(f14299,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3791,f446,f2300,f10273,f346])).
fof(f17533,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14300,f14458])).
fof(f14300,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2667,f438,f2300,f10273,f346])).
fof(f17532,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f17531,f941])).
fof(f17531,plain,(
  sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14301,f16152])).
fof(f14301,plain,(
  sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f443,f438,f2300,f10273,f346])).
fof(f17530,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f17529,f16178])).
fof(f17529,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14302,f16152])).
fof(f14302,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2297,f438,f2300,f10273,f346])).
fof(f17528,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f17527,f17361])).
fof(f17361,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(forward_demodulation,[],[f17360,f15769])).
fof(f17360,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(forward_demodulation,[],[f14608,f16152])).
fof(f14608,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)))) ) | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6025,f10273,f363])).
fof(f17527,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14303,f16152])).
fof(f14303,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3583,f438,f2300,f10273,f346])).
fof(f17526,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14304,f14458])).
fof(f14304,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3804,f2286,f2300,f10273,f346])).
fof(f17525,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f17524,f16598])).
fof(f16598,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16225,f16152])).
fof(f16225,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f15848])).
fof(f15848,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl5 | $spl8 | $spl110)),
  inference(backward_demodulation,[],[f15769,f5924])).
fof(f5924,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | $spl5),
  inference(unit_resulting_resolution,[],[f244,f2604,f363])).
fof(f17524,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14305,f16152])).
fof(f14305,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f443,f2286,f2300,f10273,f346])).
fof(f17520,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f17519,f16598])).
fof(f17519,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f17518,f16152])).
fof(f17518,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14308,f16152])).
fof(f14308,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f443,f3570,f2300,f10273,f346])).
fof(f17512,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14311,f14458])).
fof(f14311,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2667,f441,f2300,f10273,f346])).
fof(f17511,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14312,f14458])).
fof(f14312,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3804,f441,f2300,f10273,f346])).
fof(f17506,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14317,f14396])).
fof(f14317,plain,(
  ( ! [X0] : (~sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6035,f443,f3586,f10273,f346])).
fof(f3586,plain,(
  ( ! [X2,X0,X1] : (sP2('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),X1,X2)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2255,f360])).
fof(f17505,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14318,f14396])).
fof(f14318,plain,(
  ( ! [X0] : (~sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6035,f446,f3586,f10273,f346])).
fof(f17504,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14319,f14458])).
fof(f14319,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6141,f438,f3586,f10273,f346])).
fof(f17503,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f17502,f941])).
fof(f17502,plain,(
  sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14320,f16152])).
fof(f14320,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f443,f438,f3586,f10273,f346])).
fof(f17501,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f17500,f17361])).
fof(f17500,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14321,f16152])).
fof(f14321,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2297,f438,f3586,f10273,f346])).
fof(f17499,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f17498,f17361])).
fof(f17498,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14322,f16152])).
fof(f14322,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3583,f438,f3586,f10273,f346])).
fof(f17497,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f17496,f16598])).
fof(f17496,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f17495,f16152])).
fof(f17495,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14323,f16152])).
fof(f14323,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f443,f2286,f3586,f10273,f346])).
fof(f17489,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f17488,f16598])).
fof(f17488,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f17487,f16152])).
fof(f17487,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14326,f16152])).
fof(f14326,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f443,f3570,f3586,f10273,f346])).
fof(f17480,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14329,f14458])).
fof(f14329,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6141,f441,f3586,f10273,f346])).
fof(f17475,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14334,f10273])).
fof(f14334,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f439,f434,f10273,f349])).
fof(f17474,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14335,f10273])).
fof(f14335,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f512,f434,f10273,f349])).
fof(f17473,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f14336,f10273])).
fof(f14336,plain,(
  ~p(s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f9777,f434,f10273,f349])).
fof(f9777,plain,(
  sP4('T_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f9536,f344])).
fof(f9536,plain,(
  sP5('T_0','F_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl61),
  inference(unit_resulting_resolution,[],[f244,f7425,f362])).
fof(f17472,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f14337,f10273])).
fof(f14337,plain,(
  ~p(s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f12496,f434,f10273,f349])).
fof(f12496,plain,(
  sP4('T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | $spl65),
  inference(unit_resulting_resolution,[],[f12275,f344])).
fof(f12275,plain,(
  sP5('T_0','F_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl65),
  inference(unit_resulting_resolution,[],[f244,f11587,f362])).
fof(f17471,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14338,f10273])).
fof(f14338,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f441,f434,f10273,f349])).
fof(f17470,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14339,f10273])).
fof(f14339,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f513,f434,f10273,f349])).
fof(f17469,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14362,f14396])).
fof(f14362,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f434,f10273,f10273,f349])).
fof(f17468,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14364,f10273])).
fof(f14364,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f439,f434,f10273,f349])).
fof(f17467,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14365,f10273])).
fof(f14365,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f512,f434,f10273,f349])).
fof(f17466,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f14366,f244])).
fof(f14366,plain,(
  ~p(s(bool,'T_0')) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f9777,f434,f10273,f349])).
fof(f17465,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f14367,f244])).
fof(f14367,plain,(
  ~p(s(bool,'T_0')) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f12496,f434,f10273,f349])).
fof(f17464,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14368,f10273])).
fof(f14368,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f441,f434,f10273,f349])).
fof(f17463,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14369,f10273])).
fof(f14369,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f513,f434,f10273,f349])).
fof(f17462,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14380,f14396])).
fof(f14380,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f434,f244,f10273,f349])).
fof(f17461,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14394,f14396])).
fof(f14394,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f434,f10273,f10273,f349])).
fof(f17460,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14398,f10273])).
fof(f14398,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f444,f434,f10273,f353])).
fof(f17459,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14399,f10273])).
fof(f14399,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f514,f434,f10273,f353])).
fof(f17458,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f14400,f10273])).
fof(f14400,plain,(
  ~p(s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7900,f434,f10273,f353])).
fof(f7900,plain,(
  sP3('T_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f7434,f343])).
fof(f7434,plain,(
  sP5('F_0','T_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl61),
  inference(unit_resulting_resolution,[],[f434,f4960,f361])).
fof(f17457,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f14401,f10273])).
fof(f14401,plain,(
  ~p(s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11885,f434,f10273,f353])).
fof(f11885,plain,(
  sP3('T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | $spl65),
  inference(unit_resulting_resolution,[],[f11596,f343])).
fof(f11596,plain,(
  sP5('F_0','T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl65),
  inference(unit_resulting_resolution,[],[f434,f4963,f361])).
fof(f17456,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14402,f10273])).
fof(f14402,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f446,f434,f10273,f353])).
fof(f17455,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f14403,f244])).
fof(f14403,plain,(
  ~p(s(bool,'T_0')) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7889,f434,f10273,f353])).
fof(f7889,plain,(
  sP3(f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'T_0','F_0') | $spl61),
  inference(unit_resulting_resolution,[],[f7430,f343])).
fof(f7430,plain,(
  sP5('F_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'T_0') | $spl61),
  inference(unit_resulting_resolution,[],[f434,f4960,f361])).
fof(f17454,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f14404,f244])).
fof(f14404,plain,(
  ~p(s(bool,'T_0')) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11828,f434,f10273,f353])).
fof(f11828,plain,(
  sP3(f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0','F_0') | $spl65),
  inference(unit_resulting_resolution,[],[f11592,f343])).
fof(f11592,plain,(
  sP5('F_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0') | $spl65),
  inference(unit_resulting_resolution,[],[f434,f4963,f361])).
fof(f17453,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14405,f10273])).
fof(f14405,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f515,f434,f10273,f353])).
fof(f17452,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14426,f14458])).
fof(f14426,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f434,f10273,f10273,f353])).
fof(f17451,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14428,f10273])).
fof(f14428,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f444,f434,f10273,f353])).
fof(f17450,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14429,f10273])).
fof(f14429,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f514,f434,f10273,f353])).
fof(f17449,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f14430,f244])).
fof(f14430,plain,(
  ~p(s(bool,'T_0')) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7900,f434,f10273,f353])).
fof(f17448,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f14431,f244])).
fof(f14431,plain,(
  ~p(s(bool,'T_0')) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11885,f434,f10273,f353])).
fof(f17447,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14432,f10273])).
fof(f14432,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f446,f434,f10273,f353])).
fof(f17446,plain,(
  $false | ($spl61 | $spl110)),
  inference(subsumption_resolution,[],[f14433,f10273])).
fof(f14433,plain,(
  ~p(s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7889,f434,f10273,f353])).
fof(f17445,plain,(
  $false | ($spl65 | $spl110)),
  inference(subsumption_resolution,[],[f14434,f10273])).
fof(f14434,plain,(
  ~p(s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11828,f434,f10273,f353])).
fof(f17444,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14435,f10273])).
fof(f14435,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f515,f434,f10273,f353])).
fof(f17443,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14443,f14458])).
fof(f14443,plain,(
  ( ! [X0] : (~sP3('T_0',X0,'F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f434,f244,f10273,f353])).
fof(f17442,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14456,f14458])).
fof(f14456,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f434,f10273,f10273,f353])).
fof(f17441,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14468,f16016])).
fof(f14468,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5975,f10273,f362])).
fof(f17440,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f14469,f16016])).
fof(f14469,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6058,f10273,f362])).
fof(f17439,plain,(
  $false | ($spl5 | $spl8 | $spl61 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14470,f16016])).
fof(f14470,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7435,f10273,f362])).
fof(f17438,plain,(
  $false | ($spl5 | $spl8 | $spl65 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14471,f16016])).
fof(f14471,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11597,f10273,f362])).
fof(f17433,plain,(
  $false | ($spl5 | $spl8 | $spl61 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14477,f16016])).
fof(f14477,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7431,f10273,f362])).
fof(f17432,plain,(
  $false | ($spl5 | $spl8 | $spl65 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14478,f16016])).
fof(f14478,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11593,f10273,f362])).
fof(f17431,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14480,f16016])).
fof(f14480,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5970,f10273,f362])).
fof(f17430,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f14481,f16016])).
fof(f14481,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6053,f10273,f362])).
fof(f17427,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f14485,f16016])).
fof(f14485,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f8592,f10273,f362])).
fof(f8592,plain,(
  ~sP5('F_0','T_0',f(s(bool,'T_0'))) | $spl105),
  inference(unit_resulting_resolution,[],[f434,f6844,f364])).
fof(f17426,plain,(
  $false | ($spl5 | $spl8 | $spl10 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14486,f16016])).
fof(f14486,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | ($spl10 | $spl110)),
  inference(unit_resulting_resolution,[],[f6901,f10273,f362])).
fof(f17425,plain,(
  $false | ($spl5 | $spl8 | $spl14 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14487,f16016])).
fof(f14487,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl14 | $spl110)),
  inference(unit_resulting_resolution,[],[f11222,f10273,f362])).
fof(f17424,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14488,f16016])).
fof(f14488,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2667,f10273,f362])).
fof(f17423,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14489,f16016])).
fof(f14489,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6141,f10273,f362])).
fof(f17417,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl105 | $spl110)),
  inference(subsumption_resolution,[],[f14496,f16016])).
fof(f14496,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f8589,f10273,f362])).
fof(f8589,plain,(
  ~sP5('F_0',f(s(bool,'T_0')),'T_0') | $spl105),
  inference(unit_resulting_resolution,[],[f434,f6844,f364])).
fof(f17416,plain,(
  $false | ($spl5 | $spl8 | $spl10 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14497,f16016])).
fof(f14497,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | ($spl10 | $spl110)),
  inference(unit_resulting_resolution,[],[f6898,f10273,f362])).
fof(f6898,plain,(
  ~sP5('F_0',g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f434,f3920,f364])).
fof(f17415,plain,(
  $false | ($spl5 | $spl8 | $spl14 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14498,f16016])).
fof(f14498,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl14 | $spl110)),
  inference(unit_resulting_resolution,[],[f11220,f10273,f362])).
fof(f11220,plain,(
  ~sP5('F_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f434,f3975,f364])).
fof(f17413,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14500,f16016])).
fof(f14500,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2662,f10273,f362])).
fof(f17412,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f14501,f16016])).
fof(f14501,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6093,f10273,f362])).
fof(f6093,plain,(
  ( ! [X0] : (~sP5('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3571,f343])).
fof(f3571,plain,(
  ( ! [X0] : (~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f2255,f353])).
fof(f17405,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f17404,f17355])).
fof(f17404,plain,(
  ( ! [X0] : (sP5(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14523,f16152])).
fof(f14523,plain,(
  ( ! [X0] : (sP5(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2260,f10273,f362])).
fof(f17403,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f17402,f17355])).
fof(f17402,plain,(
  ( ! [X0] : (sP5(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14525,f16152])).
fof(f14525,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3536,f10273,f362])).
fof(f3536,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f2255,f311])).
fof(f17401,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f17400,f16985])).
fof(f16985,plain,(
  ( ! [X0] : (~sP5(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f15627,f16152])).
fof(f15627,plain,(
  ( ! [X0] : (~sP5(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f10273,f3187])).
fof(f17400,plain,(
  ( ! [X0] : (sP5(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f17399,f15769])).
fof(f17399,plain,(
  ( ! [X0] : (sP5(X0,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14529,f16152])).
fof(f14529,plain,(
  ( ! [X0] : (sP5(X0,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'F_0')) ) | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5967,f10273,f362])).
fof(f17398,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f17397,f16985])).
fof(f17397,plain,(
  ( ! [X0] : (sP5(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(forward_demodulation,[],[f17396,f15769])).
fof(f17396,plain,(
  ( ! [X0] : (sP5(X0,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(forward_demodulation,[],[f14530,f16152])).
fof(f14530,plain,(
  ( ! [X0] : (sP5(X0,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)),'F_0')) ) | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6050,f10273,f362])).
fof(f17391,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f17390,f17361])).
fof(f17390,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14541,f16152])).
fof(f14541,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2260,f10273,f362])).
fof(f17389,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f17388,f17361])).
fof(f17388,plain,(
  ( ! [X0] : (sP5(X0,'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f14543,f16152])).
fof(f14543,plain,(
  ( ! [X0,X1] : (sP5(X0,'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3536,f10273,f362])).
fof(f17351,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14643,f14936])).
fof(f14936,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f10273,f10273,f10273,f378])).
fof(f14643,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f434,f10273,f366])).
fof(f17350,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14662,f14936])).
fof(f14662,plain,(
  ( ! [X0,X1] : (~sP7(X0,'F_0',X1)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f434,f10273,f367])).
fof(f17349,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14664,f10273])).
fof(f14664,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f452,f540,f10273,f368])).
fof(f17348,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14665,f10273])).
fof(f14665,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f454,f545,f10273,f368])).
fof(f17347,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14666,f10273])).
fof(f14666,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f451,f545,f10273,f368])).
fof(f17346,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14667,f10273])).
fof(f14667,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f452,f545,f10273,f368])).
fof(f17345,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14668,f10273])).
fof(f14668,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f516,f545,f10273,f368])).
fof(f17344,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14669,f10273])).
fof(f14669,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f452,f2307,f10273,f368])).
fof(f2307,plain,(
  ( ! [X0] : (~sP7('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f425,f366])).
fof(f17343,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14670,f10273])).
fof(f14670,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f454,f2308,f10273,f368])).
fof(f2308,plain,(
  ( ! [X0] : (~sP7('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f425,f367])).
fof(f17342,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14671,f10273])).
fof(f14671,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2325,f2308,f10273,f368])).
fof(f2325,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f425,f373])).
fof(f17341,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14672,f10273])).
fof(f14672,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f452,f2308,f10273,f368])).
fof(f17340,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14673,f10273])).
fof(f14673,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f516,f2308,f10273,f368])).
fof(f17339,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14674,f10273])).
fof(f14674,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f452,f3595,f10273,f368])).
fof(f3595,plain,(
  ( ! [X0,X1] : (~sP7('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),X1,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f2255,f366])).
fof(f17338,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14675,f10273])).
fof(f14675,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f454,f3596,f10273,f368])).
fof(f3596,plain,(
  ( ! [X0,X1] : (~sP7('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0',X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f2255,f367])).
fof(f17337,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14676,f10273])).
fof(f14676,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3621,f3596,f10273,f368])).
fof(f3621,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2255,f373])).
fof(f17336,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14677,f10273])).
fof(f14677,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f452,f3596,f10273,f368])).
fof(f17335,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14678,f10273])).
fof(f14678,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f516,f3596,f10273,f368])).
fof(f17334,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14684,f14812])).
fof(f14812,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f10273,f373])).
fof(f14684,plain,(
  ~sP6('F_0','T_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f545,f244,f10273,f368])).
fof(f17333,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14689,f14812])).
fof(f14689,plain,(
  ~sP6('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2308,f244,f10273,f368])).
fof(f17332,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14691,f14812])).
fof(f14691,plain,(
  ( ! [X0] : (~sP6('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3596,f244,f10273,f368])).
fof(f17331,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14700,f14812])).
fof(f14700,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f540,f10273,f10273,f368])).
fof(f17330,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14702,f14812])).
fof(f14702,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f545,f10273,f10273,f368])).
fof(f17329,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14707,f14812])).
fof(f14707,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2344,f10273,f10273,f368])).
fof(f2344,plain,(
  ~sP7('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f425,f381])).
fof(f17328,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14708,f14812])).
fof(f14708,plain,(
  ~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2345,f10273,f10273,f368])).
fof(f2345,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f425,f434,f425,f381])).
fof(f17327,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14709,f14812])).
fof(f14709,plain,(
  ( ! [X0] : (~sP6('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3640,f10273,f10273,f368])).
fof(f3640,plain,(
  ( ! [X0] : (~sP7('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f2255,f381])).
fof(f17326,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14710,f14812])).
fof(f14710,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2307,f10273,f10273,f368])).
fof(f17325,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14712,f14812])).
fof(f14712,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2308,f10273,f10273,f368])).
fof(f17324,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14714,f14812])).
fof(f14714,plain,(
  ( ! [X0,X1] : (~sP6(X0,'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3595,f10273,f10273,f368])).
fof(f17323,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14716,f14812])).
fof(f14716,plain,(
  ( ! [X0,X1] : (~sP6('F_0',X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3596,f10273,f10273,f368])).
fof(f17322,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14718,f10273])).
fof(f14718,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f452,f540,f10273,f368])).
fof(f17321,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14719,f10273])).
fof(f14719,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f451,f540,f10273,f368])).
fof(f17320,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14720,f10273])).
fof(f14720,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f454,f540,f10273,f368])).
fof(f17319,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14721,f10273])).
fof(f14721,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f517,f540,f10273,f368])).
fof(f17318,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14722,f10273])).
fof(f14722,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f454,f545,f10273,f368])).
fof(f17317,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14723,f10273])).
fof(f14723,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f452,f2307,f10273,f368])).
fof(f17316,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14724,f10273])).
fof(f14724,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2325,f2307,f10273,f368])).
fof(f17315,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14725,f10273])).
fof(f14725,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f454,f2307,f10273,f368])).
fof(f17314,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14726,f10273])).
fof(f14726,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f517,f2307,f10273,f368])).
fof(f17313,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14727,f10273])).
fof(f14727,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f454,f2308,f10273,f368])).
fof(f17312,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14728,f10273])).
fof(f14728,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f452,f3595,f10273,f368])).
fof(f17311,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14729,f10273])).
fof(f14729,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3621,f3595,f10273,f368])).
fof(f17310,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14730,f10273])).
fof(f14730,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f454,f3595,f10273,f368])).
fof(f17309,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14731,f10273])).
fof(f14731,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f517,f3595,f10273,f368])).
fof(f17308,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14732,f10273])).
fof(f14732,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f454,f3596,f10273,f368])).
fof(f17307,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14738,f14812])).
fof(f14738,plain,(
  ~sP6('T_0','F_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f540,f244,f10273,f368])).
fof(f17306,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14741,f14812])).
fof(f14741,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2344,f244,f10273,f368])).
fof(f17305,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14742,f14812])).
fof(f14742,plain,(
  ( ! [X0] : (~sP6('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3640,f244,f10273,f368])).
fof(f17304,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14743,f14812])).
fof(f14743,plain,(
  ~sP6('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2307,f244,f10273,f368])).
fof(f17303,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14745,f14812])).
fof(f14745,plain,(
  ( ! [X0] : (~sP6('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3595,f244,f10273,f368])).
fof(f17302,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14754,f14812])).
fof(f14754,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f540,f10273,f10273,f368])).
fof(f17301,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14756,f14812])).
fof(f14756,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f545,f10273,f10273,f368])).
fof(f17300,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14761,f14812])).
fof(f14761,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2344,f10273,f10273,f368])).
fof(f17299,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14762,f14812])).
fof(f14762,plain,(
  ~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2345,f10273,f10273,f368])).
fof(f17298,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14763,f14812])).
fof(f14763,plain,(
  ( ! [X0] : (~sP6('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3640,f10273,f10273,f368])).
fof(f17297,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14764,f14812])).
fof(f14764,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2307,f10273,f10273,f368])).
fof(f17296,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14766,f14812])).
fof(f14766,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2308,f10273,f10273,f368])).
fof(f17295,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14768,f14812])).
fof(f14768,plain,(
  ( ! [X0,X1] : (~sP6(X0,'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3595,f10273,f10273,f368])).
fof(f17294,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14770,f14812])).
fof(f14770,plain,(
  ( ! [X0,X1] : (~sP6('F_0',X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3596,f10273,f10273,f368])).
fof(f17293,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14772,f10273])).
fof(f14772,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f452,f434,f10273,f372])).
fof(f17292,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14773,f10273])).
fof(f14773,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f516,f434,f10273,f372])).
fof(f17291,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14774,f10273])).
fof(f14774,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f454,f434,f10273,f372])).
fof(f17290,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14775,f10273])).
fof(f14775,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f517,f434,f10273,f372])).
fof(f17289,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14790,f14812])).
fof(f14790,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f434,f10273,f10273,f372])).
fof(f17288,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14792,f10273])).
fof(f14792,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f452,f434,f10273,f372])).
fof(f17287,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14793,f10273])).
fof(f14793,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f516,f434,f10273,f372])).
fof(f17286,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14794,f10273])).
fof(f14794,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f454,f434,f10273,f372])).
fof(f17285,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14795,f10273])).
fof(f14795,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f517,f434,f10273,f372])).
fof(f17284,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14801,f14812])).
fof(f14801,plain,(
  ( ! [X0] : (~sP6('T_0',X0,'F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f434,f244,f10273,f372])).
fof(f17283,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14810,f14812])).
fof(f14810,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f434,f10273,f10273,f372])).
fof(f17282,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14814,f10273])).
fof(f14814,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f540,f244,f10273,f378])).
fof(f17281,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14815,f10273])).
fof(f14815,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f545,f244,f10273,f378])).
fof(f17280,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14816,f10273])).
fof(f14816,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f540,f10273,f10273,f378])).
fof(f17279,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14817,f10273])).
fof(f14817,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f545,f10273,f10273,f378])).
fof(f17278,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14818,f244])).
fof(f14818,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f868,f10273,f10273,f378])).
fof(f17277,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14819,f244])).
fof(f14819,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2340,f10273,f10273,f378])).
fof(f2340,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f425,f381])).
fof(f17276,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14820,f244])).
fof(f14820,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3636,f10273,f10273,f378])).
fof(f3636,plain,(
  ( ! [X0] : (~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f2255,f381])).
fof(f17275,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14821,f10273])).
fof(f14821,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2344,f10273,f10273,f378])).
fof(f17274,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14822,f10273])).
fof(f14822,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2345,f10273,f10273,f378])).
fof(f17273,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14823,f10273])).
fof(f14823,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3640,f10273,f10273,f378])).
fof(f17272,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14824,f10273])).
fof(f14824,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2307,f10273,f10273,f378])).
fof(f17271,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14825,f10273])).
fof(f14825,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2308,f10273,f10273,f378])).
fof(f17270,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14826,f10273])).
fof(f14826,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3595,f10273,f10273,f378])).
fof(f17269,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14827,f10273])).
fof(f14827,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3596,f10273,f10273,f378])).
fof(f17268,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14828,f244])).
fof(f14828,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f545,f244,f10273,f378])).
fof(f17267,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14829,f10273])).
fof(f14829,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f868,f244,f10273,f378])).
fof(f17266,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14830,f10273])).
fof(f14830,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2340,f244,f10273,f378])).
fof(f17265,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14831,f10273])).
fof(f14831,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3636,f244,f10273,f378])).
fof(f17264,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14832,f10273])).
fof(f14832,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2308,f244,f10273,f378])).
fof(f17263,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14833,f10273])).
fof(f14833,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3596,f244,f10273,f378])).
fof(f17262,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14838,f244])).
fof(f14838,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f540,f10273,f10273,f378])).
fof(f17261,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14839,f244])).
fof(f14839,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f545,f10273,f10273,f378])).
fof(f17260,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14840,f10273])).
fof(f14840,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f868,f10273,f10273,f378])).
fof(f17259,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14841,f10273])).
fof(f14841,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2340,f10273,f10273,f378])).
fof(f17258,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14842,f10273])).
fof(f14842,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3636,f10273,f10273,f378])).
fof(f17257,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14843,f10273])).
fof(f14843,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2344,f10273,f10273,f378])).
fof(f17256,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14844,f10273])).
fof(f14844,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2345,f10273,f10273,f378])).
fof(f17255,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14845,f10273])).
fof(f14845,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3640,f10273,f10273,f378])).
fof(f17254,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14846,f10273])).
fof(f14846,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2307,f10273,f10273,f378])).
fof(f17253,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14847,f10273])).
fof(f14847,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2308,f10273,f10273,f378])).
fof(f17252,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14848,f10273])).
fof(f14848,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3595,f10273,f10273,f378])).
fof(f17251,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14849,f10273])).
fof(f14849,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3596,f10273,f10273,f378])).
fof(f17250,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14854,f10273])).
fof(f14854,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f540,f244,f10273,f378])).
fof(f17249,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14855,f10273])).
fof(f14855,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f545,f244,f10273,f378])).
fof(f17248,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14856,f10273])).
fof(f14856,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f540,f10273,f10273,f378])).
fof(f17247,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14857,f10273])).
fof(f14857,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f545,f10273,f10273,f378])).
fof(f17246,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14858,f244])).
fof(f14858,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f868,f10273,f10273,f378])).
fof(f17245,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14859,f10273])).
fof(f14859,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2340,f10273,f10273,f378])).
fof(f17244,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14860,f10273])).
fof(f14860,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3636,f10273,f10273,f378])).
fof(f17243,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14861,f244])).
fof(f14861,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2344,f10273,f10273,f378])).
fof(f17242,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14862,f10273])).
fof(f14862,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2345,f10273,f10273,f378])).
fof(f17241,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14863,f244])).
fof(f14863,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3640,f10273,f10273,f378])).
fof(f17240,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14864,f10273])).
fof(f14864,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2307,f10273,f10273,f378])).
fof(f17239,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14865,f10273])).
fof(f14865,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2308,f10273,f10273,f378])).
fof(f17238,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14866,f10273])).
fof(f14866,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3595,f10273,f10273,f378])).
fof(f17237,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14867,f10273])).
fof(f14867,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3596,f10273,f10273,f378])).
fof(f17236,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14868,f244])).
fof(f14868,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f540,f244,f10273,f378])).
fof(f17235,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14869,f10273])).
fof(f14869,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f868,f244,f10273,f378])).
fof(f17234,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14870,f10273])).
fof(f14870,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2344,f244,f10273,f378])).
fof(f17233,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14871,f10273])).
fof(f14871,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3640,f244,f10273,f378])).
fof(f17232,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14872,f10273])).
fof(f14872,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2307,f244,f10273,f378])).
fof(f17231,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14873,f10273])).
fof(f14873,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3595,f244,f10273,f378])).
fof(f17230,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14878,f244])).
fof(f14878,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f540,f10273,f10273,f378])).
fof(f17229,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14879,f244])).
fof(f14879,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f545,f10273,f10273,f378])).
fof(f17228,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14880,f10273])).
fof(f14880,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f868,f10273,f10273,f378])).
fof(f17227,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14881,f10273])).
fof(f14881,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2340,f10273,f10273,f378])).
fof(f17226,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14882,f10273])).
fof(f14882,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3636,f10273,f10273,f378])).
fof(f17225,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14883,f10273])).
fof(f14883,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2344,f10273,f10273,f378])).
fof(f17224,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14884,f10273])).
fof(f14884,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2345,f10273,f10273,f378])).
fof(f17223,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14885,f10273])).
fof(f14885,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3640,f10273,f10273,f378])).
fof(f17222,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14886,f10273])).
fof(f14886,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2307,f10273,f10273,f378])).
fof(f17221,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14887,f10273])).
fof(f14887,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2308,f10273,f10273,f378])).
fof(f17220,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14888,f10273])).
fof(f14888,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3595,f10273,f10273,f378])).
fof(f17219,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14889,f10273])).
fof(f14889,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3596,f10273,f10273,f378])).
fof(f17218,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14894,f10273])).
fof(f14894,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f545,f244,f10273,f378])).
fof(f17217,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14895,f244])).
fof(f14895,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f868,f244,f10273,f378])).
fof(f17216,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14896,f10273])).
fof(f14896,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2340,f244,f10273,f378])).
fof(f17215,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14897,f10273])).
fof(f14897,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3636,f244,f10273,f378])).
fof(f17214,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14898,f10273])).
fof(f14898,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2308,f244,f10273,f378])).
fof(f17213,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14899,f10273])).
fof(f14899,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3596,f244,f10273,f378])).
fof(f17212,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14900,f10273])).
fof(f14900,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f540,f10273,f10273,f378])).
fof(f17211,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14901,f10273])).
fof(f14901,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f545,f10273,f10273,f378])).
fof(f17210,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14902,f244])).
fof(f14902,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f868,f10273,f10273,f378])).
fof(f17209,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14903,f10273])).
fof(f14903,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2340,f10273,f10273,f378])).
fof(f17208,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14904,f10273])).
fof(f14904,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3636,f10273,f10273,f378])).
fof(f17207,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14905,f244])).
fof(f14905,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2344,f10273,f10273,f378])).
fof(f17206,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14906,f10273])).
fof(f14906,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2345,f10273,f10273,f378])).
fof(f17205,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14907,f244])).
fof(f14907,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3640,f10273,f10273,f378])).
fof(f17204,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14908,f10273])).
fof(f14908,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2307,f10273,f10273,f378])).
fof(f17203,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14909,f10273])).
fof(f14909,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2308,f10273,f10273,f378])).
fof(f17202,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14910,f10273])).
fof(f14910,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3595,f10273,f10273,f378])).
fof(f17201,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14911,f10273])).
fof(f14911,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3596,f10273,f10273,f378])).
fof(f17200,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14912,f10273])).
fof(f14912,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f540,f244,f10273,f378])).
fof(f17199,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14913,f244])).
fof(f14913,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f868,f244,f10273,f378])).
fof(f17198,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14914,f10273])).
fof(f14914,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2344,f244,f10273,f378])).
fof(f17197,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14915,f10273])).
fof(f14915,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3640,f244,f10273,f378])).
fof(f17196,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14916,f10273])).
fof(f14916,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2307,f244,f10273,f378])).
fof(f17195,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14917,f10273])).
fof(f14917,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3595,f244,f10273,f378])).
fof(f17194,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14922,f10273])).
fof(f14922,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f540,f10273,f10273,f378])).
fof(f17193,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14923,f10273])).
fof(f14923,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f545,f10273,f10273,f378])).
fof(f17192,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14924,f244])).
fof(f14924,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f868,f10273,f10273,f378])).
fof(f17191,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14925,f244])).
fof(f14925,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2340,f10273,f10273,f378])).
fof(f17190,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14926,f244])).
fof(f14926,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3636,f10273,f10273,f378])).
fof(f17189,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14927,f10273])).
fof(f14927,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2344,f10273,f10273,f378])).
fof(f17188,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14928,f10273])).
fof(f14928,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2345,f10273,f10273,f378])).
fof(f17187,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14929,f10273])).
fof(f14929,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3640,f10273,f10273,f378])).
fof(f17186,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14930,f10273])).
fof(f14930,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2307,f10273,f10273,f378])).
fof(f17185,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14931,f10273])).
fof(f14931,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2308,f10273,f10273,f378])).
fof(f17184,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14932,f10273])).
fof(f14932,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3595,f10273,f10273,f378])).
fof(f17183,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f14933,f10273])).
fof(f14933,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3596,f10273,f10273,f378])).
fof(f17182,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14938,f10273])).
fof(f14938,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f552,f434,f10273,f381])).
fof(f17181,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14939,f10273])).
fof(f14939,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f731,f434,f10273,f381])).
fof(f17180,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14940,f10273])).
fof(f14940,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f563,f434,f10273,f381])).
fof(f17179,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14941,f10273])).
fof(f14941,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f878,f434,f10273,f381])).
fof(f17178,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14972,f10273])).
fof(f14972,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f552,f434,f10273,f381])).
fof(f17177,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14973,f10273])).
fof(f14973,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f731,f434,f10273,f381])).
fof(f17176,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14974,f10273])).
fof(f14974,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f563,f434,f10273,f381])).
fof(f17175,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f14975,f10273])).
fof(f14975,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f878,f434,f10273,f381])).
fof(f17174,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15026,f15624])).
fof(f15624,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f458,f10273,f417])).
fof(f15026,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f434,f10273,f382])).
fof(f17173,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15048,f15624])).
fof(f15048,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f434,f10273,f383])).
fof(f17172,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15077,f15139])).
fof(f15139,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f10273,f391])).
fof(f15077,plain,(
  ~sP8('F_0','T_0','F_0') | $spl110),
  inference(unit_resulting_resolution,[],[f573,f434,f10273,f386])).
fof(f17171,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15079,f15139])).
fof(f15079,plain,(
  ~sP8('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2350,f434,f10273,f386])).
fof(f2350,plain,(
  ( ! [X0] : (~sP9('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f425,f383])).
fof(f17170,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15081,f15139])).
fof(f15081,plain,(
  ( ! [X0] : (~sP8('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3647,f434,f10273,f386])).
fof(f3647,plain,(
  ( ! [X0,X1] : (~sP9('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f2255,f383])).
fof(f17169,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15084,f15139])).
fof(f15084,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2369,f434,f10273,f386])).
fof(f2369,plain,(
  ~sP9('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f434,f434,f425,f396])).
fof(f17168,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15085,f15139])).
fof(f15085,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3673,f434,f10273,f386])).
fof(f3673,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f434,f2255,f396])).
fof(f17167,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15110,f15139])).
fof(f15110,plain,(
  ~sP8('F_0','F_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f572,f434,f10273,f387])).
fof(f17166,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15111,f15139])).
fof(f15111,plain,(
  ~sP8('F_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2348,f434,f10273,f387])).
fof(f2348,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f425,f382])).
fof(f17165,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15112,f15139])).
fof(f15112,plain,(
  ( ! [X0] : (~sP8('F_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3644,f434,f10273,f387])).
fof(f3644,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f2255,f382])).
fof(f17164,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15114,f15139])).
fof(f15114,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2369,f434,f10273,f387])).
fof(f17163,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15115,f15139])).
fof(f15115,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3673,f434,f10273,f387])).
fof(f17162,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15135,f15139])).
fof(f15135,plain,(
  ( ! [X0] : (~sP8(X0,'F_0','F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f434,f434,f10273,f389])).
fof(f17161,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15141,f10273])).
fof(f15141,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f572,f10273,f394])).
fof(f17160,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15142,f10273])).
fof(f15142,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2348,f10273,f394])).
fof(f17159,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15143,f10273])).
fof(f15143,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3644,f10273,f394])).
fof(f17158,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15144,f10273])).
fof(f15144,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f573,f10273,f394])).
fof(f17157,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15145,f10273])).
fof(f15145,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2350,f10273,f394])).
fof(f17156,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15146,f10273])).
fof(f15146,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3647,f10273,f394])).
fof(f17155,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15147,f244])).
fof(f15147,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f975,f10273,f394])).
fof(f17154,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15148,f10273])).
fof(f15148,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2369,f10273,f394])).
fof(f17153,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15149,f10273])).
fof(f15149,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3673,f10273,f394])).
fof(f17152,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15154,f244])).
fof(f15154,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f572,f10273,f394])).
fof(f17151,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15155,f10273])).
fof(f15155,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2348,f10273,f394])).
fof(f17150,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15156,f10273])).
fof(f15156,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3644,f10273,f394])).
fof(f17149,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15157,f10273])).
fof(f15157,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f573,f10273,f394])).
fof(f17148,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15158,f10273])).
fof(f15158,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2350,f10273,f394])).
fof(f17147,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15159,f10273])).
fof(f15159,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3647,f10273,f394])).
fof(f17146,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15160,f10273])).
fof(f15160,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f975,f10273,f394])).
fof(f17145,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15161,f10273])).
fof(f15161,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2369,f10273,f394])).
fof(f17144,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15162,f10273])).
fof(f15162,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3673,f10273,f394])).
fof(f17143,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15167,f10273])).
fof(f15167,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f572,f10273,f395])).
fof(f17142,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15168,f10273])).
fof(f15168,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2348,f10273,f395])).
fof(f17141,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15169,f10273])).
fof(f15169,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3644,f10273,f395])).
fof(f17140,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15170,f10273])).
fof(f15170,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f573,f10273,f395])).
fof(f17139,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15171,f10273])).
fof(f15171,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2350,f10273,f395])).
fof(f17138,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15172,f10273])).
fof(f15172,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3647,f10273,f395])).
fof(f17137,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15173,f244])).
fof(f15173,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f975,f10273,f395])).
fof(f17136,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15174,f10273])).
fof(f15174,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2369,f10273,f395])).
fof(f17135,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15175,f10273])).
fof(f15175,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3673,f10273,f395])).
fof(f17134,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15180,f10273])).
fof(f15180,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f572,f10273,f395])).
fof(f17133,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15181,f10273])).
fof(f15181,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2348,f10273,f395])).
fof(f17132,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15182,f10273])).
fof(f15182,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3644,f10273,f395])).
fof(f17131,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15183,f244])).
fof(f15183,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f573,f10273,f395])).
fof(f17130,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15184,f10273])).
fof(f15184,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2350,f10273,f395])).
fof(f17129,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15185,f10273])).
fof(f15185,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3647,f10273,f395])).
fof(f17128,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15186,f10273])).
fof(f15186,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f975,f10273,f395])).
fof(f17127,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15187,f10273])).
fof(f15187,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2369,f10273,f395])).
fof(f17126,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15188,f10273])).
fof(f15188,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3673,f10273,f395])).
fof(f17125,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15239,f15600])).
fof(f15600,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f463,f10273,f416])).
fof(f15239,plain,(
  ( ! [X0,X1] : (~sP11('F_0',X0,X1)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f434,f10273,f400])).
fof(f17124,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15241,f10273])).
fof(f15241,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f463,f588,f10273,f403])).
fof(f17123,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15242,f10273])).
fof(f15242,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f463,f589,f10273,f403])).
fof(f17122,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15243,f10273])).
fof(f15243,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f462,f589,f10273,f403])).
fof(f17121,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15244,f10273])).
fof(f15244,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f460,f589,f10273,f403])).
fof(f17120,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15245,f10273])).
fof(f15245,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f519,f589,f10273,f403])).
fof(f17119,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15246,f10273])).
fof(f15246,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f463,f2370,f10273,f403])).
fof(f2370,plain,(
  ( ! [X0] : (~sP11('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f425,f400])).
fof(f17118,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15247,f10273])).
fof(f15247,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2395,f2370,f10273,f403])).
fof(f2395,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f425,f408])).
fof(f17117,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15248,f10273])).
fof(f15248,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f460,f2370,f10273,f403])).
fof(f17116,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15249,f10273])).
fof(f15249,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f519,f2370,f10273,f403])).
fof(f17115,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15250,f10273])).
fof(f15250,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f463,f3674,f10273,f403])).
fof(f3674,plain,(
  ( ! [X0,X1] : (~sP11('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f2255,f400])).
fof(f17114,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15251,f10273])).
fof(f15251,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3709,f3674,f10273,f403])).
fof(f3709,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))),X2)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2255,f408])).
fof(f17113,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15252,f10273])).
fof(f15252,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f460,f3674,f10273,f403])).
fof(f17112,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15253,f10273])).
fof(f15253,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f519,f3674,f10273,f403])).
fof(f17111,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15259,f15430])).
fof(f15430,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f10273,f408])).
fof(f15259,plain,(
  ~sP10('F_0','T_0','T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f589,f244,f10273,f403])).
fof(f17110,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15261,f15430])).
fof(f15261,plain,(
  ~sP10('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2370,f244,f10273,f403])).
fof(f17109,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15263,f15430])).
fof(f15263,plain,(
  ( ! [X0] : (~sP10('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3674,f244,f10273,f403])).
fof(f17108,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15275,f15430])).
fof(f15275,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f588,f10273,f10273,f403])).
fof(f17107,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15277,f15430])).
fof(f15277,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f589,f10273,f10273,f403])).
fof(f17106,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15279,f15430])).
fof(f15279,plain,(
  ( ! [X0] : (~sP10('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X0)) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2370,f10273,f10273,f403])).
fof(f17105,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15281,f15430])).
fof(f15281,plain,(
  ( ! [X0,X1] : (~sP10('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),X1)) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3674,f10273,f10273,f403])).
fof(f17104,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15284,f15430])).
fof(f15284,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2404,f10273,f10273,f403])).
fof(f2404,plain,(
  ~sP11('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f425,f413])).
fof(f17103,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15285,f15430])).
fof(f15285,plain,(
  ( ! [X0] : (~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3719,f10273,f10273,f403])).
fof(f3719,plain,(
  ( ! [X0] : (~sP11('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f2255,f413])).
fof(f17102,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15287,f15430])).
fof(f15287,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2408,f10273,f10273,f403])).
fof(f2408,plain,(
  ~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f425,f434,f425,f413])).
fof(f17101,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15289,f10273])).
fof(f15289,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f588,f460,f10273,f403])).
fof(f17100,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15290,f10273])).
fof(f15290,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f589,f460,f10273,f403])).
fof(f17099,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15291,f10273])).
fof(f15291,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2370,f460,f10273,f403])).
fof(f17098,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15292,f10273])).
fof(f15292,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3674,f460,f10273,f403])).
fof(f17097,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15293,f10273])).
fof(f15293,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f589,f519,f10273,f403])).
fof(f17096,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15294,f10273])).
fof(f15294,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2370,f519,f10273,f403])).
fof(f17095,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15295,f10273])).
fof(f15295,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3674,f519,f10273,f403])).
fof(f17094,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15296,f10273])).
fof(f15296,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f589,f462,f10273,f403])).
fof(f17093,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15297,f10273])).
fof(f15297,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f588,f462,f10273,f403])).
fof(f17092,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15298,f10273])).
fof(f15298,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2370,f2395,f10273,f403])).
fof(f17091,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15299,f10273])).
fof(f15299,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f588,f2395,f10273,f403])).
fof(f17090,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15300,f10273])).
fof(f15300,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3674,f3709,f10273,f403])).
fof(f17089,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15301,f10273])).
fof(f15301,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f588,f3709,f10273,f403])).
fof(f17088,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15302,f10273])).
fof(f15302,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f588,f463,f10273,f403])).
fof(f17087,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15304,f15430])).
fof(f15304,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2404,f244,f10273,f403])).
fof(f17086,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15305,f15430])).
fof(f15305,plain,(
  ( ! [X0] : (~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3719,f244,f10273,f403])).
fof(f17085,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15311,f15430])).
fof(f15311,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f588,f10273,f10273,f403])).
fof(f17084,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15312,f15430])).
fof(f15312,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f589,f10273,f10273,f403])).
fof(f17083,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15313,f15430])).
fof(f15313,plain,(
  ( ! [X0] : (~sP10('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X0)) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2370,f10273,f10273,f403])).
fof(f17082,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15314,f15430])).
fof(f15314,plain,(
  ( ! [X0,X1] : (~sP10('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),X1)) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3674,f10273,f10273,f403])).
fof(f17081,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15316,f15430])).
fof(f15316,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2404,f10273,f10273,f403])).
fof(f17080,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15317,f15430])).
fof(f15317,plain,(
  ( ! [X0] : (~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3719,f10273,f10273,f403])).
fof(f17079,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15319,f15430])).
fof(f15319,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2408,f10273,f10273,f403])).
fof(f17078,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15349,f15430])).
fof(f15349,plain,(
  ~sP10('F_0','F_0','F_0') | $spl110),
  inference(unit_resulting_resolution,[],[f588,f434,f10273,f404])).
fof(f17077,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15351,f15430])).
fof(f15351,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2404,f434,f10273,f404])).
fof(f17076,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15352,f15430])).
fof(f15352,plain,(
  ( ! [X0] : (~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3719,f434,f10273,f404])).
fof(f17075,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15354,f15430])).
fof(f15354,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2408,f434,f10273,f404])).
fof(f17074,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15377,f15430])).
fof(f15377,plain,(
  ~sP10('F_0','F_0','F_0') | $spl110),
  inference(unit_resulting_resolution,[],[f588,f434,f10273,f405])).
fof(f17073,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15379,f15430])).
fof(f15379,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2404,f434,f10273,f405])).
fof(f17072,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15380,f15430])).
fof(f15380,plain,(
  ( ! [X0] : (~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3719,f434,f10273,f405])).
fof(f17071,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15382,f15430])).
fof(f15382,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2408,f434,f10273,f405])).
fof(f17070,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15402,f10273])).
fof(f15402,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f460,f434,f10273,f406])).
fof(f17069,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15403,f10273])).
fof(f15403,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f519,f434,f10273,f406])).
fof(f17068,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15404,f10273])).
fof(f15404,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f463,f434,f10273,f406])).
fof(f17067,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15405,f10273])).
fof(f15405,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f520,f434,f10273,f406])).
fof(f17066,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15408,f15430])).
fof(f15408,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f10273,f434,f10273,f406])).
fof(f17065,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15422,f10273])).
fof(f15422,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f460,f434,f10273,f406])).
fof(f17064,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15423,f10273])).
fof(f15423,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f519,f434,f10273,f406])).
fof(f17063,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15424,f10273])).
fof(f15424,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f463,f434,f10273,f406])).
fof(f17062,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15425,f10273])).
fof(f15425,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f520,f434,f10273,f406])).
fof(f17061,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15426,f15430])).
fof(f15426,plain,(
  ( ! [X0] : (~sP10('T_0','F_0',X0)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f244,f434,f10273,f406])).
fof(f17060,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15428,f15430])).
fof(f15428,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f10273,f434,f10273,f406])).
fof(f17059,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15457,f10273])).
fof(f15457,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f588,f10273,f412])).
fof(f17058,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15458,f10273])).
fof(f15458,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f589,f10273,f412])).
fof(f17057,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15459,f10273])).
fof(f15459,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2370,f10273,f412])).
fof(f17056,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15460,f10273])).
fof(f15460,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3674,f10273,f412])).
fof(f17055,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15461,f244])).
fof(f15461,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f1149,f10273,f412])).
fof(f17054,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15462,f244])).
fof(f15462,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2404,f10273,f412])).
fof(f17053,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15463,f244])).
fof(f15463,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3719,f10273,f412])).
fof(f17052,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15464,f10273])).
fof(f15464,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2407,f10273,f412])).
fof(f2407,plain,(
  ~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f425,f413])).
fof(f17051,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15465,f10273])).
fof(f15465,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2408,f10273,f412])).
fof(f17050,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15466,f10273])).
fof(f15466,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3723,f10273,f412])).
fof(f3723,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f2255,f413])).
fof(f17049,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15471,f10273])).
fof(f15471,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f588,f10273,f412])).
fof(f17048,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15472,f244])).
fof(f15472,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f589,f10273,f412])).
fof(f17047,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15473,f10273])).
fof(f15473,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2370,f10273,f412])).
fof(f17046,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15474,f10273])).
fof(f15474,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3674,f10273,f412])).
fof(f17045,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15475,f10273])).
fof(f15475,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f1149,f10273,f412])).
fof(f17044,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15476,f10273])).
fof(f15476,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2404,f10273,f412])).
fof(f17043,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15477,f10273])).
fof(f15477,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3719,f10273,f412])).
fof(f17042,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15478,f10273])).
fof(f15478,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2407,f10273,f412])).
fof(f17041,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15479,f10273])).
fof(f15479,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2408,f10273,f412])).
fof(f17040,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15480,f10273])).
fof(f15480,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3723,f10273,f412])).
fof(f17039,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15485,f244])).
fof(f15485,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f592,f434,f10273,f413])).
fof(f17038,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15486,f244])).
fof(f15486,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f900,f434,f10273,f413])).
fof(f17037,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15487,f10273])).
fof(f15487,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f1123,f434,f10273,f413])).
fof(f17036,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15488,f10273])).
fof(f15488,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2396,f434,f10273,f413])).
fof(f2396,plain,(
  sP11('F_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f434,f425,f410])).
fof(f17035,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15489,f10273])).
fof(f15489,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3710,f434,f10273,f413])).
fof(f3710,plain,(
  ( ! [X0] : (sP11('F_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f434,f2255,f410])).
fof(f17034,plain,(
  $false | ($spl105 | $spl110)),
  inference(subsumption_resolution,[],[f15490,f10273])).
fof(f15490,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f12733,f434,f10273,f413])).
fof(f12733,plain,(
  ( ! [X0] : (sP11(f(s(bool,'T_0')),X0,'F_0')) ) | $spl105),
  inference(unit_resulting_resolution,[],[f434,f8739])).
fof(f8739,plain,(
  ( ! [X28,X27] : (sP11(f(s(bool,'T_0')),X28,X27) | p(s(bool,X27))) ) | $spl105),
  inference(subsumption_resolution,[],[f8626,f244])).
fof(f8626,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP11(f(s(bool,'T_0')),X28,X27)) ) | $spl105),
  inference(superposition,[],[f411,f6844])).
fof(f17033,plain,(
  $false | ($spl10 | $spl110)),
  inference(subsumption_resolution,[],[f15491,f10273])).
fof(f15491,plain,(
  ~p(s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) | ($spl10 | $spl110)),
  inference(unit_resulting_resolution,[],[f13219,f434,f10273,f413])).
fof(f13219,plain,(
  ( ! [X0] : (sP11(g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),X0,'F_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f434,f7038])).
fof(f7038,plain,(
  ( ! [X28,X27] : (sP11(g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),X28,X27) | p(s(bool,X27))) ) | $spl10),
  inference(subsumption_resolution,[],[f6929,f244])).
fof(f6929,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP11(g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),X28,X27)) ) | $spl10),
  inference(superposition,[],[f411,f3920])).
fof(f17032,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15492,f10273])).
fof(f15492,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f1125,f434,f10273,f413])).
fof(f17031,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15493,f10273])).
fof(f15493,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2384,f434,f10273,f413])).
fof(f2384,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f520,f425,f405])).
fof(f17030,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15494,f10273])).
fof(f15494,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3696,f434,f10273,f413])).
fof(f3696,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f520,f2255,f405])).
fof(f17029,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15495,f10273])).
fof(f15495,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2409,f434,f10273,f413])).
fof(f2409,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f460,f425,f416])).
fof(f17028,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15496,f10273])).
fof(f15496,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2410,f434,f10273,f413])).
fof(f2410,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f519,f425,f416])).
fof(f17027,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15497,f10273])).
fof(f15497,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3725,f434,f10273,f413])).
fof(f3725,plain,(
  ( ! [X0,X1] : (sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),X1,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f460,f2255,f416])).
fof(f17026,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15498,f10273])).
fof(f15498,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3726,f434,f10273,f413])).
fof(f3726,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f519,f2255,f416])).
fof(f17025,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15530,f10273])).
fof(f15530,plain,(
  ~p(s(bool,'F_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f592,f434,f10273,f413])).
fof(f17024,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15531,f10273])).
fof(f15531,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl110),
  inference(unit_resulting_resolution,[],[f900,f434,f10273,f413])).
fof(f17023,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15532,f244])).
fof(f15532,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f1123,f434,f10273,f413])).
fof(f17022,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15533,f10273])).
fof(f15533,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2396,f434,f10273,f413])).
fof(f17021,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15534,f10273])).
fof(f15534,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3710,f434,f10273,f413])).
fof(f17020,plain,(
  $false | ($spl105 | $spl110)),
  inference(subsumption_resolution,[],[f15535,f10273])).
fof(f15535,plain,(
  ~p(s(bool,'F_0')) | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f12733,f434,f10273,f413])).
fof(f17019,plain,(
  $false | ($spl10 | $spl110)),
  inference(subsumption_resolution,[],[f15536,f10273])).
fof(f15536,plain,(
  ~p(s(bool,'F_0')) | ($spl10 | $spl110)),
  inference(unit_resulting_resolution,[],[f13219,f434,f10273,f413])).
fof(f17018,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15537,f244])).
fof(f15537,plain,(
  ~p(s(bool,'T_0')) | $spl110),
  inference(unit_resulting_resolution,[],[f1125,f434,f10273,f413])).
fof(f17017,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15538,f10273])).
fof(f15538,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2384,f434,f10273,f413])).
fof(f17016,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15539,f10273])).
fof(f15539,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14))))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3696,f434,f10273,f413])).
fof(f17015,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15540,f10273])).
fof(f15540,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2409,f434,f10273,f413])).
fof(f17014,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15541,f10273])).
fof(f15541,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2410,f434,f10273,f413])).
fof(f17013,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15542,f10273])).
fof(f15542,plain,(
  ~p(s(bool,'F_0')) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3725,f434,f10273,f413])).
fof(f17012,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15543,f10273])).
fof(f15543,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3726,f434,f10273,f413])).
fof(f17011,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15580,f460])).
fof(f15580,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f588,f10273,f416])).
fof(f17010,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15581,f462])).
fof(f15581,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f589,f10273,f416])).
fof(f17009,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15582,f463])).
fof(f15582,plain,(
  ( ! [X0] : (~sP10('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X0)) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2370,f10273,f416])).
fof(f17008,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15583,f463])).
fof(f15583,plain,(
  ( ! [X0,X1] : (~sP10('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),X1)) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3674,f10273,f416])).
fof(f17007,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f17006,f519])).
fof(f17006,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f15585,f16152])).
fof(f15585,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2404,f10273,f416])).
fof(f17005,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f17004,f519])).
fof(f17004,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f15586,f16152])).
fof(f15586,plain,(
  ( ! [X0] : (~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3719,f10273,f416])).
fof(f17003,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f17002,f519])).
fof(f17002,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f15588,f16152])).
fof(f15588,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2408,f10273,f416])).
fof(f16999,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15603,f456])).
fof(f15603,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl110),
  inference(unit_resulting_resolution,[],[f572,f10273,f417])).
fof(f16998,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15604,f458])).
fof(f15604,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2348,f10273,f417])).
fof(f16997,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15605,f458])).
fof(f15605,plain,(
  ( ! [X0,X1] : (~sP8('F_0',X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3644,f10273,f417])).
fof(f16996,plain,(
  $false | $spl110),
  inference(subsumption_resolution,[],[f15606,f457])).
fof(f15606,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl110),
  inference(unit_resulting_resolution,[],[f573,f10273,f417])).
fof(f16995,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15607,f458])).
fof(f15607,plain,(
  ( ! [X0] : (~sP8('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X0)) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2350,f10273,f417])).
fof(f16994,plain,(
  $false | ($spl0 | $spl110)),
  inference(subsumption_resolution,[],[f15608,f458])).
fof(f15608,plain,(
  ( ! [X0,X1] : (~sP8('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),X1)) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3647,f10273,f417])).
fof(f16993,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16992,f518])).
fof(f16992,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f15610,f16152])).
fof(f15610,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','F_0') | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2369,f10273,f417])).
fof(f16991,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16990,f518])).
fof(f16990,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f15611,f16152])).
fof(f15611,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3673,f10273,f417])).
fof(f16984,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16983,f512])).
fof(f16983,plain,(
  ( ! [X0] : (~sP4(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f15635,f16152])).
fof(f15635,plain,(
  ( ! [X0] : (~sP4(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f10273,f3188])).
fof(f3188,plain,(
  ( ! [X10,X11] : (~sP4(X10,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X11) | p(s(bool,X11)) | ~p(s(bool,X10))) ) | $spl0),
  inference(subsumption_resolution,[],[f3116,f244])).
fof(f3116,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | p(s(bool,X11)) | ~sP4(X10,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X11)) ) | $spl0),
  inference(superposition,[],[f349,f2260])).
fof(f16982,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16981,f514])).
fof(f16981,plain,(
  ( ! [X0] : (~sP3(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f15643,f16152])).
fof(f15643,plain,(
  ( ! [X0] : (~sP3(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f10273,f3189])).
fof(f3189,plain,(
  ( ! [X12,X13] : (~sP3(X12,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X13) | p(s(bool,X13)) | ~p(s(bool,X12))) ) | $spl0),
  inference(subsumption_resolution,[],[f3117,f244])).
fof(f3117,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X12)) | p(s(bool,X13)) | ~sP3(X12,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X13)) ) | $spl0),
  inference(superposition,[],[f353,f2260])).
fof(f16980,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16979,f516])).
fof(f16979,plain,(
  ( ! [X0] : (~sP6(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f15651,f16152])).
fof(f15651,plain,(
  ( ! [X0] : (~sP6(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0')) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f10273,f3190])).
fof(f3190,plain,(
  ( ! [X14,X15] : (~sP6(X14,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X15) | p(s(bool,X15)) | ~p(s(bool,X14))) ) | $spl0),
  inference(subsumption_resolution,[],[f3118,f244])).
fof(f3118,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | p(s(bool,X15)) | ~sP6(X14,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X15)) ) | $spl0),
  inference(superposition,[],[f372,f2260])).
fof(f16978,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16977,f731])).
fof(f16977,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f15659,f16152])).
fof(f15659,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f10273,f3191])).
fof(f3191,plain,(
  ( ! [X17,X16] : (~sP7(X17,X16,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | p(s(bool,X17)) | ~p(s(bool,X16))) ) | $spl0),
  inference(subsumption_resolution,[],[f3119,f244])).
fof(f3119,plain,(
  ( ! [X17,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | p(s(bool,X17)) | ~sP7(X17,X16,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | $spl0),
  inference(superposition,[],[f381,f2260])).
fof(f16974,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16973,f519])).
fof(f16973,plain,(
  ( ! [X0] : (~sP10(X0,'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f15672,f16152])).
fof(f15672,plain,(
  ( ! [X0] : (~sP10(X0,'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f10273,f3196])).
fof(f3196,plain,(
  ( ! [X26,X27] : (~sP10(X27,X26,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ~p(s(bool,X27)) | p(s(bool,X26))) ) | $spl0),
  inference(subsumption_resolution,[],[f3124,f244])).
fof(f3124,plain,(
  ( ! [X26,X27] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | ~p(s(bool,X27)) | ~sP10(X27,X26,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | $spl0),
  inference(superposition,[],[f406,f2260])).
fof(f16971,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16970,f16666])).
fof(f16666,plain,(
  ( ! [X57,X56] : (sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),X56,X57)) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16337,f10273])).
fof(f16337,plain,(
  ( ! [X57,X56] : (sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),X56,X57) | ~p(s(bool,X56))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f7811])).
fof(f7811,plain,(
  ( ! [X57,X56,X55] : (sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X55),s(fun(sK12,bool),sK14)))),X56,X57) | ~p(s(bool,X56))) ) | $spl0),
  inference(subsumption_resolution,[],[f7703,f244])).
fof(f7703,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X55),s(fun(sK12,bool),sK14)))),X56,X57)) ) | $spl0),
  inference(superposition,[],[f412,f3536])).
fof(f16970,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0',X0)) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f15682,f16152])).
fof(f15682,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0',X0)) ) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f10273,f3199])).
fof(f3199,plain,(
  ( ! [X33,X32] : (~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X33,X32) | p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl0),
  inference(subsumption_resolution,[],[f3127,f244])).
fof(f3127,plain,(
  ( ! [X33,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X33,X32)) ) | $spl0),
  inference(superposition,[],[f413,f2260])).
fof(f16960,plain,(
  $false | ($spl5 | $spl8 | $spl61 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f15955,f16016])).
fof(f15955,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK13))) | ($spl8 | $spl61 | $spl110)),
  inference(backward_demodulation,[],[f15769,f4960])).
fof(f16954,plain,(
  $false | ($spl5 | $spl8 | $spl65 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f15903,f16016])).
fof(f15903,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK13))) | ($spl8 | $spl65 | $spl110)),
  inference(backward_demodulation,[],[f15769,f4963])).
fof(f16953,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f16952,f1254])).
fof(f16952,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(forward_demodulation,[],[f16951,f16152])).
fof(f16951,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(forward_demodulation,[],[f15897,f16152])).
fof(f15897,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl80 | $spl110)),
  inference(backward_demodulation,[],[f15769,f8866])).
fof(f8866,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)),'T_0')) ) | ($spl0 | $spl5 | $spl80)),
  inference(unit_resulting_resolution,[],[f6025,f3531,f363])).
fof(f16946,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f16945,f16563])).
fof(f16563,plain,(
  ( ! [X4,X0,X3,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16162,f16152])).
fof(f16162,plain,(
  ( ! [X4,X0,X5,X3,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2048])).
fof(f16945,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(forward_demodulation,[],[f16944,f16152])).
fof(f16944,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(forward_demodulation,[],[f15894,f16152])).
fof(f15894,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl80 | $spl110)),
  inference(backward_demodulation,[],[f15769,f8847])).
fof(f8847,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)),'F_0')) ) | ($spl0 | $spl5 | $spl80)),
  inference(unit_resulting_resolution,[],[f6050,f3531,f362])).
fof(f16929,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl80 | $spl110)),
  inference(subsumption_resolution,[],[f15878,f16016])).
fof(f15878,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl5 | $spl8 | $spl80 | $spl110)),
  inference(backward_demodulation,[],[f15769,f6025])).
fof(f16920,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16919,f1254])).
fof(f16919,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16918,f16152])).
fof(f16918,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f15872,f16152])).
fof(f15872,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f15769,f8865])).
fof(f16913,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16912,f16563])).
fof(f16912,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16911,f16152])).
fof(f16911,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f15869,f16152])).
fof(f15869,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f15769,f8846])).
fof(f16898,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16554,f16167])).
fof(f16167,plain,(
  ( ! [X2,X0,X3,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2063])).
fof(f16554,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3961])).
fof(f3961,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f2298,f2297,f2306,f346])).
fof(f16897,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16553,f16321])).
fof(f16321,plain,(
  ( ! [X0,X1] : (sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f6272])).
fof(f6272,plain,(
  ( ! [X2,X0,X1] : (sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f515,f441,f3586,f348])).
fof(f16553,plain,(
  ~sP5('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3804])).
fof(f16896,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16552,f16176])).
fof(f16176,plain,(
  ( ! [X4,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2083])).
fof(f16552,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3791])).
fof(f16891,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16533,f1037])).
fof(f16533,plain,(
  sP5('F_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2712])).
fof(f16890,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16530,f1035])).
fof(f16530,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2692])).
fof(f16889,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16526,f942])).
fof(f16526,plain,(
  ~sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2667])).
fof(f16888,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16525,f939])).
fof(f16525,plain,(
  ~sP5('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2662])).
fof(f16887,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16524,f938])).
fof(f16524,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2654])).
fof(f16886,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16523,f1038])).
fof(f16523,plain,(
  ~sP5('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2645])).
fof(f16885,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16517,f16208])).
fof(f16208,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2174])).
fof(f16517,plain,(
  ~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2408])).
fof(f16884,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16516,f1125])).
fof(f16516,plain,(
  ~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2407])).
fof(f16883,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16515,f900])).
fof(f16515,plain,(
  ~sP11('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2404])).
fof(f16882,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16510,f519])).
fof(f16510,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2391])).
fof(f2391,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f425,f425,f406])).
fof(f16881,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16509,f519])).
fof(f16509,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2390])).
fof(f2390,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f244,f425,f406])).
fof(f16880,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16508,f520])).
fof(f16508,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2385])).
fof(f2385,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f434,f244,f425,f406])).
fof(f16879,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16505,f956])).
fof(f16505,plain,(
  ~sP9('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2369])).
fof(f16878,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16498,f518])).
fof(f16498,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2354])).
fof(f2354,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f434,f434,f425,f389])).
fof(f16877,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16495,f731])).
fof(f16495,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2345])).
fof(f16876,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16494,f731])).
fof(f16494,plain,(
  ~sP7('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2344])).
fof(f16875,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16493,f878])).
fof(f16493,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2340])).
fof(f16874,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16484,f516])).
fof(f16484,plain,(
  ~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2321])).
fof(f2321,plain,(
  ~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f425,f434,f425,f372])).
fof(f16873,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16483,f516])).
fof(f16483,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2320])).
fof(f2320,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f425,f372])).
fof(f16872,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16482,f517])).
fof(f16482,plain,(
  ~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2315])).
fof(f2315,plain,(
  ~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f425,f372])).
fof(f16871,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16479,f1254])).
fof(f16479,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2306])).
fof(f16870,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16478,f1252])).
fof(f16478,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2305])).
fof(f2305,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f465,f425,f363])).
fof(f16869,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16474,f16563])).
fof(f16474,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2301])).
fof(f16868,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16469,f514])).
fof(f16469,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2293])).
fof(f16867,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16468,f514])).
fof(f16468,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2292])).
fof(f16866,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16467,f515])).
fof(f16467,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2287])).
fof(f16865,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16465,f512])).
fof(f16465,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2282])).
fof(f16864,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16464,f512])).
fof(f16464,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2281])).
fof(f16863,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16463,f513])).
fof(f16463,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2276])).
fof(f2276,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f425,f349])).
fof(f16862,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16462,f941])).
fof(f16462,plain,(
  sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2275])).
fof(f16861,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16460,f1036])).
fof(f16460,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2272])).
fof(f16860,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16459,f508])).
fof(f16459,plain,(
  sP1('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2269])).
fof(f16859,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16458,f506])).
fof(f16458,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2268])).
fof(f16858,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16456,f510])).
fof(f16456,plain,(
  ~sP1('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2266])).
fof(f16857,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16455,f511])).
fof(f16455,plain,(
  ~sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2264])).
fof(f16856,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16451,f16167])).
fof(f16451,plain,(
  ( ! [X2,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9478])).
fof(f9478,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f3584,f3583,f3594,f346])).
fof(f16855,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl52 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16854,f900])).
fof(f16854,plain,(
  ~sP11('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl52 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16450,f16152])).
fof(f16450,plain,(
  ( ! [X0] : (~sP11('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl0 | $spl5 | $spl8 | $spl52 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9094])).
fof(f9094,plain,(
  ( ! [X0,X1] : (~sP11('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))))) ) | ($spl0 | $spl52)),
  inference(unit_resulting_resolution,[],[f434,f3531,f4930])).
fof(f16852,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl46 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16851,f520])).
fof(f16851,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','T_0') | ($spl0 | $spl5 | $spl8 | $spl46 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16448,f16152])).
fof(f16448,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0','T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl46 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9089])).
fof(f9089,plain,(
  ( ! [X0,X1] : (~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'F_0','T_0')) ) | ($spl0 | $spl46)),
  inference(unit_resulting_resolution,[],[f434,f3531,f4921])).
fof(f16848,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl36 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16847,f878])).
fof(f16847,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0') | ($spl0 | $spl5 | $spl8 | $spl36 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16445,f16152])).
fof(f16445,plain,(
  ( ! [X0] : (~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl36 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9086])).
fof(f9086,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'T_0')) ) | ($spl0 | $spl36)),
  inference(unit_resulting_resolution,[],[f434,f3531,f4906])).
fof(f16846,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl34 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16845,f517])).
fof(f16845,plain,(
  ~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl34 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16444,f16152])).
fof(f16444,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'T_0','F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl34 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9082])).
fof(f9082,plain,(
  ( ! [X0,X1] : (~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'T_0','F_0')) ) | ($spl0 | $spl34)),
  inference(unit_resulting_resolution,[],[f434,f3531,f4903])).
fof(f16844,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl32 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16843,f515])).
fof(f16843,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl32 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16443,f16152])).
fof(f16443,plain,(
  ( ! [X0] : (~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'T_0','F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl32 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9078])).
fof(f9078,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'T_0','F_0')) ) | ($spl0 | $spl32)),
  inference(unit_resulting_resolution,[],[f434,f3531,f4900])).
fof(f16842,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl30 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16841,f513])).
fof(f16841,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl30 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16442,f16152])).
fof(f16442,plain,(
  ( ! [X0] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'T_0','F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl30 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9074])).
fof(f9074,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'T_0','F_0')) ) | ($spl0 | $spl30)),
  inference(unit_resulting_resolution,[],[f434,f3531,f4897])).
fof(f16839,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16838,f16208])).
fof(f16838,plain,(
  ~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16837,f16152])).
fof(f16837,plain,(
  ~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16440,f16152])).
fof(f16440,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9070])).
fof(f9070,plain,(
  ( ! [X0,X1] : (~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f3531,f3199])).
fof(f16834,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16833,f519])).
fof(f16833,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16438,f16152])).
fof(f16438,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9065])).
fof(f9065,plain,(
  ( ! [X0,X1] : (~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f3531,f3196])).
fof(f16828,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16827,f731])).
fof(f16827,plain,(
  ( ! [X0] : (~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16435,f16152])).
fof(f16435,plain,(
  ( ! [X0] : (~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9062])).
fof(f9062,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f3531,f3191])).
fof(f16826,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16825,f516])).
fof(f16825,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16434,f16152])).
fof(f16434,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9058])).
fof(f9058,plain,(
  ( ! [X0,X1] : (~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f3531,f3190])).
fof(f16824,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16823,f514])).
fof(f16823,plain,(
  ( ! [X0] : (~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16433,f16152])).
fof(f16433,plain,(
  ( ! [X0] : (~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9054])).
fof(f9054,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f3531,f3189])).
fof(f16822,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16821,f512])).
fof(f16821,plain,(
  ( ! [X0] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16432,f16152])).
fof(f16432,plain,(
  ( ! [X0] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9050])).
fof(f9050,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f3531,f3188])).
fof(f16801,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16800,f16208])).
fof(f16800,plain,(
  ~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16799,f16152])).
fof(f16799,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16419,f16152])).
fof(f16419,plain,(
  ( ! [X2,X0,X3] : (~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X3),s(fun(sK12,bool),sK14)))))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9037])).
fof(f9037,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X3),s(fun(sK12,bool),sK14)))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3531,f434,f3531,f413])).
fof(f16798,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16797,f1125])).
fof(f16797,plain,(
  ~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16418,f16152])).
fof(f16418,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0','T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9036])).
fof(f9036,plain,(
  ( ! [X0,X1] : (~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'F_0','T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f3531,f413])).
fof(f16792,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16791,f519])).
fof(f16791,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16414,f16152])).
fof(f16414,plain,(
  ( ! [X2,X0,X3] : (~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X3),s(fun(sK12,bool),sK14)))))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9013])).
fof(f9013,plain,(
  ( ! [X2,X0,X3,X1] : (~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X3),s(fun(sK12,bool),sK14)))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f3531,f3531,f406])).
fof(f16790,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16789,f519])).
fof(f16789,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16413,f16152])).
fof(f16413,plain,(
  ( ! [X0] : (~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f9012])).
fof(f9012,plain,(
  ( ! [X0,X1] : (~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f244,f3531,f406])).
fof(f16786,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16785,f956])).
fof(f16785,plain,(
  ~sP9('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16410,f16152])).
fof(f16410,plain,(
  ( ! [X0] : (~sP9('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0','F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8971])).
fof(f8971,plain,(
  ( ! [X0,X1] : (~sP9('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'F_0','F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f434,f3531,f396])).
fof(f16778,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16777,f518])).
fof(f16777,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16405,f16152])).
fof(f16405,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0','F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8948])).
fof(f8948,plain,(
  ( ! [X0,X1] : (~sP8('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'F_0','F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f434,f3531,f389])).
fof(f16774,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16773,f731])).
fof(f16773,plain,(
  ( ! [X0] : (~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16402,f16152])).
fof(f16402,plain,(
  ( ! [X2,X0,X3] : (~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X3),s(fun(sK12,bool),sK14)))))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8925])).
fof(f8925,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X3),s(fun(sK12,bool),sK14)))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3531,f434,f3531,f381])).
fof(f16772,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16771,f731])).
fof(f16771,plain,(
  ~sP7('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16401,f16152])).
fof(f16401,plain,(
  ( ! [X0] : (~sP7('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8924])).
fof(f8924,plain,(
  ( ! [X0,X1] : (~sP7('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f3531,f381])).
fof(f16757,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16756,f516])).
fof(f16756,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16392,f16152])).
fof(f16392,plain,(
  ( ! [X2,X0,X3] : (~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X3),s(fun(sK12,bool),sK14)))))),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8898])).
fof(f8898,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X3),s(fun(sK12,bool),sK14)))))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3531,f434,f3531,f372])).
fof(f16755,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16754,f516])).
fof(f16754,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16391,f16152])).
fof(f16391,plain,(
  ( ! [X0] : (~sP6('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8897])).
fof(f8897,plain,(
  ( ! [X0,X1] : (~sP6('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f3531,f372])).
fof(f16740,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16739,f1254])).
fof(f16739,plain,(
  ( ! [X2,X3] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16382,f16152])).
fof(f16382,plain,(
  ( ! [X2,X0,X3] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8864])).
fof(f8864,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f504,f3531,f363])).
fof(f16738,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16737,f1252])).
fof(f16737,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16381,f16152])).
fof(f16381,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0','T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8863])).
fof(f8863,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'F_0','T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f465,f3531,f363])).
fof(f16734,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16733,f16178])).
fof(f16733,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16732,f16152])).
fof(f16732,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16378,f16152])).
fof(f16378,plain,(
  ( ! [X2,X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8853])).
fof(f8853,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3536,f3531,f362])).
fof(f16731,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16730,f16178])).
fof(f16730,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16729,f16152])).
fof(f16729,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16377,f16152])).
fof(f16377,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8852])).
fof(f16727,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16726,f16563])).
fof(f16726,plain,(
  ( ! [X2,X3] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16375,f16152])).
fof(f16375,plain,(
  ( ! [X2,X0,X3] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8845])).
fof(f8845,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f505,f3531,f362])).
fof(f16716,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16715,f514])).
fof(f16715,plain,(
  ( ! [X0] : (~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16367,f16152])).
fof(f16367,plain,(
  ( ! [X2,X0,X3] : (~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X3),s(fun(sK12,bool),sK14)))))),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8828])).
fof(f8828,plain,(
  ( ! [X2,X0,X3,X1] : (~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X3),s(fun(sK12,bool),sK14)))))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3531,f434,f3531,f353])).
fof(f16714,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16713,f514])).
fof(f16713,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16366,f16152])).
fof(f16366,plain,(
  ( ! [X0] : (~sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8827])).
fof(f8827,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f3531,f353])).
fof(f16711,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16710,f512])).
fof(f16710,plain,(
  ( ! [X0] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16364,f16152])).
fof(f16364,plain,(
  ( ! [X2,X0,X3] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X3),s(fun(sK12,bool),sK14)))))),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8811])).
fof(f8811,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X3),s(fun(sK12,bool),sK14)))))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3531,f434,f3531,f349])).
fof(f16709,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16708,f512])).
fof(f16708,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16363,f16152])).
fof(f16363,plain,(
  ( ! [X0] : (~sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8810])).
fof(f8810,plain,(
  ( ! [X0,X1] : (~sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f3531,f349])).
fof(f16703,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16702,f16598])).
fof(f16702,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16360,f16152])).
fof(f16360,plain,(
  ( ! [X2,X1] : (sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8799])).
fof(f8799,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f443,f3570,f3585,f3531,f346])).
fof(f16697,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16696,f16598])).
fof(f16696,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16695,f16152])).
fof(f16695,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16357,f16152])).
fof(f16357,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8796])).
fof(f8796,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f443,f2286,f2299,f3531,f346])).
fof(f16694,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16693,f16178])).
fof(f16693,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16356,f16152])).
fof(f16356,plain,(
  ( ! [X2,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8795])).
fof(f8795,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3583,f438,f449,f3531,f346])).
fof(f16692,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16691,f16178])).
fof(f16691,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16690,f16152])).
fof(f16690,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16355,f16152])).
fof(f16355,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8794])).
fof(f8794,plain,(
  ( ! [X0,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2297,f438,f449,f3531,f346])).
fof(f16689,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16688,f941])).
fof(f16688,plain,(
  sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16354,f16152])).
fof(f16354,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8793])).
fof(f8793,plain,(
  ( ! [X0,X1] : (sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f443,f438,f449,f3531,f346])).
fof(f16685,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16684,f1036])).
fof(f16684,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16352,f16152])).
fof(f16352,plain,(
  ( ! [X0] : (~sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8782])).
fof(f8782,plain,(
  ( ! [X0,X1] : (~sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f3531,f345])).
fof(f16683,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16682,f508])).
fof(f16682,plain,(
  sP1('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16351,f16152])).
fof(f16351,plain,(
  ( ! [X0] : (sP1('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8779])).
fof(f8779,plain,(
  ( ! [X0,X1] : (sP1('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f3531,f336])).
fof(f16681,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16680,f506])).
fof(f16680,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16350,f16152])).
fof(f16350,plain,(
  ( ! [X0] : (sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8778])).
fof(f8778,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f3531,f333])).
fof(f16677,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16676,f510])).
fof(f16676,plain,(
  ~sP1('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16348,f16152])).
fof(f16348,plain,(
  ( ! [X0] : (~sP1('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8776])).
fof(f8776,plain,(
  ( ! [X0,X1] : (~sP1('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f3531,f332])).
fof(f16675,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16674,f511])).
fof(f16674,plain,(
  ~sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16347,f16152])).
fof(f16347,plain,(
  ( ! [X0] : (~sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f8774])).
fof(f8774,plain,(
  ( ! [X0,X1] : (~sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))),'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f3531,f332])).
fof(f16664,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16322,f1037])).
fof(f16322,plain,(
  sP5('F_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f6273])).
fof(f6273,plain,(
  ( ! [X0] : (sP5('F_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f446,f441,f3586,f347])).
fof(f16663,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16319,f1035])).
fof(f16319,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f6211])).
fof(f16661,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16315,f942])).
fof(f16315,plain,(
  ~sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f6141])).
fof(f16660,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16314,f939])).
fof(f16314,plain,(
  ~sP5('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f6093])).
fof(f16658,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16312,f938])).
fof(f16312,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f6081])).
fof(f6081,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f443,f450,f3570,f346])).
fof(f16657,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16311,f1038])).
fof(f16311,plain,(
  ~sP5('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f6035])).
fof(f16653,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16652,f16208])).
fof(f16652,plain,(
  ~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16302,f16152])).
fof(f16302,plain,(
  ( ! [X1] : (~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3724])).
fof(f3724,plain,(
  ( ! [X0,X1] : (~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2255,f434,f2255,f413])).
fof(f16651,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16301,f1125])).
fof(f16301,plain,(
  ~sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3723])).
fof(f16650,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16300,f900])).
fof(f16300,plain,(
  ~sP11('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3719])).
fof(f16648,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16295,f520])).
fof(f16295,plain,(
  ( ! [X1] : (~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3704])).
fof(f3704,plain,(
  ( ! [X0,X1] : (~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f2255,f2255,f406])).
fof(f16647,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16294,f519])).
fof(f16294,plain,(
  ~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3703])).
fof(f3703,plain,(
  ( ! [X0] : (~sP10('T_0','F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f244,f2255,f406])).
fof(f16646,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16293,f520])).
fof(f16293,plain,(
  ~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3697])).
fof(f3697,plain,(
  ( ! [X0] : (~sP10('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f244,f2255,f406])).
fof(f16645,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16290,f956])).
fof(f16290,plain,(
  ~sP9('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3673])).
fof(f16642,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16283,f518])).
fof(f16283,plain,(
  ~sP8('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3654])).
fof(f3654,plain,(
  ( ! [X0] : (~sP8('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f434,f434,f2255,f389])).
fof(f16641,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16280,f878])).
fof(f16280,plain,(
  ( ! [X1] : (~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3641])).
fof(f3641,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2255,f434,f2255,f381])).
fof(f16640,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16279,f731])).
fof(f16279,plain,(
  ~sP7('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3640])).
fof(f16639,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16278,f878])).
fof(f16278,plain,(
  ~sP7('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3636])).
fof(f16633,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16269,f517])).
fof(f16269,plain,(
  ( ! [X1] : (~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3616])).
fof(f3616,plain,(
  ( ! [X0,X1] : (~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2255,f434,f2255,f372])).
fof(f16632,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16268,f516])).
fof(f16268,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3615])).
fof(f3615,plain,(
  ( ! [X0] : (~sP6('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f2255,f372])).
fof(f16631,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16267,f517])).
fof(f16267,plain,(
  ~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3609])).
fof(f3609,plain,(
  ( ! [X0] : (~sP6('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f2255,f372])).
fof(f16626,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16261,f1254])).
fof(f16261,plain,(
  ( ! [X2,X1] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3594])).
fof(f16625,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16260,f1252])).
fof(f16260,plain,(
  ~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0','T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3593])).
fof(f3593,plain,(
  ( ! [X0] : (~sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'F_0','T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f465,f2255,f363])).
fof(f16624,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16623,f16178])).
fof(f16623,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16257,f16152])).
fof(f16257,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3590])).
fof(f3590,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2260,f2255,f362])).
fof(f16621,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16254,f16563])).
fof(f16254,plain,(
  ( ! [X2,X1] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3587])).
fof(f16620,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16249,f515])).
fof(f16249,plain,(
  ( ! [X1] : (~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3578])).
fof(f3578,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2255,f434,f2255,f353])).
fof(f16619,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16248,f514])).
fof(f16248,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3577])).
fof(f16618,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16247,f515])).
fof(f16247,plain,(
  ~sP3('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3571])).
fof(f16617,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16245,f513])).
fof(f16245,plain,(
  ( ! [X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))),'F_0')) ) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3565])).
fof(f3565,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))),'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2255,f434,f2255,f349])).
fof(f16616,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16244,f512])).
fof(f16244,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3564])).
fof(f16615,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16243,f513])).
fof(f16243,plain,(
  ~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3558])).
fof(f3558,plain,(
  ( ! [X0] : (~sP4('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0','F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f434,f2255,f349])).
fof(f16613,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16612,f16598])).
fof(f16612,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16241,f16152])).
fof(f16241,plain,(
  sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3556])).
fof(f3556,plain,(
  ( ! [X0] : (sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f443,f2286,f2299,f2255,f346])).
fof(f16611,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16610,f16178])).
fof(f16610,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16240,f16152])).
fof(f16240,plain,(
  sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3555])).
fof(f3555,plain,(
  ( ! [X0] : (sP5('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2297,f438,f449,f2255,f346])).
fof(f16609,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16239,f941])).
fof(f16239,plain,(
  sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3554])).
fof(f3554,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f443,f438,f449,f2255,f346])).
fof(f16607,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16237,f1036])).
fof(f16237,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3548])).
fof(f16606,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16236,f508])).
fof(f16236,plain,(
  sP1('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3545])).
fof(f16605,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16235,f506])).
fof(f16235,plain,(
  sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3544])).
fof(f16603,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16233,f510])).
fof(f16233,plain,(
  ~sP1('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3542])).
fof(f3542,plain,(
  ( ! [X0] : (~sP1('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f2255,f332])).
fof(f16602,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16232,f511])).
fof(f16232,plain,(
  ~sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),'T_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f3540])).
fof(f3540,plain,(
  ( ! [X0] : (~sP1('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))),'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f244,f2255,f332])).
fof(f16597,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16224,f16152])).
fof(f16224,plain,(
  s(fun(sK12,bool),'const.pred_set.EMPTY_0') != s(fun(sK12,bool),sK14) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f430])).
fof(f16596,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16595,f16208])).
fof(f16595,plain,(
  ( ! [X4,X0,X3,X1] : (~sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f16218,f16152])).
fof(f16218,plain,(
  ( ! [X4,X0,X5,X3,X1] : (~sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2200])).
fof(f16594,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16217,f1125])).
fof(f16217,plain,(
  ( ! [X0,X1] : (~sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','T_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2199])).
fof(f16593,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16216,f900])).
fof(f16216,plain,(
  ( ! [X0,X1] : (~sP11('T_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2196])).
fof(f16591,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16211,f520])).
fof(f16211,plain,(
  ( ! [X4,X0,X5,X3,X1] : (~sP10('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2182])).
fof(f16590,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16210,f519])).
fof(f16210,plain,(
  ( ! [X0,X1] : (~sP10('T_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2181])).
fof(f16589,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16209,f520])).
fof(f16209,plain,(
  ( ! [X0,X1] : (~sP10('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','T_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2175])).
fof(f16588,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16206,f956])).
fof(f16206,plain,(
  ( ! [X0,X1] : (~sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2156])).
fof(f16585,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16199,f518])).
fof(f16199,plain,(
  ( ! [X0,X1] : (~sP8('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2141])).
fof(f16584,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16196,f878])).
fof(f16196,plain,(
  ( ! [X4,X0,X5,X3,X1] : (~sP7('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2131])).
fof(f16583,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16195,f731])).
fof(f16195,plain,(
  ( ! [X0,X1] : (~sP7('F_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2130])).
fof(f16582,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16194,f878])).
fof(f16194,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2125])).
fof(f16576,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16185,f517])).
fof(f16185,plain,(
  ( ! [X4,X0,X5,X3,X1] : (~sP6('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2105])).
fof(f16575,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16184,f516])).
fof(f16184,plain,(
  ( ! [X0,X1] : (~sP6('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2104])).
fof(f16574,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16183,f517])).
fof(f16183,plain,(
  ( ! [X0,X1] : (~sP6('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2098])).
fof(f16573,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16180,f1254])).
fof(f16180,plain,(
  ( ! [X4,X0,X3,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2089])).
fof(f16572,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16179,f1252])).
fof(f16179,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','T_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2087])).
fof(f16571,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16175,f16563])).
fof(f16175,plain,(
  ( ! [X4,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2081])).
fof(f16570,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16170,f515])).
fof(f16170,plain,(
  ( ! [X4,X0,X5,X3,X1] : (~sP3('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2071])).
fof(f16569,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16169,f514])).
fof(f16169,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2070])).
fof(f16568,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16168,f515])).
fof(f16168,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2064])).
fof(f16567,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16166,f513])).
fof(f16166,plain,(
  ( ! [X4,X0,X5,X3,X1] : (~sP4('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2058])).
fof(f16566,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16165,f512])).
fof(f16165,plain,(
  ( ! [X0,X1] : (~sP4('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2057])).
fof(f16565,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16164,f513])).
fof(f16164,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0','F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2051])).
fof(f16564,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16163,f941])).
fof(f16163,plain,(
  ( ! [X0,X1] : (sP5('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2050])).
fof(f16562,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16161,f1036])).
fof(f16161,plain,(
  ( ! [X0,X1] : (~sP5('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2047])).
fof(f16561,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16160,f508])).
fof(f16160,plain,(
  ( ! [X0,X1] : (sP1('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2043])).
fof(f16560,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16159,f506])).
fof(f16159,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2041])).
fof(f16558,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16157,f510])).
fof(f16157,plain,(
  ( ! [X0,X1] : (~sP1('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2038])).
fof(f16557,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16156,f511])).
fof(f16156,plain,(
  ( ! [X0,X1] : (~sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'T_0')) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16152,f2035])).
fof(f16151,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16019,f244])).
fof(f16019,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16016,f434])).
fof(f16150,plain,(
  $false | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16149,f10273])).
fof(f16149,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16017,f16016])).
fof(f16017,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f16016,f263])).
fof(f16015,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f15858,f2667])).
fof(f15858,plain,(
  sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f15769,f5976])).
fof(f16014,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f15857,f2275])).
fof(f15857,plain,(
  ~sP5('T_0','T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f15769,f5975])).
fof(f16013,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f15856,f2662])).
fof(f15856,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f15769,f5971])).
fof(f16012,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f15855,f2654])).
fof(f15855,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f15769,f5970])).
fof(f16011,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f16010,f465])).
fof(f16010,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(forward_demodulation,[],[f15854,f2260])).
fof(f15854,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f15769,f5967])).
fof(f16009,plain,(
  $false | ($spl0 | $spl5 | $spl8 | $spl70 | $spl110)),
  inference(subsumption_resolution,[],[f15853,f2260])).
fof(f15853,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl5 | $spl8 | $spl70 | $spl110)),
  inference(backward_demodulation,[],[f15769,f5942])).
fof(f15763,plain,(
  $false | ($spl52 | $spl110)),
  inference(unit_resulting_resolution,[],[f592,f434,f10273,f4930])).
fof(f15756,plain,(
  $false | ($spl50 | $spl110)),
  inference(unit_resulting_resolution,[],[f1149,f10273,f4927])).
fof(f15754,plain,(
  $false | ($spl46 | $spl110)),
  inference(unit_resulting_resolution,[],[f463,f434,f10273,f4921])).
fof(f15747,plain,(
  $false | ($spl42 | $spl110)),
  inference(unit_resulting_resolution,[],[f975,f10273,f4915])).
fof(f15745,plain,(
  $false | ($spl40 | $spl110)),
  inference(unit_resulting_resolution,[],[f975,f10273,f4912])).
fof(f15743,plain,(
  $false | ($spl36 | $spl110)),
  inference(unit_resulting_resolution,[],[f563,f434,f10273,f4906])).
fof(f15732,plain,(
  $false | ($spl34 | $spl110)),
  inference(unit_resulting_resolution,[],[f454,f434,f10273,f4903])).
fof(f15725,plain,(
  $false | ($spl32 | $spl110)),
  inference(unit_resulting_resolution,[],[f446,f434,f10273,f4900])).
fof(f15716,plain,(
  $false | ($spl30 | $spl110)),
  inference(unit_resulting_resolution,[],[f441,f434,f10273,f4897])).
fof(f15706,plain,(
  $false | ($spl28 | $spl110)),
  inference(unit_resulting_resolution,[],[f739,f434,f10273,f4894])).
fof(f15686,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f2410,f10273,f3199])).
fof(f15684,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f2409,f10273,f3199])).
fof(f15681,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2407,f10273,f3198])).
fof(f3198,plain,(
  ( ! [X30,X31] : (sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X30,X31) | ~p(s(bool,X30))) ) | $spl0),
  inference(subsumption_resolution,[],[f3126,f244])).
fof(f3126,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | ~p(s(bool,X30)) | sP11('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X30,X31)) ) | $spl0),
  inference(superposition,[],[f412,f2260])).
fof(f15679,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f520,f10273,f3196])).
fof(f15677,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f463,f10273,f3196])).
fof(f15671,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2369,f10273,f3194])).
fof(f3194,plain,(
  ( ! [X23,X22] : (sP9('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X22,X23) | ~p(s(bool,X22))) ) | $spl0),
  inference(subsumption_resolution,[],[f3122,f244])).
fof(f3122,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X22)) | sP9('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X22,X23)) ) | $spl0),
  inference(superposition,[],[f395,f2260])).
fof(f15669,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2369,f10273,f3193])).
fof(f3193,plain,(
  ( ! [X21,X20] : (sP9('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X21,X20) | ~p(s(bool,X20))) ) | $spl0),
  inference(subsumption_resolution,[],[f3121,f244])).
fof(f3121,plain,(
  ( ! [X21,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | sP9('const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X21,X20)) ) | $spl0),
  inference(superposition,[],[f394,f2260])).
fof(f15663,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f878,f10273,f3191])).
fof(f15661,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f563,f10273,f3191])).
fof(f15658,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f517,f10273,f3190])).
fof(f15656,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f454,f10273,f3190])).
fof(f15650,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f515,f10273,f3189])).
fof(f15648,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f446,f10273,f3189])).
fof(f15642,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f513,f10273,f3188])).
fof(f15640,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f441,f10273,f3188])).
fof(f15629,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f434,f2692,f10273,f3187])).
fof(f15625,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f572,f458,f10273,f417])).
fof(f15623,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3647,f3660,f10273,f417])).
fof(f3660,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X1),s(fun(sK12,bool),sK14)))),X2)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2255,f391])).
fof(f15621,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2350,f2358,f10273,f417])).
fof(f2358,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f425,f391])).
fof(f15619,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f573,f457,f10273,f417])).
fof(f15617,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3644,f3659,f10273,f417])).
fof(f3659,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,X2),s(fun(sK12,bool),sK14)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f2255,f390])).
fof(f15615,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2348,f2357,f10273,f417])).
fof(f2357,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f425,f390])).
fof(f15613,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f572,f456,f10273,f417])).
fof(f15601,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f588,f463,f10273,f416])).
fof(f15599,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3674,f3709,f10273,f416])).
fof(f15597,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2370,f2395,f10273,f416])).
fof(f15595,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f589,f462,f10273,f416])).
fof(f15593,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f589,f519,f10273,f416])).
fof(f15591,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f588,f460,f10273,f416])).
fof(f15579,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f592,f434,f10273,f10273,f413])).
fof(f15555,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f1123,f434,f244,f10273,f413])).
fof(f15529,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f592,f434,f10273,f10273,f413])).
fof(f15505,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f592,f434,f244,f10273,f413])).
fof(f15484,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f588,f10273,f10273,f412])).
fof(f15482,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f589,f244,f10273,f412])).
fof(f15470,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f588,f10273,f10273,f412])).
fof(f15468,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f1149,f244,f10273,f412])).
fof(f15456,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f588,f434,f10273,f411])).
fof(f15444,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f588,f434,f434,f10273,f410])).
fof(f15431,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f1137,f10273,f408])).
fof(f15429,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f460,f10273,f434,f10273,f406])).
fof(f15427,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f460,f244,f434,f10273,f406])).
fof(f15409,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f460,f10273,f434,f10273,f406])).
fof(f15407,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f463,f244,f434,f10273,f406])).
fof(f15388,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f588,f463,f434,f10273,f405])).
fof(f15385,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f588,f460,f434,f10273,f405])).
fof(f15360,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f588,f463,f434,f10273,f404])).
fof(f15357,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f588,f460,f434,f10273,f404])).
fof(f15332,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f588,f463,f10273,f10273,f403])).
fof(f15330,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3674,f3709,f10273,f10273,f403])).
fof(f15328,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2370,f2395,f10273,f10273,f403])).
fof(f15326,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f589,f462,f10273,f10273,f403])).
fof(f15324,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f589,f519,f10273,f10273,f403])).
fof(f15322,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f588,f460,f10273,f10273,f403])).
fof(f15282,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f463,f3674,f10273,f10273,f403])).
fof(f15280,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f463,f2370,f10273,f10273,f403])).
fof(f15278,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f463,f589,f10273,f10273,f403])).
fof(f15276,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f463,f588,f10273,f10273,f403])).
fof(f15264,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f463,f3674,f244,f10273,f403])).
fof(f15262,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f463,f2370,f244,f10273,f403])).
fof(f15260,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f463,f589,f244,f10273,f403])).
fof(f15240,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f1123,f434,f10273,f400])).
fof(f15216,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f954,f434,f434,f10273,f396])).
fof(f15192,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f572,f10273,f10273,f395])).
fof(f15190,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f573,f244,f10273,f395])).
fof(f15179,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f572,f10273,f10273,f395])).
fof(f15177,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f975,f244,f10273,f395])).
fof(f15166,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f572,f10273,f10273,f394])).
fof(f15164,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f572,f244,f10273,f394])).
fof(f15153,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f572,f10273,f10273,f394])).
fof(f15151,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f975,f244,f10273,f394])).
fof(f15140,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f966,f10273,f391])).
fof(f15138,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f966,f10273,f390])).
fof(f15136,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f458,f434,f434,f10273,f389])).
fof(f15123,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f572,f458,f434,f10273,f387])).
fof(f15121,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3644,f3659,f434,f10273,f387])).
fof(f15119,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2348,f2357,f434,f10273,f387])).
fof(f15117,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f572,f456,f434,f10273,f387])).
fof(f15082,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f458,f3647,f434,f10273,f386])).
fof(f15080,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f458,f2350,f434,f10273,f386])).
fof(f15078,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f458,f573,f434,f10273,f386])).
fof(f15049,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f954,f434,f10273,f383])).
fof(f15027,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f954,f434,f10273,f382])).
fof(f15005,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f552,f434,f10273,f10273,f381])).
fof(f14986,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f552,f434,f244,f10273,f381])).
fof(f14971,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f552,f434,f10273,f10273,f381])).
fof(f14952,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f563,f434,f244,f10273,f381])).
fof(f14937,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f540,f10273,f10273,f10273,f378])).
fof(f14935,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f545,f244,f10273,f10273,f378])).
fof(f14921,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f540,f10273,f244,f10273,f378])).
fof(f14919,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f868,f244,f244,f10273,f378])).
fof(f14893,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f540,f10273,f10273,f10273,f378])).
fof(f14891,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f540,f244,f10273,f10273,f378])).
fof(f14877,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f540,f10273,f244,f10273,f378])).
fof(f14875,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f540,f244,f244,f10273,f378])).
fof(f14853,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f540,f10273,f10273,f10273,f378])).
fof(f14851,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f540,f244,f10273,f10273,f378])).
fof(f14837,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f545,f10273,f244,f10273,f378])).
fof(f14835,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f545,f244,f244,f10273,f378])).
fof(f14813,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f858,f10273,f373])).
fof(f14811,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f452,f434,f10273,f10273,f372])).
fof(f14802,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f452,f434,f244,f10273,f372])).
fof(f14791,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f452,f434,f10273,f10273,f372])).
fof(f14782,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f454,f434,f244,f10273,f372])).
fof(f14771,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f454,f3596,f10273,f10273,f368])).
fof(f14769,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f452,f3595,f10273,f10273,f368])).
fof(f14767,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f454,f2308,f10273,f10273,f368])).
fof(f14765,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f452,f2307,f10273,f10273,f368])).
fof(f14757,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f454,f545,f10273,f10273,f368])).
fof(f14755,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f452,f540,f10273,f10273,f368])).
fof(f14746,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f452,f3595,f244,f10273,f368])).
fof(f14744,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f452,f2307,f244,f10273,f368])).
fof(f14739,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f452,f540,f244,f10273,f368])).
fof(f14717,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f454,f3596,f10273,f10273,f368])).
fof(f14715,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f452,f3595,f10273,f10273,f368])).
fof(f14713,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f454,f2308,f10273,f10273,f368])).
fof(f14711,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f452,f2307,f10273,f10273,f368])).
fof(f14703,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f454,f545,f10273,f10273,f368])).
fof(f14701,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f452,f540,f10273,f10273,f368])).
fof(f14692,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f454,f3596,f244,f10273,f368])).
fof(f14690,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f454,f2308,f244,f10273,f368])).
fof(f14685,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f454,f545,f244,f10273,f368])).
fof(f14663,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f563,f434,f10273,f367])).
fof(f14644,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f552,f434,f10273,f366])).
fof(f14625,plain,(
  $false | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11592,f4963,f10273,f363])).
fof(f14623,plain,(
  $false | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7430,f4960,f10273,f363])).
fof(f14621,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6054,f6025,f10273,f363])).
fof(f6054,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)),'T_0') | ($spl5 | $spl80)),
  inference(unit_resulting_resolution,[],[f434,f6025,f361])).
fof(f14619,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5971,f5942,f10273,f363])).
fof(f14617,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f939,f504,f10273,f363])).
fof(f14615,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f737,f465,f10273,f363])).
fof(f14613,plain,(
  $false | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f11596,f4963,f10273,f363])).
fof(f14611,plain,(
  $false | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7434,f4960,f10273,f363])).
fof(f14609,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f6059,f6025,f10273,f363])).
fof(f6059,plain,(
  sP5('F_0','T_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl80)),
  inference(unit_resulting_resolution,[],[f434,f6025,f361])).
fof(f14607,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f5976,f5942,f10273,f363])).
fof(f14605,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f942,f504,f10273,f363])).
fof(f14603,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f739,f465,f10273,f363])).
fof(f14546,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f1037,f505,f10273,f362])).
fof(f14544,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6141,f3536,f10273,f362])).
fof(f14542,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2667,f2260,f10273,f362])).
fof(f14540,plain,(
  $false | ($spl14 | $spl110)),
  inference(unit_resulting_resolution,[],[f11222,f3975,f10273,f362])).
fof(f14538,plain,(
  $false | ($spl10 | $spl110)),
  inference(unit_resulting_resolution,[],[f6901,f3920,f10273,f362])).
fof(f14536,plain,(
  $false | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f12274,f11587,f10273,f362])).
fof(f14534,plain,(
  $false | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f9535,f7425,f10273,f362])).
fof(f14532,plain,(
  $false | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f8592,f6844,f10273,f362])).
fof(f14528,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f1035,f505,f10273,f362])).
fof(f14526,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f6093,f3536,f10273,f362])).
fof(f14524,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2662,f2260,f10273,f362])).
fof(f14522,plain,(
  $false | ($spl14 | $spl110)),
  inference(unit_resulting_resolution,[],[f11220,f3975,f10273,f362])).
fof(f14520,plain,(
  $false | ($spl10 | $spl110)),
  inference(unit_resulting_resolution,[],[f6898,f3920,f10273,f362])).
fof(f14518,plain,(
  $false | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f12269,f11587,f10273,f362])).
fof(f14516,plain,(
  $false | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f9530,f7425,f10273,f362])).
fof(f14514,plain,(
  $false | ($spl105 | $spl110)),
  inference(unit_resulting_resolution,[],[f8589,f6844,f10273,f362])).
fof(f15847,plain,(
  $false | ($spl0 | $spl110)),
  inference(trivial_inequality_removal,[],[f14502])).
fof(f14502,plain,(
  s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) != s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3804,f10273,f362])).
fof(f15846,plain,(
  $false | $spl110),
  inference(trivial_inequality_removal,[],[f14490])).
fof(f14490,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl110),
  inference(unit_resulting_resolution,[],[f842,f10273,f362])).
fof(f15845,plain,(
  $false | $spl110),
  inference(trivial_inequality_removal,[],[f14484])).
fof(f14484,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl110),
  inference(unit_resulting_resolution,[],[f833,f10273,f362])).
fof(f14465,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f837,f10273,f360])).
fof(f14463,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f837,f10273,f359])).
fof(f14461,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f837,f10273,f358])).
fof(f14459,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f822,f10273,f354])).
fof(f14457,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f444,f434,f10273,f10273,f353])).
fof(f14444,plain,(
  $false | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7900,f434,f244,f10273,f353])).
fof(f14427,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f444,f434,f10273,f10273,f353])).
fof(f14414,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f446,f434,f244,f10273,f353])).
fof(f14397,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f801,f10273,f350])).
fof(f14395,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f439,f434,f10273,f10273,f349])).
fof(f14381,plain,(
  $false | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f9777,f434,f244,f10273,f349])).
fof(f14363,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f439,f434,f10273,f10273,f349])).
fof(f14349,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f441,f434,f244,f10273,f349])).
fof(f14239,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f443,f439,f3585,f10273,f346])).
fof(f14237,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f444,f439,f3585,f10273,f346])).
fof(f14233,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f444,f3570,f3585,f10273,f346])).
fof(f14229,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f443,f3570,f3585,f10273,f346])).
fof(f14212,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f2297,f439,f2299,f10273,f346])).
fof(f14210,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f443,f439,f2299,f10273,f346])).
fof(f14208,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f444,f439,f2299,f10273,f346])).
fof(f14203,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f444,f2286,f2299,f10273,f346])).
fof(f14200,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f2297,f2286,f2299,f10273,f346])).
fof(f14198,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f443,f2286,f2299,f10273,f346])).
fof(f14176,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f3583,f512,f449,f10273,f346])).
fof(f14174,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f2297,f512,f449,f10273,f346])).
fof(f14172,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f941,f443,f512,f449,f10273,f346])).
fof(f14170,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f941,f514,f512,f449,f10273,f346])).
fof(f14165,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f3583,f439,f449,f10273,f346])).
fof(f14163,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f2297,f439,f449,f10273,f346])).
fof(f14161,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f760,f443,f439,f449,f10273,f346])).
fof(f14159,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f760,f444,f439,f449,f10273,f346])).
fof(f14154,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f941,f514,f438,f449,f10273,f346])).
fof(f14152,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f760,f444,f438,f449,f10273,f346])).
fof(f14150,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f3583,f438,f449,f10273,f346])).
fof(f14148,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f2297,f438,f449,f10273,f346])).
fof(f14146,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f760,f443,f438,f449,f10273,f346])).
fof(f14108,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f3583,f512,f3584,f10273,f346])).
fof(f14106,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f514,f512,f3584,f10273,f346])).
fof(f14101,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f3583,f439,f3584,f10273,f346])).
fof(f14099,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f444,f439,f3584,f10273,f346])).
fof(f14090,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3592,f514,f438,f3584,f10273,f346])).
fof(f14088,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f444,f438,f3584,f10273,f346])).
fof(f14086,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3591,f3583,f438,f3584,f10273,f346])).
fof(f14072,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f2297,f512,f2298,f10273,f346])).
fof(f14070,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f514,f512,f2298,f10273,f346])).
fof(f14065,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f2297,f439,f2298,f10273,f346])).
fof(f14063,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f444,f439,f2298,f10273,f346])).
fof(f14055,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f444,f2286,f2298,f10273,f346])).
fof(f14053,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2273,f2297,f2286,f2298,f10273,f346])).
fof(f14050,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2304,f514,f438,f2298,f10273,f346])).
fof(f14048,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f444,f438,f2298,f10273,f346])).
fof(f14046,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2303,f2297,f438,f2298,f10273,f346])).
fof(f14028,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f941,f443,f512,f448,f10273,f346])).
fof(f14026,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f941,f514,f512,f448,f10273,f346])).
fof(f14021,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f760,f443,f439,f448,f10273,f346])).
fof(f14019,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f760,f444,f439,f448,f10273,f346])).
fof(f14012,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f444,f3570,f448,f10273,f346])).
fof(f14010,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f3548,f443,f3570,f448,f10273,f346])).
fof(f14006,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f444,f2286,f448,f10273,f346])).
fof(f14004,plain,(
  $false | ($spl0 | $spl110)),
  inference(unit_resulting_resolution,[],[f2272,f443,f2286,f448,f10273,f346])).
fof(f14001,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f941,f514,f438,f448,f10273,f346])).
fof(f13999,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f760,f444,f438,f448,f10273,f346])).
fof(f13997,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f760,f443,f438,f448,f10273,f346])).
fof(f13896,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f739,f434,f10273,f10273,f345])).
fof(f13841,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f1325,f434,f244,f10273,f345])).
fof(f13808,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f739,f434,f10273,f10273,f345])).
fof(f13753,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f739,f434,f244,f10273,f345])).
fof(f13723,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f485,f434,f10273,f336])).
fof(f13711,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f485,f434,f10273,f333])).
fof(f13699,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f489,f10273,f10273,f332])).
fof(f13697,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f489,f244,f10273,f332])).
fof(f13673,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f489,f10273,f10273,f332])).
fof(f13671,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f495,f244,f10273,f332])).
fof(f13647,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f499,f10273,f318])).
fof(f13645,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f499,f10273,f317])).
fof(f13643,plain,(
  $false | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f10273,f4963,f10273,f311])).
fof(f13641,plain,(
  $false | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f10273,f4960,f10273,f311])).
fof(f13639,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f10273,f6025,f10273,f311])).
fof(f13637,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f10273,f5942,f10273,f311])).
fof(f13635,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f10273,f504,f10273,f311])).
fof(f13633,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f10273,f465,f10273,f311])).
fof(f13631,plain,(
  $false | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f244,f4963,f10273,f311])).
fof(f13629,plain,(
  $false | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f244,f4960,f10273,f311])).
fof(f13627,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f244,f6025,f10273,f311])).
fof(f13625,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f244,f5942,f10273,f311])).
fof(f13623,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f244,f504,f10273,f311])).
fof(f13621,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f244,f465,f10273,f311])).
fof(f13617,plain,(
  $false | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f10273,f4963,f10273,f311])).
fof(f13615,plain,(
  $false | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f10273,f4960,f10273,f311])).
fof(f13613,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f10273,f6025,f10273,f311])).
fof(f13611,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f10273,f5942,f10273,f311])).
fof(f13609,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f10273,f504,f10273,f311])).
fof(f13607,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f10273,f465,f10273,f311])).
fof(f13605,plain,(
  $false | ($spl65 | $spl110)),
  inference(unit_resulting_resolution,[],[f244,f4963,f10273,f311])).
fof(f13603,plain,(
  $false | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f244,f4960,f10273,f311])).
fof(f13601,plain,(
  $false | ($spl5 | $spl80 | $spl110)),
  inference(unit_resulting_resolution,[],[f244,f6025,f10273,f311])).
fof(f13599,plain,(
  $false | ($spl5 | $spl70 | $spl110)),
  inference(unit_resulting_resolution,[],[f244,f5942,f10273,f311])).
fof(f13597,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f244,f504,f10273,f311])).
fof(f13595,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f244,f465,f10273,f311])).
fof(f13591,plain,(
  $false | ($spl61 | $spl110)),
  inference(unit_resulting_resolution,[],[f7425,f10273,f263])).
fof(f13589,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f465,f10273,f260])).
fof(f13587,plain,(
  $false | $spl110),
  inference(unit_resulting_resolution,[],[f434,f10273])).
fof(f13586,plain,(
  ( ! [X178] : (sP1('const.bool.IN_2'(s(sK12,X178),s(fun(sK12,bool),sK14)),'T_0') | s(sK12,sK13) = s(sK12,X178)) ) | $spl172),
  inference(cnf_transformation,[],[f13586_D])).
fof(f13586_D,plain,(
  ( ! [X178] : (sP1('const.bool.IN_2'(s(sK12,X178),s(fun(sK12,bool),sK14)),'T_0') | s(sK12,sK13) = s(sK12,X178)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f10273,plain,(
  ( ! [X115] : (p(s(bool,X115))) ) | $spl110),
  inference(cnf_transformation,[],[f10273_D])).
fof(f10273_D,plain,(
  ( ! [X115] : (p(s(bool,X115))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f13585,plain,(
  ( ! [X178,X179] : (p(s(bool,X179)) | sP1('const.bool.IN_2'(s(sK12,X178),s(fun(sK12,bool),sK14)),'T_0') | s(sK12,sK13) = s(sK12,X178)) ) | ($spl8 | $spl26)),
  inference(subsumption_resolution,[],[f13451,f259])).
fof(f13451,plain,(
  ( ! [X178,X179] : (p(s(bool,X179)) | sP1('const.bool.IN_2'(s(sK12,X178),s(fun(sK12,bool),sK14)),'T_0') | s(sK12,sK13) = s(sK12,X178) | s(bool,'T_0') = s(bool,X179)) ) | ($spl8 | $spl26)),
  inference(superposition,[],[f4891,f4966])).
fof(f13584,plain,(
  ( ! [X176] : (sP1('T_0','const.bool.IN_2'(s(sK12,X176),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X176)) ) | $spl170),
  inference(cnf_transformation,[],[f13584_D])).
fof(f13584_D,plain,(
  ( ! [X176] : (sP1('T_0','const.bool.IN_2'(s(sK12,X176),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X176)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f13583,plain,(
  ( ! [X177,X176] : (p(s(bool,X177)) | sP1('T_0','const.bool.IN_2'(s(sK12,X176),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X176)) ) | ($spl8 | $spl22)),
  inference(subsumption_resolution,[],[f13450,f259])).
fof(f13450,plain,(
  ( ! [X177,X176] : (p(s(bool,X177)) | sP1('T_0','const.bool.IN_2'(s(sK12,X176),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X176) | s(bool,'T_0') = s(bool,X177)) ) | ($spl8 | $spl22)),
  inference(superposition,[],[f4887,f4966])).
fof(f13576,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK14))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) | s(sK12,sK13) = s(sK12,X1)) ) | $spl168),
  inference(cnf_transformation,[],[f13576_D])).
fof(f13576_D,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK14))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) | s(sK12,sK13) = s(sK12,X1)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f13575,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK14))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) | s(sK12,sK13) = s(sK12,X1)) ) | $spl8),
  inference(condensation,[],[f13574])).
fof(f13574,plain,(
  ( ! [X142,X140,X143,X141] : (p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X140),s(fun(sK12,bool),sK14))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X143)))))) | s(bool,'T_0') = s(bool,X142) | s(sK12,sK13) = s(sK12,X140) | s(bool,'T_0') = s(bool,X141)) ) | $spl8),
  inference(subsumption_resolution,[],[f13441,f635])).
fof(f13441,plain,(
  ( ! [X142,X140,X143,X141] : (s(bool,X141) != s(bool,X142) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X140),s(fun(sK12,bool),sK14))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X143)))))) | s(bool,'T_0') = s(bool,X142) | s(sK12,sK13) = s(sK12,X140) | s(bool,'T_0') = s(bool,X141)) ) | $spl8),
  inference(superposition,[],[f1425,f4966])).
fof(f13571,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK14))),s(fun(bool,bool),X2)))))) | s(sK12,sK13) = s(sK12,X1)) ) | $spl166),
  inference(cnf_transformation,[],[f13571_D])).
fof(f13571_D,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK14))),s(fun(bool,bool),X2)))))) | s(sK12,sK13) = s(sK12,X1)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f13570,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK14))),s(fun(bool,bool),X2)))))) | s(sK12,sK13) = s(sK12,X1)) ) | $spl8),
  inference(condensation,[],[f13569])).
fof(f13569,plain,(
  ( ! [X130,X128,X131,X129] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK12,X128),s(fun(sK12,bool),sK14))),s(fun(bool,bool),X131)))))) | s(bool,'T_0') = s(bool,X130) | s(sK12,sK13) = s(sK12,X128) | s(bool,'T_0') = s(bool,X129)) ) | $spl8),
  inference(subsumption_resolution,[],[f13438,f635])).
fof(f13438,plain,(
  ( ! [X130,X128,X131,X129] : (s(bool,X129) != s(bool,X130) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK12,X128),s(fun(sK12,bool),sK14))),s(fun(bool,bool),X131)))))) | s(bool,'T_0') = s(bool,X130) | s(sK12,sK13) = s(sK12,X128) | s(bool,'T_0') = s(bool,X129)) ) | $spl8),
  inference(superposition,[],[f1420,f4966])).
fof(f13552,plain,(
  ( ! [X152,X151,X153] : (~p(s(bool,'const.bool.IN_2'(s(bool,X152),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK12,X151),s(fun(sK12,bool),sK14))),s(fun(bool,bool),X153)))))) | s(sK12,sK13) = s(sK12,X151) | p(s(bool,'const.bool.IN_2'(s(bool,X152),s(fun(bool,bool),X153)))) | s(bool,'F_0') = s(bool,X152)) ) | $spl164),
  inference(cnf_transformation,[],[f13552_D])).
fof(f13552_D,plain,(
  ( ! [X152,X151,X153] : (~p(s(bool,'const.bool.IN_2'(s(bool,X152),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK12,X151),s(fun(sK12,bool),sK14))),s(fun(bool,bool),X153)))))) | s(sK12,sK13) = s(sK12,X151) | p(s(bool,'const.bool.IN_2'(s(bool,X152),s(fun(bool,bool),X153)))) | s(bool,'F_0') = s(bool,X152)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f13518,plain,(
  ( ! [X152,X151,X153,X150] : (~p(s(bool,'const.bool.IN_2'(s(bool,X152),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK12,X151),s(fun(sK12,bool),sK14))),s(fun(bool,bool),X153)))))) | s(bool,'F_0') = s(bool,X152) | p(s(bool,'const.bool.IN_2'(s(bool,X152),s(fun(bool,bool),X153)))) | s(bool,'T_0') = s(bool,X150) | s(sK12,sK13) = s(sK12,X151)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f13295])).
fof(f13295,plain,(
  ( ! [X152,X151,X153,X150] : (~p(s(bool,'const.bool.IN_2'(s(bool,X152),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK12,X151),s(fun(sK12,bool),sK14))),s(fun(bool,bool),X153)))))) | s(bool,'F_0') = s(bool,X152) | p(s(bool,'const.bool.IN_2'(s(bool,X152),s(fun(bool,bool),X153)))) | s(bool,'T_0') = s(bool,X150) | s(sK12,sK13) = s(sK12,X151) | s(bool,'T_0') = s(bool,X150)) ) | $spl8),
  inference(superposition,[],[f1739,f4966])).
fof(f13551,plain,(
  ( ! [X140,X141,X139] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X139),s(fun(sK12,bool),sK14))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X140),s(fun(bool,bool),X141)))))) | s(sK12,sK13) = s(sK12,X139) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X141)))) | s(bool,'F_0') = s(bool,X140)) ) | $spl162),
  inference(cnf_transformation,[],[f13551_D])).
fof(f13551_D,plain,(
  ( ! [X140,X141,X139] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X139),s(fun(sK12,bool),sK14))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X140),s(fun(bool,bool),X141)))))) | s(sK12,sK13) = s(sK12,X139) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X141)))) | s(bool,'F_0') = s(bool,X140)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f13519,plain,(
  ( ! [X140,X138,X141,X139] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X139),s(fun(sK12,bool),sK14))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X140),s(fun(bool,bool),X141)))))) | s(bool,'F_0') = s(bool,X140) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X141)))) | s(bool,'T_0') = s(bool,X138) | s(sK12,sK13) = s(sK12,X139)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f13292])).
fof(f13292,plain,(
  ( ! [X140,X138,X141,X139] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X139),s(fun(sK12,bool),sK14))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X140),s(fun(bool,bool),X141)))))) | s(bool,'F_0') = s(bool,X140) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X141)))) | s(bool,'T_0') = s(bool,X138) | s(sK12,sK13) = s(sK12,X139) | s(bool,'T_0') = s(bool,X138)) ) | $spl8),
  inference(superposition,[],[f1732,f4966])).
fof(f13550,plain,(
  ( ! [X132,X133,X131] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X131),s(fun(sK12,bool),sK14))),s(fun(bool,bool),X132)))) | s(sK12,sK13) = s(sK12,X131) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X133),s(fun(bool,bool),X132))))))) ) | $spl160),
  inference(cnf_transformation,[],[f13550_D])).
fof(f13550_D,plain,(
  ( ! [X132,X133,X131] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X131),s(fun(sK12,bool),sK14))),s(fun(bool,bool),X132)))) | s(sK12,sK13) = s(sK12,X131) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X133),s(fun(bool,bool),X132))))))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f13520,plain,(
  ( ! [X132,X130,X133,X131] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X131),s(fun(sK12,bool),sK14))),s(fun(bool,bool),X132)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X133),s(fun(bool,bool),X132)))))) | s(bool,'T_0') = s(bool,X130) | s(sK12,sK13) = s(sK12,X131)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f13290])).
fof(f13290,plain,(
  ( ! [X132,X130,X133,X131] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X131),s(fun(sK12,bool),sK14))),s(fun(bool,bool),X132)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X133),s(fun(bool,bool),X132)))))) | s(bool,'T_0') = s(bool,X130) | s(sK12,sK13) = s(sK12,X131) | s(bool,'T_0') = s(bool,X130)) ) | $spl8),
  inference(superposition,[],[f1458,f4966])).
fof(f13549,plain,(
  ( ! [X121,X120,X119] : (s(bool,X120) != s(bool,'const.bool.IN_2'(s(sK12,X119),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X119) | p(s(bool,'const.bool.IN_2'(s(bool,X120),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X121))))))) ) | $spl158),
  inference(cnf_transformation,[],[f13549_D])).
fof(f13549_D,plain,(
  ( ! [X121,X120,X119] : (s(bool,X120) != s(bool,'const.bool.IN_2'(s(sK12,X119),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X119) | p(s(bool,'const.bool.IN_2'(s(bool,X120),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X121))))))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f13521,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,X120) != s(bool,'const.bool.IN_2'(s(sK12,X119),s(fun(sK12,bool),sK14))) | p(s(bool,'const.bool.IN_2'(s(bool,X120),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X121)))))) | s(bool,'T_0') = s(bool,X118) | s(sK12,sK13) = s(sK12,X119)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f13287])).
fof(f13287,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,X120) != s(bool,'const.bool.IN_2'(s(sK12,X119),s(fun(sK12,bool),sK14))) | p(s(bool,'const.bool.IN_2'(s(bool,X120),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X121)))))) | s(bool,'T_0') = s(bool,X118) | s(sK12,sK13) = s(sK12,X119) | s(bool,'T_0') = s(bool,X118)) ) | $spl8),
  inference(superposition,[],[f1425,f4966])).
fof(f13548,plain,(
  ( ! [X109,X107,X108] : (s(bool,X108) != s(bool,'const.bool.IN_2'(s(sK12,X107),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X107) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X108),s(fun(bool,bool),X109))))))) ) | $spl156),
  inference(cnf_transformation,[],[f13548_D])).
fof(f13548_D,plain,(
  ( ! [X109,X107,X108] : (s(bool,X108) != s(bool,'const.bool.IN_2'(s(sK12,X107),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X107) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X108),s(fun(bool,bool),X109))))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f13522,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,X108) != s(bool,'const.bool.IN_2'(s(sK12,X107),s(fun(sK12,bool),sK14))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X108),s(fun(bool,bool),X109)))))) | s(bool,'T_0') = s(bool,X106) | s(sK12,sK13) = s(sK12,X107)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f13284])).
fof(f13284,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,X108) != s(bool,'const.bool.IN_2'(s(sK12,X107),s(fun(sK12,bool),sK14))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X108),s(fun(bool,bool),X109)))))) | s(bool,'T_0') = s(bool,X106) | s(sK12,sK13) = s(sK12,X107) | s(bool,'T_0') = s(bool,X106)) ) | $spl8),
  inference(superposition,[],[f1420,f4966])).
fof(f13544,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X3),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X3)) ) | $spl154),
  inference(cnf_transformation,[],[f13544_D])).
fof(f13544_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X3),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X3)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f13543,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X3),s(fun(sK12,bool),sK14))) | p(s(bool,X2)) | s(sK12,sK13) = s(sK12,X3)) ) | $spl8),
  inference(subsumption_resolution,[],[f13253,f259])).
fof(f13253,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X3),s(fun(sK12,bool),sK14))) | p(s(bool,X2)) | s(sK12,sK13) = s(sK12,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl8),
  inference(superposition,[],[f259,f4966])).
fof(f13541,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X16),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X16)) ) | $spl152),
  inference(cnf_transformation,[],[f13541_D])).
fof(f13541_D,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X16),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X16)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f13534,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X16),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X16) | s(bool,'T_0') = s(bool,X15)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f13230])).
fof(f13230,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X16),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X16) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl8),
  inference(superposition,[],[f4966,f251])).
fof(f13540,plain,(
  ( ! [X2,X1] : (s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK14))) = s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X1) | s(sK12,sK13) = s(sK12,X2)) ) | $spl150),
  inference(cnf_transformation,[],[f13540_D])).
fof(f13540_D,plain,(
  ( ! [X2,X1] : (s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK14))) = s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X1) | s(sK12,sK13) = s(sK12,X2)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f13539,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK14))) = s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X2) | s(bool,'T_0') = s(bool,X0) | s(sK12,sK13) = s(sK12,X1)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f13225])).
fof(f13225,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),sK14))) = s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK14))) | s(sK12,sK13) = s(sK12,X2) | s(bool,'T_0') = s(bool,X0) | s(sK12,sK13) = s(sK12,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f4966,f4966])).
fof(f10469,plain,(
  $false | ($spl5 | $spl14 | $spl20 | $spl65)),
  inference(subsumption_resolution,[],[f4963,f10378])).
fof(f10378,plain,(
  ( ! [X108] : (s(bool,'T_0') = s(bool,X108)) ) | ($spl5 | $spl14 | $spl20)),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f430,f2605,f2603,f3917,f3916,f3915,f3914,f3918,f3972,f3971,f3970,f3969,f3973,f4020,f4075,f4076,f4081,f4082,f4136,f4170,f4171,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4343,f4344,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4515,f4516,f3975,f2604,f2606,f1981,f1977,f1975,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f434,f264,f5928,f5927,f5925,f5924,f5929,f5930,f5931,f5932,f5933,f6011,f6010,f6008,f6007,f6012,f6013,f6014,f6015,f6016,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6466,f6467,f6481,f6496,f6508,f6729,f6737,f6740,f6759,f6767,f7135,f7134,f7132,f7131,f7257,f7258,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7274,f7275,f5095,f7364,f7363,f7362,f7360,f7359,f7374,f7375,f7354,f7373,f7393,f7395,f7401,f7492,f7491,f7487,f7486,f7840,f7845,f7853,f7913,f7930,f7940,f9504,f9509,f4402,f4263,f4356,f4355,f4354,f4261,f4088,f4087,f4264,f6665,f4353,f4352,f4265,f4262,f4135,f4089,f4086,f4085,f304,f303,f307,f306,f305,f4964,f4965,f2026,f4134,f4401,f6757,f6755,f6727,f6725,f6710,f6709,f6705,f6670,f6604,f6603,f6597,f6560,f6559,f6555,f6506,f6504,f6503,f6494,f6492,f6491,f6455,f6424,f1979,f1973,f505,f263,f4874])).
fof(f9509,plain,(
  sP1('F_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl14),
  inference(unit_resulting_resolution,[],[f434,f7260])).
fof(f9504,plain,(
  sP1(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f434,f7259])).
fof(f7940,plain,(
  sP4('T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f7492,f344])).
fof(f7930,plain,(
  ~sP2(f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0','F_0') | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f434,f446,f441,f7491,f347])).
fof(f7913,plain,(
  ~sP2('F_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f434,f444,f441,f7486,f348])).
fof(f7853,plain,(
  ~sP3('T_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f434,f449,f441,f7134,f348])).
fof(f7845,plain,(
  sP4('T_0','T_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl14),
  inference(unit_resulting_resolution,[],[f7132,f344])).
fof(f7840,plain,(
  ~sP3(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0','F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f244,f450,f441,f7131,f346])).
fof(f7486,plain,(
  ~sP5('F_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f434,f7354,f364])).
fof(f7487,plain,(
  sP5('T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f244,f7354,f362])).
fof(f7491,plain,(
  ~sP5('F_0','F_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f434,f7354,f364])).
fof(f7492,plain,(
  sP5('T_0','F_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f244,f7354,f362])).
fof(f7401,plain,(
  sP3('T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'F_0') | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f7363,f343])).
fof(f7395,plain,(
  ~sP4('T_0','T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f244,f450,f443,f7360,f346])).
fof(f7393,plain,(
  sP3(f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0','F_0') | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f7359,f343])).
fof(f7373,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl5 | $spl14)),
  inference(trivial_inequality_removal,[],[f7367])).
fof(f7367,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl5 | $spl14)),
  inference(superposition,[],[f5095,f635])).
fof(f7354,plain,(
  s(bool,'F_0') = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f434,f5095,f704])).
fof(f7375,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2))))))) ) | ($spl5 | $spl14)),
  inference(forward_demodulation,[],[f7356,f505])).
fof(f7356,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))),s(fun(bool,bool),X2))))))) ) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f505,f5095,f1421])).
fof(f7374,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X2))))))) ) | ($spl5 | $spl14)),
  inference(forward_demodulation,[],[f7357,f505])).
fof(f7357,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X2))))))) ) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f505,f5095,f1426])).
fof(f7359,plain,(
  sP5('F_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0') | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f434,f5095,f361])).
fof(f7360,plain,(
  ~sP5('T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0') | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f244,f5095,f363])).
fof(f7362,plain,(
  ~p(s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f244,f5095,f311])).
fof(f7363,plain,(
  sP5('F_0','T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f434,f5095,f361])).
fof(f7364,plain,(
  ~sP5('T_0','T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl14)),
  inference(unit_resulting_resolution,[],[f244,f5095,f363])).
fof(f5095,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl14)),
  inference(forward_demodulation,[],[f2605,f3975])).
fof(f7275,plain,(
  ( ! [X35,X36] : (~sP8(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X35,X36) | sP9(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X35,X36)) ) | $spl14),
  inference(subsumption_resolution,[],[f7166,f244])).
fof(f7166,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | ~sP8(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X35,X36) | sP9(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X35,X36)) ) | $spl14),
  inference(superposition,[],[f417,f3975])).
fof(f7274,plain,(
  ( ! [X33,X34] : (~sP10(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X33,X34) | sP11(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X33,X34)) ) | $spl14),
  inference(subsumption_resolution,[],[f7165,f244])).
fof(f7165,plain,(
  ( ! [X33,X34] : (~p(s(bool,'T_0')) | ~sP10(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X33,X34) | sP11(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X33,X34)) ) | $spl14),
  inference(superposition,[],[f416,f3975])).
fof(f7273,plain,(
  ( ! [X31,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | ~sP11(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X32,X31)) ) | $spl14),
  inference(subsumption_resolution,[],[f7164,f244])).
fof(f7164,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | p(s(bool,X32)) | ~sP11(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X32,X31)) ) | $spl14),
  inference(superposition,[],[f413,f3975])).
fof(f7272,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | sP11(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X29,X30)) ) | $spl14),
  inference(subsumption_resolution,[],[f7163,f244])).
fof(f7163,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | sP11(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X29,X30)) ) | $spl14),
  inference(superposition,[],[f412,f3975])).
fof(f7271,plain,(
  ( ! [X28,X27] : (p(s(bool,X27)) | sP11(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X28,X27)) ) | $spl14),
  inference(subsumption_resolution,[],[f7162,f244])).
fof(f7162,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP11(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X28,X27)) ) | $spl14),
  inference(superposition,[],[f411,f3975])).
fof(f7270,plain,(
  ( ! [X26,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | ~sP10(X26,X25,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl14),
  inference(subsumption_resolution,[],[f7161,f244])).
fof(f7161,plain,(
  ( ! [X26,X25] : (~p(s(bool,'T_0')) | p(s(bool,X25)) | ~p(s(bool,X26)) | ~sP10(X26,X25,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl14),
  inference(superposition,[],[f406,f3975])).
fof(f7269,plain,(
  ( ! [X24,X23] : (p(s(bool,X23)) | p(s(bool,X24)) | ~sP9(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X24,X23)) ) | $spl14),
  inference(subsumption_resolution,[],[f7160,f244])).
fof(f7160,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | p(s(bool,X23)) | p(s(bool,X24)) | ~sP9(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X24,X23)) ) | $spl14),
  inference(superposition,[],[f396,f3975])).
fof(f7268,plain,(
  ( ! [X21,X22] : (~p(s(bool,X21)) | sP9(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X21,X22)) ) | $spl14),
  inference(subsumption_resolution,[],[f7159,f244])).
fof(f7159,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | sP9(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X21,X22)) ) | $spl14),
  inference(superposition,[],[f395,f3975])).
fof(f7267,plain,(
  ( ! [X19,X20] : (~p(s(bool,X19)) | sP9(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X20,X19)) ) | $spl14),
  inference(subsumption_resolution,[],[f7158,f244])).
fof(f7158,plain,(
  ( ! [X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X19)) | sP9(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X20,X19)) ) | $spl14),
  inference(superposition,[],[f394,f3975])).
fof(f7266,plain,(
  ( ! [X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | ~sP8(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X17,X18)) ) | $spl14),
  inference(subsumption_resolution,[],[f7157,f244])).
fof(f7157,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,X17)) | p(s(bool,X18)) | ~sP8(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X17,X18)) ) | $spl14),
  inference(superposition,[],[f389,f3975])).
fof(f7265,plain,(
  ( ! [X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | ~sP7(X16,X15,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl14),
  inference(subsumption_resolution,[],[f7156,f244])).
fof(f7156,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X15)) | p(s(bool,X16)) | ~sP7(X16,X15,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl14),
  inference(superposition,[],[f381,f3975])).
fof(f7264,plain,(
  ( ! [X14,X13] : (~p(s(bool,X13)) | p(s(bool,X14)) | ~sP6(X13,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X14)) ) | $spl14),
  inference(subsumption_resolution,[],[f7155,f244])).
fof(f7155,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | p(s(bool,X14)) | ~sP6(X13,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X14)) ) | $spl14),
  inference(superposition,[],[f372,f3975])).
fof(f7263,plain,(
  ( ! [X12,X11] : (~p(s(bool,X11)) | p(s(bool,X12)) | ~sP3(X11,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X12)) ) | $spl14),
  inference(subsumption_resolution,[],[f7154,f244])).
fof(f7154,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | p(s(bool,X12)) | ~sP3(X11,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X12)) ) | $spl14),
  inference(superposition,[],[f353,f3975])).
fof(f7262,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | p(s(bool,X10)) | ~sP4(X9,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X10)) ) | $spl14),
  inference(subsumption_resolution,[],[f7153,f244])).
fof(f7153,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | p(s(bool,X10)) | ~sP4(X9,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X10)) ) | $spl14),
  inference(superposition,[],[f349,f3975])).
fof(f7261,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | ~sP5(X7,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X8)) ) | $spl14),
  inference(subsumption_resolution,[],[f7152,f244])).
fof(f7152,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP5(X7,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X8)) ) | $spl14),
  inference(superposition,[],[f345,f3975])).
fof(f7260,plain,(
  ( ! [X6] : (sP1(X6,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | p(s(bool,X6))) ) | $spl14),
  inference(subsumption_resolution,[],[f7151,f244])).
fof(f7151,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1(X6,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl14),
  inference(superposition,[],[f336,f3975])).
fof(f7259,plain,(
  ( ! [X5] : (sP1(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X5) | p(s(bool,X5))) ) | $spl14),
  inference(subsumption_resolution,[],[f7150,f244])).
fof(f7150,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,X5)) | sP1(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X5)) ) | $spl14),
  inference(superposition,[],[f333,f3975])).
fof(f7258,plain,(
  ( ! [X4] : (sP0(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X4)) ) | $spl14),
  inference(subsumption_resolution,[],[f7149,f244])).
fof(f7149,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | sP0(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),X4)) ) | $spl14),
  inference(superposition,[],[f318,f3975])).
fof(f7257,plain,(
  ( ! [X3] : (sP0(X3,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl14),
  inference(subsumption_resolution,[],[f7148,f244])).
fof(f7148,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | sP0(X3,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl14),
  inference(superposition,[],[f317,f3975])).
fof(f7131,plain,(
  ~sP5('F_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f434,f3975,f364])).
fof(f7132,plain,(
  sP5('T_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f244,f3975,f362])).
fof(f7134,plain,(
  ~sP5('F_0','T_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl14),
  inference(unit_resulting_resolution,[],[f434,f3975,f364])).
fof(f7135,plain,(
  sP5('T_0','T_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl14),
  inference(unit_resulting_resolution,[],[f244,f3975,f362])).
fof(f6007,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | $spl5),
  inference(unit_resulting_resolution,[],[f244,f2606,f363])).
fof(f6008,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | $spl5),
  inference(unit_resulting_resolution,[],[f434,f2606,f361])).
fof(f6010,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | $spl5),
  inference(unit_resulting_resolution,[],[f244,f2606,f363])).
fof(f6011,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | $spl5),
  inference(unit_resulting_resolution,[],[f434,f2606,f361])).
fof(f5925,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))),'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | $spl5),
  inference(unit_resulting_resolution,[],[f434,f2604,f361])).
fof(f5927,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | $spl5),
  inference(unit_resulting_resolution,[],[f244,f2604,f363])).
fof(f5928,plain,(
  sP5('F_0','const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14)),'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | $spl5),
  inference(unit_resulting_resolution,[],[f434,f2604,f361])).
fof(f3969,plain,(
  ~sP5('T_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl5),
  inference(unit_resulting_resolution,[],[f244,f2605,f363])).
fof(f3970,plain,(
  sP5('F_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl5),
  inference(unit_resulting_resolution,[],[f434,f2605,f361])).
fof(f3971,plain,(
  ~sP5('T_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl5),
  inference(unit_resulting_resolution,[],[f244,f2605,f363])).
fof(f3972,plain,(
  sP5('F_0',f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl5),
  inference(unit_resulting_resolution,[],[f434,f2605,f361])).
fof(f3914,plain,(
  ~sP5('T_0',g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl5),
  inference(unit_resulting_resolution,[],[f244,f2603,f363])).
fof(f3915,plain,(
  sP5('F_0',g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl5),
  inference(unit_resulting_resolution,[],[f434,f2603,f361])).
fof(f3916,plain,(
  ~sP5('T_0',f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))),g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl5),
  inference(unit_resulting_resolution,[],[f244,f2603,f363])).
fof(f2605,plain,(
  s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) != s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl5),
  inference(unit_resulting_resolution,[],[f430,f341])).
fof(f10468,plain,(
  $false | ($spl5 | $spl14 | $spl20)),
  inference(subsumption_resolution,[],[f10467,f244])).
fof(f10467,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl14 | $spl20)),
  inference(forward_demodulation,[],[f434,f10378])).
fof(f10465,plain,(
  $false | ($spl5 | $spl14 | $spl20)),
  inference(subsumption_resolution,[],[f10464,f10378])).
fof(f10464,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl14 | $spl20)),
  inference(forward_demodulation,[],[f2605,f10378])).
fof(f10463,plain,(
  $false | ($spl5 | $spl8 | $spl14 | $spl20)),
  inference(global_subsumption,[],[f10462,f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f430,f2605,f2603,f3917,f3916,f3915,f3914,f3918,f3972,f3971,f3970,f3969,f3973,f4020,f4075,f4076,f4081,f4082,f4136,f4170,f4171,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4343,f4344,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4515,f4516,f3975,f2604,f2606,f1981,f1977,f1975,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f434,f264,f5928,f5927,f5925,f5924,f5929,f5930,f5931,f5932,f5933,f6011,f6010,f6008,f6007,f6012,f6013,f6014,f6015,f6016,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6466,f6467,f6481,f6496,f6508,f6729,f6737,f6740,f6759,f6767,f7135,f7134,f7132,f7131,f7257,f7258,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7274,f7275,f5095,f7364,f7363,f7362,f7360,f7359,f7374,f7375,f7354,f7373,f7393,f7395,f7401,f7492,f7491,f7487,f7486,f7840,f7845,f7853,f7913,f7930,f7940,f9504,f9509,f4402,f4263,f4356,f4355,f4354,f4261,f4088,f4087,f4264,f6665,f4353,f4352,f4265,f4262,f4135,f4089,f4086,f4085,f304,f303,f307,f306,f305,f4964,f4965,f2026,f4134,f4401,f6757,f6755,f6727,f6725,f6710,f6709,f6705,f6670,f6604,f6603,f6597,f6560,f6559,f6555,f6506,f6504,f6503,f6494,f6492,f6491,f6455,f6424,f1979,f1973,f505,f263,f4874,f10378,f10452,f10454,f10456,f10460])).
fof(f10460,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl5 | $spl14 | $spl20)),
  inference(subsumption_resolution,[],[f10418,f10378])).
fof(f10418,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK16(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl5 | $spl14 | $spl20)),
  inference(backward_demodulation,[],[f10378,f304])).
fof(f10456,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl5 | $spl14 | $spl20)),
  inference(subsumption_resolution,[],[f10382,f10378])).
fof(f10382,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK16('const.pred_set.EMPTY_0',X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl5 | $spl14 | $spl20)),
  inference(backward_demodulation,[],[f10378,f1979])).
fof(f10454,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl5 | $spl14 | $spl20)),
  inference(subsumption_resolution,[],[f10381,f10378])).
fof(f10381,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK16(X1,'const.pred_set.EMPTY_0',X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl5 | $spl14 | $spl20)),
  inference(backward_demodulation,[],[f10378,f1973])).
fof(f10452,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl5 | $spl14 | $spl20)),
  inference(subsumption_resolution,[],[f10379,f10378])).
fof(f10379,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl5 | $spl14 | $spl20)),
  inference(backward_demodulation,[],[f10378,f263])).
fof(f10462,plain,(
  ( ! [X4] : (s(sK12,sK13) = s(sK12,X4)) ) | ($spl5 | $spl8 | $spl14 | $spl20)),
  inference(subsumption_resolution,[],[f10437,f244])).
fof(f10437,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(sK12,sK13) = s(sK12,X4)) ) | ($spl5 | $spl8 | $spl14 | $spl20)),
  inference(backward_demodulation,[],[f10378,f433])).
fof(f10461,plain,(
  $false | ($spl5 | $spl14 | $spl20)),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f430,f2605,f2603,f3917,f3916,f3915,f3914,f3918,f3972,f3971,f3970,f3969,f3973,f4020,f4075,f4076,f4081,f4082,f4136,f4170,f4171,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4343,f4344,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4515,f4516,f3975,f2604,f2606,f1981,f1977,f1975,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f434,f264,f5928,f5927,f5925,f5924,f5929,f5930,f5931,f5932,f5933,f6011,f6010,f6008,f6007,f6012,f6013,f6014,f6015,f6016,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6466,f6467,f6481,f6496,f6508,f6729,f6737,f6740,f6759,f6767,f7135,f7134,f7132,f7131,f7257,f7258,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7274,f7275,f5095,f7364,f7363,f7362,f7360,f7359,f7374,f7375,f7354,f7373,f7393,f7395,f7401,f7492,f7491,f7487,f7486,f7840,f7845,f7853,f7913,f7930,f7940,f9504,f9509,f4402,f4263,f4356,f4355,f4354,f4261,f4088,f4087,f4264,f6665,f4353,f4352,f4265,f4262,f4135,f4089,f4086,f4085,f304,f303,f307,f306,f305,f4964,f4965,f2026,f4134,f4401,f6757,f6755,f6727,f6725,f6710,f6709,f6705,f6670,f6604,f6603,f6597,f6560,f6559,f6555,f6506,f6504,f6503,f6494,f6492,f6491,f6455,f6424,f1979,f1973,f505,f263,f4874,f10378,f10452,f10454,f10456,f10460])).
fof(f10457,plain,(
  $false | ($spl5 | $spl14 | $spl20)),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f430,f2605,f2603,f3917,f3916,f3915,f3914,f3918,f3972,f3971,f3970,f3969,f3973,f4020,f4075,f4076,f4081,f4082,f4136,f4170,f4171,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4343,f4344,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4515,f4516,f3975,f2604,f2606,f1981,f1977,f1975,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f434,f264,f5928,f5927,f5925,f5924,f5929,f5930,f5931,f5932,f5933,f6011,f6010,f6008,f6007,f6012,f6013,f6014,f6015,f6016,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6466,f6467,f6481,f6496,f6508,f6729,f6737,f6740,f6759,f6767,f7135,f7134,f7132,f7131,f7257,f7258,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7274,f7275,f5095,f7364,f7363,f7362,f7360,f7359,f7374,f7375,f7354,f7373,f7393,f7395,f7401,f7492,f7491,f7487,f7486,f7840,f7845,f7853,f7913,f7930,f7940,f9504,f9509,f4402,f4263,f4356,f4355,f4354,f4261,f4088,f4087,f4264,f6665,f4353,f4352,f4265,f4262,f4135,f4089,f4086,f4085,f304,f303,f307,f306,f305,f4964,f4965,f2026,f4134,f4401,f6757,f6755,f6727,f6725,f6710,f6709,f6705,f6670,f6604,f6603,f6597,f6560,f6559,f6555,f6506,f6504,f6503,f6494,f6492,f6491,f6455,f6424,f1979,f1973,f505,f263,f4874,f10378,f10452,f10454,f10456])).
fof(f10455,plain,(
  $false | ($spl5 | $spl14 | $spl20)),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f430,f2605,f2603,f3917,f3916,f3915,f3914,f3918,f3972,f3971,f3970,f3969,f3973,f4020,f4075,f4076,f4081,f4082,f4136,f4170,f4171,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4343,f4344,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4515,f4516,f3975,f2604,f2606,f1981,f1977,f1975,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f434,f264,f5928,f5927,f5925,f5924,f5929,f5930,f5931,f5932,f5933,f6011,f6010,f6008,f6007,f6012,f6013,f6014,f6015,f6016,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6466,f6467,f6481,f6496,f6508,f6729,f6737,f6740,f6759,f6767,f7135,f7134,f7132,f7131,f7257,f7258,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7274,f7275,f5095,f7364,f7363,f7362,f7360,f7359,f7374,f7375,f7354,f7373,f7393,f7395,f7401,f7492,f7491,f7487,f7486,f7840,f7845,f7853,f7913,f7930,f7940,f9504,f9509,f4402,f4263,f4356,f4355,f4354,f4261,f4088,f4087,f4264,f6665,f4353,f4352,f4265,f4262,f4135,f4089,f4086,f4085,f304,f303,f307,f306,f305,f4964,f4965,f2026,f4134,f4401,f6757,f6755,f6727,f6725,f6710,f6709,f6705,f6670,f6604,f6603,f6597,f6560,f6559,f6555,f6506,f6504,f6503,f6494,f6492,f6491,f6455,f6424,f1979,f1973,f505,f263,f4874,f10378,f10452,f10454])).
fof(f10453,plain,(
  $false | ($spl5 | $spl14 | $spl20)),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f430,f2605,f2603,f3917,f3916,f3915,f3914,f3918,f3972,f3971,f3970,f3969,f3973,f4020,f4075,f4076,f4081,f4082,f4136,f4170,f4171,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4343,f4344,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4515,f4516,f3975,f2604,f2606,f1981,f1977,f1975,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f434,f264,f5928,f5927,f5925,f5924,f5929,f5930,f5931,f5932,f5933,f6011,f6010,f6008,f6007,f6012,f6013,f6014,f6015,f6016,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6466,f6467,f6481,f6496,f6508,f6729,f6737,f6740,f6759,f6767,f7135,f7134,f7132,f7131,f7257,f7258,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7274,f7275,f5095,f7364,f7363,f7362,f7360,f7359,f7374,f7375,f7354,f7373,f7393,f7395,f7401,f7492,f7491,f7487,f7486,f7840,f7845,f7853,f7913,f7930,f7940,f9504,f9509,f4402,f4263,f4356,f4355,f4354,f4261,f4088,f4087,f4264,f6665,f4353,f4352,f4265,f4262,f4135,f4089,f4086,f4085,f304,f303,f307,f306,f305,f4964,f4965,f2026,f4134,f4401,f6757,f6755,f6727,f6725,f6710,f6709,f6705,f6670,f6604,f6603,f6597,f6560,f6559,f6555,f6506,f6504,f6503,f6494,f6492,f6491,f6455,f6424,f1979,f1973,f505,f263,f4874,f10378,f10452])).
fof(f10349,plain,(
  ( ! [X76,X75] : (s(bool,'F_0') != s(bool,X75) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X75),s(fun(bool,bool),X76))))))) ) | $spl148),
  inference(cnf_transformation,[],[f10349_D])).
fof(f10349_D,plain,(
  ( ! [X76,X75] : (s(bool,'F_0') != s(bool,X75) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X75),s(fun(bool,bool),X76))))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f10348,plain,(
  ( ! [X76,X74,X75] : (s(bool,'F_0') != s(bool,X75) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X75),s(fun(bool,bool),X76)))))) | s(bool,'T_0') = s(bool,X74)) ) | ($spl5 | $spl62 | $spl80)),
  inference(forward_demodulation,[],[f10094,f10310])).
fof(f10310,plain,(
  s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl62 | $spl80)),
  inference(subsumption_resolution,[],[f9898,f6025])).
fof(f9898,plain,(
  s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl62 | $spl80)),
  inference(superposition,[],[f6050,f4961])).
fof(f10094,plain,(
  ( ! [X76,X74,X75] : (s(bool,X75) != s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X75),s(fun(bool,bool),X76)))))) | s(bool,'T_0') = s(bool,X74)) ) | $spl62),
  inference(duplicate_literal_removal,[],[f9869])).
fof(f9869,plain,(
  ( ! [X76,X74,X75] : (s(bool,X75) != s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X75),s(fun(bool,bool),X76)))))) | s(bool,'T_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) ) | $spl62),
  inference(superposition,[],[f1420,f4961])).
fof(f10347,plain,(
  ( ! [X85,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,'const.bool.IN_2'(s(bool,X84),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X85))))))) ) | $spl146),
  inference(cnf_transformation,[],[f10347_D])).
fof(f10347_D,plain,(
  ( ! [X85,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,'const.bool.IN_2'(s(bool,X84),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X85))))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f10346,plain,(
  ( ! [X85,X83,X84] : (s(bool,'F_0') != s(bool,X84) | p(s(bool,'const.bool.IN_2'(s(bool,X84),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X85)))))) | s(bool,'T_0') = s(bool,X83)) ) | ($spl5 | $spl62 | $spl80)),
  inference(forward_demodulation,[],[f10093,f10310])).
fof(f10093,plain,(
  ( ! [X85,X83,X84] : (s(bool,X84) != s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | p(s(bool,'const.bool.IN_2'(s(bool,X84),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X85)))))) | s(bool,'T_0') = s(bool,X83)) ) | $spl62),
  inference(duplicate_literal_removal,[],[f9872])).
fof(f9872,plain,(
  ( ! [X85,X83,X84] : (s(bool,X84) != s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | p(s(bool,'const.bool.IN_2'(s(bool,X84),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X85)))))) | s(bool,'T_0') = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) ) | $spl62),
  inference(superposition,[],[f1425,f4961])).
fof(f10344,plain,(
  ( ! [X94,X93] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X93)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X94),s(fun(bool,bool),X93))))))) ) | $spl144),
  inference(cnf_transformation,[],[f10344_D])).
fof(f10344_D,plain,(
  ( ! [X94,X93] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X93)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X94),s(fun(bool,bool),X93))))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f10343,plain,(
  ( ! [X94,X92,X93] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X93)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X94),s(fun(bool,bool),X93)))))) | s(bool,'T_0') = s(bool,X92)) ) | ($spl5 | $spl62 | $spl80)),
  inference(forward_demodulation,[],[f10092,f10310])).
fof(f10092,plain,(
  ( ! [X94,X92,X93] : (~p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X93)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X94),s(fun(bool,bool),X93)))))) | s(bool,'T_0') = s(bool,X92)) ) | $spl62),
  inference(duplicate_literal_removal,[],[f9875])).
fof(f9875,plain,(
  ( ! [X94,X92,X93] : (~p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X93)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X94),s(fun(bool,bool),X93)))))) | s(bool,'T_0') = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) ) | $spl62),
  inference(superposition,[],[f1458,f4961])).
fof(f10342,plain,(
  ( ! [X99,X100] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X99),s(fun(bool,bool),X100)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X100)))) | s(bool,'F_0') = s(bool,X99)) ) | $spl142),
  inference(cnf_transformation,[],[f10342_D])).
fof(f10342_D,plain,(
  ( ! [X99,X100] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X99),s(fun(bool,bool),X100)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X100)))) | s(bool,'F_0') = s(bool,X99)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f10341,plain,(
  ( ! [X99,X100,X98] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X99),s(fun(bool,bool),X100)))))) | s(bool,'F_0') = s(bool,X99) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X100)))) | s(bool,'T_0') = s(bool,X98)) ) | ($spl5 | $spl62 | $spl80)),
  inference(forward_demodulation,[],[f10091,f10310])).
fof(f10091,plain,(
  ( ! [X99,X100,X98] : (~p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X99),s(fun(bool,bool),X100)))))) | s(bool,'F_0') = s(bool,X99) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X100)))) | s(bool,'T_0') = s(bool,X98)) ) | $spl62),
  inference(duplicate_literal_removal,[],[f9877])).
fof(f9877,plain,(
  ( ! [X99,X100,X98] : (~p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X99),s(fun(bool,bool),X100)))))) | s(bool,'F_0') = s(bool,X99) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X100)))) | s(bool,'T_0') = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) ) | $spl62),
  inference(superposition,[],[f1732,f4961])).
fof(f10340,plain,(
  ( ! [X109,X108] : (~p(s(bool,'const.bool.IN_2'(s(bool,X108),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X109)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X108),s(fun(bool,bool),X109)))) | s(bool,'F_0') = s(bool,X108)) ) | $spl140),
  inference(cnf_transformation,[],[f10340_D])).
fof(f10340_D,plain,(
  ( ! [X109,X108] : (~p(s(bool,'const.bool.IN_2'(s(bool,X108),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X109)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X108),s(fun(bool,bool),X109)))) | s(bool,'F_0') = s(bool,X108)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f10339,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,'const.bool.IN_2'(s(bool,X108),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X109)))))) | s(bool,'F_0') = s(bool,X108) | p(s(bool,'const.bool.IN_2'(s(bool,X108),s(fun(bool,bool),X109)))) | s(bool,'T_0') = s(bool,X107)) ) | ($spl5 | $spl62 | $spl80)),
  inference(forward_demodulation,[],[f10090,f10310])).
fof(f10090,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,'const.bool.IN_2'(s(bool,X108),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X109)))))) | s(bool,'F_0') = s(bool,X108) | p(s(bool,'const.bool.IN_2'(s(bool,X108),s(fun(bool,bool),X109)))) | s(bool,'T_0') = s(bool,X107)) ) | $spl62),
  inference(duplicate_literal_removal,[],[f9880])).
fof(f9880,plain,(
  ( ! [X109,X107,X108] : (~p(s(bool,'const.bool.IN_2'(s(bool,X108),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X109)))))) | s(bool,'F_0') = s(bool,X108) | p(s(bool,'const.bool.IN_2'(s(bool,X108),s(fun(bool,bool),X109)))) | s(bool,'T_0') = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) ) | $spl62),
  inference(superposition,[],[f1739,f4961])).
fof(f10308,plain,(
  ( ! [X146,X144,X145] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X146),s(fun(bool,bool),X145)))))) | s(bool,'T_0') = s(bool,X144) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X144),s(fun(bool,bool),X145)))) ) | $spl138),
  inference(cnf_transformation,[],[f10308_D])).
fof(f10308_D,plain,(
  ( ! [X146,X144,X145] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X146),s(fun(bool,bool),X145)))))) | s(bool,'T_0') = s(bool,X144) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X144),s(fun(bool,bool),X145)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f10303,plain,(
  ~p(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | $spl129),
  inference(cnf_transformation,[],[f10303_D])).
fof(f10303_D,plain,(
  ~p(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f9901,plain,(
  ( ! [X146,X144,X145] : (~p(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X146),s(fun(bool,bool),X145)))))) | s(bool,'T_0') = s(bool,X144) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X144),s(fun(bool,bool),X145)))) ) | $spl62),
  inference(superposition,[],[f1458,f4961])).
fof(f10307,plain,(
  ( ! [X149,X147,X148] : (s(bool,'F_0') = s(bool,X147) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X147),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X148),s(fun(bool,bool),X149))))) | s(bool,'T_0') = s(bool,X148) | p(s(bool,'const.bool.IN_2'(s(bool,X147),s(fun(bool,bool),X149))))) ) | $spl136),
  inference(cnf_transformation,[],[f10307_D])).
fof(f10307_D,plain,(
  ( ! [X149,X147,X148] : (s(bool,'F_0') = s(bool,X147) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X147),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X148),s(fun(bool,bool),X149))))) | s(bool,'T_0') = s(bool,X148) | p(s(bool,'const.bool.IN_2'(s(bool,X147),s(fun(bool,bool),X149))))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f9902,plain,(
  ( ! [X149,X147,X148] : (~p(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | s(bool,'F_0') = s(bool,X147) | p(s(bool,'const.bool.IN_2'(s(bool,X147),s(fun(bool,bool),X149)))) | s(bool,'T_0') = s(bool,X148) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X147),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X148),s(fun(bool,bool),X149)))))) ) | $spl62),
  inference(superposition,[],[f1739,f4961])).
fof(f10306,plain,(
  ( ! [X152,X151,X150] : (s(bool,'F_0') = s(bool,X151) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X150),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X151),s(fun(bool,bool),X152))))) | s(bool,'T_0') = s(bool,X150) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X152))))) ) | $spl134),
  inference(cnf_transformation,[],[f10306_D])).
fof(f10306_D,plain,(
  ( ! [X152,X151,X150] : (s(bool,'F_0') = s(bool,X151) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X150),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X151),s(fun(bool,bool),X152))))) | s(bool,'T_0') = s(bool,X150) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X152))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f9903,plain,(
  ( ! [X152,X151,X150] : (~p(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | s(bool,'F_0') = s(bool,X151) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X152)))) | s(bool,'T_0') = s(bool,X150) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X150),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X151),s(fun(bool,bool),X152)))))) ) | $spl62),
  inference(superposition,[],[f1732,f4961])).
fof(f10305,plain,(
  ( ! [X154,X155,X153] : (s(bool,X153) = s(bool,X155) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X153),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X154))))) | s(bool,'T_0') = s(bool,X155) | p(s(bool,'const.bool.IN_2'(s(bool,X153),s(fun(bool,bool),X154))))) ) | $spl132),
  inference(cnf_transformation,[],[f10305_D])).
fof(f10305_D,plain,(
  ( ! [X154,X155,X153] : (s(bool,X153) = s(bool,X155) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X153),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X154))))) | s(bool,'T_0') = s(bool,X155) | p(s(bool,'const.bool.IN_2'(s(bool,X153),s(fun(bool,bool),X154))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f9904,plain,(
  ( ! [X154,X155,X153] : (~p(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | s(bool,X153) = s(bool,X155) | p(s(bool,'const.bool.IN_2'(s(bool,X153),s(fun(bool,bool),X154)))) | s(bool,'T_0') = s(bool,X155) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X153),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X154)))))) ) | $spl62),
  inference(superposition,[],[f1741,f4961])).
fof(f10304,plain,(
  ( ! [X161,X162,X160] : (p(s(bool,'const.bool.IN_2'(s(bool,X161),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X162),s(fun(bool,bool),X160)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X160))) | s(bool,'T_0') = s(bool,X161)) ) | $spl130),
  inference(cnf_transformation,[],[f10304_D])).
fof(f10304_D,plain,(
  ( ! [X161,X162,X160] : (p(s(bool,'const.bool.IN_2'(s(bool,X161),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X162),s(fun(bool,bool),X160)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X160))) | s(bool,'T_0') = s(bool,X161)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f9907,plain,(
  ( ! [X161,X162,X160] : (~p(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | p(s(bool,'const.bool.IN_2'(s(bool,X161),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X162),s(fun(bool,bool),X160)))))) | s(bool,'T_0') = s(bool,X161) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X160)))) ) | $spl62),
  inference(superposition,[],[f1459,f4961])).
fof(f10301,plain,(
  ( ! [X165,X163,X164] : (s(bool,X163) = s(bool,X165) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X163),s(fun(bool,bool),X164))))) | s(bool,'T_0') = s(bool,X165) | p(s(bool,'const.bool.IN_2'(s(bool,X165),s(fun(bool,bool),X164))))) ) | $spl126),
  inference(cnf_transformation,[],[f10301_D])).
fof(f10301_D,plain,(
  ( ! [X165,X163,X164] : (s(bool,X163) = s(bool,X165) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X163),s(fun(bool,bool),X164))))) | s(bool,'T_0') = s(bool,X165) | p(s(bool,'const.bool.IN_2'(s(bool,X165),s(fun(bool,bool),X164))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f9908,plain,(
  ( ! [X165,X163,X164] : (~p(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) | s(bool,X163) = s(bool,X165) | p(s(bool,'const.bool.IN_2'(s(bool,X165),s(fun(bool,bool),X164)))) | s(bool,'T_0') = s(bool,X165) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X163),s(fun(bool,bool),X164)))))) ) | $spl62),
  inference(superposition,[],[f1734,f4961])).
fof(f10300,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | $spl124),
  inference(cnf_transformation,[],[f10300_D])).
fof(f10300_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f9915,plain,(
  ( ! [X182] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X182)) ) | $spl62),
  inference(superposition,[],[f505,f4961])).
fof(f10280,plain,(
  s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl114),
  inference(cnf_transformation,[],[f10280_D])).
fof(f10280_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f10254,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl62),
  inference(subsumption_resolution,[],[f9979,f259])).
fof(f9979,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X5)) ) | $spl62),
  inference(superposition,[],[f264,f4961])).
fof(f10297,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X1))))))) ) | $spl122),
  inference(cnf_transformation,[],[f10297_D])).
fof(f10297_D,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X1))))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f10296,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X1))))))) ) | $spl62),
  inference(condensation,[],[f10295])).
fof(f10295,plain,(
  ( ! [X80,X78,X79] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X80)))))) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X78)) ) | $spl62),
  inference(subsumption_resolution,[],[f10008,f635])).
fof(f10008,plain,(
  ( ! [X80,X78,X79] : (s(bool,X78) != s(bool,X79) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X80)))))) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X78)) ) | $spl62),
  inference(superposition,[],[f1420,f4961])).
fof(f10291,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1))))))) ) | $spl120),
  inference(cnf_transformation,[],[f10291_D])).
fof(f10291_D,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1))))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f10290,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1))))))) ) | $spl62),
  inference(condensation,[],[f10289])).
fof(f10289,plain,(
  ( ! [X88,X87,X89] : (p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X89)))))) | s(bool,'T_0') = s(bool,X88) | s(bool,'T_0') = s(bool,X87)) ) | $spl62),
  inference(subsumption_resolution,[],[f10011,f635])).
fof(f10011,plain,(
  ( ! [X88,X87,X89] : (s(bool,X87) != s(bool,X88) | p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X89)))))) | s(bool,'T_0') = s(bool,X88) | s(bool,'T_0') = s(bool,X87)) ) | $spl62),
  inference(superposition,[],[f1425,f4961])).
fof(f10285,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1))))) ) | $spl118),
  inference(cnf_transformation,[],[f10285_D])).
fof(f10285_D,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f10284,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1))))) ) | $spl62),
  inference(condensation,[],[f10283])).
fof(f10283,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') = s(bool,X96) | s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X98)))) | s(bool,'T_0') = s(bool,X97)) ) | $spl62),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4263,f4402,f4352,f4265,f4135,f4264,f4353,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6424,f6455,f6466,f6467,f6481,f6491,f6492,f6494,f6496,f6503,f6504,f6506,f6508,f6555,f6559,f6560,f6603,f6597,f6604,f6665,f6670,f6709,f6705,f6710,f6725,f6727,f6729,f6737,f6740,f6755,f6757,f6759,f6767,f4961,f10082,f10029,f10027,f10026,f10025,f10024,f10023,f10018,f10276,f10017,f10278,f10015,f10282,f10014])).
fof(f10014,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'const.bool.IN_2'(s(bool,X97),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X96),s(fun(bool,bool),X98)))))) | s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X98)))) | s(bool,'T_0') = s(bool,X97) | s(bool,'T_0') = s(bool,X96)) ) | $spl62),
  inference(superposition,[],[f1732,f4961])).
fof(f10282,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') = s(bool,X99) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X101)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X99),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X100),s(fun(bool,bool),X101)))))) | s(bool,'F_0') = s(bool,X100)) ) | $spl62),
  inference(global_subsumption,[],[f10015,f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4263,f4402,f4352,f4265,f4135,f4264,f4353,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6424,f6455,f6466,f6467,f6481,f6491,f6492,f6494,f6496,f6503,f6504,f6506,f6508,f6555,f6559,f6560,f6603,f6597,f6604,f6665,f6670,f6709,f6705,f6710,f6725,f6727,f6729,f6737,f6740,f6755,f6757,f6759,f6767])).
fof(f10015,plain,(
  ( ! [X101,X99,X100] : (~p(s(bool,'const.bool.IN_2'(s(bool,X99),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X100),s(fun(bool,bool),X101)))))) | s(bool,'F_0') = s(bool,X100) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X101)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X99)) ) | $spl62),
  inference(superposition,[],[f1732,f4961])).
fof(f10017,plain,(
  ( ! [X107,X105,X106] : (~p(s(bool,'const.bool.IN_2'(s(bool,X105),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X106),s(fun(bool,bool),X107)))))) | s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X107)))) | s(bool,'T_0') = s(bool,X106) | s(bool,'T_0') = s(bool,X105)) ) | $spl62),
  inference(superposition,[],[f1739,f4961])).
fof(f10276,plain,(
  ( ! [X109,X110,X108] : (p(s(bool,'const.bool.IN_2'(s(bool,X109),s(fun(bool,bool),X110)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X109),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X108),s(fun(bool,bool),X110)))))) | s(bool,'F_0') = s(bool,X109) | s(bool,'T_0') = s(bool,X108)) ) | $spl62),
  inference(global_subsumption,[],[f10018,f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4263,f4402,f4352,f4265,f4135,f4264,f4353,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6424,f6455,f6466,f6467,f6481,f6491,f6492,f6494,f6496,f6503,f6504,f6506,f6508,f6555,f6559,f6560,f6603,f6597,f6604,f6665,f6670,f6709,f6705,f6710,f6725,f6727,f6729,f6737,f6740,f6755,f6757,f6759,f6767])).
fof(f10018,plain,(
  ( ! [X109,X110,X108] : (~p(s(bool,'const.bool.IN_2'(s(bool,X109),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X108),s(fun(bool,bool),X110)))))) | s(bool,'F_0') = s(bool,X109) | p(s(bool,'const.bool.IN_2'(s(bool,X109),s(fun(bool,bool),X110)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X108)) ) | $spl62),
  inference(superposition,[],[f1739,f4961])).
fof(f10023,plain,(
  ( ! [X118,X119,X117] : (~p(s(bool,'const.bool.IN_2'(s(bool,X117),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X118),s(fun(bool,bool),X119)))))) | s(bool,X118) = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X119)))) | s(bool,'T_0') = s(bool,X117)) ) | $spl62),
  inference(superposition,[],[f305,f4961])).
fof(f10024,plain,(
  ( ! [X121,X122,X120] : (~p(s(bool,'const.bool.IN_2'(s(bool,X121),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X120),s(fun(bool,bool),X122)))))) | s(bool,X121) = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | p(s(bool,'const.bool.IN_2'(s(bool,X121),s(fun(bool,bool),X122)))) | s(bool,'T_0') = s(bool,X120)) ) | $spl62),
  inference(superposition,[],[f305,f4961])).
fof(f10025,plain,(
  ( ! [X125,X123,X124] : (s(bool,X123) != s(bool,X124) | p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X124),s(fun(bool,bool),X125)))))) | s(bool,'T_0') = s(bool,X123)) ) | $spl62),
  inference(superposition,[],[f306,f4961])).
fof(f10026,plain,(
  ( ! [X127,X128,X126] : (s(bool,X126) != s(bool,X127) | p(s(bool,'const.bool.IN_2'(s(bool,X127),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X128)))))) | s(bool,'T_0') = s(bool,X126)) ) | $spl62),
  inference(superposition,[],[f306,f4961])).
fof(f10027,plain,(
  ( ! [X130,X131,X129] : (~p(s(bool,'const.bool.IN_2'(s(bool,X129),s(fun(bool,bool),X130)))) | p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X131),s(fun(bool,bool),X130)))))) | s(bool,'T_0') = s(bool,X129)) ) | $spl62),
  inference(superposition,[],[f307,f4961])).
fof(f10029,plain,(
  ( ! [X133] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,X133),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X133)) ) | $spl62),
  inference(superposition,[],[f505,f4961])).
fof(f10281,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X1))))) ) | $spl116),
  inference(cnf_transformation,[],[f10281_D])).
fof(f10281_D,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X1))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f10278,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X1)))) | s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl62),
  inference(condensation,[],[f10277])).
fof(f10277,plain,(
  ( ! [X107,X105,X106] : (p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))),s(fun(bool,bool),X107)))) | s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X105) | s(bool,'T_0') = s(bool,X106)) ) | $spl62),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4263,f4402,f4352,f4265,f4135,f4264,f4353,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6424,f6455,f6466,f6467,f6481,f6491,f6492,f6494,f6496,f6503,f6504,f6506,f6508,f6555,f6559,f6560,f6603,f6597,f6604,f6665,f6670,f6709,f6705,f6710,f6725,f6727,f6729,f6737,f6740,f6755,f6757,f6759,f6767,f4961,f10082,f10029,f10027,f10026,f10025,f10024,f10023,f10018,f10276,f10017])).
fof(f10275,plain,(
  sP1('T_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl112),
  inference(cnf_transformation,[],[f10275_D])).
fof(f10275_D,plain,(
  sP1('T_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f10263,plain,(
  ( ! [X114] : (p(s(bool,X114)) | sP1('T_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | ($spl22 | $spl62)),
  inference(subsumption_resolution,[],[f10020,f259])).
fof(f10020,plain,(
  ( ! [X114] : (p(s(bool,X114)) | sP1('T_0',g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X114)) ) | ($spl22 | $spl62)),
  inference(superposition,[],[f4887,f4961])).
fof(f10272,plain,(
  sP1(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0') | $spl108),
  inference(cnf_transformation,[],[f10272_D])).
fof(f10272_D,plain,(
  sP1(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0') <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f10265,plain,(
  ( ! [X115] : (p(s(bool,X115)) | sP1(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0')) ) | ($spl26 | $spl62)),
  inference(subsumption_resolution,[],[f10021,f259])).
fof(f10021,plain,(
  ( ! [X115] : (p(s(bool,X115)) | sP1(g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))),'T_0') | s(bool,'T_0') = s(bool,X115)) ) | ($spl26 | $spl62)),
  inference(superposition,[],[f4891,f4961])).
fof(f10269,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl106),
  inference(cnf_transformation,[],[f10269_D])).
fof(f10269_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f3974,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl15),
  inference(cnf_transformation,[],[f3974_D])).
fof(f3974_D,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f10082,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,X0) = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl62),
  inference(equality_factoring,[],[f4961])).
fof(f10268,plain,(
  $false | ($spl14 | $spl62)),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f3975,f4263,f4402,f4352,f4265,f4135,f4264,f4353,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6424,f6455,f6466,f6467,f6481,f6491,f6492,f6494,f6496,f6503,f6504,f6506,f6508,f6555,f6559,f6560,f6603,f6597,f6604,f6665,f6670,f6709,f6705,f6710,f6725,f6727,f6729,f6737,f6740,f6755,f6757,f6759,f6767,f7135,f7134,f7132,f7131,f7257,f7258,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7274,f7275,f4961,f7840,f7845,f7853,f9504,f9509,f10105,f10236,f10239,f10242,f10247,f10252,f10267])).
fof(f10267,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | ($spl14 | $spl62)),
  inference(subsumption_resolution,[],[f10082,f10105])).
fof(f10252,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | ($spl14 | $spl62)),
  inference(subsumption_resolution,[],[f10251,f10105])).
fof(f10251,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | p(s(bool,X1))) ) | $spl62),
  inference(subsumption_resolution,[],[f9838,f259])).
fof(f9838,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl62),
  inference(superposition,[],[f259,f4961])).
fof(f10247,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl14 | $spl62)),
  inference(subsumption_resolution,[],[f10200,f10105])).
fof(f10200,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK16(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl14 | $spl62)),
  inference(backward_demodulation,[],[f10105,f304])).
fof(f10242,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl14 | $spl62)),
  inference(subsumption_resolution,[],[f10156,f10105])).
fof(f10156,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK16('const.pred_set.EMPTY_0',X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl14 | $spl62)),
  inference(backward_demodulation,[],[f10105,f1979])).
fof(f10239,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl14 | $spl62)),
  inference(subsumption_resolution,[],[f10153,f10105])).
fof(f10153,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK16(X1,'const.pred_set.EMPTY_0',X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | ($spl14 | $spl62)),
  inference(backward_demodulation,[],[f10105,f1973])).
fof(f10236,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl14 | $spl62)),
  inference(subsumption_resolution,[],[f10137,f10105])).
fof(f10137,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl14 | $spl62)),
  inference(backward_demodulation,[],[f10105,f263])).
fof(f10105,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl14 | $spl62)),
  inference(duplicate_literal_removal,[],[f9830])).
fof(f9830,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl14 | $spl62)),
  inference(superposition,[],[f4961,f3975])).
fof(f10266,plain,(
  $false | ($spl14 | $spl26 | $spl62)),
  inference(global_subsumption,[],[f10265,f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f3975,f4263,f4402,f4352,f4265,f4135,f4264,f4353,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6424,f6455,f6466,f6467,f6481,f6491,f6492,f6494,f6496,f6503,f6504,f6506,f6508,f6555,f6559,f6560,f6603,f6597,f6604,f6665,f6670,f6709,f6705,f6710,f6725,f6727,f6729,f6737,f6740,f6755,f6757,f6759,f6767,f7135,f7134,f7132,f7131,f7257,f7258,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7274,f7275,f4961,f7840,f7845,f7853,f9504,f9509,f10105,f10236,f10239,f10242,f10247,f10252])).
fof(f10264,plain,(
  $false | ($spl14 | $spl22 | $spl62)),
  inference(global_subsumption,[],[f10263,f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f3975,f4263,f4402,f4352,f4265,f4135,f4264,f4353,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6424,f6455,f6466,f6467,f6481,f6491,f6492,f6494,f6496,f6503,f6504,f6506,f6508,f6555,f6559,f6560,f6603,f6597,f6604,f6665,f6670,f6709,f6705,f6710,f6725,f6727,f6729,f6737,f6740,f6755,f6757,f6759,f6767,f7135,f7134,f7132,f7131,f7257,f7258,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7274,f7275,f4961,f7840,f7845,f7853,f9504,f9509,f10105,f10236,f10239,f10242,f10247,f10252])).
fof(f10253,plain,(
  $false | ($spl14 | $spl62)),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f3975,f4263,f4402,f4352,f4265,f4135,f4264,f4353,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6424,f6455,f6466,f6467,f6481,f6491,f6492,f6494,f6496,f6503,f6504,f6506,f6508,f6555,f6559,f6560,f6603,f6597,f6604,f6665,f6670,f6709,f6705,f6710,f6725,f6727,f6729,f6737,f6740,f6755,f6757,f6759,f6767,f7135,f7134,f7132,f7131,f7257,f7258,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7274,f7275,f4961,f7840,f7845,f7853,f9504,f9509,f10105,f10236,f10239,f10242,f10247,f10252])).
fof(f10250,plain,(
  $false | ($spl8 | $spl14 | $spl62)),
  inference(global_subsumption,[],[f10249,f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f3975,f4263,f4402,f4352,f4265,f4135,f4264,f4353,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6424,f6455,f6466,f6467,f6481,f6491,f6492,f6494,f6496,f6503,f6504,f6506,f6508,f6555,f6559,f6560,f6603,f6597,f6604,f6665,f6670,f6709,f6705,f6710,f6725,f6727,f6729,f6737,f6740,f6755,f6757,f6759,f6767,f7135,f7134,f7132,f7131,f7257,f7258,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7274,f7275,f4961,f7840,f7845,f7853,f9504,f9509,f10105,f10236,f10239,f10242,f10247])).
fof(f10249,plain,(
  ( ! [X4] : (s(sK12,sK13) = s(sK12,X4)) ) | ($spl8 | $spl14 | $spl62)),
  inference(subsumption_resolution,[],[f10220,f244])).
fof(f10220,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(sK12,sK13) = s(sK12,X4)) ) | ($spl8 | $spl14 | $spl62)),
  inference(backward_demodulation,[],[f10105,f433])).
fof(f10248,plain,(
  $false | ($spl14 | $spl62)),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f3975,f4263,f4402,f4352,f4265,f4135,f4264,f4353,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6424,f6455,f6466,f6467,f6481,f6491,f6492,f6494,f6496,f6503,f6504,f6506,f6508,f6555,f6559,f6560,f6603,f6597,f6604,f6665,f6670,f6709,f6705,f6710,f6725,f6727,f6729,f6737,f6740,f6755,f6757,f6759,f6767,f7135,f7134,f7132,f7131,f7257,f7258,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7274,f7275,f4961,f7840,f7845,f7853,f9504,f9509,f10105,f10236,f10239,f10242,f10247])).
fof(f10243,plain,(
  $false | ($spl14 | $spl62)),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f3975,f4263,f4402,f4352,f4265,f4135,f4264,f4353,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6424,f6455,f6466,f6467,f6481,f6491,f6492,f6494,f6496,f6503,f6504,f6506,f6508,f6555,f6559,f6560,f6603,f6597,f6604,f6665,f6670,f6709,f6705,f6710,f6725,f6727,f6729,f6737,f6740,f6755,f6757,f6759,f6767,f7135,f7134,f7132,f7131,f7257,f7258,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7274,f7275,f4961,f7840,f7845,f7853,f9504,f9509,f10105,f10236,f10239,f10242])).
fof(f10240,plain,(
  $false | ($spl14 | $spl62)),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f3975,f4263,f4402,f4352,f4265,f4135,f4264,f4353,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6424,f6455,f6466,f6467,f6481,f6491,f6492,f6494,f6496,f6503,f6504,f6506,f6508,f6555,f6559,f6560,f6603,f6597,f6604,f6665,f6670,f6709,f6705,f6710,f6725,f6727,f6729,f6737,f6740,f6755,f6757,f6759,f6767,f7135,f7134,f7132,f7131,f7257,f7258,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7274,f7275,f4961,f7840,f7845,f7853,f9504,f9509,f10105,f10236,f10239])).
fof(f10238,plain,(
  $false | ($spl14 | $spl62)),
  inference(subsumption_resolution,[],[f10139,f244])).
fof(f10139,plain,(
  ~p(s(bool,'T_0')) | ($spl14 | $spl62)),
  inference(backward_demodulation,[],[f10105,f434])).
fof(f10237,plain,(
  $false | ($spl14 | $spl62)),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f251,f260,f259,f302,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1427,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f4020,f4081,f4082,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4349,f4350,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f3975,f4263,f4402,f4352,f4265,f4135,f4264,f4353,f4356,f4355,f4354,f4262,f4261,f4089,f4088,f4087,f4086,f4085,f304,f303,f2026,f307,f306,f305,f4964,f4965,f4134,f4401,f1981,f1979,f1977,f1975,f1973,f1741,f1739,f1734,f1732,f1670,f1668,f1459,f1458,f1426,f1425,f1421,f1420,f505,f434,f264,f263,f6129,f6113,f6114,f6189,f6175,f6371,f6372,f6373,f6374,f6375,f6424,f6455,f6466,f6467,f6481,f6491,f6492,f6494,f6496,f6503,f6504,f6506,f6508,f6555,f6559,f6560,f6603,f6597,f6604,f6665,f6670,f6709,f6705,f6710,f6725,f6727,f6729,f6737,f6740,f6755,f6757,f6759,f6767,f7135,f7134,f7132,f7131,f7257,f7258,f7259,f7260,f7261,f7262,f7263,f7264,f7265,f7266,f7267,f7268,f7269,f7270,f7271,f7272,f7273,f7274,f7275,f4961,f7840,f7845,f7853,f9504,f9509,f10105,f10236])).
fof(f10136,plain,(
  $false | ($spl14 | $spl62)),
  inference(subsumption_resolution,[],[f10135,f465])).
fof(f10135,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl14 | $spl62)),
  inference(forward_demodulation,[],[f10134,f3975])).
fof(f10134,plain,(
  s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl62),
  inference(subsumption_resolution,[],[f9828,f504])).
fof(f9828,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X15,X16),s(fun(X15,bool),'const.pred_set.EMPTY_0')))) ) | $spl62),
  inference(superposition,[],[f4961,f505])).
fof(f10133,plain,(
  $false | ($spl5 | $spl14 | $spl62 | $spl80)),
  inference(subsumption_resolution,[],[f10132,f465])).
fof(f10132,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl5 | $spl14 | $spl62 | $spl80)),
  inference(forward_demodulation,[],[f10131,f3975])).
fof(f10131,plain,(
  s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl62 | $spl80)),
  inference(subsumption_resolution,[],[f9827,f6025])).
fof(f9827,plain,(
  s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl62 | $spl80)),
  inference(superposition,[],[f4961,f6050])).
fof(f10130,plain,(
  $false | ($spl5 | $spl14 | $spl62 | $spl70)),
  inference(subsumption_resolution,[],[f10129,f465])).
fof(f10129,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl5 | $spl14 | $spl62 | $spl70)),
  inference(forward_demodulation,[],[f10128,f3975])).
fof(f10128,plain,(
  s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl62 | $spl70)),
  inference(subsumption_resolution,[],[f9826,f5942])).
fof(f9826,plain,(
  s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl62 | $spl70)),
  inference(superposition,[],[f4961,f5967])).
fof(f10127,plain,(
  $false | ($spl14 | $spl62)),
  inference(subsumption_resolution,[],[f10126,f465])).
fof(f10126,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl14 | $spl62)),
  inference(condensation,[],[f10125])).
fof(f10125,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X9)) ) | ($spl14 | $spl62)),
  inference(forward_demodulation,[],[f10106,f3975])).
fof(f10106,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X9)) ) | $spl62),
  inference(duplicate_literal_removal,[],[f9823])).
fof(f9823,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl62),
  inference(superposition,[],[f4961,f251])).
fof(f10124,plain,(
  $false | ($spl14 | $spl62)),
  inference(subsumption_resolution,[],[f10123,f465])).
fof(f10123,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl14 | $spl62)),
  inference(forward_demodulation,[],[f9813,f3975])).
fof(f9813,plain,(
  s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl62),
  inference(unit_resulting_resolution,[],[f465,f4961])).
fof(f10122,plain,(
  $false | ($spl14 | $spl62)),
  inference(subsumption_resolution,[],[f10121,f504])).
fof(f10121,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl14 | $spl62)),
  inference(forward_demodulation,[],[f9814,f3975])).
fof(f9814,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))) = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl62),
  inference(unit_resulting_resolution,[],[f504,f4961])).
fof(f10120,plain,(
  $false | ($spl5 | $spl14 | $spl62 | $spl70)),
  inference(subsumption_resolution,[],[f10119,f465])).
fof(f10119,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl5 | $spl14 | $spl62 | $spl70)),
  inference(forward_demodulation,[],[f10118,f3975])).
fof(f10118,plain,(
  s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl62 | $spl70)),
  inference(forward_demodulation,[],[f9815,f5967])).
fof(f9815,plain,(
  s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl62 | $spl70)),
  inference(unit_resulting_resolution,[],[f5942,f4961])).
fof(f10117,plain,(
  $false | ($spl5 | $spl14 | $spl62 | $spl80)),
  inference(subsumption_resolution,[],[f10116,f465])).
fof(f10116,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl5 | $spl14 | $spl62 | $spl80)),
  inference(forward_demodulation,[],[f10115,f3975])).
fof(f10115,plain,(
  s(bool,'F_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl62 | $spl80)),
  inference(forward_demodulation,[],[f9816,f6050])).
fof(f9816,plain,(
  s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | ($spl5 | $spl62 | $spl80)),
  inference(unit_resulting_resolution,[],[f6025,f4961])).
fof(f10114,plain,(
  $false | ($spl5 | $spl14 | $spl62)),
  inference(subsumption_resolution,[],[f10113,f5095])).
fof(f10113,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl14 | $spl62)),
  inference(forward_demodulation,[],[f9817,f3975])).
fof(f9817,plain,(
  s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl14 | $spl62)),
  inference(unit_resulting_resolution,[],[f5095,f4961])).
fof(f10112,plain,(
  $false | ($spl14 | $spl61 | $spl62)),
  inference(subsumption_resolution,[],[f10111,f4960])).
fof(f10111,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | ($spl14 | $spl61 | $spl62)),
  inference(forward_demodulation,[],[f9818,f3975])).
fof(f9818,plain,(
  s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl61 | $spl62)),
  inference(unit_resulting_resolution,[],[f4960,f4961])).
fof(f7335,plain,(
  $false | ($spl5 | $spl14 | $spl17)),
  inference(subsumption_resolution,[],[f5095,f7319])).
fof(f7319,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl17),
  inference(trivial_inequality_removal,[],[f7318])).
fof(f7318,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl17),
  inference(superposition,[],[f3977,f251])).
fof(f7334,plain,(
  $false | ($spl5 | $spl14 | $spl17)),
  inference(subsumption_resolution,[],[f7333,f430])).
fof(f7333,plain,(
  s(fun(sK12,bool),sK14) = s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | ($spl14 | $spl17)),
  inference(subsumption_resolution,[],[f7185,f7319])).
fof(f7185,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(fun(sK12,bool),sK14) = s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | $spl14),
  inference(superposition,[],[f341,f3975])).
fof(f7330,plain,(
  $false | ($spl17 | $spl65)),
  inference(subsumption_resolution,[],[f7319,f4963])).
fof(f7329,plain,(
  $false | ($spl17 | $spl65)),
  inference(subsumption_resolution,[],[f7328,f4963])).
fof(f7328,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl17),
  inference(condensation,[],[f7327])).
fof(f7327,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl17),
  inference(subsumption_resolution,[],[f7317,f251])).
fof(f7317,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl17),
  inference(superposition,[],[f3977,f635])).
fof(f7326,plain,(
  $false | ($spl17 | $spl65)),
  inference(subsumption_resolution,[],[f7325,f4963])).
fof(f7325,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl17),
  inference(condensation,[],[f7324])).
fof(f7324,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl17),
  inference(subsumption_resolution,[],[f7316,f251])).
fof(f7316,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl17),
  inference(superposition,[],[f3977,f635])).
fof(f7323,plain,(
  $false | ($spl17 | $spl65)),
  inference(subsumption_resolution,[],[f7303,f4963])).
fof(f7303,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl17),
  inference(unit_resulting_resolution,[],[f3977,f251])).
fof(f7322,plain,(
  $false | ($spl17 | $spl65)),
  inference(subsumption_resolution,[],[f7307,f4963])).
fof(f7307,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl17),
  inference(unit_resulting_resolution,[],[f465,f3977,f635])).
fof(f7321,plain,(
  $false | ($spl17 | $spl65)),
  inference(subsumption_resolution,[],[f7312,f4963])).
fof(f7312,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl17),
  inference(unit_resulting_resolution,[],[f465,f3977,f635])).
fof(f7320,plain,(
  $false | ($spl17 | $spl65)),
  inference(subsumption_resolution,[],[f7313,f4963])).
fof(f7313,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl17),
  inference(unit_resulting_resolution,[],[f434,f3977,f704])).
fof(f7106,plain,(
  $false | ($spl13 | $spl61)),
  inference(subsumption_resolution,[],[f7095,f4960])).
fof(f7095,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl13),
  inference(trivial_inequality_removal,[],[f7094])).
fof(f7094,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl13),
  inference(superposition,[],[f3922,f251])).
fof(f7105,plain,(
  $false | ($spl13 | $spl61)),
  inference(subsumption_resolution,[],[f7104,f4960])).
fof(f7104,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl13),
  inference(condensation,[],[f7103])).
fof(f7103,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) ) | $spl13),
  inference(subsumption_resolution,[],[f7093,f251])).
fof(f7093,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) ) | $spl13),
  inference(superposition,[],[f3922,f635])).
fof(f7102,plain,(
  $false | ($spl13 | $spl61)),
  inference(subsumption_resolution,[],[f7101,f4960])).
fof(f7101,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl13),
  inference(condensation,[],[f7100])).
fof(f7100,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl13),
  inference(subsumption_resolution,[],[f7092,f251])).
fof(f7092,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl13),
  inference(superposition,[],[f3922,f635])).
fof(f7099,plain,(
  $false | ($spl13 | $spl61)),
  inference(subsumption_resolution,[],[f7079,f4960])).
fof(f7079,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl13),
  inference(unit_resulting_resolution,[],[f3922,f251])).
fof(f7098,plain,(
  $false | ($spl13 | $spl61)),
  inference(subsumption_resolution,[],[f7083,f4960])).
fof(f7083,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl13),
  inference(unit_resulting_resolution,[],[f465,f3922,f635])).
fof(f7097,plain,(
  $false | ($spl13 | $spl61)),
  inference(subsumption_resolution,[],[f7088,f4960])).
fof(f7088,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl13),
  inference(unit_resulting_resolution,[],[f465,f3922,f635])).
fof(f7096,plain,(
  $false | ($spl13 | $spl61)),
  inference(subsumption_resolution,[],[f7089,f4960])).
fof(f7089,plain,(
  s(bool,'T_0') = s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl13),
  inference(unit_resulting_resolution,[],[f434,f3922,f704])).
fof(f6477,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK17(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X0,X1,bool,bool))))) ) | $spl102),
  inference(cnf_transformation,[],[f6477_D])).
fof(f6477_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK17(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X0,X1,bool,bool))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f6479,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl105),
  inference(cnf_transformation,[],[f6479_D])).
fof(f6479_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f6467,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK17(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1670,f635])).
fof(f6466,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK17(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK17(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1670,f635])).
fof(f6381,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl96),
  inference(cnf_transformation,[],[f6381_D])).
fof(f6381_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f6377,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK17(X0,X1,bool,bool))) ) | $spl92),
  inference(cnf_transformation,[],[f6377_D])).
fof(f6377_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK17(X0,X1,bool,bool))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f6384,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl101),
  inference(cnf_transformation,[],[f6384_D])).
fof(f6384_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f6375,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK17(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1668,f251])).
fof(f6382,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl98),
  inference(cnf_transformation,[],[f6382_D])).
fof(f6382_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f6374,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK17(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1668,f635])).
fof(f6373,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK17(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1668,f635])).
fof(f6376,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl90),
  inference(cnf_transformation,[],[f6376_D])).
fof(f6376_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f6379,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl95),
  inference(cnf_transformation,[],[f6379_D])).
fof(f6379_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f6372,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK17(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1668,f635])).
fof(f6371,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK17(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1668,f635])).
fof(f6130,plain,(
  ( ! [X21] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X21))))))) ) | $spl86),
  inference(cnf_transformation,[],[f6130_D])).
fof(f6130_D,plain,(
  ( ! [X21] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X21))))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f6131,plain,(
  ( ! [X19,X22,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X22) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X20)) ) | $spl88),
  inference(cnf_transformation,[],[f6131_D])).
fof(f6131_D,plain,(
  ( ! [X19,X22,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X22) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X20)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f6114,plain,(
  ( ! [X26,X24,X23,X25] : (s(bool,X23) != s(bool,X24) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))))) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f1420,f635])).
fof(f6113,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X19) != s(bool,X20) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X21)))))) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f1420,f635])).
fof(f6021,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | $spl80),
  inference(cnf_transformation,[],[f6021_D])).
fof(f6021_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f6024,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | $spl85),
  inference(cnf_transformation,[],[f6024_D])).
fof(f6024_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f6016,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | $spl5),
  inference(superposition,[],[f2606,f251])).
fof(f6022,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X2)) ) | $spl82),
  inference(cnf_transformation,[],[f6022_D])).
fof(f6022_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f6015,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))))) ) | $spl5),
  inference(superposition,[],[f2606,f635])).
fof(f6014,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl5),
  inference(superposition,[],[f2606,f635])).
fof(f6017,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl76),
  inference(cnf_transformation,[],[f6017_D])).
fof(f6017_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f6019,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | $spl79),
  inference(cnf_transformation,[],[f6019_D])).
fof(f6019_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f6013,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl5),
  inference(superposition,[],[f2606,f635])).
fof(f6012,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),sK14))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl5),
  inference(superposition,[],[f2606,f635])).
fof(f5938,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | $spl70),
  inference(cnf_transformation,[],[f5938_D])).
fof(f5938_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f5941,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | $spl75),
  inference(cnf_transformation,[],[f5941_D])).
fof(f5941_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f5933,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | $spl5),
  inference(superposition,[],[f2604,f251])).
fof(f5939,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X2)) ) | $spl72),
  inference(cnf_transformation,[],[f5939_D])).
fof(f5939_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f5932,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))))) ) | $spl5),
  inference(superposition,[],[f2604,f635])).
fof(f5931,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl5),
  inference(superposition,[],[f2604,f635])).
fof(f5934,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl66),
  inference(cnf_transformation,[],[f5934_D])).
fof(f5934_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f5936,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | $spl69),
  inference(cnf_transformation,[],[f5936_D])).
fof(f5936_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f5930,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl5),
  inference(superposition,[],[f2604,f635])).
fof(f5929,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),sK14))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl5),
  inference(superposition,[],[f2604,f635])).
fof(f5093,plain,(
  $false | ($spl5 | $spl24)),
  inference(subsumption_resolution,[],[f5059,f4888])).
fof(f5059,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | ($spl5 | $spl24)),
  inference(backward_demodulation,[],[f4888,f2604])).
fof(f5092,plain,(
  $false | ($spl5 | $spl24)),
  inference(subsumption_resolution,[],[f5058,f4888])).
fof(f5058,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,sK16(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK12)),s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) | ($spl5 | $spl24)),
  inference(backward_demodulation,[],[f4888,f2606])).
fof(f5091,plain,(
  $false | ($spl8 | $spl24)),
  inference(global_subsumption,[],[f5090,f243,f242,f241,f240,f244,f434,f251,f260,f259,f264,f263,f302,f304,f303,f307,f306,f305,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f505,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1420,f1421,f1425,f1426,f1427,f1458,f1459,f1973,f1979,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f2026,f1981,f1977,f1975,f1741,f1739,f1734,f1732,f1670,f1668,f4020,f4081,f4082,f4085,f4086,f4087,f4088,f4089,f4134,f4135,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4261,f4262,f4263,f4264,f4265,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4964,f4965,f4349,f4350,f4352,f4353,f4354,f4355,f4356,f4401,f4402,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4888,f5077,f5080,f5082,f5088])).
fof(f5088,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl24),
  inference(subsumption_resolution,[],[f5039,f4888])).
fof(f5039,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK16(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl24),
  inference(backward_demodulation,[],[f4888,f304])).
fof(f5082,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | $spl24),
  inference(subsumption_resolution,[],[f5024,f4888])).
fof(f5024,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK16('const.pred_set.EMPTY_0',X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | $spl24),
  inference(backward_demodulation,[],[f4888,f1979])).
fof(f5080,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | $spl24),
  inference(subsumption_resolution,[],[f5021,f4888])).
fof(f5021,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK16(X1,'const.pred_set.EMPTY_0',X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | $spl24),
  inference(backward_demodulation,[],[f4888,f1973])).
fof(f5077,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl24),
  inference(subsumption_resolution,[],[f5005,f4888])).
fof(f5005,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl24),
  inference(backward_demodulation,[],[f4888,f263])).
fof(f5090,plain,(
  ( ! [X4] : (s(sK12,sK13) = s(sK12,X4)) ) | ($spl8 | $spl24)),
  inference(subsumption_resolution,[],[f5057,f244])).
fof(f5057,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s(sK12,sK13) = s(sK12,X4)) ) | ($spl8 | $spl24)),
  inference(backward_demodulation,[],[f4888,f433])).
fof(f5089,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f434,f251,f260,f259,f264,f263,f302,f304,f303,f307,f306,f305,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f505,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1420,f1421,f1425,f1426,f1427,f1458,f1459,f1973,f1979,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f2026,f1981,f1977,f1975,f1741,f1739,f1734,f1732,f1670,f1668,f4020,f4081,f4082,f4085,f4086,f4087,f4088,f4089,f4134,f4135,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4261,f4262,f4263,f4264,f4265,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4964,f4965,f4349,f4350,f4352,f4353,f4354,f4355,f4356,f4401,f4402,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4888,f5077,f5080,f5082,f5088])).
fof(f5085,plain,(
  $false | ($spl5 | $spl24)),
  inference(subsumption_resolution,[],[f5031,f4888])).
fof(f5031,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | ($spl5 | $spl24)),
  inference(backward_demodulation,[],[f4888,f2603])).
fof(f5084,plain,(
  $false | ($spl5 | $spl24)),
  inference(subsumption_resolution,[],[f5030,f4888])).
fof(f5030,plain,(
  s(bool,'T_0') != s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | ($spl5 | $spl24)),
  inference(backward_demodulation,[],[f4888,f2605])).
fof(f5083,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f434,f251,f260,f259,f264,f263,f302,f304,f303,f307,f306,f305,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f505,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1420,f1421,f1425,f1426,f1427,f1458,f1459,f1973,f1979,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f2026,f1981,f1977,f1975,f1741,f1739,f1734,f1732,f1670,f1668,f4020,f4081,f4082,f4085,f4086,f4087,f4088,f4089,f4134,f4135,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4261,f4262,f4263,f4264,f4265,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4964,f4965,f4349,f4350,f4352,f4353,f4354,f4355,f4356,f4401,f4402,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4888,f5077,f5080,f5082])).
fof(f5081,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f434,f251,f260,f259,f264,f263,f302,f304,f303,f307,f306,f305,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f505,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1420,f1421,f1425,f1426,f1427,f1458,f1459,f1973,f1979,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f2026,f1981,f1977,f1975,f1741,f1739,f1734,f1732,f1670,f1668,f4020,f4081,f4082,f4085,f4086,f4087,f4088,f4089,f4134,f4135,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4261,f4262,f4263,f4264,f4265,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4964,f4965,f4349,f4350,f4352,f4353,f4354,f4355,f4356,f4401,f4402,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4888,f5077,f5080])).
fof(f5079,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f5007,f244])).
fof(f5007,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(backward_demodulation,[],[f4888,f434])).
fof(f5078,plain,(
  $false | $spl24),
  inference(global_subsumption,[],[f243,f242,f241,f240,f244,f434,f251,f260,f259,f264,f263,f302,f304,f303,f307,f306,f305,f311,f308,f318,f317,f435,f336,f333,f332,f331,f341,f348,f347,f346,f345,f344,f343,f342,f352,f351,f350,f349,f356,f355,f354,f353,f360,f359,f358,f357,f364,f363,f362,f361,f418,f368,f367,f366,f365,f375,f374,f373,f372,f381,f378,f377,f376,f388,f387,f386,f417,f384,f383,f382,f392,f391,f390,f389,f396,f395,f394,f393,f405,f404,f403,f416,f401,f400,f399,f409,f408,f407,f406,f413,f412,f411,f410,f436,f437,f438,f439,f441,f443,f444,f446,f448,f449,f450,f451,f452,f454,f456,f457,f458,f460,f462,f463,f465,f485,f486,f489,f495,f499,f525,f524,f521,f520,f519,f518,f517,f516,f515,f514,f513,f512,f511,f510,f509,f508,f506,f505,f504,f540,f545,f552,f563,f572,f573,f575,f583,f588,f589,f592,f599,f635,f704,f731,f739,f737,f760,f756,f801,f822,f833,f837,f842,f846,f858,f868,f878,f900,f942,f941,f939,f938,f956,f954,f966,f975,f1038,f1037,f1036,f1035,f1042,f1125,f1123,f1137,f1149,f1211,f1263,f1260,f1254,f1252,f1325,f1322,f1380,f1428,f1420,f1421,f1425,f1426,f1427,f1458,f1459,f1973,f1979,f2206,f2205,f2204,f2203,f2201,f2200,f2199,f2196,f2195,f2191,f2188,f2187,f2182,f2181,f2175,f2174,f2157,f2156,f2155,f2151,f2150,f2146,f2145,f2144,f2141,f2137,f2135,f2131,f2130,f2125,f2124,f2123,f2122,f2121,f2118,f2116,f2111,f2110,f2105,f2104,f2098,f2091,f2090,f2089,f2087,f2086,f2084,f2083,f2081,f2079,f2078,f2077,f2076,f2071,f2070,f2064,f2063,f2058,f2057,f2051,f2050,f2048,f2047,f2043,f2041,f2039,f2038,f2035,f2034,f2033,f2032,f2031,f2026,f1981,f1977,f1975,f1741,f1739,f1734,f1732,f1670,f1668,f4020,f4081,f4082,f4085,f4086,f4087,f4088,f4089,f4134,f4135,f4136,f4232,f4886,f4890,f4893,f4896,f4899,f4902,f4905,f4908,f4911,f4914,f4917,f4920,f4923,f4926,f4929,f4932,f4935,f4261,f4262,f4263,f4264,f4265,f4266,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4964,f4965,f4349,f4350,f4352,f4353,f4354,f4355,f4356,f4401,f4402,f4403,f4972,f4974,f4976,f4978,f4980,f4982,f4984,f4986,f4988,f4990,f4992,f4994,f4996,f4998,f5000,f5002,f5004,f4888,f5077])).
fof(f4961,plain,(
  ( ! [X104] : (s(bool,X104) = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X104)) ) | $spl62),
  inference(cnf_transformation,[],[f4961_D])).
fof(f4961_D,plain,(
  ( ! [X104] : (s(bool,X104) = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X104)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f4963,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl65),
  inference(cnf_transformation,[],[f4963_D])).
fof(f4963_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f4516,plain,(
  ( ! [X104] : (s(bool,'T_0') != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,X104) = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X104)) ) | $spl5),
  inference(superposition,[],[f2605,f635])).
fof(f4958,plain,(
  ( ! [X103] : (s(bool,X103) = s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | s(bool,'T_0') = s(bool,X103)) ) | $spl58),
  inference(cnf_transformation,[],[f4958_D])).
fof(f4958_D,plain,(
  ( ! [X103] : (s(bool,X103) = s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | s(bool,'T_0') = s(bool,X103)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f4960,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl61),
  inference(cnf_transformation,[],[f4960_D])).
fof(f4960_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f4515,plain,(
  ( ! [X103] : (s(bool,'T_0') != s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | s(bool,X103) = s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | s(bool,'T_0') = s(bool,X103)) ) | $spl5),
  inference(superposition,[],[f2603,f635])).
fof(f4936,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl56),
  inference(cnf_transformation,[],[f4936_D])).
fof(f4936_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f5004,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f5003])).
fof(f5003,plain,(
  ( ! [X94,X97,X95,X96] : (~sP8('T_0',X95,X96) | sP9('T_0',X95,X96) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f4431,f704])).
fof(f4431,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,X94)) | ~sP8('T_0',X95,X96) | sP9('T_0',X95,X96) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f417,f635])).
fof(f4933,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl54),
  inference(cnf_transformation,[],[f4933_D])).
fof(f4933_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f5002,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f5001])).
fof(f5001,plain,(
  ( ! [X92,X90,X93,X91] : (~sP10('T_0',X91,X92) | sP11('T_0',X91,X92) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f4430,f704])).
fof(f4430,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,X90)) | ~sP10('T_0',X91,X92) | sP11('T_0',X91,X92) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f416,f635])).
fof(f4930,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl52),
  inference(cnf_transformation,[],[f4930_D])).
fof(f4930_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f5000,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4999])).
fof(f4999,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,X87)) | p(s(bool,X88)) | ~sP11('T_0',X88,X87) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f4429,f704])).
fof(f4429,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,X86)) | ~p(s(bool,X87)) | p(s(bool,X88)) | ~sP11('T_0',X88,X87) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f413,f635])).
fof(f4927,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl50),
  inference(cnf_transformation,[],[f4927_D])).
fof(f4927_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4998,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f4997])).
fof(f4997,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,X83)) | sP11('T_0',X83,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f4428,f704])).
fof(f4428,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,X82)) | ~p(s(bool,X83)) | sP11('T_0',X83,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f412,f635])).
fof(f4924,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl48),
  inference(cnf_transformation,[],[f4924_D])).
fof(f4924_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f4996,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4995])).
fof(f4995,plain,(
  ( ! [X80,X78,X81,X79] : (p(s(bool,X79)) | sP11('T_0',X80,X79) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f4427,f704])).
fof(f4427,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X78)) | p(s(bool,X79)) | sP11('T_0',X80,X79) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f411,f635])).
fof(f4921,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl46),
  inference(cnf_transformation,[],[f4921_D])).
fof(f4921_D,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4994,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f4993])).
fof(f4993,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X75)) | ~p(s(bool,X76)) | ~sP10(X76,X75,'T_0') | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f4426,f704])).
fof(f4426,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X74)) | p(s(bool,X75)) | ~p(s(bool,X76)) | ~sP10(X76,X75,'T_0') | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f406,f635])).
fof(f4918,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl44),
  inference(cnf_transformation,[],[f4918_D])).
fof(f4918_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4992,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4991])).
fof(f4991,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,X71)) | p(s(bool,X72)) | ~sP9('T_0',X72,X71) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f4425,f704])).
fof(f4425,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X70)) | p(s(bool,X71)) | p(s(bool,X72)) | ~sP9('T_0',X72,X71) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f396,f635])).
fof(f4915,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl42),
  inference(cnf_transformation,[],[f4915_D])).
fof(f4915_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4990,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4989])).
fof(f4989,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X67)) | sP9('T_0',X67,X68) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f4424,f704])).
fof(f4424,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X66)) | ~p(s(bool,X67)) | sP9('T_0',X67,X68) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f395,f635])).
fof(f4912,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl40),
  inference(cnf_transformation,[],[f4912_D])).
fof(f4912_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4988,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4987])).
fof(f4987,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X63)) | sP9('T_0',X64,X63) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f4423,f704])).
fof(f4423,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X62)) | ~p(s(bool,X63)) | sP9('T_0',X64,X63) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f394,f635])).
fof(f4909,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl38),
  inference(cnf_transformation,[],[f4909_D])).
fof(f4909_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4986,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f4985])).
fof(f4985,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,X59)) | p(s(bool,X60)) | ~sP8('T_0',X59,X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f4422,f704])).
fof(f4422,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X58)) | p(s(bool,X59)) | p(s(bool,X60)) | ~sP8('T_0',X59,X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f389,f635])).
fof(f4906,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl36),
  inference(cnf_transformation,[],[f4906_D])).
fof(f4906_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4984,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f4983])).
fof(f4983,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | ~sP7(X56,X55,'T_0') | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f4421,f704])).
fof(f4421,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X54)) | ~p(s(bool,X55)) | p(s(bool,X56)) | ~sP7(X56,X55,'T_0') | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f381,f635])).
fof(f4903,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl34),
  inference(cnf_transformation,[],[f4903_D])).
fof(f4903_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4982,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f4981])).
fof(f4981,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP6(X51,'T_0',X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f4420,f704])).
fof(f4420,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP6(X51,'T_0',X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f372,f635])).
fof(f4900,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl32),
  inference(cnf_transformation,[],[f4900_D])).
fof(f4900_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4980,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f4979])).
fof(f4979,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | p(s(bool,X48)) | ~sP3(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f4419,f704])).
fof(f4419,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP3(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f353,f635])).
fof(f4897,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl30),
  inference(cnf_transformation,[],[f4897_D])).
fof(f4897_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4978,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f4977])).
fof(f4977,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X43)) | p(s(bool,X44)) | ~sP4(X43,'T_0',X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f4418,f704])).
fof(f4418,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X42)) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP4(X43,'T_0',X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f349,f635])).
fof(f4894,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl28),
  inference(cnf_transformation,[],[f4894_D])).
fof(f4894_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f4976,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f4975])).
fof(f4975,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X39)) | p(s(bool,X40)) | ~sP5(X39,'T_0',X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f4417,f704])).
fof(f4417,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X38)) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~sP5(X39,'T_0',X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f345,f635])).
fof(f4891,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl26),
  inference(cnf_transformation,[],[f4891_D])).
fof(f4891_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4974,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4973])).
fof(f4973,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X36)) | sP1(X36,'T_0') | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f4416,f704])).
fof(f4416,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X35)) | p(s(bool,X36)) | sP1(X36,'T_0') | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f336,f635])).
fof(f4887,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl22),
  inference(cnf_transformation,[],[f4887_D])).
fof(f4887_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f4972,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4971])).
fof(f4971,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1('T_0',X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4415,f704])).
fof(f4415,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1('T_0',X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f333,f635])).
fof(f4344,plain,(
  ( ! [X104] : (s(bool,'T_0') != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,X104) = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X104)) ) | $spl5),
  inference(superposition,[],[f2605,f635])).
fof(f4343,plain,(
  ( ! [X103] : (s(bool,'T_0') != s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | s(bool,X103) = s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | s(bool,'T_0') = s(bool,X103)) ) | $spl5),
  inference(superposition,[],[f2603,f635])).
fof(f4935,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4934])).
fof(f4934,plain,(
  ( ! [X94,X97,X95,X96] : (~sP8('T_0',X95,X96) | sP9('T_0',X95,X96) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f4259,f704])).
fof(f4259,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,X94)) | ~sP8('T_0',X95,X96) | sP9('T_0',X95,X96) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f417,f635])).
fof(f4932,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f4931])).
fof(f4931,plain,(
  ( ! [X92,X90,X93,X91] : (~sP10('T_0',X91,X92) | sP11('T_0',X91,X92) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f4258,f704])).
fof(f4258,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,X90)) | ~sP10('T_0',X91,X92) | sP11('T_0',X91,X92) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f416,f635])).
fof(f4929,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4928])).
fof(f4928,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,X87)) | p(s(bool,X88)) | ~sP11('T_0',X88,X87) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f4257,f704])).
fof(f4257,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,X86)) | ~p(s(bool,X87)) | p(s(bool,X88)) | ~sP11('T_0',X88,X87) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f413,f635])).
fof(f4926,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f4925])).
fof(f4925,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,X83)) | sP11('T_0',X83,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f4256,f704])).
fof(f4256,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,X82)) | ~p(s(bool,X83)) | sP11('T_0',X83,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f412,f635])).
fof(f4923,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4922])).
fof(f4922,plain,(
  ( ! [X80,X78,X81,X79] : (p(s(bool,X79)) | sP11('T_0',X80,X79) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f4255,f704])).
fof(f4255,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X78)) | p(s(bool,X79)) | sP11('T_0',X80,X79) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f411,f635])).
fof(f4920,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f4919])).
fof(f4919,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X75)) | ~p(s(bool,X76)) | ~sP10(X76,X75,'T_0') | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f4254,f704])).
fof(f4254,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X74)) | p(s(bool,X75)) | ~p(s(bool,X76)) | ~sP10(X76,X75,'T_0') | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f406,f635])).
fof(f4917,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4916])).
fof(f4916,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,X71)) | p(s(bool,X72)) | ~sP9('T_0',X72,X71) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f4253,f704])).
fof(f4253,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X70)) | p(s(bool,X71)) | p(s(bool,X72)) | ~sP9('T_0',X72,X71) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f396,f635])).
fof(f4914,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4913])).
fof(f4913,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X67)) | sP9('T_0',X67,X68) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f4252,f704])).
fof(f4252,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X66)) | ~p(s(bool,X67)) | sP9('T_0',X67,X68) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f395,f635])).
fof(f4911,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4910])).
fof(f4910,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X63)) | sP9('T_0',X64,X63) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f4251,f704])).
fof(f4251,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X62)) | ~p(s(bool,X63)) | sP9('T_0',X64,X63) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f394,f635])).
fof(f4908,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f4907])).
fof(f4907,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,X59)) | p(s(bool,X60)) | ~sP8('T_0',X59,X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f4250,f704])).
fof(f4250,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X58)) | p(s(bool,X59)) | p(s(bool,X60)) | ~sP8('T_0',X59,X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f389,f635])).
fof(f4905,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f4904])).
fof(f4904,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | ~sP7(X56,X55,'T_0') | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f4249,f704])).
fof(f4249,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X54)) | ~p(s(bool,X55)) | p(s(bool,X56)) | ~sP7(X56,X55,'T_0') | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f381,f635])).
fof(f4902,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f4901])).
fof(f4901,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP6(X51,'T_0',X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f4248,f704])).
fof(f4248,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP6(X51,'T_0',X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f372,f635])).
fof(f4899,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f4898])).
fof(f4898,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | p(s(bool,X48)) | ~sP3(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f4247,f704])).
fof(f4247,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP3(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f353,f635])).
fof(f4896,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f4895])).
fof(f4895,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X43)) | p(s(bool,X44)) | ~sP4(X43,'T_0',X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f4246,f704])).
fof(f4246,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X42)) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP4(X43,'T_0',X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f349,f635])).
fof(f4893,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f4892])).
fof(f4892,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X39)) | p(s(bool,X40)) | ~sP5(X39,'T_0',X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f4245,f704])).
fof(f4245,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X38)) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~sP5(X39,'T_0',X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f345,f635])).
fof(f4890,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4889])).
fof(f4889,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X36)) | sP1(X36,'T_0') | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f4244,f704])).
fof(f4244,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X35)) | p(s(bool,X36)) | sP1(X36,'T_0') | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f336,f635])).
fof(f4886,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4885])).
fof(f4885,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1('T_0',X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4243,f704])).
fof(f4243,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1('T_0',X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f333,f635])).
fof(f3975,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl14),
  inference(cnf_transformation,[],[f3975_D])).
fof(f3975_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f4874,plain,(
  ( ! [X108] : (s(bool,X108) != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X108)) ) | $spl20),
  inference(cnf_transformation,[],[f4874_D])).
fof(f4874_D,plain,(
  ( ! [X108] : (s(bool,X108) != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X108)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f4171,plain,(
  ( ! [X108] : (s(bool,X108) != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X108)) ) | $spl5),
  inference(superposition,[],[f2605,f635])).
fof(f3920,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl10),
  inference(cnf_transformation,[],[f3920_D])).
fof(f3920_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f4873,plain,(
  ( ! [X107] : (s(bool,X107) != s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | s(bool,'T_0') = s(bool,X107)) ) | $spl18),
  inference(cnf_transformation,[],[f4873_D])).
fof(f4873_D,plain,(
  ( ! [X107] : (s(bool,X107) != s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | s(bool,'T_0') = s(bool,X107)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4170,plain,(
  ( ! [X107] : (s(bool,X107) != s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | s(bool,'T_0') = s(bool,X107)) ) | $spl5),
  inference(superposition,[],[f2603,f635])).
fof(f4076,plain,(
  ( ! [X108] : (s(bool,X108) != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,X108) | s(bool,'T_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12))))) ) | $spl5),
  inference(superposition,[],[f2605,f635])).
fof(f4075,plain,(
  ( ! [X107] : (s(bool,X107) != s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | s(bool,'T_0') = s(bool,X107) | s(bool,'T_0') = s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12))))) ) | $spl5),
  inference(superposition,[],[f2603,f635])).
fof(f3977,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl17),
  inference(cnf_transformation,[],[f3977_D])).
fof(f3977_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f3973,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK17(sK14,'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),bool,sK12)))) | $spl5),
  inference(superposition,[],[f2605,f251])).
fof(f3922,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl13),
  inference(cnf_transformation,[],[f3922_D])).
fof(f3922_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f3918,plain,(
  s(bool,'F_0') != s(bool,f(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | s(bool,'T_0') = s(bool,g(s(sK12,sK17('const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')),sK14,bool,sK12)))) | $spl5),
  inference(superposition,[],[f2603,f251])).
fof(f2614,plain,(
  $false | ($spl3 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f2613,f244])).
fof(f2613,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f2611,f2432])).
fof(f2432,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK14))) | $spl6),
  inference(unit_resulting_resolution,[],[f244,f432,f311])).
fof(f2611,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f428,f433])).
fof(f2254,plain,(
  $false | ($spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f2253,f432])).
fof(f2253,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK14)))) | ($spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f428,f1762])).
fof(f1762,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))) | s(sK12,sK13) = s(sK12,X0)) ) | $spl4),
  inference(subsumption_resolution,[],[f1761,f434])).
fof(f1761,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))) | s(sK12,sK13) = s(sK12,X0)) ) | $spl4),
  inference(forward_demodulation,[],[f1743,f505])).
fof(f1743,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),sK14)))) | s(sK12,sK13) = s(sK12,X0) | p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))) ) | $spl4),
  inference(superposition,[],[f305,f429])).
fof(f2252,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f2251,f434])).
fof(f2251,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f2215,f473])).
fof(f473,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))) | $spl1),
  inference(unit_resulting_resolution,[],[f426,f264])).
fof(f2215,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | $spl4),
  inference(superposition,[],[f1427,f429])).
fof(f425,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | $spl0),
  inference(cnf_transformation,[],[f425_D])).
fof(f425_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f429,plain,(
  s(fun(sK12,bool),sK14) = s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | $spl4),
  inference(cnf_transformation,[],[f429_D])).
fof(f429_D,plain,(
  s(fun(sK12,bool),sK14) = s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f240,plain,(
  s(fun(sK12,bool),sK14) = s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))))),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ((p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK14)))) | s(sK12,sK13) = s(sK12,X4))) | s(fun(sK12,bool),sK14) = s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0')))) & (~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | (p(s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK14)))) & s(sK12,sK13) != s(sK12,sK15)) | s(fun(sK12,bool),sK14) != s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f159])).
fof(f159,plain,(
  ? [X0,X1,X2] : (((p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) | s(X0,X1) = s(X0,X4))) | s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & s(X0,X1) != s(X0,X3)) | s(fun(X0,bool),X2) != s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))))),
  inference(rectify,[],[f158])).
fof(f158,plain,(
  ? [X0,X1,X2] : (((p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) & ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | s(X0,X1) = s(X0,X3))) | s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & s(X0,X1) != s(X0,X3)) | s(fun(X0,bool),X2) != s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ? [X0,X1,X2] : (((p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) & ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | s(X0,X1) = s(X0,X3))) | s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & s(X0,X1) != s(X0,X3))) | s(fun(X0,bool),X2) != s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))))),
  inference(nnf_transformation,[],[f111])).
fof(f111,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) & ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | s(X0,X1) = s(X0,X3))) <~> s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ~! [X0,X1,X2] : ((p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) & ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) => s(X0,X1) = s(X0,X3))) <=> s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ~! [X0] : ! [X1,X2] : ((p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) & ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) => s(X0,X1) = s(X0,X3))) <=> s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ((p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8)))) & ! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8)))) => s(X6,X7) = s(X6,X9))) <=> s(fun(X6,bool),X8) = s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8)))) & ! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8)))) => s(X6,X7) = s(X6,X9))) <=> s(fun(X6,bool),X8) = s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.UNIQUE_MEMBER_SING','thm.pred_set.UNIQUE_MEMBER_SING')).
fof(f433,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK14)))) | s(sK12,sK13) = s(sK12,X4)) ) | $spl8),
  inference(cnf_transformation,[],[f433_D])).
fof(f433_D,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK14)))) | s(sK12,sK13) = s(sK12,X4)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f241,plain,(
  ( ! [X4] : (s(fun(sK12,bool),sK14) = s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | s(sK12,sK13) = s(sK12,X4) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),sK14))))) )),
  inference(cnf_transformation,[],[f160])).
fof(f426,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) | $spl1),
  inference(cnf_transformation,[],[f426_D])).
fof(f426_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f432,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK14)))) | $spl6),
  inference(cnf_transformation,[],[f432_D])).
fof(f432_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK14)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f430,plain,(
  s(fun(sK12,bool),sK14) != s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | $spl5),
  inference(cnf_transformation,[],[f430_D])).
fof(f430_D,plain,(
  s(fun(sK12,bool),sK14) != s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f242,plain,(
  s(fun(sK12,bool),sK14) != s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK15),s(fun(sK12,bool),sK14)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))))),
  inference(cnf_transformation,[],[f160])).
fof(f428,plain,(
  s(sK12,sK13) != s(sK12,sK15) | $spl3),
  inference(cnf_transformation,[],[f428_D])).
fof(f428_D,plain,(
  s(sK12,sK13) != s(sK12,sK15) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f243,plain,(
  s(fun(sK12,bool),sK14) != s(fun(sK12,bool),'const.pred_set.INSERT_2'(s(sK12,sK13),s(fun(sK12,bool),'const.pred_set.EMPTY_0'))) | s(sK12,sK13) != s(sK12,sK15) | ~p(s(bool,'const.bool.IN_2'(s(sK12,sK13),s(fun(sK12,bool),sK14))))),
  inference(cnf_transformation,[],[f160])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1057
% Passive clauses: 2656
% Generated clauses: 26897
% Final active clauses: 525
% Final passive clauses: 448
% Input formulas: 44
% Initial clauses: 176
% 
% Duplicate literals: 355
% Trivial inequalities: 133
% Fw subsumption resolutions: 1631
% Fw demodulations: 3813
% Bw demodulations: 1209
% Condensations: 649
% Global subsumptions: 38
% 
% Simple tautologies: 155
% Equational tautologies: 168
% Forward subsumptions: 15938
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 7
% 
% Unit resulting resolution: 9512
% Factoring: 1
% Forward superposition: 2690
% Backward superposition: 5172
% Self superposition: 42
% Equality factoring: 38
% Equality resolution: 29
% 
% Split clauses: 100
% Split components: 209
% SAT solver clauses: 4155
% SAT solver unit clauses: 603
% SAT solver binary clauses: 1810
% 
% Sat splits: 100
% Sat splitting refutations: 1259
% 
% Memory used [KB]: 11001
% Time elapsed: 0.784 s
% ------------------------------
% ------------------------------
% Success in time 2.102 s
