Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numRing
% (87907)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 159
% Passive clauses: 3945
% Generated clauses: 20167
% Final active clauses: 151
% Final passive clauses: 3652
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 553
% Trivial inequalities: 331
% Fw subsumption resolutions: 940
% Fw demodulations: 1182
% Bw demodulations: 4
% 
% Simple tautologies: 328
% Equational tautologies: 22
% Forward subsumptions: 12518
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 31
% Forward superposition: 9479
% Backward superposition: 7216
% Self superposition: 211
% Equality factoring: 54
% Equality resolution: 28
% 
% Split clauses: 228
% Split components: 456
% SAT solver clauses: 452
% SAT solver unit clauses: 5
% SAT solver binary clauses: 100
% 
% Sat splits: 228
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 452
% TWLsolver calls for satisfiability: 58
% 
% Memory used [KB]: 20596
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.numRing
% (87935)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 144
% Passive clauses: 1569
% Generated clauses: 17846
% Final active clauses: 132
% Final passive clauses: 1281
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 631
% Trivial inequalities: 211
% Fw subsumption resolutions: 592
% Fw demodulations: 1061
% Bw demodulations: 121
% Condensations: 934
% 
% Simple tautologies: 185
% Equational tautologies: 26
% Forward subsumptions: 9935
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 26
% 
% Binary resolution: 43
% Forward superposition: 6931
% Backward superposition: 5287
% Self superposition: 221
% Equality factoring: 76
% Equality resolution: 20
% 
% Split clauses: 118
% Split components: 236
% SAT solver clauses: 2340
% SAT solver unit clauses: 235
% SAT solver binary clauses: 51
% 
% Sat splits: 118
% Sat splitting refutations: 30
% 
% TWLsolver clauses: 2340
% TWLsolver calls for satisfiability: 278
% 
% Memory used [KB]: 24690
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.numRing
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numRing
% SZS output start Proof for thm.numRing
fof(f133641,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f133489])).
fof(f133489,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')),
  inference(superposition,[],[f307,f133176])).
fof(f133176,plain,(
  ( ! [X2] : (s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_prod_2'(s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'('type.num.num'),X2)))) )),
  inference(superposition,[],[f30459,f336])).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X0),X2)))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X7] : ! [X12,X6] : s('type.canonical.canonical_sum'(X7),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X7),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X7),X12),s('type.canonical.canonical_sum'(X7),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'(X7),X6)))),
  file('thm.numRing.num_ring_thms|split|34','thm.canonical.canonical_sum_prod_def|split|2')).
fof(f30459,plain,(
  ( ! [X0,X1] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_prod_2'(s('type.canonical.canonical_sum'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.canonical.canonical_sum'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1)))) )),
  inference(forward_demodulation,[],[f30416,f422])).
fof(f422,plain,(
  ( ! [X0,X1] : (s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_prod_2'(s('type.canonical.canonical_sum'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_prod_0'),s('type.canonical.canonical_sum'('type.num.num'),X0))),s('type.canonical.canonical_sum'('type.num.num'),X1)))) )),
  inference(cnf_transformation,[],[f288])).
fof(f288,plain,(
  ! [X0,X1] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_prod_2'(s('type.canonical.canonical_sum'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_prod_0'),s('type.canonical.canonical_sum'('type.num.num'),X0))),s('type.canonical.canonical_sum'('type.num.num'),X1)))),
  inference(rectify,[],[f102])).
fof(f102,axiom,(
  ! [X45,X46] : s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_prod_2'(s('type.canonical.canonical_sum'('type.num.num'),X45),s('type.canonical.canonical_sum'('type.num.num'),X46))) = s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_prod_0'),s('type.canonical.canonical_sum'('type.num.num'),X45))),s('type.canonical.canonical_sum'('type.num.num'),X46)))),
  file('thm.numRing.num_ring_thms|split|34','ARITY_0_2_const.numRing.num_canonical_sum_prod')).
fof(f30416,plain,(
  ( ! [X0,X1] : (s('type.canonical.canonical_sum'('type.num.num'),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num')),happ(s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_prod_0'),s('type.canonical.canonical_sum'('type.num.num'),X0))),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.canonical.canonical_sum'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1)))) )),
  inference(superposition,[],[f366,f320])).
fof(f320,plain,(
  s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_prod_0') = s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  inference(cnf_transformation,[],[f68])).
fof(f68,axiom,(
  s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.numRing.num_canonical_sum_prod_0') = s(fun('type.canonical.canonical_sum'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.canonical.canonical_sum'('type.num.num'))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  file('thm.numRing.num_ring_thms|split|34','thm.numRing.num_canonical_sum_prod_def')).
fof(f366,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),happ(s(fun('type.canonical.canonical_sum'(X0),fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X7] : ! [X45,X46,X47] : s('type.canonical.canonical_sum'(X7),'const.canonical.canonical_sum_prod_3'(s('type.semi_ring.semi_ring'(X7),X45),s('type.canonical.canonical_sum'(X7),X46),s('type.canonical.canonical_sum'(X7),X47))) = s('type.canonical.canonical_sum'(X7),happ(s(fun('type.canonical.canonical_sum'(X7),'type.canonical.canonical_sum'(X7)),happ(s(fun('type.canonical.canonical_sum'(X7),fun('type.canonical.canonical_sum'(X7),'type.canonical.canonical_sum'(X7))),'const.canonical.canonical_sum_prod_1'(s('type.semi_ring.semi_ring'(X7),X45))),s('type.canonical.canonical_sum'(X7),X46))),s('type.canonical.canonical_sum'(X7),X47)))),
  file('thm.numRing.num_ring_thms|split|34','ARITY_1_3_const.canonical.canonical_sum_prod')).
fof(f307,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_prod_2'(s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'('type.num.num'),sK0)))),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_prod_2'(s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'('type.num.num'),sK0)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f298])).
fof(f298,plain,(
  ? [X0] : s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_prod_2'(s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'('type.num.num'),X0)))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ~! [X0] : s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_prod_2'(s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'('type.num.num'),X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_prod_2'(s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'('type.num.num'),X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'('type.num.num'),'const.numRing.num_canonical_sum_prod_2'(s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0'),s('type.canonical.canonical_sum'('type.num.num'),X6)))),
  file('thm.numRing.num_ring_thms|split|34','thm.numRing.num_ring_thms|split|34')).
% SZS output end Proof for thm.numRing
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 564
% Passive clauses: 10354
% Generated clauses: 136211
% Final active clauses: 241
% Final passive clauses: 1613
% Input formulas: 118
% Initial clauses: 123
% Discarded non-redundant clauses: 6761
% 
% Duplicate literals: 3761
% Trivial inequalities: 1233
% Fw subsumption resolutions: 3527
% Fw demodulations: 15539
% Bw demodulations: 2110
% Condensations: 5227
% Global subsumptions: 207
% 
% Simple tautologies: 633
% Equational tautologies: 271
% Forward subsumptions: 94024
% Fw demodulations to eq. taut.: 25
% Bw demodulations to eq. taut.: 212
% 
% Unit resulting resolution: 2807
% Forward superposition: 50832
% Backward superposition: 45506
% Self superposition: 780
% Equality factoring: 233
% Equality resolution: 51
% 
% Split clauses: 343
% Split components: 691
% SAT solver clauses: 11483
% SAT solver unit clauses: 1271
% SAT solver binary clauses: 2790
% 
% Sat splits: 343
% Sat splitting refutations: 234
% 
% Memory used [KB]: 162982
% Time elapsed: 7.661 s
% ------------------------------
% ------------------------------
% Success in time 8.998 s
