Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.sum
% (50989)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 140
% Passive clauses: 3725
% Generated clauses: 14602
% Final active clauses: 136
% Final passive clauses: 3419
% Input formulas: 37
% Initial clauses: 88
% 
% Duplicate literals: 558
% Trivial inequalities: 228
% Fw subsumption resolutions: 1706
% Fw demodulations: 1147
% Bw demodulations: 1
% 
% Simple tautologies: 291
% Equational tautologies: 42
% Forward subsumptions: 6306
% 
% Binary resolution: 868
% Factoring: 2
% Forward superposition: 5728
% Backward superposition: 3855
% Self superposition: 146
% Equality factoring: 98
% Equality resolution: 31
% 
% Split clauses: 471
% Split components: 982
% SAT solver clauses: 768
% SAT solver unit clauses: 28
% SAT solver binary clauses: 184
% 
% Sat splits: 471
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 766
% TWLsolver calls for satisfiability: 190
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.sum
% (51009)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 127
% Passive clauses: 1076
% Generated clauses: 12888
% Final active clauses: 101
% Final passive clauses: 873
% Input formulas: 37
% Initial clauses: 88
% 
% Duplicate literals: 292
% Trivial inequalities: 221
% Fw subsumption resolutions: 898
% Fw demodulations: 2045
% Bw demodulations: 68
% Condensations: 885
% Global subsumptions: 8
% 
% Simple tautologies: 151
% Equational tautologies: 29
% Forward subsumptions: 5186
% Fw demodulations to eq. taut.: 28
% Bw demodulations to eq. taut.: 18
% 
% Binary resolution: 124
% Factoring: 1
% Forward superposition: 2211
% Backward superposition: 4643
% Self superposition: 149
% Equality factoring: 69
% Equality resolution: 39
% 
% Split clauses: 154
% Split components: 311
% SAT solver clauses: 2119
% SAT solver unit clauses: 107
% SAT solver binary clauses: 70
% SAT solver learnt clauses: 16
% 
% Sat splits: 154
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 2087
% TWLsolver calls for satisfiability: 620
% 
% Memory used [KB]: 13304
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.sum
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.sum
% SZS output start Proof for thm.sum
fof(f9869,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f417,f421,f419,f709,f714,f711,f2305,f712,f2396,f2567,f2568,f3342,f3371,f3369,f3364,f3374,f3372,f9166,f9170,f9168,f9164,f9180,f9178,f9224,f9228,f9226,f9266,f9270,f9268,f9324,f9328,f9326,f9329,f9333,f9331,f9291,f9338,f9337,f9396,f9400,f9398,f9441,f9454,f8499,f9459,f8496,f9461,f8495,f9462,f8494,f9460,f9463,f8487,f9465,f8480,f9469,f8537,f8536,f8524,f9466,f9472,f9474,f9470,f2569,f9581,f9456,f9582,f9698,f9769,f9855])).
fof(f9855,plain,(
  $false | $spl28),
  inference(trivial_inequality_removal,[],[f9854])).
fof(f9854,plain,(
  ( ! [X0] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(sK2,sK1,sK0)))))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(sK2,sK1,sK0))))))) ) | $spl28),
  inference(forward_demodulation,[],[f9853,f9783])).
fof(f9783,plain,(
  s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(sK2,sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(sK2,sK1,sK0)))) | $spl28),
  inference(global_subsumption,[],[f156,f155,f154,f153,f259,f260,f157,f262,f163,f166,f165,f170,f169,f193,f194,f196,f195,f204,f201,f200,f205,f206,f207,f209,f208,f210,f211,f214,f213,f212,f218,f217,f216,f215,f226,f225,f227,f228,f232,f231,f230,f229,f263,f272,f274,f306,f325,f351,f363,f417,f422,f423,f425,f460,f461,f458,f522,f523,f526,f527,f635,f637,f769,f890,f1033,f1079,f1060,f1189,f1241,f1219,f1313,f1314,f1315,f1316,f1322,f1324,f1330,f1331,f1332,f1337,f1339,f1472,f1479,f1542,f1515,f1634,f1641,f1709,f1680,f1857,f1976,f1977,f1981,f1982,f1984,f1985,f1986,f1987,f1988,f1989,f1990,f2009,f2010,f2013,f2019,f2020,f2022,f2023,f2024,f2025,f2026,f2027,f2028,f2049,f2051,f2690,f2692,f2694,f2696,f2698,f2700,f2703,f2831,f2832,f2833,f2834,f2835,f2836,f2837,f2838,f2839,f2840,f2841,f2842,f2844,f2845,f2846,f2847,f2848,f2849,f2850,f2851,f2853,f2854,f3123,f3124,f3125,f3126,f3127,f3128,f3266,f3267,f3268,f3269,f3270,f3271,f3387,f3388,f3389,f3390,f3391,f3392,f3463,f3464,f3482,f3483,f3484,f3485,f3486,f3487,f3472,f3473,f3493,f3496,f3499,f3502,f3505,f3508,f3480,f3665,f3666,f3667,f3668,f3669,f3670,f3673,f3674,f3675,f3676,f3677,f3678,f3664,f4033,f4034,f4035,f4036,f4037,f4038,f4042,f4043,f4044,f4045,f4046,f4047,f4032,f4171,f4172,f4173,f4174,f4175,f4176,f4180,f4181,f4182,f4183,f4184,f4185,f4170,f5193,f5194,f5196,f5197,f5198,f5199,f5200,f5201,f5235,f5259,f5260,f5261,f5262,f5263,f5264,f5265,f5266,f5396,f5398,f5402,f5403,f5413,f5417,f5418,f5419,f5420,f5421,f5422,f5427,f5428,f5429,f5430,f5432,f5433,f5434,f5435,f5436,f5437,f5578,f5622,f6373,f6374,f6375,f6376,f6378,f6379,f6380,f6381,f6382,f6383,f5729,f5731,f5735,f5736,f5746,f5750,f5751,f5752,f5753,f5754,f5755,f5760,f5761,f5762,f5763,f5764,f5765,f5766,f5767,f5768,f5769,f5770,f6775,f6776,f6777,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6962,f6803,f6805,f6809,f6810,f6821,f6822,f6823,f6824,f6825,f6826,f6828,f6829,f6831,f6836,f6837,f6860,f6861,f6866,f6867,f6868,f6869,f6870,f6871,f6876,f6877,f6878,f6879,f6880,f6881,f6882,f6883,f6884,f6885,f6886,f6966,f6967,f6968,f6969,f6930,f6932,f6936,f6937,f7117,f7148,f7149,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7299,f7301,f7303,f7304,f7250,f7309,f7256,f7257,f7482,f7483,f7487,f7488,f7489,f7490,f7491,f7492,f7493,f7494,f7578,f7579,f7580,f7593,f7595,f7599,f7600,f7612,f7613,f7614,f7615,f7616,f7617,f7619,f7620,f7622,f7627,f7628,f7629,f7656,f7657,f7662,f7663,f7664,f7665,f7666,f7667,f7668,f7673,f7674,f7675,f7676,f7677,f7678,f7682,f7683,f7684,f7685,f7992,f8048,f8049,f8055,f8056,f8057,f8058,f8059,f8060,f8061,f8062,f8178,f8123,f8129,f8130,f8463,f8492,f8493,f8500,f8501,f8502,f8503,f8504,f8505,f8506,f8507,f8508,f9166,f9164,f9181,f9224,f9188,f9189,f9190,f9191,f9192,f9231,f9234,f9235,f9266,f9350,f9178,f9396,f9359,f9360,f9361,f9362,f9363,f9403,f9406,f9407,f8496,f8495,f8494,f8524,f8402,f8412,f9472,f9474,f5570,f9476,f3490,f3481,f9163,f9263,f9162,f9262,f9176,f9276,f9175,f9275,f9174,f9274,f9173,f9273,f9154,f9254,f9172,f9272,f9171,f9271,f9155,f9255,f7686,f9629,f9630,f9622,f9627,f9626,f9625,f9623,f261,f2843,f9264])).
fof(f9264,plain,(
  s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(sK2,sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(sK2,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))))),
  inference(equality_resolution,[],[f1330])).
fof(f2843,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X2),s(fun(sK1,bool),X3),s('type.sum.sum'(sK0,sK1),sK2)))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X0),s(fun(sK1,bool),X1),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK15(X0,X2,X1,X3,sK3,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK14(X0,X2,X1,X3,sK3,sK1,sK0)))))) )),
  inference(forward_demodulation,[],[f2815,f153])).
fof(f2815,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X0),s(fun(sK1,bool),X1),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X2),s(fun(sK1,bool),X3),s('type.sum.sum'(sK0,sK1),sK3)))) | p(s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK15(X0,X2,X1,X3,sK3,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK14(X0,X2,X1,X3,sK3,sK1,sK0)))))) )),
  inference(superposition,[],[f215,f153])).
fof(f261,plain,(
  s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK4),s(fun(sK1,bool),sK6),s('type.sum.sum'(sK0,sK1),sK2))) != s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2)))),
  inference(backward_demodulation,[],[f153,f156])).
fof(f9623,plain,(
  ( ! [X2,X0,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2))))) != s('type.sum.sum'(X0,bool),'const.bool.COND_3'(s(bool,X1),s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK4),s(fun(sK1,bool),sK6),s('type.sum.sum'(sK0,sK1),sK2))))),s('type.sum.sum'(X0,bool),'const.sum.INL_1'(s(X0,X2)))))) )),
  inference(unit_resulting_resolution,[],[f261,f213])).
fof(f9625,plain,(
  ( ! [X0,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INL_1'(s(X0,X1))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK4),s(fun(sK1,bool),sK6),s('type.sum.sum'(sK0,sK1),sK2)))))) )),
  inference(unit_resulting_resolution,[],[f261,f2051])).
fof(f9626,plain,(
  ( ! [X2,X0,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK4),s(fun(sK1,bool),sK6),s('type.sum.sum'(sK0,sK1),sK2))))) != s('type.sum.sum'(X0,bool),'const.bool.COND_3'(s(bool,X1),s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2))))),s('type.sum.sum'(X0,bool),'const.sum.INL_1'(s(X0,X2)))))) )),
  inference(unit_resulting_resolution,[],[f261,f213])).
fof(f9627,plain,(
  ( ! [X0] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK4),s(fun(sK1,bool),sK6),s('type.sum.sum'(sK0,sK1),sK2))))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2)))))) )),
  inference(unit_resulting_resolution,[],[f261,f208])).
fof(f9622,plain,(
  ( ! [X0,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INL_1'(s(X0,X1))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2)))))) )),
  inference(unit_resulting_resolution,[],[f261,f2051])).
fof(f9630,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK4),s(fun(sK1,bool),sK6),s('type.sum.sum'(sK0,sK1),sK2))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2)))) )),
  inference(superposition,[],[f261,f272])).
fof(f9629,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK4),s(fun(sK1,bool),sK6),s('type.sum.sum'(sK0,sK1),sK2))) | s(bool,X0) = s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f261,f272])).
fof(f7686,plain,(
  ( ! [X551,X552,X550] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X552))) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,'const.bool.COND_3'(s(bool,X550),s(sK1,X551),s(sK1,X552))))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,'const.bool.COND_3'(s(bool,X550),s(sK1,X551),s(sK1,X552))))) | s(bool,'T_0') = s(bool,X550)) )),
  inference(superposition,[],[f260,f351])).
fof(f9255,plain,(
  ( ! [X1] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X1) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(X1,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(X1,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(X1,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(X1,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),X1) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(X1,sK1,sK0))))) )),
  inference(superposition,[],[f1330,f205])).
fof(f9155,plain,(
  ( ! [X1] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X1) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(X1,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(X1,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(X1,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(X1,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),X1) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(X1,sK1,sK0))))) )),
  inference(superposition,[],[f1313,f205])).
fof(f9271,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s('type.sum.sum'(sK0,sK1),happ(s(fun(X2,'type.sum.sum'(sK0,sK1)),X5),s(X2,X6))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(happ(s(fun('type.sum.sum'(X2,X3),'type.sum.sum'(sK0,sK1)),sK16(X4,X5,X3,'type.sum.sum'(sK0,sK1),X2)),s('type.sum.sum'(X2,X3),'const.sum.INL_1'(s(X2,X6)))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X2,'type.sum.sum'(sK0,sK1)),X5),s(X2,X6))) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(happ(s(fun('type.sum.sum'(X2,X3),'type.sum.sum'(sK0,sK1)),sK16(X4,X5,X3,'type.sum.sum'(sK0,sK1),X2)),s('type.sum.sum'(X2,X3),'const.sum.INL_1'(s(X2,X6)))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(happ(s(fun('type.sum.sum'(X2,X3),'type.sum.sum'(sK0,sK1)),sK16(X4,X5,X3,'type.sum.sum'(sK0,sK1),X2)),s('type.sum.sum'(X2,X3),'const.sum.INL_1'(s(X2,X6)))),sK1,sK0))))) )),
  inference(forward_demodulation,[],[f9256,f225])).
fof(f9256,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X2,'type.sum.sum'(sK0,sK1)),X5),s(X2,X6))) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(happ(s(fun('type.sum.sum'(X2,X3),'type.sum.sum'(sK0,sK1)),sK16(X4,X5,X3,'type.sum.sum'(sK0,sK1),X2)),s('type.sum.sum'(X2,X3),'const.sum.INL_1'(s(X2,X6)))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(happ(s(fun('type.sum.sum'(X2,X3),'type.sum.sum'(sK0,sK1)),sK16(X4,X5,X3,'type.sum.sum'(sK0,sK1),X2)),s('type.sum.sum'(X2,X3),'const.sum.INL_1'(s(X2,X6)))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),happ(s(fun('type.sum.sum'(X2,X3),'type.sum.sum'(sK0,sK1)),sK16(X4,X5,X3,'type.sum.sum'(sK0,sK1),X2)),s('type.sum.sum'(X2,X3),'const.sum.INL_1'(s(X2,X6))))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(happ(s(fun('type.sum.sum'(X2,X3),'type.sum.sum'(sK0,sK1)),sK16(X4,X5,X3,'type.sum.sum'(sK0,sK1),X2)),s('type.sum.sum'(X2,X3),'const.sum.INL_1'(s(X2,X6)))),sK1,sK0))))) )),
  inference(superposition,[],[f1330,f225])).
fof(f9171,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s('type.sum.sum'(sK0,sK1),happ(s(fun(X2,'type.sum.sum'(sK0,sK1)),X5),s(X2,X6))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(happ(s(fun('type.sum.sum'(X2,X3),'type.sum.sum'(sK0,sK1)),sK16(X4,X5,X3,'type.sum.sum'(sK0,sK1),X2)),s('type.sum.sum'(X2,X3),'const.sum.INL_1'(s(X2,X6)))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X2,'type.sum.sum'(sK0,sK1)),X5),s(X2,X6))) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(happ(s(fun('type.sum.sum'(X2,X3),'type.sum.sum'(sK0,sK1)),sK16(X4,X5,X3,'type.sum.sum'(sK0,sK1),X2)),s('type.sum.sum'(X2,X3),'const.sum.INL_1'(s(X2,X6)))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(happ(s(fun('type.sum.sum'(X2,X3),'type.sum.sum'(sK0,sK1)),sK16(X4,X5,X3,'type.sum.sum'(sK0,sK1),X2)),s('type.sum.sum'(X2,X3),'const.sum.INL_1'(s(X2,X6)))),sK1,sK0))))) )),
  inference(forward_demodulation,[],[f9156,f225])).
fof(f9156,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X2,'type.sum.sum'(sK0,sK1)),X5),s(X2,X6))) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(happ(s(fun('type.sum.sum'(X2,X3),'type.sum.sum'(sK0,sK1)),sK16(X4,X5,X3,'type.sum.sum'(sK0,sK1),X2)),s('type.sum.sum'(X2,X3),'const.sum.INL_1'(s(X2,X6)))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(happ(s(fun('type.sum.sum'(X2,X3),'type.sum.sum'(sK0,sK1)),sK16(X4,X5,X3,'type.sum.sum'(sK0,sK1),X2)),s('type.sum.sum'(X2,X3),'const.sum.INL_1'(s(X2,X6)))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),happ(s(fun('type.sum.sum'(X2,X3),'type.sum.sum'(sK0,sK1)),sK16(X4,X5,X3,'type.sum.sum'(sK0,sK1),X2)),s('type.sum.sum'(X2,X3),'const.sum.INL_1'(s(X2,X6))))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(happ(s(fun('type.sum.sum'(X2,X3),'type.sum.sum'(sK0,sK1)),sK16(X4,X5,X3,'type.sum.sum'(sK0,sK1),X2)),s('type.sum.sum'(X2,X3),'const.sum.INL_1'(s(X2,X6)))),sK1,sK0))))) )),
  inference(superposition,[],[f1313,f225])).
fof(f9272,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s('type.sum.sum'(sK0,sK1),happ(s(fun(X8,'type.sum.sum'(sK0,sK1)),X9),s(X8,X11))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(happ(s(fun('type.sum.sum'(X7,X8),'type.sum.sum'(sK0,sK1)),sK16(X9,X10,X8,'type.sum.sum'(sK0,sK1),X7)),s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X11)))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X8,'type.sum.sum'(sK0,sK1)),X9),s(X8,X11))) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(happ(s(fun('type.sum.sum'(X7,X8),'type.sum.sum'(sK0,sK1)),sK16(X9,X10,X8,'type.sum.sum'(sK0,sK1),X7)),s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X11)))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(happ(s(fun('type.sum.sum'(X7,X8),'type.sum.sum'(sK0,sK1)),sK16(X9,X10,X8,'type.sum.sum'(sK0,sK1),X7)),s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X11)))),sK1,sK0))))) )),
  inference(forward_demodulation,[],[f9257,f226])).
fof(f9257,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X8,'type.sum.sum'(sK0,sK1)),X9),s(X8,X11))) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(happ(s(fun('type.sum.sum'(X7,X8),'type.sum.sum'(sK0,sK1)),sK16(X9,X10,X8,'type.sum.sum'(sK0,sK1),X7)),s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X11)))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(happ(s(fun('type.sum.sum'(X7,X8),'type.sum.sum'(sK0,sK1)),sK16(X9,X10,X8,'type.sum.sum'(sK0,sK1),X7)),s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X11)))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),happ(s(fun('type.sum.sum'(X7,X8),'type.sum.sum'(sK0,sK1)),sK16(X9,X10,X8,'type.sum.sum'(sK0,sK1),X7)),s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X11))))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(happ(s(fun('type.sum.sum'(X7,X8),'type.sum.sum'(sK0,sK1)),sK16(X9,X10,X8,'type.sum.sum'(sK0,sK1),X7)),s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X11)))),sK1,sK0))))) )),
  inference(superposition,[],[f1330,f226])).
fof(f9172,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s('type.sum.sum'(sK0,sK1),happ(s(fun(X8,'type.sum.sum'(sK0,sK1)),X9),s(X8,X11))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(happ(s(fun('type.sum.sum'(X7,X8),'type.sum.sum'(sK0,sK1)),sK16(X9,X10,X8,'type.sum.sum'(sK0,sK1),X7)),s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X11)))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X8,'type.sum.sum'(sK0,sK1)),X9),s(X8,X11))) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(happ(s(fun('type.sum.sum'(X7,X8),'type.sum.sum'(sK0,sK1)),sK16(X9,X10,X8,'type.sum.sum'(sK0,sK1),X7)),s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X11)))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(happ(s(fun('type.sum.sum'(X7,X8),'type.sum.sum'(sK0,sK1)),sK16(X9,X10,X8,'type.sum.sum'(sK0,sK1),X7)),s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X11)))),sK1,sK0))))) )),
  inference(forward_demodulation,[],[f9157,f226])).
fof(f9157,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X8,'type.sum.sum'(sK0,sK1)),X9),s(X8,X11))) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(happ(s(fun('type.sum.sum'(X7,X8),'type.sum.sum'(sK0,sK1)),sK16(X9,X10,X8,'type.sum.sum'(sK0,sK1),X7)),s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X11)))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(happ(s(fun('type.sum.sum'(X7,X8),'type.sum.sum'(sK0,sK1)),sK16(X9,X10,X8,'type.sum.sum'(sK0,sK1),X7)),s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X11)))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),happ(s(fun('type.sum.sum'(X7,X8),'type.sum.sum'(sK0,sK1)),sK16(X9,X10,X8,'type.sum.sum'(sK0,sK1),X7)),s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X11))))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(happ(s(fun('type.sum.sum'(X7,X8),'type.sum.sum'(sK0,sK1)),sK16(X9,X10,X8,'type.sum.sum'(sK0,sK1),X7)),s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X11)))),sK1,sK0))))) )),
  inference(superposition,[],[f1313,f226])).
fof(f9254,plain,(
  ( ! [X0] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X0) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(X0,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(X0,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(X0,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(X0,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),X0) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(X0,sK1,sK0))))) )),
  inference(superposition,[],[f1330,f205])).
fof(f9154,plain,(
  ( ! [X0] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X0) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(X0,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(X0,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(X0,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(X0,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),X0) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(X0,sK1,sK0))))) )),
  inference(superposition,[],[f1313,f205])).
fof(f9273,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.sum.sum'(sK0,sK1),happ(s(fun(X13,'type.sum.sum'(sK0,sK1)),X16),s(X13,X14))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,'type.sum.sum'(sK0,sK1)),X15),s(fun(X13,'type.sum.sum'(sK0,sK1)),X16)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X13,'type.sum.sum'(sK0,sK1)),X16),s(X13,X14))) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,'type.sum.sum'(sK0,sK1)),X15),s(fun(X13,'type.sum.sum'(sK0,sK1)),X16)),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,'type.sum.sum'(sK0,sK1)),X15),s(fun(X13,'type.sum.sum'(sK0,sK1)),X16)),sK1,sK0))))) )),
  inference(forward_demodulation,[],[f9258,f227])).
fof(f9258,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X13,'type.sum.sum'(sK0,sK1)),X16),s(X13,X14))) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,'type.sum.sum'(sK0,sK1)),X15),s(fun(X13,'type.sum.sum'(sK0,sK1)),X16)),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,'type.sum.sum'(sK0,sK1)),X15),s(fun(X13,'type.sum.sum'(sK0,sK1)),X16)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,'type.sum.sum'(sK0,sK1)),X15),s(fun(X13,'type.sum.sum'(sK0,sK1)),X16))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,'type.sum.sum'(sK0,sK1)),X15),s(fun(X13,'type.sum.sum'(sK0,sK1)),X16)),sK1,sK0))))) )),
  inference(superposition,[],[f1330,f227])).
fof(f9173,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.sum.sum'(sK0,sK1),happ(s(fun(X13,'type.sum.sum'(sK0,sK1)),X16),s(X13,X14))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,'type.sum.sum'(sK0,sK1)),X15),s(fun(X13,'type.sum.sum'(sK0,sK1)),X16)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X13,'type.sum.sum'(sK0,sK1)),X16),s(X13,X14))) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,'type.sum.sum'(sK0,sK1)),X15),s(fun(X13,'type.sum.sum'(sK0,sK1)),X16)),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,'type.sum.sum'(sK0,sK1)),X15),s(fun(X13,'type.sum.sum'(sK0,sK1)),X16)),sK1,sK0))))) )),
  inference(forward_demodulation,[],[f9158,f227])).
fof(f9158,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X13,'type.sum.sum'(sK0,sK1)),X16),s(X13,X14))) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,'type.sum.sum'(sK0,sK1)),X15),s(fun(X13,'type.sum.sum'(sK0,sK1)),X16)),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,'type.sum.sum'(sK0,sK1)),X15),s(fun(X13,'type.sum.sum'(sK0,sK1)),X16)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,'type.sum.sum'(sK0,sK1)),X15),s(fun(X13,'type.sum.sum'(sK0,sK1)),X16))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,'type.sum.sum'(sK0,sK1)),X15),s(fun(X13,'type.sum.sum'(sK0,sK1)),X16)),sK1,sK0))))) )),
  inference(superposition,[],[f1313,f227])).
fof(f9274,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s('type.sum.sum'(sK0,sK1),happ(s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(X17,X19))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INL_1'(s(X17,X19))),s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(fun(X18,'type.sum.sum'(sK0,sK1)),X21)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(X17,X19))) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INL_1'(s(X17,X19))),s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(fun(X18,'type.sum.sum'(sK0,sK1)),X21)),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INL_1'(s(X17,X19))),s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(fun(X18,'type.sum.sum'(sK0,sK1)),X21)),sK1,sK0))))) )),
  inference(forward_demodulation,[],[f9259,f228])).
fof(f9259,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(X17,X19))) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INL_1'(s(X17,X19))),s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(fun(X18,'type.sum.sum'(sK0,sK1)),X21)),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INL_1'(s(X17,X19))),s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(fun(X18,'type.sum.sum'(sK0,sK1)),X21)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INL_1'(s(X17,X19))),s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(fun(X18,'type.sum.sum'(sK0,sK1)),X21))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INL_1'(s(X17,X19))),s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(fun(X18,'type.sum.sum'(sK0,sK1)),X21)),sK1,sK0))))) )),
  inference(superposition,[],[f1330,f228])).
fof(f9174,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s('type.sum.sum'(sK0,sK1),happ(s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(X17,X19))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INL_1'(s(X17,X19))),s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(fun(X18,'type.sum.sum'(sK0,sK1)),X21)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(X17,X19))) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INL_1'(s(X17,X19))),s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(fun(X18,'type.sum.sum'(sK0,sK1)),X21)),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INL_1'(s(X17,X19))),s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(fun(X18,'type.sum.sum'(sK0,sK1)),X21)),sK1,sK0))))) )),
  inference(forward_demodulation,[],[f9159,f228])).
fof(f9159,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),happ(s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(X17,X19))) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INL_1'(s(X17,X19))),s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(fun(X18,'type.sum.sum'(sK0,sK1)),X21)),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INL_1'(s(X17,X19))),s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(fun(X18,'type.sum.sum'(sK0,sK1)),X21)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INL_1'(s(X17,X19))),s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(fun(X18,'type.sum.sum'(sK0,sK1)),X21))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INL_1'(s(X17,X19))),s(fun(X17,'type.sum.sum'(sK0,sK1)),X20),s(fun(X18,'type.sum.sum'(sK0,sK1)),X21)),sK1,sK0))))) )),
  inference(superposition,[],[f1313,f228])).
fof(f9275,plain,(
  ( ! [X23,X22] : (s('type.sum.sum'(sK0,sK1),X22) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK0,sK1),X22),s('type.sum.sum'(sK0,sK1),X23)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X22) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK0,sK1),X22),s('type.sum.sum'(sK0,sK1),X23)),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK0,sK1),X22),s('type.sum.sum'(sK0,sK1),X23)),sK1,sK0))))) )),
  inference(forward_demodulation,[],[f9260,f193])).
fof(f9260,plain,(
  ( ! [X23,X22] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X22) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK0,sK1),X22),s('type.sum.sum'(sK0,sK1),X23)),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK0,sK1),X22),s('type.sum.sum'(sK0,sK1),X23)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK0,sK1),X22),s('type.sum.sum'(sK0,sK1),X23))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK0,sK1),X22),s('type.sum.sum'(sK0,sK1),X23)),sK1,sK0))))) )),
  inference(superposition,[],[f1330,f193])).
fof(f9175,plain,(
  ( ! [X23,X22] : (s('type.sum.sum'(sK0,sK1),X22) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK0,sK1),X22),s('type.sum.sum'(sK0,sK1),X23)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X22) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK0,sK1),X22),s('type.sum.sum'(sK0,sK1),X23)),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK0,sK1),X22),s('type.sum.sum'(sK0,sK1),X23)),sK1,sK0))))) )),
  inference(forward_demodulation,[],[f9160,f193])).
fof(f9160,plain,(
  ( ! [X23,X22] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X22) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK0,sK1),X22),s('type.sum.sum'(sK0,sK1),X23)),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK0,sK1),X22),s('type.sum.sum'(sK0,sK1),X23)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK0,sK1),X22),s('type.sum.sum'(sK0,sK1),X23))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(sK0,sK1),X22),s('type.sum.sum'(sK0,sK1),X23)),sK1,sK0))))) )),
  inference(superposition,[],[f1313,f193])).
fof(f9276,plain,(
  ( ! [X24,X25] : (s('type.sum.sum'(sK0,sK1),X25) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),X24),s('type.sum.sum'(sK0,sK1),X25)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X25) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),X24),s('type.sum.sum'(sK0,sK1),X25)),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),X24),s('type.sum.sum'(sK0,sK1),X25)),sK1,sK0))))) )),
  inference(forward_demodulation,[],[f9261,f194])).
fof(f9261,plain,(
  ( ! [X24,X25] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X25) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),X24),s('type.sum.sum'(sK0,sK1),X25)),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),X24),s('type.sum.sum'(sK0,sK1),X25)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),X24),s('type.sum.sum'(sK0,sK1),X25))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),X24),s('type.sum.sum'(sK0,sK1),X25)),sK1,sK0))))) )),
  inference(superposition,[],[f1330,f194])).
fof(f9176,plain,(
  ( ! [X24,X25] : (s('type.sum.sum'(sK0,sK1),X25) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),X24),s('type.sum.sum'(sK0,sK1),X25)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X25) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),X24),s('type.sum.sum'(sK0,sK1),X25)),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),X24),s('type.sum.sum'(sK0,sK1),X25)),sK1,sK0))))) )),
  inference(forward_demodulation,[],[f9161,f194])).
fof(f9161,plain,(
  ( ! [X24,X25] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X25) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),X24),s('type.sum.sum'(sK0,sK1),X25)),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),X24),s('type.sum.sum'(sK0,sK1),X25)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),X24),s('type.sum.sum'(sK0,sK1),X25))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(sK0,sK1),X24),s('type.sum.sum'(sK0,sK1),X25)),sK1,sK0))))) )),
  inference(superposition,[],[f1313,f194])).
fof(f9262,plain,(
  ( ! [X28,X26,X27] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X28) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.bool.COND_3'(s(bool,X26),s('type.sum.sum'(sK0,sK1),X27),s('type.sum.sum'(sK0,sK1),X28)),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.bool.COND_3'(s(bool,X26),s('type.sum.sum'(sK0,sK1),X27),s('type.sum.sum'(sK0,sK1),X28)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,X26),s('type.sum.sum'(sK0,sK1),X27),s('type.sum.sum'(sK0,sK1),X28))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.bool.COND_3'(s(bool,X26),s('type.sum.sum'(sK0,sK1),X27),s('type.sum.sum'(sK0,sK1),X28)),sK1,sK0)))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f1330,f351])).
fof(f9162,plain,(
  ( ! [X28,X26,X27] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X28) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.bool.COND_3'(s(bool,X26),s('type.sum.sum'(sK0,sK1),X27),s('type.sum.sum'(sK0,sK1),X28)),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.bool.COND_3'(s(bool,X26),s('type.sum.sum'(sK0,sK1),X27),s('type.sum.sum'(sK0,sK1),X28)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,X26),s('type.sum.sum'(sK0,sK1),X27),s('type.sum.sum'(sK0,sK1),X28))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.bool.COND_3'(s(bool,X26),s('type.sum.sum'(sK0,sK1),X27),s('type.sum.sum'(sK0,sK1),X28)),sK1,sK0)))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f1313,f351])).
fof(f9263,plain,(
  ( ! [X30,X29] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X29) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(sK8(X29,X30,'type.sum.sum'(sK0,sK1)),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(sK8(X29,X30,'type.sum.sum'(sK0,sK1)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK8(X29,X30,'type.sum.sum'(sK0,sK1))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK8(X29,X30,'type.sum.sum'(sK0,sK1)),sK1,sK0)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.sum.sum'(sK0,sK1),bool),X30)))) )),
  inference(superposition,[],[f1330,f425])).
fof(f9163,plain,(
  ( ! [X30,X29] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X29) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK8(X29,X30,'type.sum.sum'(sK0,sK1)),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(sK8(X29,X30,'type.sum.sum'(sK0,sK1)),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK8(X29,X30,'type.sum.sum'(sK0,sK1))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK8(X29,X30,'type.sum.sum'(sK0,sK1)),sK1,sK0)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.sum.sum'(sK0,sK1),bool),X30)))) )),
  inference(superposition,[],[f1313,f425])).
fof(f3481,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK2),s(fun(sK0,X5),X2),s(fun(sK1,X5),X1))) = s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X0),s(fun(sK0,X5),X3),s(fun(sK1,X5),X4))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X0) | s('type.sum.sum'(sK0,sK1),X0) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X1,X2,X0,sK3,X3,X4,X5,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),X0) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X1,X2,X0,sK3,X3,X4,X5,sK1,sK0))))) )),
  inference(forward_demodulation,[],[f3462,f153])).
fof(f3462,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X0) | s('type.sum.sum'(sK0,sK1),X0) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X1,X2,X0,sK3,X3,X4,X5,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),X0) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X1,X2,X0,sK3,X3,X4,X5,sK1,sK0)))) | s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK3),s(fun(sK0,X5),X2),s(fun(sK1,X5),X1))) = s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X0),s(fun(sK0,X5),X3),s(fun(sK1,X5),X4)))) )),
  inference(superposition,[],[f229,f153])).
fof(f3490,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK2),s(fun(sK0,X5),X3),s(fun(sK1,X5),X4))) = s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X0),s(fun(sK0,X5),X2),s(fun(sK1,X5),X1))) | s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X1,X2,sK3,X0,X3,X4,X5,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X1,X2,sK3,X0,X3,X4,X5,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X0)) )),
  inference(forward_demodulation,[],[f3489,f153])).
fof(f3489,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X1,X2,sK3,X0,X3,X4,X5,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X1,X2,sK3,X0,X3,X4,X5,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X0) | s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK3),s(fun(sK0,X5),X3),s(fun(sK1,X5),X4))) = s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X0),s(fun(sK0,X5),X2),s(fun(sK1,X5),X1)))) )),
  inference(forward_demodulation,[],[f3488,f153])).
fof(f3488,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X1,X2,sK3,X0,X3,X4,X5,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X0) | s('type.sum.sum'(sK0,sK1),sK3) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X1,X2,sK3,X0,X3,X4,X5,sK1,sK0)))) | s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK3),s(fun(sK0,X5),X3),s(fun(sK1,X5),X4))) = s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X0),s(fun(sK0,X5),X2),s(fun(sK1,X5),X1)))) )),
  inference(forward_demodulation,[],[f3471,f153])).
fof(f3471,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X0) | s('type.sum.sum'(sK0,sK1),sK3) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X1,X2,sK3,X0,X3,X4,X5,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK3) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X1,X2,sK3,X0,X3,X4,X5,sK1,sK0)))) | s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK3),s(fun(sK0,X5),X3),s(fun(sK1,X5),X4))) = s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X0),s(fun(sK0,X5),X2),s(fun(sK1,X5),X1)))) )),
  inference(superposition,[],[f229,f153])).
fof(f9476,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.sum.sum'(X1,X2),'const.sum.INR_1'(s(X2,X3))) != s('type.sum.sum'(X1,X2),'const.bool.COND_3'(s(bool,X0),s('type.sum.sum'(X1,X2),'const.sum.INR_1'(s(X2,X4))),s('type.sum.sum'(X1,X2),'const.sum.INL_1'(s(X1,X5)))))) )),
  inference(condensation,[],[f9475])).
fof(f9475,plain,(
  ( ! [X28,X26,X24,X23,X27,X25,X22] : (s('type.sum.sum'(X24,X25),'const.sum.INR_1'(s(X25,X26))) != s('type.sum.sum'(X24,X25),'const.bool.COND_3'(s(bool,X22),s('type.sum.sum'(X24,X25),'const.sum.INR_1'(s(X25,X27))),s('type.sum.sum'(X24,X25),'const.sum.INL_1'(s(X24,X28))))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f5447,f166])).
fof(f5447,plain,(
  ( ! [X28,X26,X24,X23,X27,X25,X22] : (p(s(bool,X23)) | s('type.sum.sum'(X24,X25),'const.sum.INR_1'(s(X25,X26))) != s('type.sum.sum'(X24,X25),'const.bool.COND_3'(s(bool,X22),s('type.sum.sum'(X24,X25),'const.sum.INR_1'(s(X25,X27))),s('type.sum.sum'(X24,X25),'const.sum.INL_1'(s(X24,X28))))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f212,f272])).
fof(f5570,plain,(
  ( ! [X478,X480,X477,X479] : (s(X479,f(s(bool,X480))) != s(X479,g(s(bool,X480))) | s(fun(bool,X479),X477) = s(fun(bool,X479),X478) | s(bool,'T_0') = s(bool,sK13(X477,X478,X479,bool)) | s(bool,'T_0') = s(bool,X480)) )),
  inference(superposition,[],[f207,f272])).
fof(f8412,plain,(
  ( ! [X0,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X0,bool),'const.sum.INL_1'(s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f263,f2051])).
fof(f8402,plain,(
  ( ! [X0,X1] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'F_0'))) != s('type.sum.sum'(X0,bool),'const.sum.INL_1'(s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f263,f2051])).
fof(f9407,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK2),s(fun(sK0,X24),X21),s(fun(sK1,X24),X20))) = s(X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X19),s(fun(sK0,X24),X22),s(fun(sK1,X24),X23))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X19) | s('type.sum.sum'(sK0,sK1),X19) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X20,X21,X19,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X22,X23,X24,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),X19) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X20,X21,X19,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X22,X23,X24,sK1,sK0))))) ) | $spl28),
  inference(forward_demodulation,[],[f9365,f9178])).
fof(f9365,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X19) | s('type.sum.sum'(sK0,sK1),X19) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X20,X21,X19,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X22,X23,X24,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),X19) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X20,X21,X19,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X22,X23,X24,sK1,sK0)))) | s(X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X19),s(fun(sK0,X24),X22),s(fun(sK1,X24),X23))) = s(X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))),s(fun(sK0,X24),X21),s(fun(sK1,X24),X20)))) ) | $spl28),
  inference(superposition,[],[f229,f9178])).
fof(f9406,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK2),s(fun(sK0,X18),X16),s(fun(sK1,X18),X17))) = s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X13),s(fun(sK0,X18),X15),s(fun(sK1,X18),X14))) | s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X13)) ) | $spl28),
  inference(forward_demodulation,[],[f9405,f9178])).
fof(f9405,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X13) | s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X13),s(fun(sK0,X18),X15),s(fun(sK1,X18),X14))) = s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))),s(fun(sK0,X18),X16),s(fun(sK1,X18),X17)))) ) | $spl28),
  inference(forward_demodulation,[],[f9404,f9178])).
fof(f9404,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X13) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X13),s(fun(sK0,X18),X15),s(fun(sK1,X18),X14))) = s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))),s(fun(sK0,X18),X16),s(fun(sK1,X18),X17)))) ) | $spl28),
  inference(forward_demodulation,[],[f9364,f9178])).
fof(f9364,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X13) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X13),s(fun(sK0,X18),X15),s(fun(sK1,X18),X14))) = s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))),s(fun(sK0,X18),X16),s(fun(sK1,X18),X17)))) ) | $spl28),
  inference(superposition,[],[f229,f9178])).
fof(f9403,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK9(sK2,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK9(sK2,sK1,sK0))))) | p(s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK15(X0,X2,X1,X3,sK3,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK14(X0,X2,X1,X3,sK3,sK1,sK0)))))) ) | $spl28),
  inference(forward_demodulation,[],[f9402,f9363])).
fof(f9402,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK9(sK2,sK1,sK0))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X2),s(fun(sK1,bool),X3),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK15(X0,X2,X1,X3,sK3,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK14(X0,X2,X1,X3,sK3,sK1,sK0)))))) ) | $spl28),
  inference(backward_demodulation,[],[f9363,f2843])).
fof(f9363,plain,(
  ( ! [X12,X11] : (s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK9(sK2,sK1,sK0)))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X11),s(fun(sK1,bool),X12),s('type.sum.sum'(sK0,sK1),sK2)))) ) | $spl28),
  inference(superposition,[],[f211,f9178])).
fof(f9362,plain,(
  ( ! [X10,X8,X9] : (s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9(sK2,sK1,sK0)))) = s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK2),s(fun(sK0,X8),X9),s(fun(sK1,X8),X10)))) ) | $spl28),
  inference(superposition,[],[f228,f9178])).
fof(f9361,plain,(
  ( ! [X6,X7,X5] : (s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9(sK2,sK1,sK0)))) = s(X5,happ(s(fun('type.sum.sum'(sK0,sK1),X5),sK16(X7,X6,sK1,X5,sK0)),s('type.sum.sum'(sK0,sK1),sK2)))) ) | $spl28),
  inference(superposition,[],[f225,f9178])).
fof(f9360,plain,(
  ( ! [X4,X2,X3] : (s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X2))) != s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,X3),s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X4))),s('type.sum.sum'(sK0,sK1),sK2))) | s(sK1,X2) = s(sK1,X4)) ) | $spl28),
  inference(superposition,[],[f213,f9178])).
fof(f9359,plain,(
  ( ! [X0,X1] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X0))) | s(sK1,X0) = s(sK1,X1)) ) | $spl28),
  inference(superposition,[],[f2051,f9178])).
fof(f9350,plain,(
  s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(sK2,sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(sK2,sK1,sK0)))) | $spl28),
  inference(global_subsumption,[],[f156,f155,f154,f153,f259,f260,f261,f157,f262,f163,f166,f165,f170,f169,f193,f194,f196,f195,f204,f201,f200,f205,f206,f207,f209,f208,f210,f211,f214,f213,f212,f218,f217,f216,f215,f226,f225,f227,f228,f232,f231,f230,f229,f263,f272,f274,f306,f325,f351,f363,f417,f422,f423,f425,f460,f461,f458,f522,f523,f526,f527,f635,f637,f769,f890,f1033,f1079,f1060,f1189,f1241,f1219,f1313,f1314,f1315,f1316,f1322,f1324,f1330,f1331,f1332,f1337,f1339,f1472,f1479,f1542,f1515,f1634,f1641,f1709,f1680,f1857,f1976,f1977,f1981,f1982,f1984,f1985,f1986,f1987,f1988,f1989,f1990,f2009,f2010,f2013,f2019,f2020,f2022,f2023,f2024,f2025,f2026,f2027,f2028,f2049,f2051,f2690,f2692,f2694,f2696,f2698,f2700,f2703,f2831,f2832,f2833,f2834,f2835,f2836,f2837,f2838,f2839,f2840,f2841,f2842,f2844,f2845,f2846,f2847,f2848,f2849,f2850,f2851,f2853,f2854,f3123,f3124,f3125,f3126,f3127,f3128,f3266,f3267,f3268,f3269,f3270,f3271,f3387,f3388,f3389,f3390,f3391,f3392,f3481,f3463,f3464,f3482,f3483,f3484,f3485,f3486,f3487,f3490,f3472,f3473,f3493,f3496,f3499,f3502,f3505,f3508,f3480,f3665,f3666,f3667,f3668,f3669,f3670,f3673,f3674,f3675,f3676,f3677,f3678,f3664,f4033,f4034,f4035,f4036,f4037,f4038,f4042,f4043,f4044,f4045,f4046,f4047,f4032,f4171,f4172,f4173,f4174,f4175,f4176,f4180,f4181,f4182,f4183,f4184,f4185,f4170,f5193,f5194,f5196,f5197,f5198,f5199,f5200,f5201,f5235,f5259,f5260,f5261,f5262,f5263,f5264,f5265,f5266,f5396,f5398,f5402,f5403,f5413,f5417,f5418,f5419,f5420,f5421,f5422,f5427,f5428,f5429,f5430,f5432,f5433,f5434,f5435,f5436,f5437,f5578,f5622,f6373,f6374,f6375,f6376,f6378,f6379,f6380,f6381,f6382,f6383,f5729,f5731,f5735,f5736,f5746,f5750,f5751,f5752,f5753,f5754,f5755,f5760,f5761,f5762,f5763,f5764,f5765,f5766,f5767,f5768,f5769,f5770,f6775,f6776,f6777,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6962,f6803,f6805,f6809,f6810,f6821,f6822,f6823,f6824,f6825,f6826,f6828,f6829,f6831,f6836,f6837,f6860,f6861,f6866,f6867,f6868,f6869,f6870,f6871,f6876,f6877,f6878,f6879,f6880,f6881,f6882,f6883,f6884,f6885,f6886,f6966,f6967,f6968,f6969,f6930,f6932,f6936,f6937,f7117,f7148,f7149,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7299,f7301,f7303,f7304,f7250,f7309,f7256,f7257,f7482,f7483,f7487,f7488,f7489,f7490,f7491,f7492,f7493,f7494,f7578,f7579,f7580,f7593,f7595,f7599,f7600,f7612,f7613,f7614,f7615,f7616,f7617,f7619,f7620,f7622,f7627,f7628,f7629,f7656,f7657,f7662,f7663,f7664,f7665,f7666,f7667,f7668,f7673,f7674,f7675,f7676,f7677,f7678,f7682,f7683,f7684,f7685,f7686,f7992,f8048,f8049,f8055,f8056,f8057,f8058,f8059,f8060,f8061,f8062,f8178,f8123,f8129,f8130,f8463,f8492,f8493,f8500,f8501,f8502,f8503,f8504,f8505,f8506,f8507,f8508,f9166,f9154,f9155,f9171,f9172,f9173,f9174,f9175,f9176,f9162,f9163,f9164,f9178,f9181,f9224,f9188,f9189,f9190,f9191,f9192,f9231,f9234,f9235,f9266,f9254,f9255,f9271,f9272,f9273,f9274,f9275,f9276,f9262,f9263,f2843,f9264])).
fof(f9235,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK2),s(fun(sK0,X24),X21),s(fun(sK1,X24),X20))) = s(X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X19),s(fun(sK0,X24),X22),s(fun(sK1,X24),X23))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X19) | s('type.sum.sum'(sK0,sK1),X19) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X20,X21,X19,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X22,X23,X24,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),X19) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X20,X21,X19,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X22,X23,X24,sK1,sK0))))) ) | $spl28),
  inference(forward_demodulation,[],[f9194,f9178])).
fof(f9194,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X19) | s('type.sum.sum'(sK0,sK1),X19) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X20,X21,X19,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X22,X23,X24,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),X19) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X20,X21,X19,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X22,X23,X24,sK1,sK0)))) | s(X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X19),s(fun(sK0,X24),X22),s(fun(sK1,X24),X23))) = s(X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))),s(fun(sK0,X24),X21),s(fun(sK1,X24),X20)))) ) | $spl28),
  inference(superposition,[],[f229,f9178])).
fof(f9234,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK2),s(fun(sK0,X18),X16),s(fun(sK1,X18),X17))) = s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X13),s(fun(sK0,X18),X15),s(fun(sK1,X18),X14))) | s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X13)) ) | $spl28),
  inference(forward_demodulation,[],[f9233,f9178])).
fof(f9233,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X13) | s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X13),s(fun(sK0,X18),X15),s(fun(sK1,X18),X14))) = s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))),s(fun(sK0,X18),X16),s(fun(sK1,X18),X17)))) ) | $spl28),
  inference(forward_demodulation,[],[f9232,f9178])).
fof(f9232,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X13) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X13),s(fun(sK0,X18),X15),s(fun(sK1,X18),X14))) = s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))),s(fun(sK0,X18),X16),s(fun(sK1,X18),X17)))) ) | $spl28),
  inference(forward_demodulation,[],[f9193,f9178])).
fof(f9193,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X13) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X14,X15,'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),X13,X16,X17,X18,sK1,sK0)))) | s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X13),s(fun(sK0,X18),X15),s(fun(sK1,X18),X14))) = s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))),s(fun(sK0,X18),X16),s(fun(sK1,X18),X17)))) ) | $spl28),
  inference(superposition,[],[f229,f9178])).
fof(f9231,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK0,bool),X2),s(sK0,sK9(sK2,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK9(sK2,sK1,sK0))))) | p(s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK15(X0,X2,X1,X3,sK3,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK14(X0,X2,X1,X3,sK3,sK1,sK0)))))) ) | $spl28),
  inference(forward_demodulation,[],[f9230,f9192])).
fof(f9230,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK9(sK2,sK1,sK0))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X2),s(fun(sK1,bool),X3),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK15(X0,X2,X1,X3,sK3,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK14(X0,X2,X1,X3,sK3,sK1,sK0)))))) ) | $spl28),
  inference(backward_demodulation,[],[f9192,f2843])).
fof(f9192,plain,(
  ( ! [X12,X11] : (s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK9(sK2,sK1,sK0)))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X11),s(fun(sK1,bool),X12),s('type.sum.sum'(sK0,sK1),sK2)))) ) | $spl28),
  inference(superposition,[],[f211,f9178])).
fof(f9191,plain,(
  ( ! [X10,X8,X9] : (s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9(sK2,sK1,sK0)))) = s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK2),s(fun(sK0,X8),X9),s(fun(sK1,X8),X10)))) ) | $spl28),
  inference(superposition,[],[f228,f9178])).
fof(f9190,plain,(
  ( ! [X6,X7,X5] : (s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9(sK2,sK1,sK0)))) = s(X5,happ(s(fun('type.sum.sum'(sK0,sK1),X5),sK16(X7,X6,sK1,X5,sK0)),s('type.sum.sum'(sK0,sK1),sK2)))) ) | $spl28),
  inference(superposition,[],[f225,f9178])).
fof(f9189,plain,(
  ( ! [X4,X2,X3] : (s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X2))) != s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,X3),s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X4))),s('type.sum.sum'(sK0,sK1),sK2))) | s(sK1,X2) = s(sK1,X4)) ) | $spl28),
  inference(superposition,[],[f213,f9178])).
fof(f9188,plain,(
  ( ! [X0,X1] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X0))) | s(sK1,X0) = s(sK1,X1)) ) | $spl28),
  inference(superposition,[],[f2051,f9178])).
fof(f9181,plain,(
  s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(sK2,sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(sK2,sK1,sK0)))) | $spl28),
  inference(unit_resulting_resolution,[],[f9178,f259])).
fof(f8508,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.sum.sum'(X0,X1),X2) != s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X3))) | s(X1,X4) = s(X1,sK10(X2,X1,X0)) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK9(X2,X1,X0))))) )),
  inference(superposition,[],[f2051,f205])).
fof(f8507,plain,(
  ( ! [X218,X222,X220,X219,X223,X221] : (s('type.sum.sum'(X221,X218),'const.sum.INL_1'(s(X221,X222))) != s('type.sum.sum'(X221,X218),'const.sum.INR_1'(s(X218,X219))) | s(X218,X223) = s(X218,sK8(X219,X220,X218)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X218,bool),X220)))) )),
  inference(superposition,[],[f2051,f425])).
fof(f8506,plain,(
  ( ! [X212,X216,X215,X213,X211,X217,X214] : (s('type.sum.sum'(X215,X211),'const.sum.INL_1'(s(X215,X216))) != s('type.sum.sum'(X215,X211),'const.sum.INR_1'(s(X211,X214))) | s(X211,X217) = s(X211,'const.bool.COND_3'(s(bool,X212),s(X211,X213),s(X211,X214))) | s(bool,'T_0') = s(bool,X212)) )),
  inference(superposition,[],[f2051,f351])).
fof(f8505,plain,(
  ( ! [X210,X208,X206,X209,X207,X205] : (s('type.sum.sum'(X208,X205),'const.sum.INL_1'(s(X208,X209))) != s('type.sum.sum'(X208,X205),'const.sum.INR_1'(s(X205,X207))) | s(X205,X210) = s(X205,'const.bool.COND_3'(s(bool,'F_0'),s(X205,X206),s(X205,X207)))) )),
  inference(superposition,[],[f2051,f194])).
fof(f8504,plain,(
  ( ! [X204,X202,X200,X199,X203,X201] : (s('type.sum.sum'(X202,X199),'const.sum.INL_1'(s(X202,X203))) != s('type.sum.sum'(X202,X199),'const.sum.INR_1'(s(X199,X200))) | s(X199,X204) = s(X199,'const.bool.COND_3'(s(bool,'T_0'),s(X199,X200),s(X199,X201)))) )),
  inference(superposition,[],[f2051,f193])).
fof(f8503,plain,(
  ( ! [X191,X198,X196,X194,X192,X190,X197,X195,X193] : (s('type.sum.sum'(X196,X190),'const.sum.INL_1'(s(X196,X197))) != s('type.sum.sum'(X196,X190),'const.sum.INR_1'(s(X190,happ(s(fun(X191,X190),X194),s(X191,X193))))) | s(X190,X198) = s(X190,'const.sum.sum_CASE_3'(s('type.sum.sum'(X191,X192),'const.sum.INL_1'(s(X191,X193))),s(fun(X191,X190),X194),s(fun(X192,X190),X195)))) )),
  inference(superposition,[],[f2051,f228])).
fof(f8502,plain,(
  ( ! [X189,X187,X185,X182,X188,X186,X184,X183,X181] : (s('type.sum.sum'(X187,X181),'const.sum.INL_1'(s(X187,X188))) != s('type.sum.sum'(X187,X181),'const.sum.INR_1'(s(X181,happ(s(fun(X183,X181),X186),s(X183,X184))))) | s(X181,X189) = s(X181,'const.sum.sum_CASE_3'(s('type.sum.sum'(X182,X183),'const.sum.INR_1'(s(X183,X184))),s(fun(X182,X181),X185),s(fun(X183,X181),X186)))) )),
  inference(superposition,[],[f2051,f227])).
fof(f8501,plain,(
  ( ! [X177,X175,X173,X180,X178,X176,X174,X172,X179] : (s('type.sum.sum'(X178,X172),'const.sum.INL_1'(s(X178,X179))) != s('type.sum.sum'(X178,X172),'const.sum.INR_1'(s(X172,happ(s(fun(X174,X172),X175),s(X174,X177))))) | s(X172,X180) = s(X172,happ(s(fun('type.sum.sum'(X173,X174),X172),sK16(X175,X176,X174,X172,X173)),s('type.sum.sum'(X173,X174),'const.sum.INR_1'(s(X174,X177)))))) )),
  inference(superposition,[],[f2051,f226])).
fof(f8500,plain,(
  ( ! [X167,X165,X163,X171,X169,X166,X164,X170,X168] : (s('type.sum.sum'(X169,X163),'const.sum.INL_1'(s(X169,X170))) != s('type.sum.sum'(X169,X163),'const.sum.INR_1'(s(X163,happ(s(fun(X164,X163),X167),s(X164,X168))))) | s(X163,X171) = s(X163,happ(s(fun('type.sum.sum'(X164,X165),X163),sK16(X166,X167,X165,X163,X164)),s('type.sum.sum'(X164,X165),'const.sum.INL_1'(s(X164,X168)))))) )),
  inference(superposition,[],[f2051,f225])).
fof(f8493,plain,(
  ( ! [X127,X132,X130,X128,X126,X133,X131,X129] : (s('type.sum.sum'(X131,bool),'const.sum.INL_1'(s(X131,X132))) != s('type.sum.sum'(X131,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X126,bool),X127),s(X126,X130))))) | s(bool,X133) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X126,bool),X127),s(fun(X128,bool),X129),s('type.sum.sum'(X126,X128),'const.sum.INL_1'(s(X126,X130)))))) )),
  inference(superposition,[],[f2051,f211])).
fof(f8492,plain,(
  ( ! [X125,X123,X121,X118,X124,X122,X120,X119] : (s('type.sum.sum'(X123,bool),'const.sum.INL_1'(s(X123,X124))) != s('type.sum.sum'(X123,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X120,bool),X121),s(X120,X122))))) | s(bool,X125) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X118,bool),X119),s(fun(X120,bool),X121),s('type.sum.sum'(X118,X120),'const.sum.INR_1'(s(X120,X122)))))) )),
  inference(superposition,[],[f2051,f210])).
fof(f8463,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.sum.sum'(X0,X1),X2) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) | s(X1,X3) = s(X1,X4) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK10(X2,X1,X0))))) )),
  inference(superposition,[],[f2051,f205])).
fof(f8130,plain,(
  ( ! [X185,X182,X186,X184,X183,X181] : (s('type.sum.sum'(X183,bool),'const.sum.INR_1'(s(bool,X184))) != s('type.sum.sum'(X183,bool),'const.bool.COND_3'(s(bool,X185),s('type.sum.sum'(X183,bool),'const.sum.INR_1'(s(bool,'T_0'))),s('type.sum.sum'(X183,bool),'const.sum.INL_1'(s(X183,X186))))) | s(bool,X184) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X181),s(bool,X182))) | s(bool,'F_0') = s(bool,X182)) )),
  inference(superposition,[],[f213,f363])).
fof(f8129,plain,(
  ( ! [X177,X175,X180,X178,X176,X179] : (s('type.sum.sum'(X177,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X177,bool),'const.bool.COND_3'(s(bool,X178),s('type.sum.sum'(X177,bool),'const.sum.INR_1'(s(bool,X179))),s('type.sum.sum'(X177,bool),'const.sum.INL_1'(s(X177,X180))))) | s(bool,X179) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X175),s(bool,X176))) | s(bool,'F_0') = s(bool,X176)) )),
  inference(superposition,[],[f213,f363])).
fof(f8123,plain,(
  ( ! [X152,X151,X149,X150] : (s('type.sum.sum'(X151,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X151,bool),'const.sum.INR_1'(s(bool,X152))) | s(bool,X152) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X149),s(bool,X150))) | s(bool,'F_0') = s(bool,X150)) )),
  inference(superposition,[],[f208,f363])).
fof(f8178,plain,(
  ( ! [X54,X55,X53] : (s(bool,X55) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X53),s(bool,X54))) | s(bool,'T_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f8086,f157])).
fof(f8086,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,'T_0')) | s(bool,X55) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X53),s(bool,X54))) | s(bool,'T_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X54)) )),
  inference(superposition,[],[f274,f363])).
fof(f8062,plain,(
  ( ! [X116,X114,X115] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X116),s(bool,X114))) | s(bool,'F_0') = s(bool,sK8(X114,X115,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X115)))) )),
  inference(superposition,[],[f363,f425])).
fof(f8061,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X113),s(bool,X112))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X110),s(bool,X111),s(bool,X112))) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f363,f351])).
fof(f8060,plain,(
  ( ! [X109,X107,X108] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X109),s(bool,X108))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X107),s(bool,X108)))) )),
  inference(superposition,[],[f363,f194])).
fof(f8059,plain,(
  ( ! [X105,X106,X104] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X106),s(bool,X104))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X104),s(bool,X105)))) )),
  inference(superposition,[],[f363,f193])).
fof(f8058,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X103),s(bool,happ(s(fun(X98,bool),X101),s(X98,X100))))) | s(bool,'F_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X98,X99),'const.sum.INL_1'(s(X98,X100))),s(fun(X98,bool),X101),s(fun(X99,bool),X102)))) )),
  inference(superposition,[],[f363,f228])).
fof(f8057,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X97),s(bool,happ(s(fun(X93,bool),X96),s(X93,X94))))) | s(bool,'F_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X92,X93),'const.sum.INR_1'(s(X93,X94))),s(fun(X92,bool),X95),s(fun(X93,bool),X96)))) )),
  inference(superposition,[],[f363,f227])).
fof(f8056,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X91),s(bool,happ(s(fun(X87,bool),X88),s(X87,X90))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(X86,X87),bool),sK16(X88,X89,X87,bool,X86)),s('type.sum.sum'(X86,X87),'const.sum.INR_1'(s(X87,X90)))))) )),
  inference(superposition,[],[f363,f226])).
fof(f8055,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X85),s(bool,happ(s(fun(X80,bool),X83),s(X80,X84))))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(X80,X81),bool),sK16(X82,X83,X81,bool,X80)),s('type.sum.sum'(X80,X81),'const.sum.INL_1'(s(X80,X84)))))) )),
  inference(superposition,[],[f363,f225])).
fof(f8049,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X64),s(bool,happ(s(fun(X59,bool),X60),s(X59,X63))))) | s(bool,'F_0') = s(bool,'const.sum.SUM_ALL_3'(s(fun(X59,bool),X60),s(fun(X61,bool),X62),s('type.sum.sum'(X59,X61),'const.sum.INL_1'(s(X59,X63)))))) )),
  inference(superposition,[],[f363,f211])).
fof(f8048,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X58),s(bool,happ(s(fun(X55,bool),X56),s(X55,X57))))) | s(bool,'F_0') = s(bool,'const.sum.SUM_ALL_3'(s(fun(X53,bool),X54),s(fun(X55,bool),X56),s('type.sum.sum'(X53,X55),'const.sum.INR_1'(s(X55,X57)))))) )),
  inference(superposition,[],[f363,f210])).
fof(f7992,plain,(
  ( ! [X24,X23,X25] : (s(bool,X23) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X23),s(bool,X24))) | s(bool,'F_0') = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f363,f272])).
fof(f7685,plain,(
  ( ! [X549,X547,X548] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,X549))) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,'const.bool.COND_3'(s(bool,X547),s(sK0,X548),s(sK0,X549))))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,'const.bool.COND_3'(s(bool,X547),s(sK0,X548),s(sK0,X549))))) | s(bool,'T_0') = s(bool,X547)) )),
  inference(superposition,[],[f259,f351])).
fof(f7684,plain,(
  ( ! [X546,X542,X540,X538,X545,X541,X543,X539,X544] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X538,bool),X543),s(fun(X539,bool),X544),s('type.sum.sum'(X538,X539),X542)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X538,bool),X545),s(fun(X539,bool),X546),s('type.sum.sum'(X538,X539),'const.bool.COND_3'(s(bool,X540),s('type.sum.sum'(X538,X539),X541),s('type.sum.sum'(X538,X539),X542)))))) | p(s(bool,happ(s(fun(X539,bool),X544),s(X539,sK15(X543,X545,X544,X546,'const.bool.COND_3'(s(bool,X540),s('type.sum.sum'(X538,X539),X541),s('type.sum.sum'(X538,X539),X542)),X539,X538))))) | p(s(bool,happ(s(fun(X538,bool),X543),s(X538,sK14(X543,X545,X544,X546,'const.bool.COND_3'(s(bool,X540),s('type.sum.sum'(X538,X539),X541),s('type.sum.sum'(X538,X539),X542)),X539,X538))))) | s(bool,'T_0') = s(bool,X540)) )),
  inference(superposition,[],[f215,f351])).
fof(f7683,plain,(
  ( ! [X536,X532,X534,X528,X530,X537,X533,X535,X529,X531,X527] : (s('type.sum.sum'(X527,X528),X531) != s('type.sum.sum'(X527,X528),X532) | s('type.sum.sum'(X527,X528),X532) = s('type.sum.sum'(X527,X528),'const.sum.INR_1'(s(X528,sK18(X533,X534,X532,'const.bool.COND_3'(s(bool,X529),s('type.sum.sum'(X527,X528),X530),s('type.sum.sum'(X527,X528),X531)),X535,X536,X537,X528,X527)))) | s('type.sum.sum'(X527,X528),X532) = s('type.sum.sum'(X527,X528),'const.sum.INL_1'(s(X527,sK17(X533,X534,X532,'const.bool.COND_3'(s(bool,X529),s('type.sum.sum'(X527,X528),X530),s('type.sum.sum'(X527,X528),X531)),X535,X536,X537,X528,X527)))) | s(X537,'const.sum.sum_CASE_3'(s('type.sum.sum'(X527,X528),X532),s(fun(X527,X537),X535),s(fun(X528,X537),X536))) = s(X537,'const.sum.sum_CASE_3'(s('type.sum.sum'(X527,X528),'const.bool.COND_3'(s(bool,X529),s('type.sum.sum'(X527,X528),X530),s('type.sum.sum'(X527,X528),X531))),s(fun(X527,X537),X534),s(fun(X528,X537),X533))) | s(bool,'T_0') = s(bool,X529)) )),
  inference(superposition,[],[f229,f351])).
fof(f7682,plain,(
  ( ! [X524,X526,X520,X522,X516,X518,X525,X521,X523,X517,X519] : (s('type.sum.sum'(X516,X517),X520) != s('type.sum.sum'(X516,X517),X521) | s('type.sum.sum'(X516,X517),'const.bool.COND_3'(s(bool,X518),s('type.sum.sum'(X516,X517),X519),s('type.sum.sum'(X516,X517),X520))) = s('type.sum.sum'(X516,X517),'const.sum.INR_1'(s(X517,sK18(X522,X523,'const.bool.COND_3'(s(bool,X518),s('type.sum.sum'(X516,X517),X519),s('type.sum.sum'(X516,X517),X520)),X521,X524,X525,X526,X517,X516)))) | s('type.sum.sum'(X516,X517),'const.bool.COND_3'(s(bool,X518),s('type.sum.sum'(X516,X517),X519),s('type.sum.sum'(X516,X517),X520))) = s('type.sum.sum'(X516,X517),'const.sum.INL_1'(s(X516,sK17(X522,X523,'const.bool.COND_3'(s(bool,X518),s('type.sum.sum'(X516,X517),X519),s('type.sum.sum'(X516,X517),X520)),X521,X524,X525,X526,X517,X516)))) | s(X526,'const.sum.sum_CASE_3'(s('type.sum.sum'(X516,X517),X521),s(fun(X516,X526),X523),s(fun(X517,X526),X522))) = s(X526,'const.sum.sum_CASE_3'(s('type.sum.sum'(X516,X517),'const.bool.COND_3'(s(bool,X518),s('type.sum.sum'(X516,X517),X519),s('type.sum.sum'(X516,X517),X520))),s(fun(X516,X526),X524),s(fun(X517,X526),X525))) | s(bool,'T_0') = s(bool,X518)) )),
  inference(superposition,[],[f229,f351])).
fof(f7678,plain,(
  ( ! [X499,X494,X500,X502,X496,X498,X493,X495,X501,X503,X497] : (s(X494,happ(s(fun(X493,X494),X497),s(X493,sK18('const.bool.COND_3'(s(bool,X495),s(fun(X493,X494),X496),s(fun(X493,X494),X497)),X499,X500,X501,X502,X498,X494,X493,X503)))) != s(X494,happ(s(fun(X493,X494),X498),s(X493,sK18('const.bool.COND_3'(s(bool,X495),s(fun(X493,X494),X496),s(fun(X493,X494),X497)),X499,X500,X501,X502,X498,X494,X493,X503)))) | s(X494,'const.sum.sum_CASE_3'(s('type.sum.sum'(X503,X493),X500),s(fun(X503,X494),X502),s(fun(X493,X494),X498))) = s(X494,'const.sum.sum_CASE_3'(s('type.sum.sum'(X503,X493),X501),s(fun(X503,X494),X499),s(fun(X493,X494),'const.bool.COND_3'(s(bool,X495),s(fun(X493,X494),X496),s(fun(X493,X494),X497))))) | s(X494,happ(s(fun(X503,X494),X499),s(X503,sK17('const.bool.COND_3'(s(bool,X495),s(fun(X493,X494),X496),s(fun(X493,X494),X497)),X499,X500,X501,X502,X498,X494,X493,X503)))) != s(X494,happ(s(fun(X503,X494),X502),s(X503,sK17('const.bool.COND_3'(s(bool,X495),s(fun(X493,X494),X496),s(fun(X493,X494),X497)),X499,X500,X501,X502,X498,X494,X493,X503)))) | s('type.sum.sum'(X503,X493),X500) != s('type.sum.sum'(X503,X493),X501) | s(bool,'T_0') = s(bool,X495)) )),
  inference(superposition,[],[f232,f351])).
fof(f7677,plain,(
  ( ! [X484,X486,X482,X492,X488,X490,X485,X487,X483,X489,X491] : (s(X483,happ(s(fun(X482,X483),X486),s(X482,sK18(X487,X488,X489,X490,X491,'const.bool.COND_3'(s(bool,X484),s(fun(X482,X483),X485),s(fun(X482,X483),X486)),X483,X482,X492)))) != s(X483,happ(s(fun(X482,X483),X487),s(X482,sK18(X487,X488,X489,X490,X491,'const.bool.COND_3'(s(bool,X484),s(fun(X482,X483),X485),s(fun(X482,X483),X486)),X483,X482,X492)))) | s(X483,'const.sum.sum_CASE_3'(s('type.sum.sum'(X492,X482),X490),s(fun(X492,X483),X488),s(fun(X482,X483),X487))) = s(X483,'const.sum.sum_CASE_3'(s('type.sum.sum'(X492,X482),X489),s(fun(X492,X483),X491),s(fun(X482,X483),'const.bool.COND_3'(s(bool,X484),s(fun(X482,X483),X485),s(fun(X482,X483),X486))))) | s(X483,happ(s(fun(X492,X483),X488),s(X492,sK17(X487,X488,X489,X490,X491,'const.bool.COND_3'(s(bool,X484),s(fun(X482,X483),X485),s(fun(X482,X483),X486)),X483,X482,X492)))) != s(X483,happ(s(fun(X492,X483),X491),s(X492,sK17(X487,X488,X489,X490,X491,'const.bool.COND_3'(s(bool,X484),s(fun(X482,X483),X485),s(fun(X482,X483),X486)),X483,X482,X492)))) | s('type.sum.sum'(X492,X482),X489) != s('type.sum.sum'(X492,X482),X490) | s(bool,'T_0') = s(bool,X484)) )),
  inference(superposition,[],[f232,f351])).
fof(f7676,plain,(
  ( ! [X476,X478,X472,X474,X480,X471,X475,X477,X479,X473,X481] : (s(X472,happ(s(fun(X471,X472),X475),s(X471,sK18('const.bool.COND_3'(s(bool,X473),s(fun(X471,X472),X474),s(fun(X471,X472),X475)),X477,X478,X479,X480,X476,X472,X471,X481)))) != s(X472,happ(s(fun(X471,X472),X476),s(X471,sK18('const.bool.COND_3'(s(bool,X473),s(fun(X471,X472),X474),s(fun(X471,X472),X475)),X477,X478,X479,X480,X476,X472,X471,X481)))) | s(X472,'const.sum.sum_CASE_3'(s('type.sum.sum'(X481,X471),X478),s(fun(X481,X472),X480),s(fun(X471,X472),X476))) = s(X472,'const.sum.sum_CASE_3'(s('type.sum.sum'(X481,X471),X479),s(fun(X481,X472),X477),s(fun(X471,X472),'const.bool.COND_3'(s(bool,X473),s(fun(X471,X472),X474),s(fun(X471,X472),X475))))) | s('type.sum.sum'(X481,X471),X478) = s('type.sum.sum'(X481,X471),'const.sum.INL_1'(s(X481,sK17('const.bool.COND_3'(s(bool,X473),s(fun(X471,X472),X474),s(fun(X471,X472),X475)),X477,X478,X479,X480,X476,X472,X471,X481)))) | s('type.sum.sum'(X481,X471),X478) != s('type.sum.sum'(X481,X471),X479) | s(bool,'T_0') = s(bool,X473)) )),
  inference(superposition,[],[f231,f351])).
fof(f7675,plain,(
  ( ! [X466,X461,X463,X469,X465,X467,X460,X462,X468,X470,X464] : (s(X461,happ(s(fun(X460,X461),X464),s(X460,sK18(X465,X466,X467,X468,X469,'const.bool.COND_3'(s(bool,X462),s(fun(X460,X461),X463),s(fun(X460,X461),X464)),X461,X460,X470)))) != s(X461,happ(s(fun(X460,X461),X465),s(X460,sK18(X465,X466,X467,X468,X469,'const.bool.COND_3'(s(bool,X462),s(fun(X460,X461),X463),s(fun(X460,X461),X464)),X461,X460,X470)))) | s(X461,'const.sum.sum_CASE_3'(s('type.sum.sum'(X470,X460),X468),s(fun(X470,X461),X466),s(fun(X460,X461),X465))) = s(X461,'const.sum.sum_CASE_3'(s('type.sum.sum'(X470,X460),X467),s(fun(X470,X461),X469),s(fun(X460,X461),'const.bool.COND_3'(s(bool,X462),s(fun(X460,X461),X463),s(fun(X460,X461),X464))))) | s('type.sum.sum'(X470,X460),X467) = s('type.sum.sum'(X470,X460),'const.sum.INL_1'(s(X470,sK17(X465,X466,X467,X468,X469,'const.bool.COND_3'(s(bool,X462),s(fun(X460,X461),X463),s(fun(X460,X461),X464)),X461,X460,X470)))) | s('type.sum.sum'(X470,X460),X467) != s('type.sum.sum'(X470,X460),X468) | s(bool,'T_0') = s(bool,X462)) )),
  inference(superposition,[],[f231,f351])).
fof(f7674,plain,(
  ( ! [X453,X455,X449,X451,X457,X459,X452,X454,X450,X456,X458] : (s(X450,happ(s(fun(X449,X450),X453),s(X449,sK17(X454,X455,X456,X457,'const.bool.COND_3'(s(bool,X451),s(fun(X449,X450),X452),s(fun(X449,X450),X453)),X458,X450,X459,X449)))) != s(X450,happ(s(fun(X449,X450),X455),s(X449,sK17(X454,X455,X456,X457,'const.bool.COND_3'(s(bool,X451),s(fun(X449,X450),X452),s(fun(X449,X450),X453)),X458,X450,X459,X449)))) | s('type.sum.sum'(X449,X459),X456) = s('type.sum.sum'(X449,X459),'const.sum.INR_1'(s(X459,sK18(X454,X455,X456,X457,'const.bool.COND_3'(s(bool,X451),s(fun(X449,X450),X452),s(fun(X449,X450),X453)),X458,X450,X459,X449)))) | s(X450,'const.sum.sum_CASE_3'(s('type.sum.sum'(X449,X459),X457),s(fun(X449,X450),X455),s(fun(X459,X450),X454))) = s(X450,'const.sum.sum_CASE_3'(s('type.sum.sum'(X449,X459),X456),s(fun(X449,X450),'const.bool.COND_3'(s(bool,X451),s(fun(X449,X450),X452),s(fun(X449,X450),X453))),s(fun(X459,X450),X458))) | s('type.sum.sum'(X449,X459),X456) != s('type.sum.sum'(X449,X459),X457) | s(bool,'T_0') = s(bool,X451)) )),
  inference(superposition,[],[f230,f351])).
fof(f7673,plain,(
  ( ! [X445,X447,X441,X443,X438,X444,X446,X440,X442,X448,X439] : (s(X439,happ(s(fun(X438,X439),X442),s(X438,sK17(X444,'const.bool.COND_3'(s(bool,X440),s(fun(X438,X439),X441),s(fun(X438,X439),X442)),X445,X446,X443,X447,X439,X448,X438)))) != s(X439,happ(s(fun(X438,X439),X443),s(X438,sK17(X444,'const.bool.COND_3'(s(bool,X440),s(fun(X438,X439),X441),s(fun(X438,X439),X442)),X445,X446,X443,X447,X439,X448,X438)))) | s('type.sum.sum'(X438,X448),X445) = s('type.sum.sum'(X438,X448),'const.sum.INR_1'(s(X448,sK18(X444,'const.bool.COND_3'(s(bool,X440),s(fun(X438,X439),X441),s(fun(X438,X439),X442)),X445,X446,X443,X447,X439,X448,X438)))) | s(X439,'const.sum.sum_CASE_3'(s('type.sum.sum'(X438,X448),X445),s(fun(X438,X439),X443),s(fun(X448,X439),X447))) = s(X439,'const.sum.sum_CASE_3'(s('type.sum.sum'(X438,X448),X446),s(fun(X438,X439),'const.bool.COND_3'(s(bool,X440),s(fun(X438,X439),X441),s(fun(X438,X439),X442))),s(fun(X448,X439),X444))) | s('type.sum.sum'(X438,X448),X445) != s('type.sum.sum'(X438,X448),X446) | s(bool,'T_0') = s(bool,X440)) )),
  inference(superposition,[],[f230,f351])).
fof(f7668,plain,(
  ( ! [X405,X401,X403,X404,X400,X402] : (~p(s(bool,happ(s(fun('type.sum.sum'(X400,X401),bool),X404),s('type.sum.sum'(X400,X401),'const.sum.INR_1'(s(X401,sK12('const.bool.COND_3'(s(bool,X402),s(fun('type.sum.sum'(X400,X401),bool),X403),s(fun('type.sum.sum'(X400,X401),bool),X404)),X401,X400))))))) | p(s(bool,happ(s(fun('type.sum.sum'(X400,X401),bool),'const.bool.COND_3'(s(bool,X402),s(fun('type.sum.sum'(X400,X401),bool),X403),s(fun('type.sum.sum'(X400,X401),bool),X404))),s('type.sum.sum'(X400,X401),X405)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X400,X401),bool),'const.bool.COND_3'(s(bool,X402),s(fun('type.sum.sum'(X400,X401),bool),X403),s(fun('type.sum.sum'(X400,X401),bool),X404))),s('type.sum.sum'(X400,X401),'const.sum.INL_1'(s(X400,sK11('const.bool.COND_3'(s(bool,X402),s(fun('type.sum.sum'(X400,X401),bool),X403),s(fun('type.sum.sum'(X400,X401),bool),X404)),X401,X400))))))) | s(bool,'T_0') = s(bool,X402)) )),
  inference(superposition,[],[f206,f351])).
fof(f7667,plain,(
  ( ! [X397,X399,X395,X396,X398] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X395,bool),X398))) | s(X395,X399) = s(X395,sK8(X399,'const.bool.COND_3'(s(bool,X396),s(fun(X395,bool),X397),s(fun(X395,bool),X398)),X395)) | s(bool,'T_0') = s(bool,X396)) )),
  inference(superposition,[],[f425,f351])).
fof(f7666,plain,(
  ( ! [X389,X391,X387,X393,X388,X390,X386,X392,X394] : (~p(s(bool,happ(s(fun(X386,bool),X389),s(X386,sK15(X390,X391,X392,'const.bool.COND_3'(s(bool,X387),s(fun(X386,bool),X388),s(fun(X386,bool),X389)),X393,X386,X394))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X394,bool),X390),s(fun(X386,bool),X392),s('type.sum.sum'(X394,X386),X393)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X394,bool),X391),s(fun(X386,bool),'const.bool.COND_3'(s(bool,X387),s(fun(X386,bool),X388),s(fun(X386,bool),X389))),s('type.sum.sum'(X394,X386),X393)))) | ~p(s(bool,happ(s(fun(X394,bool),X391),s(X394,sK14(X390,X391,X392,'const.bool.COND_3'(s(bool,X387),s(fun(X386,bool),X388),s(fun(X386,bool),X389)),X393,X386,X394))))) | s(bool,'T_0') = s(bool,X387)) )),
  inference(superposition,[],[f218,f351])).
fof(f7665,plain,(
  ( ! [X383,X377,X379,X385,X380,X382,X378,X384,X381] : (~p(s(bool,happ(s(fun(X377,bool),X380),s(X377,sK15(X381,X382,X383,'const.bool.COND_3'(s(bool,X378),s(fun(X377,bool),X379),s(fun(X377,bool),X380)),X384,X377,X385))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X385,bool),X381),s(fun(X377,bool),X383),s('type.sum.sum'(X385,X377),X384)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X385,bool),X382),s(fun(X377,bool),'const.bool.COND_3'(s(bool,X378),s(fun(X377,bool),X379),s(fun(X377,bool),X380))),s('type.sum.sum'(X385,X377),X384)))) | p(s(bool,happ(s(fun(X385,bool),X381),s(X385,sK14(X381,X382,X383,'const.bool.COND_3'(s(bool,X378),s(fun(X377,bool),X379),s(fun(X377,bool),X380)),X384,X377,X385))))) | s(bool,'T_0') = s(bool,X378)) )),
  inference(superposition,[],[f217,f351])).
fof(f7664,plain,(
  ( ! [X372,X374,X368,X370,X376,X373,X375,X369,X371] : (~p(s(bool,happ(s(fun(X368,bool),X371),s(X368,sK14(X372,'const.bool.COND_3'(s(bool,X369),s(fun(X368,bool),X370),s(fun(X368,bool),X371)),X373,X374,X375,X376,X368))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X368,bool),X372),s(fun(X376,bool),X373),s('type.sum.sum'(X368,X376),X375)))) | p(s(bool,happ(s(fun(X376,bool),X373),s(X376,sK15(X372,'const.bool.COND_3'(s(bool,X369),s(fun(X368,bool),X370),s(fun(X368,bool),X371)),X373,X374,X375,X376,X368))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X368,bool),'const.bool.COND_3'(s(bool,X369),s(fun(X368,bool),X370),s(fun(X368,bool),X371))),s(fun(X376,bool),X374),s('type.sum.sum'(X368,X376),X375)))) | s(bool,'T_0') = s(bool,X369)) )),
  inference(superposition,[],[f216,f351])).
fof(f7663,plain,(
  ( ! [X364,X366,X360,X362,X359,X365,X367,X361,X363] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X359,bool),X362),s(fun(X363,bool),X364),s('type.sum.sum'(X359,X363),X365)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X359,bool),X366),s(fun(X363,bool),X367),s('type.sum.sum'(X359,X363),X365)))) | p(s(bool,happ(s(fun(X363,bool),X364),s(X363,sK15('const.bool.COND_3'(s(bool,X360),s(fun(X359,bool),X361),s(fun(X359,bool),X362)),X366,X364,X367,X365,X363,X359))))) | p(s(bool,happ(s(fun(X359,bool),'const.bool.COND_3'(s(bool,X360),s(fun(X359,bool),X361),s(fun(X359,bool),X362))),s(X359,sK14('const.bool.COND_3'(s(bool,X360),s(fun(X359,bool),X361),s(fun(X359,bool),X362)),X366,X364,X367,X365,X363,X359))))) | s(bool,'T_0') = s(bool,X360)) )),
  inference(superposition,[],[f215,f351])).
fof(f7662,plain,(
  ( ! [X350,X356,X358,X352,X354,X351,X357,X353,X355] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X354,bool),X355),s(fun(X350,bool),X353),s('type.sum.sum'(X354,X350),X356)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X354,bool),X357),s(fun(X350,bool),X358),s('type.sum.sum'(X354,X350),X356)))) | p(s(bool,happ(s(fun(X350,bool),'const.bool.COND_3'(s(bool,X351),s(fun(X350,bool),X352),s(fun(X350,bool),X353))),s(X350,sK15(X355,X357,'const.bool.COND_3'(s(bool,X351),s(fun(X350,bool),X352),s(fun(X350,bool),X353)),X358,X356,X350,X354))))) | p(s(bool,happ(s(fun(X354,bool),X355),s(X354,sK14(X355,X357,'const.bool.COND_3'(s(bool,X351),s(fun(X350,bool),X352),s(fun(X350,bool),X353)),X358,X356,X350,X354))))) | s(bool,'T_0') = s(bool,X351)) )),
  inference(superposition,[],[f215,f351])).
fof(f7657,plain,(
  ( ! [X321,X318,X320,X317,X319] : (~p(s(bool,happ(s(fun(X317,bool),X320),s(X317,sK8(X321,'const.bool.COND_3'(s(bool,X318),s(fun(X317,bool),X319),s(fun(X317,bool),X320)),X317))))) | p(s(bool,'|dollar|exists_1'(s(fun(X317,bool),'const.bool.COND_3'(s(bool,X318),s(fun(X317,bool),X319),s(fun(X317,bool),X320)))))) | s(bool,'T_0') = s(bool,X318)) )),
  inference(superposition,[],[f196,f351])).
fof(f7656,plain,(
  ( ! [X313,X315,X316,X312,X314] : (p(s(bool,'|dollar|exists_1'(s(fun(X312,bool),X315)))) | s(X312,X316) = s(X312,sK8(X316,'const.bool.COND_3'(s(bool,X313),s(fun(X312,bool),X314),s(fun(X312,bool),X315)),X312)) | s(bool,'T_0') = s(bool,X313)) )),
  inference(superposition,[],[f195,f351])).
fof(f7629,plain,(
  ( ! [X222,X220,X219,X221] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X221),s(bool,X222))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X219),s(bool,X220),s(bool,X221))) | s(bool,'T_0') = s(bool,X219)) )),
  inference(superposition,[],[f325,f351])).
fof(f7628,plain,(
  ( ! [X218,X216,X215,X217] : (p(s(bool,X217)) | s(bool,X218) = s(bool,'const.bool.COND_3'(s(bool,X215),s(bool,X216),s(bool,X217))) | s(bool,'T_0') = s(bool,X218) | s(bool,'T_0') = s(bool,X215)) )),
  inference(superposition,[],[f306,f351])).
fof(f7627,plain,(
  ( ! [X212,X213,X211,X214] : (~p(s(bool,X213)) | s(bool,X214) != s(bool,'const.bool.COND_3'(s(bool,X211),s(bool,X212),s(bool,X213))) | s(bool,'T_0') = s(bool,X214) | s(bool,'T_0') = s(bool,X211)) )),
  inference(superposition,[],[f274,f351])).
fof(f7622,plain,(
  ( ! [X191,X189,X187,X194,X192,X190,X188,X193] : (~p(s(bool,X189)) | s(X190,X191) != s(X190,X192) | s('type.sum.sum'(X193,X190),'const.sum.INR_1'(s(X190,X192))) = s('type.sum.sum'(X193,X190),'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X187),s(bool,X188),s(bool,X189))),s('type.sum.sum'(X193,X190),'const.sum.INR_1'(s(X190,X191))),s('type.sum.sum'(X193,X190),'const.sum.INL_1'(s(X193,X194))))) | s(bool,'T_0') = s(bool,X187)) )),
  inference(superposition,[],[f214,f351])).
fof(f7620,plain,(
  ( ! [X177,X175,X173,X171,X178,X176,X174,X172] : (p(s(bool,X173)) | s('type.sum.sum'(X174,X175),'const.sum.INR_1'(s(X175,X176))) != s('type.sum.sum'(X174,X175),'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X171),s(bool,X172),s(bool,X173))),s('type.sum.sum'(X174,X175),'const.sum.INR_1'(s(X175,X177))),s('type.sum.sum'(X174,X175),'const.sum.INL_1'(s(X174,X178))))) | s(bool,'T_0') = s(bool,X171)) )),
  inference(superposition,[],[f212,f351])).
fof(f7619,plain,(
  ( ! [X167,X169,X170,X168] : (~p(s(bool,X169)) | ~p(s(bool,X170)) | s(bool,X170) = s(bool,'const.bool.COND_3'(s(bool,X167),s(bool,X168),s(bool,X169))) | s(bool,'T_0') = s(bool,X167)) )),
  inference(superposition,[],[f204,f351])).
fof(f7617,plain,(
  ( ! [X161,X159,X162,X160] : (p(s(bool,X161)) | p(s(bool,X162)) | s(bool,X162) = s(bool,'const.bool.COND_3'(s(bool,X159),s(bool,X160),s(bool,X161))) | s(bool,'T_0') = s(bool,X159)) )),
  inference(superposition,[],[f201,f351])).
fof(f7616,plain,(
  ( ! [X158,X156,X157,X155] : (~p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X158) != s(bool,'const.bool.COND_3'(s(bool,X155),s(bool,X156),s(bool,X157))) | s(bool,'T_0') = s(bool,X155)) )),
  inference(superposition,[],[f200,f351])).
fof(f7615,plain,(
  ( ! [X154,X152,X153] : (p(s(bool,X154)) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X152),s(bool,X153),s(bool,X154))) | s(bool,'T_0') = s(bool,X152)) )),
  inference(superposition,[],[f170,f351])).
fof(f7614,plain,(
  ( ! [X151,X149,X150] : (s(bool,'F_0') != s(bool,X151) | ~p(s(bool,'const.bool.COND_3'(s(bool,X149),s(bool,X150),s(bool,X151)))) | s(bool,'T_0') = s(bool,X149)) )),
  inference(superposition,[],[f169,f351])).
fof(f7613,plain,(
  ( ! [X146,X147,X148] : (~p(s(bool,X148)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X146),s(bool,X147),s(bool,X148))) | s(bool,'T_0') = s(bool,X146)) )),
  inference(superposition,[],[f166,f351])).
fof(f7612,plain,(
  ( ! [X144,X145,X143] : (p(s(bool,X145)) | s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X143),s(bool,X144),s(bool,X145))) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f165,f351])).
fof(f7600,plain,(
  ( ! [X78,X76,X74,X72,X79,X77,X75,X73] : (s('type.sum.sum'(X76,X72),'const.sum.INR_1'(s(X72,X77))) != s('type.sum.sum'(X76,X72),'const.bool.COND_3'(s(bool,X78),s('type.sum.sum'(X76,X72),'const.sum.INR_1'(s(X72,X75))),s('type.sum.sum'(X76,X72),'const.sum.INL_1'(s(X76,X79))))) | s(X72,X77) = s(X72,'const.bool.COND_3'(s(bool,X73),s(X72,X74),s(X72,X75))) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f213,f351])).
fof(f7599,plain,(
  ( ! [X70,X68,X66,X64,X71,X69,X67,X65] : (s('type.sum.sum'(X68,X64),'const.sum.INR_1'(s(X64,X67))) != s('type.sum.sum'(X68,X64),'const.bool.COND_3'(s(bool,X69),s('type.sum.sum'(X68,X64),'const.sum.INR_1'(s(X64,X70))),s('type.sum.sum'(X68,X64),'const.sum.INL_1'(s(X68,X71))))) | s(X64,X70) = s(X64,'const.bool.COND_3'(s(bool,X65),s(X64,X66),s(X64,X67))) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f213,f351])).
fof(f7595,plain,(
  ( ! [X39,X43,X41,X38,X42,X40] : (s(X38,X41) != s(X38,X42) | s('type.sum.sum'(X43,X38),'const.sum.INR_1'(s(X38,X42))) = s('type.sum.sum'(X43,X38),'const.sum.INR_1'(s(X38,'const.bool.COND_3'(s(bool,X39),s(X38,X40),s(X38,X41))))) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f209,f351])).
fof(f7593,plain,(
  ( ! [X30,X28,X26,X31,X29,X27] : (s('type.sum.sum'(X30,X26),'const.sum.INR_1'(s(X26,X29))) != s('type.sum.sum'(X30,X26),'const.sum.INR_1'(s(X26,X31))) | s(X26,X31) = s(X26,'const.bool.COND_3'(s(bool,X27),s(X26,X28),s(X26,X29))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f208,f351])).
fof(f7580,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X9),s(bool,X10),s(bool,X11)))) )),
  inference(superposition,[],[f351,f272])).
fof(f7579,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X5) | s(bool,X8) = s(bool,'const.bool.COND_3'(s(bool,X5),s(bool,X6),s(bool,X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f351,f272])).
fof(f7578,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X2),s(bool,X3),s(bool,X4)))) )),
  inference(superposition,[],[f351,f163])).
fof(f7494,plain,(
  ( ! [X185,X182,X184,X183,X181] : (s(X183,X184) = s(X183,'const.bool.COND_3'(s(bool,X181),s(X183,X185),s(X183,X184))) | s(bool,'T_0') = s(bool,sK8(X181,X182,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X182)))) )),
  inference(superposition,[],[f351,f425])).
fof(f7493,plain,(
  ( ! [X177,X175,X180,X178,X176,X179] : (s(X178,X179) = s(X178,'const.bool.COND_3'(s(bool,X177),s(X178,X180),s(X178,X179))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X175),s(bool,X176),s(bool,X177))) | s(bool,'T_0') = s(bool,X175)) )),
  inference(superposition,[],[f351,f351])).
fof(f7492,plain,(
  ( ! [X173,X171,X174,X172,X170] : (s(X172,X173) = s(X172,'const.bool.COND_3'(s(bool,X171),s(X172,X174),s(X172,X173))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X170),s(bool,X171)))) )),
  inference(superposition,[],[f351,f194])).
fof(f7491,plain,(
  ( ! [X167,X165,X169,X166,X168] : (s(X167,X168) = s(X167,'const.bool.COND_3'(s(bool,X165),s(X167,X169),s(X167,X168))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X165),s(bool,X166)))) )),
  inference(superposition,[],[f351,f193])).
fof(f7490,plain,(
  ( ! [X158,X163,X161,X159,X157,X164,X162,X160] : (s(X162,X163) = s(X162,'const.bool.COND_3'(s(bool,happ(s(fun(X157,bool),X160),s(X157,X159))),s(X162,X164),s(X162,X163))) | s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X157,X158),'const.sum.INL_1'(s(X157,X159))),s(fun(X157,bool),X160),s(fun(X158,bool),X161)))) )),
  inference(superposition,[],[f351,f228])).
fof(f7489,plain,(
  ( ! [X156,X154,X152,X151,X149,X155,X153,X150] : (s(X154,X155) = s(X154,'const.bool.COND_3'(s(bool,happ(s(fun(X150,bool),X153),s(X150,X151))),s(X154,X156),s(X154,X155))) | s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X149,X150),'const.sum.INR_1'(s(X150,X151))),s(fun(X149,bool),X152),s(fun(X150,bool),X153)))) )),
  inference(superposition,[],[f351,f227])).
fof(f7488,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141,X148] : (s(X146,X147) = s(X146,'const.bool.COND_3'(s(bool,happ(s(fun(X142,bool),X143),s(X142,X145))),s(X146,X148),s(X146,X147))) | s(bool,'T_0') = s(bool,happ(s(fun('type.sum.sum'(X141,X142),bool),sK16(X143,X144,X142,bool,X141)),s('type.sum.sum'(X141,X142),'const.sum.INR_1'(s(X142,X145)))))) )),
  inference(superposition,[],[f351,f226])).
fof(f7487,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X139,X137] : (s(X138,X139) = s(X138,'const.bool.COND_3'(s(bool,happ(s(fun(X133,bool),X136),s(X133,X137))),s(X138,X140),s(X138,X139))) | s(bool,'T_0') = s(bool,happ(s(fun('type.sum.sum'(X133,X134),bool),sK16(X135,X136,X134,bool,X133)),s('type.sum.sum'(X133,X134),'const.sum.INL_1'(s(X133,X137)))))) )),
  inference(superposition,[],[f351,f225])).
fof(f7483,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X115,X113] : (s(X114,X115) = s(X114,'const.bool.COND_3'(s(bool,happ(s(fun(X109,bool),X110),s(X109,X113))),s(X114,X116),s(X114,X115))) | s(bool,'T_0') = s(bool,'const.sum.SUM_ALL_3'(s(fun(X109,bool),X110),s(fun(X111,bool),X112),s('type.sum.sum'(X109,X111),'const.sum.INL_1'(s(X109,X113)))))) )),
  inference(superposition,[],[f351,f211])).
fof(f7482,plain,(
  ( ! [X103,X101,X107,X105,X102,X108,X106,X104] : (s(X106,X107) = s(X106,'const.bool.COND_3'(s(bool,happ(s(fun(X103,bool),X104),s(X103,X105))),s(X106,X108),s(X106,X107))) | s(bool,'T_0') = s(bool,'const.sum.SUM_ALL_3'(s(fun(X101,bool),X102),s(fun(X103,bool),X104),s('type.sum.sum'(X101,X103),'const.sum.INR_1'(s(X103,X105)))))) )),
  inference(superposition,[],[f351,f210])).
fof(f7257,plain,(
  ( ! [X167,X165,X163,X166,X164,X162] : (s('type.sum.sum'(X164,bool),'const.sum.INR_1'(s(bool,X165))) != s('type.sum.sum'(X164,bool),'const.bool.COND_3'(s(bool,X166),s('type.sum.sum'(X164,bool),'const.sum.INR_1'(s(bool,'T_0'))),s('type.sum.sum'(X164,bool),'const.sum.INL_1'(s(X164,X167))))) | s(bool,X165) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X162),s(bool,X163))) | s(bool,'F_0') = s(bool,X162)) )),
  inference(superposition,[],[f213,f325])).
fof(f7256,plain,(
  ( ! [X158,X156,X161,X159,X157,X160] : (s('type.sum.sum'(X158,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X158,bool),'const.bool.COND_3'(s(bool,X159),s('type.sum.sum'(X158,bool),'const.sum.INR_1'(s(bool,X160))),s('type.sum.sum'(X158,bool),'const.sum.INL_1'(s(X158,X161))))) | s(bool,X160) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X156),s(bool,X157))) | s(bool,'F_0') = s(bool,X156)) )),
  inference(superposition,[],[f213,f325])).
fof(f7309,plain,(
  ( ! [X140,X138,X141] : (s('type.sum.sum'(X141,bool),'const.sum.INR_1'(s(bool,X138))) = s('type.sum.sum'(X141,bool),'const.sum.INR_1'(s(bool,X140))) | s(bool,'T_0') != s(bool,X140) | s(bool,'F_0') = s(bool,X138)) )),
  inference(forward_demodulation,[],[f7252,f193])).
fof(f7252,plain,(
  ( ! [X140,X138,X141,X139] : (s(bool,'T_0') != s(bool,X140) | s('type.sum.sum'(X141,bool),'const.sum.INR_1'(s(bool,X140))) = s('type.sum.sum'(X141,bool),'const.sum.INR_1'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X138),s(bool,X139))))) | s(bool,'F_0') = s(bool,X138)) )),
  inference(superposition,[],[f209,f325])).
fof(f7250,plain,(
  ( ! [X132,X130,X133,X131] : (s('type.sum.sum'(X132,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X132,bool),'const.sum.INR_1'(s(bool,X133))) | s(bool,X133) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X130),s(bool,X131))) | s(bool,'F_0') = s(bool,X130)) )),
  inference(superposition,[],[f208,f325])).
fof(f7304,plain,(
  ( ! [X54,X55,X53] : (s(bool,X55) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X53),s(bool,X54))) | s(bool,'T_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f7219,f157])).
fof(f7219,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,'T_0')) | s(bool,X55) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X53),s(bool,X54))) | s(bool,'T_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X53)) )),
  inference(superposition,[],[f274,f325])).
fof(f7303,plain,(
  ( ! [X39,X41,X38,X36,X42,X40] : (s('type.sum.sum'(X41,X38),'const.sum.INR_1'(s(X38,X40))) = s('type.sum.sum'(X41,X38),'const.bool.COND_3'(s(bool,X36),s('type.sum.sum'(X41,X38),'const.sum.INR_1'(s(X38,X39))),s('type.sum.sum'(X41,X38),'const.sum.INL_1'(s(X41,X42))))) | s(X38,X39) != s(X38,X40) | s(bool,'F_0') = s(bool,X36)) )),
  inference(forward_demodulation,[],[f7302,f193])).
fof(f7302,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(X38,X39) != s(X38,X40) | s('type.sum.sum'(X41,X38),'const.sum.INR_1'(s(X38,X40))) = s('type.sum.sum'(X41,X38),'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X36),s(bool,X37))),s('type.sum.sum'(X41,X38),'const.sum.INR_1'(s(X38,X39))),s('type.sum.sum'(X41,X38),'const.sum.INL_1'(s(X41,X42))))) | s(bool,'F_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f7216,f157])).
fof(f7216,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (~p(s(bool,'T_0')) | s(X38,X39) != s(X38,X40) | s('type.sum.sum'(X41,X38),'const.sum.INR_1'(s(X38,X40))) = s('type.sum.sum'(X41,X38),'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X36),s(bool,X37))),s('type.sum.sum'(X41,X38),'const.sum.INR_1'(s(X38,X39))),s('type.sum.sum'(X41,X38),'const.sum.INL_1'(s(X41,X42))))) | s(bool,'F_0') = s(bool,X36)) )),
  inference(superposition,[],[f214,f325])).
fof(f7301,plain,(
  ( ! [X21,X19] : (s(bool,X19) = s(bool,X21) | ~p(s(bool,X21)) | s(bool,'F_0') = s(bool,X19)) )),
  inference(forward_demodulation,[],[f7300,f193])).
fof(f7300,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X21)) | s(bool,X21) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X19),s(bool,X20))) | s(bool,'F_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f7213,f157])).
fof(f7213,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | s(bool,X21) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X19),s(bool,X20))) | s(bool,'F_0') = s(bool,X19)) )),
  inference(superposition,[],[f204,f325])).
fof(f7299,plain,(
  ( ! [X12,X10] : (s(bool,X10) != s(bool,X12) | p(s(bool,X12)) | s(bool,'F_0') = s(bool,X10)) )),
  inference(forward_demodulation,[],[f7298,f193])).
fof(f7298,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f7210,f157])).
fof(f7210,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X12)) | s(bool,X12) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11))) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f200,f325])).
fof(f7159,plain,(
  ( ! [X107,X105,X106] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X105),s(bool,X107))) | s(bool,'F_0') = s(bool,sK8(X105,X106,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X106)))) )),
  inference(superposition,[],[f325,f425])).
fof(f7158,plain,(
  ( ! [X103,X102,X104] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X103),s(bool,X104))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X102),s(bool,X103)))) )),
  inference(superposition,[],[f325,f194])).
fof(f7157,plain,(
  ( ! [X101,X99,X100] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X99),s(bool,X101))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X99),s(bool,X100)))) )),
  inference(superposition,[],[f325,f193])).
fof(f7156,plain,(
  ( ! [X94,X97,X95,X93,X98,X96] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X93,bool),X96),s(X93,X95))),s(bool,X98))) | s(bool,'F_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X93,X94),'const.sum.INL_1'(s(X93,X95))),s(fun(X93,bool),X96),s(fun(X94,bool),X97)))) )),
  inference(superposition,[],[f325,f228])).
fof(f7155,plain,(
  ( ! [X92,X90,X88,X87,X91,X89] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X88,bool),X91),s(X88,X89))),s(bool,X92))) | s(bool,'F_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X87,X88),'const.sum.INR_1'(s(X88,X89))),s(fun(X87,bool),X90),s(fun(X88,bool),X91)))) )),
  inference(superposition,[],[f325,f227])).
fof(f7154,plain,(
  ( ! [X85,X83,X81,X86,X84,X82] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X82,bool),X83),s(X82,X85))),s(bool,X86))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(X81,X82),bool),sK16(X83,X84,X82,bool,X81)),s('type.sum.sum'(X81,X82),'const.sum.INR_1'(s(X82,X85)))))) )),
  inference(superposition,[],[f325,f226])).
fof(f7153,plain,(
  ( ! [X80,X78,X76,X79,X77,X75] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X75,bool),X78),s(X75,X79))),s(bool,X80))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(X75,X76),bool),sK16(X77,X78,X76,bool,X75)),s('type.sum.sum'(X75,X76),'const.sum.INL_1'(s(X75,X79)))))) )),
  inference(superposition,[],[f325,f225])).
fof(f7149,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X59,bool),X60),s(X59,X63))),s(bool,X64))) | s(bool,'F_0') = s(bool,'const.sum.SUM_ALL_3'(s(fun(X59,bool),X60),s(fun(X61,bool),X62),s('type.sum.sum'(X59,X61),'const.sum.INL_1'(s(X59,X63)))))) )),
  inference(superposition,[],[f325,f211])).
fof(f7148,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,happ(s(fun(X55,bool),X56),s(X55,X57))),s(bool,X58))) | s(bool,'F_0') = s(bool,'const.sum.SUM_ALL_3'(s(fun(X53,bool),X54),s(fun(X55,bool),X56),s('type.sum.sum'(X53,X55),'const.sum.INR_1'(s(X55,X57)))))) )),
  inference(superposition,[],[f325,f210])).
fof(f7117,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f325,f272])).
fof(f6937,plain,(
  ( ! [X212,X210,X215,X213,X211,X209,X214] : (s('type.sum.sum'(X211,bool),'const.sum.INR_1'(s(bool,X212))) != s('type.sum.sum'(X211,bool),'const.bool.COND_3'(s(bool,X213),s('type.sum.sum'(X211,bool),'const.sum.INR_1'(s(bool,'T_0'))),s('type.sum.sum'(X211,bool),'const.sum.INL_1'(s(X211,X214))))) | s(bool,X212) = s(bool,'|dollar|exists_1'(s(fun(X209,bool),X210))) | s(X209,X215) = s(X209,sK8(X215,X210,X209))) )),
  inference(superposition,[],[f213,f425])).
fof(f6936,plain,(
  ( ! [X208,X206,X204,X202,X207,X205,X203] : (s('type.sum.sum'(X204,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X204,bool),'const.bool.COND_3'(s(bool,X205),s('type.sum.sum'(X204,bool),'const.sum.INR_1'(s(bool,X206))),s('type.sum.sum'(X204,bool),'const.sum.INL_1'(s(X204,X207))))) | s(bool,X206) = s(bool,'|dollar|exists_1'(s(fun(X202,bool),X203))) | s(X202,X208) = s(X202,sK8(X208,X203,X202))) )),
  inference(superposition,[],[f213,f425])).
fof(f6932,plain,(
  ( ! [X182,X180,X184,X183,X181] : (s(bool,'T_0') != s(bool,X182) | s('type.sum.sum'(X183,bool),'const.sum.INR_1'(s(bool,X182))) = s('type.sum.sum'(X183,bool),'const.sum.INR_1'(s(bool,'|dollar|exists_1'(s(fun(X180,bool),X181))))) | s(X180,X184) = s(X180,sK8(X184,X181,X180))) )),
  inference(superposition,[],[f209,f425])).
fof(f6930,plain,(
  ( ! [X173,X171,X174,X172,X170] : (s('type.sum.sum'(X172,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X172,bool),'const.sum.INR_1'(s(bool,X173))) | s(bool,X173) = s(bool,'|dollar|exists_1'(s(fun(X170,bool),X171))) | s(X170,X174) = s(X170,sK8(X174,X171,X170))) )),
  inference(superposition,[],[f208,f425])).
fof(f6969,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,X73) != s(bool,'|dollar|exists_1'(s(fun(X71,bool),X72))) | s(bool,'T_0') = s(bool,X73) | s(X71,X74) = s(X71,sK8(X74,X72,X71))) )),
  inference(subsumption_resolution,[],[f6902,f157])).
fof(f6902,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | s(bool,X73) != s(bool,'|dollar|exists_1'(s(fun(X71,bool),X72))) | s(bool,'T_0') = s(bool,X73) | s(X71,X74) = s(X71,sK8(X74,X72,X71))) )),
  inference(superposition,[],[f274,f425])).
fof(f6968,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (s(X53,X54) != s(X53,X55) | s('type.sum.sum'(X56,X53),'const.sum.INR_1'(s(X53,X55))) = s('type.sum.sum'(X56,X53),'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))),s('type.sum.sum'(X56,X53),'const.sum.INR_1'(s(X53,X54))),s('type.sum.sum'(X56,X53),'const.sum.INL_1'(s(X56,X57))))) | s(X51,X58) = s(X51,sK8(X58,X52,X51))) )),
  inference(subsumption_resolution,[],[f6899,f157])).
fof(f6899,plain,(
  ( ! [X57,X54,X52,X58,X56,X55,X53,X51] : (~p(s(bool,'T_0')) | s(X53,X54) != s(X53,X55) | s('type.sum.sum'(X56,X53),'const.sum.INR_1'(s(X53,X55))) = s('type.sum.sum'(X56,X53),'const.bool.COND_3'(s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52))),s('type.sum.sum'(X56,X53),'const.sum.INR_1'(s(X53,X54))),s('type.sum.sum'(X56,X53),'const.sum.INL_1'(s(X56,X57))))) | s(X51,X58) = s(X51,sK8(X58,X52,X51))) )),
  inference(superposition,[],[f214,f425])).
fof(f6967,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK8(X34,X32,X31))) )),
  inference(subsumption_resolution,[],[f6896,f157])).
fof(f6896,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | s(bool,X33) = s(bool,'|dollar|exists_1'(s(fun(X31,bool),X32))) | s(X31,X34) = s(X31,sK8(X34,X32,X31))) )),
  inference(superposition,[],[f204,f425])).
fof(f6966,plain,(
  ( ! [X21,X19,X22,X20] : (p(s(bool,X21)) | s(bool,X21) != s(bool,'|dollar|exists_1'(s(fun(X19,bool),X20))) | s(X19,X22) = s(X19,sK8(X22,X20,X19))) )),
  inference(subsumption_resolution,[],[f6893,f157])).
fof(f6893,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'T_0')) | p(s(bool,X21)) | s(bool,X21) != s(bool,'|dollar|exists_1'(s(fun(X19,bool),X20))) | s(X19,X22) = s(X19,sK8(X22,X20,X19))) )),
  inference(superposition,[],[f200,f425])).
fof(f6886,plain,(
  ( ! [X424,X423] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X423))) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK8(X423,X424,sK1)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK8(X423,X424,sK1)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK1,bool),X424)))) )),
  inference(superposition,[],[f260,f425])).
fof(f6885,plain,(
  ( ! [X422,X421] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,X421))) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK8(X421,X422,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK8(X421,X422,sK0)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK0,bool),X422)))) )),
  inference(superposition,[],[f259,f425])).
fof(f6884,plain,(
  ( ! [X414,X420,X416,X418,X413,X415,X417,X419] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X413,bool),X417),s(fun(X414,bool),X418),s('type.sum.sum'(X413,X414),X415)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X413,bool),X419),s(fun(X414,bool),X420),s('type.sum.sum'(X413,X414),sK8(X415,X416,'type.sum.sum'(X413,X414)))))) | p(s(bool,happ(s(fun(X414,bool),X418),s(X414,sK15(X417,X419,X418,X420,sK8(X415,X416,'type.sum.sum'(X413,X414)),X414,X413))))) | p(s(bool,happ(s(fun(X413,bool),X417),s(X413,sK14(X417,X419,X418,X420,sK8(X415,X416,'type.sum.sum'(X413,X414)),X414,X413))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.sum.sum'(X413,X414),bool),X416)))) )),
  inference(superposition,[],[f215,f425])).
fof(f6883,plain,(
  ( ! [X412,X408,X410,X405,X407,X403,X409,X411,X404,X406] : (s('type.sum.sum'(X403,X404),X405) != s('type.sum.sum'(X403,X404),X407) | s('type.sum.sum'(X403,X404),X407) = s('type.sum.sum'(X403,X404),'const.sum.INR_1'(s(X404,sK18(X408,X409,X407,sK8(X405,X406,'type.sum.sum'(X403,X404)),X410,X411,X412,X404,X403)))) | s('type.sum.sum'(X403,X404),X407) = s('type.sum.sum'(X403,X404),'const.sum.INL_1'(s(X403,sK17(X408,X409,X407,sK8(X405,X406,'type.sum.sum'(X403,X404)),X410,X411,X412,X404,X403)))) | s(X412,'const.sum.sum_CASE_3'(s('type.sum.sum'(X403,X404),X407),s(fun(X403,X412),X410),s(fun(X404,X412),X411))) = s(X412,'const.sum.sum_CASE_3'(s('type.sum.sum'(X403,X404),sK8(X405,X406,'type.sum.sum'(X403,X404))),s(fun(X403,X412),X409),s(fun(X404,X412),X408))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.sum.sum'(X403,X404),bool),X406)))) )),
  inference(superposition,[],[f229,f425])).
fof(f6882,plain,(
  ( ! [X397,X399,X393,X395,X401,X396,X398,X394,X400,X402] : (s('type.sum.sum'(X393,X394),X395) != s('type.sum.sum'(X393,X394),X397) | s('type.sum.sum'(X393,X394),sK8(X395,X396,'type.sum.sum'(X393,X394))) = s('type.sum.sum'(X393,X394),'const.sum.INR_1'(s(X394,sK18(X398,X399,sK8(X395,X396,'type.sum.sum'(X393,X394)),X397,X400,X401,X402,X394,X393)))) | s('type.sum.sum'(X393,X394),sK8(X395,X396,'type.sum.sum'(X393,X394))) = s('type.sum.sum'(X393,X394),'const.sum.INL_1'(s(X393,sK17(X398,X399,sK8(X395,X396,'type.sum.sum'(X393,X394)),X397,X400,X401,X402,X394,X393)))) | s(X402,'const.sum.sum_CASE_3'(s('type.sum.sum'(X393,X394),X397),s(fun(X393,X402),X399),s(fun(X394,X402),X398))) = s(X402,'const.sum.sum_CASE_3'(s('type.sum.sum'(X393,X394),sK8(X395,X396,'type.sum.sum'(X393,X394))),s(fun(X393,X402),X400),s(fun(X394,X402),X401))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun('type.sum.sum'(X393,X394),bool),X396)))) )),
  inference(superposition,[],[f229,f425])).
fof(f6881,plain,(
  ( ! [X383,X389,X391,X385,X387,X388,X390,X384,X386,X392] : (s(X384,happ(s(fun(X383,X384),X385),s(X383,sK18(sK8(X385,X386,fun(X383,X384)),X388,X389,X390,X391,X387,X384,X383,X392)))) != s(X384,happ(s(fun(X383,X384),X387),s(X383,sK18(sK8(X385,X386,fun(X383,X384)),X388,X389,X390,X391,X387,X384,X383,X392)))) | s(X384,'const.sum.sum_CASE_3'(s('type.sum.sum'(X392,X383),X389),s(fun(X392,X384),X391),s(fun(X383,X384),X387))) = s(X384,'const.sum.sum_CASE_3'(s('type.sum.sum'(X392,X383),X390),s(fun(X392,X384),X388),s(fun(X383,X384),sK8(X385,X386,fun(X383,X384))))) | s(X384,happ(s(fun(X392,X384),X388),s(X392,sK17(sK8(X385,X386,fun(X383,X384)),X388,X389,X390,X391,X387,X384,X383,X392)))) != s(X384,happ(s(fun(X392,X384),X391),s(X392,sK17(sK8(X385,X386,fun(X383,X384)),X388,X389,X390,X391,X387,X384,X383,X392)))) | s('type.sum.sum'(X392,X383),X389) != s('type.sum.sum'(X392,X383),X390) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X383,X384),bool),X386)))) )),
  inference(superposition,[],[f232,f425])).
fof(f6880,plain,(
  ( ! [X377,X379,X374,X380,X382,X376,X378,X373,X375,X381] : (s(X374,happ(s(fun(X373,X374),X375),s(X373,sK18(X377,X378,X379,X380,X381,sK8(X375,X376,fun(X373,X374)),X374,X373,X382)))) != s(X374,happ(s(fun(X373,X374),X377),s(X373,sK18(X377,X378,X379,X380,X381,sK8(X375,X376,fun(X373,X374)),X374,X373,X382)))) | s(X374,'const.sum.sum_CASE_3'(s('type.sum.sum'(X382,X373),X380),s(fun(X382,X374),X378),s(fun(X373,X374),X377))) = s(X374,'const.sum.sum_CASE_3'(s('type.sum.sum'(X382,X373),X379),s(fun(X382,X374),X381),s(fun(X373,X374),sK8(X375,X376,fun(X373,X374))))) | s(X374,happ(s(fun(X382,X374),X378),s(X382,sK17(X377,X378,X379,X380,X381,sK8(X375,X376,fun(X373,X374)),X374,X373,X382)))) != s(X374,happ(s(fun(X382,X374),X381),s(X382,sK17(X377,X378,X379,X380,X381,sK8(X375,X376,fun(X373,X374)),X374,X373,X382)))) | s('type.sum.sum'(X382,X373),X379) != s('type.sum.sum'(X382,X373),X380) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X373,X374),bool),X376)))) )),
  inference(superposition,[],[f232,f425])).
fof(f6879,plain,(
  ( ! [X364,X366,X372,X368,X370,X365,X367,X363,X369,X371] : (s(X364,happ(s(fun(X363,X364),X365),s(X363,sK18(sK8(X365,X366,fun(X363,X364)),X368,X369,X370,X371,X367,X364,X363,X372)))) != s(X364,happ(s(fun(X363,X364),X367),s(X363,sK18(sK8(X365,X366,fun(X363,X364)),X368,X369,X370,X371,X367,X364,X363,X372)))) | s(X364,'const.sum.sum_CASE_3'(s('type.sum.sum'(X372,X363),X369),s(fun(X372,X364),X371),s(fun(X363,X364),X367))) = s(X364,'const.sum.sum_CASE_3'(s('type.sum.sum'(X372,X363),X370),s(fun(X372,X364),X368),s(fun(X363,X364),sK8(X365,X366,fun(X363,X364))))) | s('type.sum.sum'(X372,X363),X369) = s('type.sum.sum'(X372,X363),'const.sum.INL_1'(s(X372,sK17(sK8(X365,X366,fun(X363,X364)),X368,X369,X370,X371,X367,X364,X363,X372)))) | s('type.sum.sum'(X372,X363),X369) != s('type.sum.sum'(X372,X363),X370) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X363,X364),bool),X366)))) )),
  inference(superposition,[],[f231,f425])).
fof(f6878,plain,(
  ( ! [X356,X358,X354,X360,X362,X357,X359,X353,X355,X361] : (s(X354,happ(s(fun(X353,X354),X355),s(X353,sK18(X357,X358,X359,X360,X361,sK8(X355,X356,fun(X353,X354)),X354,X353,X362)))) != s(X354,happ(s(fun(X353,X354),X357),s(X353,sK18(X357,X358,X359,X360,X361,sK8(X355,X356,fun(X353,X354)),X354,X353,X362)))) | s(X354,'const.sum.sum_CASE_3'(s('type.sum.sum'(X362,X353),X360),s(fun(X362,X354),X358),s(fun(X353,X354),X357))) = s(X354,'const.sum.sum_CASE_3'(s('type.sum.sum'(X362,X353),X359),s(fun(X362,X354),X361),s(fun(X353,X354),sK8(X355,X356,fun(X353,X354))))) | s('type.sum.sum'(X362,X353),X359) = s('type.sum.sum'(X362,X353),'const.sum.INL_1'(s(X362,sK17(X357,X358,X359,X360,X361,sK8(X355,X356,fun(X353,X354)),X354,X353,X362)))) | s('type.sum.sum'(X362,X353),X359) != s('type.sum.sum'(X362,X353),X360) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X353,X354),bool),X356)))) )),
  inference(superposition,[],[f231,f425])).
fof(f6877,plain,(
  ( ! [X350,X344,X346,X352,X343,X349,X351,X345,X347,X348] : (s(X344,happ(s(fun(X343,X344),X345),s(X343,sK17(X347,X348,X349,X350,sK8(X345,X346,fun(X343,X344)),X351,X344,X352,X343)))) != s(X344,happ(s(fun(X343,X344),X348),s(X343,sK17(X347,X348,X349,X350,sK8(X345,X346,fun(X343,X344)),X351,X344,X352,X343)))) | s('type.sum.sum'(X343,X352),X349) = s('type.sum.sum'(X343,X352),'const.sum.INR_1'(s(X352,sK18(X347,X348,X349,X350,sK8(X345,X346,fun(X343,X344)),X351,X344,X352,X343)))) | s(X344,'const.sum.sum_CASE_3'(s('type.sum.sum'(X343,X352),X350),s(fun(X343,X344),X348),s(fun(X352,X344),X347))) = s(X344,'const.sum.sum_CASE_3'(s('type.sum.sum'(X343,X352),X349),s(fun(X343,X344),sK8(X345,X346,fun(X343,X344))),s(fun(X352,X344),X351))) | s('type.sum.sum'(X343,X352),X349) != s('type.sum.sum'(X343,X352),X350) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X343,X344),bool),X346)))) )),
  inference(superposition,[],[f230,f425])).
fof(f6876,plain,(
  ( ! [X333,X335,X341,X337,X339,X334,X340,X342,X336,X338] : (s(X334,happ(s(fun(X333,X334),X335),s(X333,sK17(X338,sK8(X335,X336,fun(X333,X334)),X339,X340,X337,X341,X334,X342,X333)))) != s(X334,happ(s(fun(X333,X334),X337),s(X333,sK17(X338,sK8(X335,X336,fun(X333,X334)),X339,X340,X337,X341,X334,X342,X333)))) | s('type.sum.sum'(X333,X342),X339) = s('type.sum.sum'(X333,X342),'const.sum.INR_1'(s(X342,sK18(X338,sK8(X335,X336,fun(X333,X334)),X339,X340,X337,X341,X334,X342,X333)))) | s(X334,'const.sum.sum_CASE_3'(s('type.sum.sum'(X333,X342),X339),s(fun(X333,X334),X337),s(fun(X342,X334),X341))) = s(X334,'const.sum.sum_CASE_3'(s('type.sum.sum'(X333,X342),X340),s(fun(X333,X334),sK8(X335,X336,fun(X333,X334))),s(fun(X342,X334),X338))) | s('type.sum.sum'(X333,X342),X339) != s('type.sum.sum'(X333,X342),X340) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X333,X334),bool),X336)))) )),
  inference(superposition,[],[f230,f425])).
fof(f6871,plain,(
  ( ! [X300,X302,X304,X301,X303] : (~p(s(bool,happ(s(fun('type.sum.sum'(X300,X301),bool),X302),s('type.sum.sum'(X300,X301),'const.sum.INR_1'(s(X301,sK12(sK8(X302,X303,fun('type.sum.sum'(X300,X301),bool)),X301,X300))))))) | p(s(bool,happ(s(fun('type.sum.sum'(X300,X301),bool),sK8(X302,X303,fun('type.sum.sum'(X300,X301),bool))),s('type.sum.sum'(X300,X301),X304)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X300,X301),bool),sK8(X302,X303,fun('type.sum.sum'(X300,X301),bool))),s('type.sum.sum'(X300,X301),'const.sum.INL_1'(s(X300,sK11(sK8(X302,X303,fun('type.sum.sum'(X300,X301),bool)),X301,X300))))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun('type.sum.sum'(X300,X301),bool),bool),X303)))) )),
  inference(superposition,[],[f206,f425])).
fof(f6870,plain,(
  ( ! [X292,X294,X296,X298,X293,X295,X297,X299] : (~p(s(bool,happ(s(fun(X292,bool),X293),s(X292,sK15(X295,X296,X297,sK8(X293,X294,fun(X292,bool)),X298,X292,X299))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X299,bool),X295),s(fun(X292,bool),X297),s('type.sum.sum'(X299,X292),X298)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X299,bool),X296),s(fun(X292,bool),sK8(X293,X294,fun(X292,bool))),s('type.sum.sum'(X299,X292),X298)))) | ~p(s(bool,happ(s(fun(X299,bool),X296),s(X299,sK14(X295,X296,X297,sK8(X293,X294,fun(X292,bool)),X298,X292,X299))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X292,bool),bool),X294)))) )),
  inference(superposition,[],[f218,f425])).
fof(f6869,plain,(
  ( ! [X288,X290,X285,X287,X289,X291,X284,X286] : (~p(s(bool,happ(s(fun(X284,bool),X285),s(X284,sK15(X287,X288,X289,sK8(X285,X286,fun(X284,bool)),X290,X284,X291))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X291,bool),X287),s(fun(X284,bool),X289),s('type.sum.sum'(X291,X284),X290)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X291,bool),X288),s(fun(X284,bool),sK8(X285,X286,fun(X284,bool))),s('type.sum.sum'(X291,X284),X290)))) | p(s(bool,happ(s(fun(X291,bool),X287),s(X291,sK14(X287,X288,X289,sK8(X285,X286,fun(X284,bool)),X290,X284,X291))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X284,bool),bool),X286)))) )),
  inference(superposition,[],[f217,f425])).
fof(f6868,plain,(
  ( ! [X280,X282,X277,X279,X281,X283,X276,X278] : (~p(s(bool,happ(s(fun(X276,bool),X277),s(X276,sK14(X279,sK8(X277,X278,fun(X276,bool)),X280,X281,X282,X283,X276))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X276,bool),X279),s(fun(X283,bool),X280),s('type.sum.sum'(X276,X283),X282)))) | p(s(bool,happ(s(fun(X283,bool),X280),s(X283,sK15(X279,sK8(X277,X278,fun(X276,bool)),X280,X281,X282,X283,X276))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X276,bool),sK8(X277,X278,fun(X276,bool))),s(fun(X283,bool),X281),s('type.sum.sum'(X276,X283),X282)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X276,bool),bool),X278)))) )),
  inference(superposition,[],[f216,f425])).
fof(f6867,plain,(
  ( ! [X269,X271,X273,X275,X268,X270,X272,X274] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X268,bool),X269),s(fun(X271,bool),X272),s('type.sum.sum'(X268,X271),X273)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X268,bool),X274),s(fun(X271,bool),X275),s('type.sum.sum'(X268,X271),X273)))) | p(s(bool,happ(s(fun(X271,bool),X272),s(X271,sK15(sK8(X269,X270,fun(X268,bool)),X274,X272,X275,X273,X271,X268))))) | p(s(bool,happ(s(fun(X268,bool),sK8(X269,X270,fun(X268,bool))),s(X268,sK14(sK8(X269,X270,fun(X268,bool)),X274,X272,X275,X273,X271,X268))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X268,bool),bool),X270)))) )),
  inference(superposition,[],[f215,f425])).
fof(f6866,plain,(
  ( ! [X261,X263,X265,X267,X260,X262,X264,X266] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X263,bool),X264),s(fun(X260,bool),X261),s('type.sum.sum'(X263,X260),X265)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X263,bool),X266),s(fun(X260,bool),X267),s('type.sum.sum'(X263,X260),X265)))) | p(s(bool,happ(s(fun(X260,bool),sK8(X261,X262,fun(X260,bool))),s(X260,sK15(X264,X266,sK8(X261,X262,fun(X260,bool)),X267,X265,X260,X263))))) | p(s(bool,happ(s(fun(X263,bool),X264),s(X263,sK14(X264,X266,sK8(X261,X262,fun(X260,bool)),X267,X265,X260,X263))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X260,bool),bool),X262)))) )),
  inference(superposition,[],[f215,f425])).
fof(f6861,plain,(
  ( ! [X235,X233,X234,X232] : (~p(s(bool,happ(s(fun(X232,bool),X233),s(X232,sK8(X235,sK8(X233,X234,fun(X232,bool)),X232))))) | p(s(bool,'|dollar|exists_1'(s(fun(X232,bool),sK8(X233,X234,fun(X232,bool)))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X232,bool),bool),X234)))) )),
  inference(superposition,[],[f196,f425])).
fof(f6860,plain,(
  ( ! [X231,X229,X230,X228] : (p(s(bool,'|dollar|exists_1'(s(fun(X228,bool),X229)))) | s(X228,X231) = s(X228,sK8(X231,sK8(X229,X230,fun(X228,bool)),X228)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X228,bool),bool),X230)))) )),
  inference(superposition,[],[f195,f425])).
fof(f6837,plain,(
  ( ! [X173,X174,X172] : (p(s(bool,X172)) | s(bool,X174) = s(bool,sK8(X172,X173,bool)) | s(bool,'T_0') = s(bool,X174) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X173)))) )),
  inference(superposition,[],[f306,f425])).
fof(f6836,plain,(
  ( ! [X171,X169,X170] : (~p(s(bool,X169)) | s(bool,X171) != s(bool,sK8(X169,X170,bool)) | s(bool,'T_0') = s(bool,X171) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X170)))) )),
  inference(superposition,[],[f274,f425])).
fof(f6831,plain,(
  ( ! [X156,X154,X152,X151,X155,X153,X150] : (~p(s(bool,X150)) | s(X152,X153) != s(X152,X154) | s('type.sum.sum'(X155,X152),'const.sum.INR_1'(s(X152,X154))) = s('type.sum.sum'(X155,X152),'const.bool.COND_3'(s(bool,sK8(X150,X151,bool)),s('type.sum.sum'(X155,X152),'const.sum.INR_1'(s(X152,X153))),s('type.sum.sum'(X155,X152),'const.sum.INL_1'(s(X155,X156))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X151)))) )),
  inference(superposition,[],[f214,f425])).
fof(f6829,plain,(
  ( ! [X142,X140,X138,X136,X141,X139,X137] : (p(s(bool,X136)) | s('type.sum.sum'(X138,X139),'const.sum.INR_1'(s(X139,X140))) != s('type.sum.sum'(X138,X139),'const.bool.COND_3'(s(bool,sK8(X136,X137,bool)),s('type.sum.sum'(X138,X139),'const.sum.INR_1'(s(X139,X141))),s('type.sum.sum'(X138,X139),'const.sum.INL_1'(s(X138,X142))))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X137)))) )),
  inference(superposition,[],[f212,f425])).
fof(f6828,plain,(
  ( ! [X134,X135,X133] : (~p(s(bool,X133)) | ~p(s(bool,X135)) | s(bool,X135) = s(bool,sK8(X133,X134,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X134)))) )),
  inference(superposition,[],[f204,f425])).
fof(f6826,plain,(
  ( ! [X127,X128,X129] : (p(s(bool,X127)) | p(s(bool,X129)) | s(bool,X129) = s(bool,sK8(X127,X128,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X128)))) )),
  inference(superposition,[],[f201,f425])).
fof(f6825,plain,(
  ( ! [X125,X126,X124] : (~p(s(bool,X124)) | p(s(bool,X126)) | s(bool,X126) != s(bool,sK8(X124,X125,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X125)))) )),
  inference(superposition,[],[f200,f425])).
fof(f6824,plain,(
  ( ! [X123,X122] : (p(s(bool,X122)) | s(bool,'F_0') = s(bool,sK8(X122,X123,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X123)))) )),
  inference(superposition,[],[f170,f425])).
fof(f6823,plain,(
  ( ! [X121,X120] : (s(bool,'F_0') != s(bool,X120) | ~p(s(bool,sK8(X120,X121,bool))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X121)))) )),
  inference(superposition,[],[f169,f425])).
fof(f6822,plain,(
  ( ! [X118,X119] : (~p(s(bool,X118)) | s(bool,'T_0') = s(bool,sK8(X118,X119,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X119)))) )),
  inference(superposition,[],[f166,f425])).
fof(f6821,plain,(
  ( ! [X116,X117] : (p(s(bool,X116)) | s(bool,'T_0') != s(bool,sK8(X116,X117,bool)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X117)))) )),
  inference(superposition,[],[f165,f425])).
fof(f6810,plain,(
  ( ! [X61,X59,X64,X62,X60,X58,X63] : (s('type.sum.sum'(X61,X58),'const.sum.INR_1'(s(X58,X62))) != s('type.sum.sum'(X61,X58),'const.bool.COND_3'(s(bool,X63),s('type.sum.sum'(X61,X58),'const.sum.INR_1'(s(X58,X59))),s('type.sum.sum'(X61,X58),'const.sum.INL_1'(s(X61,X64))))) | s(X58,X62) = s(X58,sK8(X59,X60,X58)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X58,bool),X60)))) )),
  inference(superposition,[],[f213,f425])).
fof(f6809,plain,(
  ( ! [X57,X54,X52,X56,X55,X53,X51] : (s('type.sum.sum'(X54,X51),'const.sum.INR_1'(s(X51,X52))) != s('type.sum.sum'(X54,X51),'const.bool.COND_3'(s(bool,X55),s('type.sum.sum'(X54,X51),'const.sum.INR_1'(s(X51,X56))),s('type.sum.sum'(X54,X51),'const.sum.INL_1'(s(X54,X57))))) | s(X51,X56) = s(X51,sK8(X52,X53,X51)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X51,bool),X53)))) )),
  inference(superposition,[],[f213,f425])).
fof(f6805,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(X29,X30) != s(X29,X32) | s('type.sum.sum'(X33,X29),'const.sum.INR_1'(s(X29,X32))) = s('type.sum.sum'(X33,X29),'const.sum.INR_1'(s(X29,sK8(X30,X31,X29)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X29,bool),X31)))) )),
  inference(superposition,[],[f209,f425])).
fof(f6803,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s('type.sum.sum'(X22,X19),'const.sum.INR_1'(s(X19,X20))) != s('type.sum.sum'(X22,X19),'const.sum.INR_1'(s(X19,X23))) | s(X19,X23) = s(X19,sK8(X20,X21,X19)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X19,bool),X21)))) )),
  inference(superposition,[],[f208,f425])).
fof(f6962,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(subsumption_resolution,[],[f6798,f166])).
fof(f6798,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f196,f425])).
fof(f6790,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X36,bool),X37))) | s(X36,X39) = s(X36,sK8(X39,sK8(X37,X38,fun(X36,bool)),X36)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(fun(X36,bool),bool),X38)))) )),
  inference(superposition,[],[f425,f425])).
fof(f6789,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X32,bool),X34))) | s(X32,X35) = s(X32,sK8(X35,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X32,bool),X33),s(fun(X32,bool),X34)),X32))) )),
  inference(superposition,[],[f425,f194])).
fof(f6788,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X28,bool),X29))) | s(X28,X31) = s(X28,sK8(X31,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X28,bool),X29),s(fun(X28,bool),X30)),X28))) )),
  inference(superposition,[],[f425,f193])).
fof(f6787,plain,(
  ( ! [X26,X24,X23,X21,X27,X25,X22] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X21,bool),happ(s(fun(X22,fun(X21,bool)),X25),s(X22,X24))))) | s(X21,X27) = s(X21,sK8(X27,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INL_1'(s(X22,X24))),s(fun(X22,fun(X21,bool)),X25),s(fun(X23,fun(X21,bool)),X26)),X21))) )),
  inference(superposition,[],[f425,f228])).
fof(f6786,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X14,bool),happ(s(fun(X16,fun(X14,bool)),X19),s(X16,X17))))) | s(X14,X20) = s(X14,sK8(X20,'const.sum.sum_CASE_3'(s('type.sum.sum'(X15,X16),'const.sum.INR_1'(s(X16,X17))),s(fun(X15,fun(X14,bool)),X18),s(fun(X16,fun(X14,bool)),X19)),X14))) )),
  inference(superposition,[],[f425,f227])).
fof(f6785,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X7,bool),happ(s(fun(X9,fun(X7,bool)),X10),s(X9,X12))))) | s(X7,X13) = s(X7,sK8(X13,happ(s(fun('type.sum.sum'(X8,X9),fun(X7,bool)),sK16(X10,X11,X9,fun(X7,bool),X8)),s('type.sum.sum'(X8,X9),'const.sum.INR_1'(s(X9,X12)))),X7))) )),
  inference(superposition,[],[f425,f226])).
fof(f6784,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X4),s(X1,X5))))) | s(X0,X6) = s(X0,sK8(X6,happ(s(fun('type.sum.sum'(X1,X2),fun(X0,bool)),sK16(X3,X4,X2,fun(X0,bool),X1)),s('type.sum.sum'(X1,X2),'const.sum.INL_1'(s(X1,X5)))),X0))) )),
  inference(superposition,[],[f425,f225])).
fof(f6777,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,sK8(X5,X6,bool))) )),
  inference(superposition,[],[f425,f272])).
fof(f6776,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X3))) | s(bool,X4) = s(bool,sK8(X2,X3,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f425,f272])).
fof(f6775,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,sK8(X0,X1,bool))) )),
  inference(superposition,[],[f425,f163])).
fof(f5770,plain,(
  ( ! [X509,X511,X505,X507,X512,X508,X510,X506,X513] : (s(X511,happ(s(fun(bool,X511),X505),s(bool,'T_0'))) != s(X511,happ(s(fun(bool,X511),X510),s(bool,'T_0'))) | s(X511,'const.sum.sum_CASE_3'(s('type.sum.sum'(X512,bool),X507),s(fun(X512,X511),X509),s(fun(bool,X511),X510))) = s(X511,'const.sum.sum_CASE_3'(s('type.sum.sum'(X512,bool),X508),s(fun(X512,X511),X506),s(fun(bool,X511),X505))) | s('type.sum.sum'(X512,bool),X507) = s('type.sum.sum'(X512,bool),'const.sum.INL_1'(s(X512,sK17(X505,X506,X507,X508,X509,X510,X511,bool,X512)))) | s('type.sum.sum'(X512,bool),X507) != s('type.sum.sum'(X512,bool),X508) | s(bool,X513) = s(bool,sK18(X505,X506,X507,X508,X509,X510,X511,bool,X512)) | s(bool,'T_0') = s(bool,X513)) )),
  inference(superposition,[],[f231,f272])).
fof(f5769,plain,(
  ( ! [X499,X500,X502,X496,X498,X504,X501,X503,X497] : (s(X502,happ(s(fun(bool,X502),X496),s(bool,'T_0'))) != s(X502,happ(s(fun(bool,X502),X501),s(bool,'T_0'))) | s(X502,'const.sum.sum_CASE_3'(s('type.sum.sum'(X503,bool),X498),s(fun(X503,X502),X500),s(fun(bool,X502),X501))) = s(X502,'const.sum.sum_CASE_3'(s('type.sum.sum'(X503,bool),X499),s(fun(X503,X502),X497),s(fun(bool,X502),X496))) | s(X502,happ(s(fun(X503,X502),X497),s(X503,sK17(X496,X497,X498,X499,X500,X501,X502,bool,X503)))) != s(X502,happ(s(fun(X503,X502),X500),s(X503,sK17(X496,X497,X498,X499,X500,X501,X502,bool,X503)))) | s('type.sum.sum'(X503,bool),X498) != s('type.sum.sum'(X503,bool),X499) | s(bool,X504) = s(bool,sK18(X496,X497,X498,X499,X500,X501,X502,bool,X503)) | s(bool,'T_0') = s(bool,X504)) )),
  inference(superposition,[],[f232,f272])).
fof(f5768,plain,(
  ( ! [X492,X494,X488,X490,X487,X493,X495,X489,X491] : (s(X493,happ(s(fun(bool,X493),X488),s(bool,'T_0'))) != s(X493,happ(s(fun(bool,X493),X491),s(bool,'T_0'))) | s('type.sum.sum'(bool,X494),X489) = s('type.sum.sum'(bool,X494),'const.sum.INR_1'(s(X494,sK18(X487,X488,X489,X490,X491,X492,X493,X494,bool)))) | s(X493,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X494),X489),s(fun(bool,X493),X491),s(fun(X494,X493),X492))) = s(X493,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X494),X490),s(fun(bool,X493),X488),s(fun(X494,X493),X487))) | s('type.sum.sum'(bool,X494),X489) != s('type.sum.sum'(bool,X494),X490) | s(bool,X495) = s(bool,sK17(X487,X488,X489,X490,X491,X492,X493,X494,bool)) | s(bool,'T_0') = s(bool,X495)) )),
  inference(superposition,[],[f230,f272])).
fof(f5767,plain,(
  ( ! [X484,X486,X480,X482,X485,X481,X483] : (~p(s(bool,happ(s(fun(bool,bool),X483),s(bool,'T_0')))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X485,bool),X480),s(fun(bool,bool),X482),s('type.sum.sum'(X485,bool),X484)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X485,bool),X481),s(fun(bool,bool),X483),s('type.sum.sum'(X485,bool),X484)))) | p(s(bool,happ(s(fun(X485,bool),X480),s(X485,sK14(X480,X481,X482,X483,X484,bool,X485))))) | s(bool,X486) = s(bool,sK15(X480,X481,X482,X483,X484,bool,X485)) | s(bool,'T_0') = s(bool,X486)) )),
  inference(superposition,[],[f217,f272])).
fof(f5766,plain,(
  ( ! [X476,X478,X474,X475,X477,X479,X473] : (~p(s(bool,happ(s(fun(bool,bool),X476),s(bool,'T_0')))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X478,bool),X473),s(fun(bool,bool),X475),s('type.sum.sum'(X478,bool),X477)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X478,bool),X474),s(fun(bool,bool),X476),s('type.sum.sum'(X478,bool),X477)))) | ~p(s(bool,happ(s(fun(X478,bool),X474),s(X478,sK14(X473,X474,X475,X476,X477,bool,X478))))) | s(bool,X479) = s(bool,sK15(X473,X474,X475,X476,X477,bool,X478)) | s(bool,'T_0') = s(bool,X479)) )),
  inference(superposition,[],[f218,f272])).
fof(f5765,plain,(
  ( ! [X466,X472,X469,X471,X467,X468,X470] : (~p(s(bool,happ(s(fun(bool,bool),X467),s(bool,'T_0')))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(bool,bool),X466),s(fun(X471,bool),X468),s('type.sum.sum'(bool,X471),X470)))) | p(s(bool,happ(s(fun(X471,bool),X468),s(X471,sK15(X466,X467,X468,X469,X470,X471,bool))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(bool,bool),X467),s(fun(X471,bool),X469),s('type.sum.sum'(bool,X471),X470)))) | s(bool,X472) = s(bool,sK14(X466,X467,X468,X469,X470,X471,bool)) | s(bool,'T_0') = s(bool,X472)) )),
  inference(superposition,[],[f216,f272])).
fof(f5764,plain,(
  ( ! [X463,X465,X462,X464] : (s(X464,f(s(bool,'T_0'))) != s(X464,g(s(bool,'T_0'))) | s(fun(bool,X464),X462) = s(fun(bool,X464),X463) | s(bool,X465) = s(bool,sK13(X462,X463,X464,bool)) | s(bool,'T_0') = s(bool,X465)) )),
  inference(superposition,[],[f207,f272])).
fof(f5763,plain,(
  ( ! [X461,X459,X460,X458] : (~p(s(bool,happ(s(fun('type.sum.sum'(X459,bool),bool),X458),s('type.sum.sum'(X459,bool),'const.sum.INR_1'(s(bool,'T_0')))))) | p(s(bool,happ(s(fun('type.sum.sum'(X459,bool),bool),X458),s('type.sum.sum'(X459,bool),X460)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X459,bool),bool),X458),s('type.sum.sum'(X459,bool),'const.sum.INL_1'(s(X459,sK11(X458,bool,X459))))))) | s(bool,X461) = s(bool,sK12(X458,bool,X459)) | s(bool,'T_0') = s(bool,X461)) )),
  inference(superposition,[],[f206,f272])).
fof(f5762,plain,(
  ( ! [X455,X457,X456] : (s('type.sum.sum'(X456,bool),X455) = s('type.sum.sum'(X456,bool),'const.sum.INR_1'(s(bool,'T_0'))) | s('type.sum.sum'(X456,bool),X455) = s('type.sum.sum'(X456,bool),'const.sum.INL_1'(s(X456,sK9(X455,bool,X456)))) | s(bool,X457) = s(bool,sK10(X455,bool,X456)) | s(bool,'T_0') = s(bool,X457)) )),
  inference(superposition,[],[f205,f272])).
fof(f5761,plain,(
  ( ! [X453,X452,X454] : (s('type.sum.sum'(bool,X453),X452) = s('type.sum.sum'(bool,X453),'const.sum.INL_1'(s(bool,'T_0'))) | s('type.sum.sum'(bool,X453),X452) = s('type.sum.sum'(bool,X453),'const.sum.INR_1'(s(X453,sK10(X452,X453,bool)))) | s(bool,X454) = s(bool,sK9(X452,X453,bool)) | s(bool,'T_0') = s(bool,X454)) )),
  inference(superposition,[],[f205,f272])).
fof(f5760,plain,(
  ( ! [X449,X451,X450] : (~p(s(bool,happ(s(fun(bool,bool),X450),s(bool,'T_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X450)))) | s(bool,X451) = s(bool,sK8(X449,X450,bool)) | s(bool,'T_0') = s(bool,X451)) )),
  inference(superposition,[],[f196,f272])).
fof(f5755,plain,(
  ( ! [X422,X428,X430,X424,X426,X423,X429,X425,X427] : (s(bool,'T_0') != s(bool,happ(s(fun(X422,bool),X428),s(X422,sK18(X423,X424,X425,X426,X427,X428,bool,X422,X429)))) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X429,X422),X425),s(fun(X429,bool),X427),s(fun(X422,bool),X428))) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X429,X422),X426),s(fun(X429,bool),X424),s(fun(X422,bool),X423))) | s('type.sum.sum'(X429,X422),X425) = s('type.sum.sum'(X429,X422),'const.sum.INL_1'(s(X429,sK17(X423,X424,X425,X426,X427,X428,bool,X422,X429)))) | s('type.sum.sum'(X429,X422),X425) != s('type.sum.sum'(X429,X422),X426) | s(bool,X430) = s(bool,happ(s(fun(X422,bool),X423),s(X422,sK18(X423,X424,X425,X426,X427,X428,bool,X422,X429)))) | s(bool,'T_0') = s(bool,X430)) )),
  inference(superposition,[],[f231,f272])).
fof(f5754,plain,(
  ( ! [X414,X420,X416,X418,X413,X415,X421,X417,X419] : (s(bool,'T_0') != s(bool,happ(s(fun(X413,bool),X419),s(X413,sK18(X414,X415,X416,X417,X418,X419,bool,X413,X420)))) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X420,X413),X416),s(fun(X420,bool),X418),s(fun(X413,bool),X419))) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X420,X413),X417),s(fun(X420,bool),X415),s(fun(X413,bool),X414))) | s(bool,happ(s(fun(X420,bool),X415),s(X420,sK17(X414,X415,X416,X417,X418,X419,bool,X413,X420)))) != s(bool,happ(s(fun(X420,bool),X418),s(X420,sK17(X414,X415,X416,X417,X418,X419,bool,X413,X420)))) | s('type.sum.sum'(X420,X413),X416) != s('type.sum.sum'(X420,X413),X417) | s(bool,X421) = s(bool,happ(s(fun(X413,bool),X414),s(X413,sK18(X414,X415,X416,X417,X418,X419,bool,X413,X420)))) | s(bool,'T_0') = s(bool,X421)) )),
  inference(superposition,[],[f232,f272])).
fof(f5753,plain,(
  ( ! [X412,X408,X410,X405,X407,X409,X411,X404,X406] : (s(bool,'T_0') != s(bool,happ(s(fun(X404,bool),X406),s(X404,sK18(X406,X407,X408,X409,X410,X405,bool,X404,X411)))) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X411,X404),X408),s(fun(X411,bool),X410),s(fun(X404,bool),X405))) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X411,X404),X409),s(fun(X411,bool),X407),s(fun(X404,bool),X406))) | s('type.sum.sum'(X411,X404),X408) = s('type.sum.sum'(X411,X404),'const.sum.INL_1'(s(X411,sK17(X406,X407,X408,X409,X410,X405,bool,X404,X411)))) | s('type.sum.sum'(X411,X404),X408) != s('type.sum.sum'(X411,X404),X409) | s(bool,X412) = s(bool,happ(s(fun(X404,bool),X405),s(X404,sK18(X406,X407,X408,X409,X410,X405,bool,X404,X411)))) | s(bool,'T_0') = s(bool,X412)) )),
  inference(superposition,[],[f231,f272])).
fof(f5752,plain,(
  ( ! [X397,X399,X395,X401,X403,X396,X398,X400,X402] : (s(bool,'T_0') != s(bool,happ(s(fun(X395,bool),X397),s(X395,sK18(X397,X398,X399,X400,X401,X396,bool,X395,X402)))) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X402,X395),X399),s(fun(X402,bool),X401),s(fun(X395,bool),X396))) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X402,X395),X400),s(fun(X402,bool),X398),s(fun(X395,bool),X397))) | s(bool,happ(s(fun(X402,bool),X398),s(X402,sK17(X397,X398,X399,X400,X401,X396,bool,X395,X402)))) != s(bool,happ(s(fun(X402,bool),X401),s(X402,sK17(X397,X398,X399,X400,X401,X396,bool,X395,X402)))) | s('type.sum.sum'(X402,X395),X399) != s('type.sum.sum'(X402,X395),X400) | s(bool,X403) = s(bool,happ(s(fun(X395,bool),X396),s(X395,sK18(X397,X398,X399,X400,X401,X396,bool,X395,X402)))) | s(bool,'T_0') = s(bool,X403)) )),
  inference(superposition,[],[f232,f272])).
fof(f5751,plain,(
  ( ! [X389,X391,X387,X393,X388,X390,X386,X392,X394] : (s(bool,'T_0') != s(bool,happ(s(fun(X386,bool),X391),s(X386,sK17(X388,X387,X389,X390,X391,X392,bool,X393,X386)))) | s('type.sum.sum'(X386,X393),X389) = s('type.sum.sum'(X386,X393),'const.sum.INR_1'(s(X393,sK18(X388,X387,X389,X390,X391,X392,bool,X393,X386)))) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X386,X393),X389),s(fun(X386,bool),X391),s(fun(X393,bool),X392))) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X386,X393),X390),s(fun(X386,bool),X387),s(fun(X393,bool),X388))) | s('type.sum.sum'(X386,X393),X389) != s('type.sum.sum'(X386,X393),X390) | s(bool,X394) = s(bool,happ(s(fun(X386,bool),X387),s(X386,sK17(X388,X387,X389,X390,X391,X392,bool,X393,X386)))) | s(bool,'T_0') = s(bool,X394)) )),
  inference(superposition,[],[f230,f272])).
fof(f5750,plain,(
  ( ! [X383,X377,X379,X385,X380,X382,X378,X384,X381] : (s(bool,'T_0') != s(bool,happ(s(fun(X377,bool),X380),s(X377,sK17(X379,X380,X381,X382,X378,X383,bool,X384,X377)))) | s('type.sum.sum'(X377,X384),X381) = s('type.sum.sum'(X377,X384),'const.sum.INR_1'(s(X384,sK18(X379,X380,X381,X382,X378,X383,bool,X384,X377)))) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X377,X384),X381),s(fun(X377,bool),X378),s(fun(X384,bool),X383))) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X377,X384),X382),s(fun(X377,bool),X380),s(fun(X384,bool),X379))) | s('type.sum.sum'(X377,X384),X381) != s('type.sum.sum'(X377,X384),X382) | s(bool,X385) = s(bool,happ(s(fun(X377,bool),X378),s(X377,sK17(X379,X380,X381,X382,X378,X383,bool,X384,X377)))) | s(bool,'T_0') = s(bool,X385)) )),
  inference(superposition,[],[f230,f272])).
fof(f5746,plain,(
  ( ! [X358,X360,X359,X361] : (s(bool,'T_0') != s(bool,f(s(X358,sK13(X359,X360,bool,X358)))) | s(fun(X358,bool),X359) = s(fun(X358,bool),X360) | s(bool,X361) = s(bool,g(s(X358,sK13(X359,X360,bool,X358)))) | s(bool,'T_0') = s(bool,X361)) )),
  inference(superposition,[],[f207,f272])).
fof(f5736,plain,(
  ( ! [X313,X310,X312,X314,X309,X311] : (s('type.sum.sum'(X310,bool),'const.sum.INR_1'(s(bool,X311))) != s('type.sum.sum'(X310,bool),'const.bool.COND_3'(s(bool,X312),s('type.sum.sum'(X310,bool),'const.sum.INR_1'(s(bool,'T_0'))),s('type.sum.sum'(X310,bool),'const.sum.INL_1'(s(X310,X313))))) | s(bool,X309) = s(bool,X311) | s(bool,X309) = s(bool,X314) | s(bool,'T_0') = s(bool,X314)) )),
  inference(superposition,[],[f213,f272])).
fof(f5735,plain,(
  ( ! [X308,X304,X306,X303,X305,X307] : (s('type.sum.sum'(X304,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X304,bool),'const.bool.COND_3'(s(bool,X305),s('type.sum.sum'(X304,bool),'const.sum.INR_1'(s(bool,X306))),s('type.sum.sum'(X304,bool),'const.sum.INL_1'(s(X304,X307))))) | s(bool,X303) = s(bool,X306) | s(bool,X303) = s(bool,X308) | s(bool,'T_0') = s(bool,X308)) )),
  inference(superposition,[],[f213,f272])).
fof(f5731,plain,(
  ( ! [X288,X285,X287,X286] : (s(bool,'T_0') != s(bool,X286) | s('type.sum.sum'(X287,bool),'const.sum.INR_1'(s(bool,X285))) = s('type.sum.sum'(X287,bool),'const.sum.INR_1'(s(bool,X286))) | s(bool,X285) = s(bool,X288) | s(bool,'T_0') = s(bool,X288)) )),
  inference(superposition,[],[f209,f272])).
fof(f5729,plain,(
  ( ! [X280,X277,X279,X278] : (s('type.sum.sum'(X278,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X278,bool),'const.sum.INR_1'(s(bool,X279))) | s(bool,X277) = s(bool,X279) | s(bool,X277) = s(bool,X280) | s(bool,'T_0') = s(bool,X280)) )),
  inference(superposition,[],[f208,f272])).
fof(f6383,plain,(
  ( ! [X255,X253,X251,X257,X254,X256,X252,X258] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X251,bool),X256),s(fun(X253,bool),X257),s('type.sum.sum'(X251,X253),X255)))) | p(s(bool,happ(s(fun(X253,bool),X254),s(X253,sK15(X252,X256,X254,X257,X255,X253,X251))))) | p(s(bool,happ(s(fun(X251,bool),X252),s(X251,sK14(X252,X256,X254,X257,X255,X253,X251))))) | s(bool,X258) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X251,bool),X252),s(fun(X253,bool),X254),s('type.sum.sum'(X251,X253),X255))) | s(bool,'T_0') = s(bool,X258)) )),
  inference(subsumption_resolution,[],[f5723,f157])).
fof(f5723,plain,(
  ( ! [X255,X253,X251,X257,X254,X256,X252,X258] : (~p(s(bool,'T_0')) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X251,bool),X256),s(fun(X253,bool),X257),s('type.sum.sum'(X251,X253),X255)))) | p(s(bool,happ(s(fun(X253,bool),X254),s(X253,sK15(X252,X256,X254,X257,X255,X253,X251))))) | p(s(bool,happ(s(fun(X251,bool),X252),s(X251,sK14(X252,X256,X254,X257,X255,X253,X251))))) | s(bool,X258) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X251,bool),X252),s(fun(X253,bool),X254),s('type.sum.sum'(X251,X253),X255))) | s(bool,'T_0') = s(bool,X258)) )),
  inference(superposition,[],[f215,f272])).
fof(f6382,plain,(
  ( ! [X231,X229,X230,X228,X232] : (p(s(bool,happ(s(fun('type.sum.sum'(X228,X229),bool),X230),s('type.sum.sum'(X228,X229),X231)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X228,X229),bool),X230),s('type.sum.sum'(X228,X229),'const.sum.INL_1'(s(X228,sK11(X230,X229,X228))))))) | s(bool,X232) = s(bool,happ(s(fun('type.sum.sum'(X228,X229),bool),X230),s('type.sum.sum'(X228,X229),'const.sum.INR_1'(s(X229,sK12(X230,X229,X228)))))) | s(bool,'T_0') = s(bool,X232)) )),
  inference(subsumption_resolution,[],[f5716,f157])).
fof(f5716,plain,(
  ( ! [X231,X229,X230,X228,X232] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.sum.sum'(X228,X229),bool),X230),s('type.sum.sum'(X228,X229),X231)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X228,X229),bool),X230),s('type.sum.sum'(X228,X229),'const.sum.INL_1'(s(X228,sK11(X230,X229,X228))))))) | s(bool,X232) = s(bool,happ(s(fun('type.sum.sum'(X228,X229),bool),X230),s('type.sum.sum'(X228,X229),'const.sum.INR_1'(s(X229,sK12(X230,X229,X228)))))) | s(bool,'T_0') = s(bool,X232)) )),
  inference(superposition,[],[f206,f272])).
fof(f6381,plain,(
  ( ! [X222,X220,X227,X225,X223,X221,X226,X224] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X226,bool),X222),s(fun(X220,bool),X224),s('type.sum.sum'(X226,X220),X225)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X226,bool),X223),s(fun(X220,bool),X221),s('type.sum.sum'(X226,X220),X225)))) | p(s(bool,happ(s(fun(X226,bool),X222),s(X226,sK14(X222,X223,X224,X221,X225,X220,X226))))) | s(bool,X227) = s(bool,happ(s(fun(X220,bool),X221),s(X220,sK15(X222,X223,X224,X221,X225,X220,X226)))) | s(bool,'T_0') = s(bool,X227)) )),
  inference(subsumption_resolution,[],[f5715,f157])).
fof(f5715,plain,(
  ( ! [X222,X220,X227,X225,X223,X221,X226,X224] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X226,bool),X222),s(fun(X220,bool),X224),s('type.sum.sum'(X226,X220),X225)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X226,bool),X223),s(fun(X220,bool),X221),s('type.sum.sum'(X226,X220),X225)))) | p(s(bool,happ(s(fun(X226,bool),X222),s(X226,sK14(X222,X223,X224,X221,X225,X220,X226))))) | s(bool,X227) = s(bool,happ(s(fun(X220,bool),X221),s(X220,sK15(X222,X223,X224,X221,X225,X220,X226)))) | s(bool,'T_0') = s(bool,X227)) )),
  inference(superposition,[],[f217,f272])).
fof(f6380,plain,(
  ( ! [X212,X218,X216,X215,X213,X219,X217,X214] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X218,bool),X214),s(fun(X212,bool),X216),s('type.sum.sum'(X218,X212),X217)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X218,bool),X215),s(fun(X212,bool),X213),s('type.sum.sum'(X218,X212),X217)))) | ~p(s(bool,happ(s(fun(X218,bool),X215),s(X218,sK14(X214,X215,X216,X213,X217,X212,X218))))) | s(bool,X219) = s(bool,happ(s(fun(X212,bool),X213),s(X212,sK15(X214,X215,X216,X213,X217,X212,X218)))) | s(bool,'T_0') = s(bool,X219)) )),
  inference(subsumption_resolution,[],[f5714,f157])).
fof(f5714,plain,(
  ( ! [X212,X218,X216,X215,X213,X219,X217,X214] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X218,bool),X214),s(fun(X212,bool),X216),s('type.sum.sum'(X218,X212),X217)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X218,bool),X215),s(fun(X212,bool),X213),s('type.sum.sum'(X218,X212),X217)))) | ~p(s(bool,happ(s(fun(X218,bool),X215),s(X218,sK14(X214,X215,X216,X213,X217,X212,X218))))) | s(bool,X219) = s(bool,happ(s(fun(X212,bool),X213),s(X212,sK15(X214,X215,X216,X213,X217,X212,X218)))) | s(bool,'T_0') = s(bool,X219)) )),
  inference(superposition,[],[f218,f272])).
fof(f6379,plain,(
  ( ! [X210,X208,X206,X204,X211,X209,X207,X205] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X204,bool),X206),s(fun(X210,bool),X207),s('type.sum.sum'(X204,X210),X209)))) | p(s(bool,happ(s(fun(X210,bool),X207),s(X210,sK15(X206,X205,X207,X208,X209,X210,X204))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X204,bool),X205),s(fun(X210,bool),X208),s('type.sum.sum'(X204,X210),X209)))) | s(bool,X211) = s(bool,happ(s(fun(X204,bool),X205),s(X204,sK14(X206,X205,X207,X208,X209,X210,X204)))) | s(bool,'T_0') = s(bool,X211)) )),
  inference(subsumption_resolution,[],[f5713,f157])).
fof(f5713,plain,(
  ( ! [X210,X208,X206,X204,X211,X209,X207,X205] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X204,bool),X206),s(fun(X210,bool),X207),s('type.sum.sum'(X204,X210),X209)))) | p(s(bool,happ(s(fun(X210,bool),X207),s(X210,sK15(X206,X205,X207,X208,X209,X210,X204))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X204,bool),X205),s(fun(X210,bool),X208),s('type.sum.sum'(X204,X210),X209)))) | s(bool,X211) = s(bool,happ(s(fun(X204,bool),X205),s(X204,sK14(X206,X205,X207,X208,X209,X210,X204)))) | s(bool,'T_0') = s(bool,X211)) )),
  inference(superposition,[],[f216,f272])).
fof(f6378,plain,(
  ( ! [X202,X200,X203,X201] : (p(s(bool,'|dollar|exists_1'(s(fun(X200,bool),X201)))) | s(bool,X203) = s(bool,happ(s(fun(X200,bool),X201),s(X200,sK8(X202,X201,X200)))) | s(bool,'T_0') = s(bool,X203)) )),
  inference(subsumption_resolution,[],[f5712,f157])).
fof(f5712,plain,(
  ( ! [X202,X200,X203,X201] : (~p(s(bool,'T_0')) | p(s(bool,'|dollar|exists_1'(s(fun(X200,bool),X201)))) | s(bool,X203) = s(bool,happ(s(fun(X200,bool),X201),s(X200,sK8(X202,X201,X200)))) | s(bool,'T_0') = s(bool,X203)) )),
  inference(superposition,[],[f196,f272])).
fof(f6376,plain,(
  ( ! [X45,X43,X44] : (s(bool,X43) != s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f5647,f157])).
fof(f5647,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'T_0')) | s(bool,X43) != s(bool,X44) | s(bool,'T_0') = s(bool,X44) | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f274,f272])).
fof(f6375,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(X37,X38) != s(X37,X39) | s('type.sum.sum'(X40,X37),'const.sum.INR_1'(s(X37,X39))) = s('type.sum.sum'(X40,X37),'const.bool.COND_3'(s(bool,X36),s('type.sum.sum'(X40,X37),'const.sum.INR_1'(s(X37,X38))),s('type.sum.sum'(X40,X37),'const.sum.INL_1'(s(X40,X41))))) | s(bool,X36) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f5646,f157])).
fof(f5646,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (~p(s(bool,'T_0')) | s(X37,X38) != s(X37,X39) | s('type.sum.sum'(X40,X37),'const.sum.INR_1'(s(X37,X39))) = s('type.sum.sum'(X40,X37),'const.bool.COND_3'(s(bool,X36),s('type.sum.sum'(X40,X37),'const.sum.INR_1'(s(X37,X38))),s('type.sum.sum'(X40,X37),'const.sum.INL_1'(s(X40,X41))))) | s(bool,X36) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f214,f272])).
fof(f6374,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X20)) | s(bool,X19) = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f5643,f157])).
fof(f5643,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | s(bool,X19) = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f204,f272])).
fof(f6373,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X11)) | s(bool,X10) != s(bool,X11) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f5640,f157])).
fof(f5640,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X11)) | s(bool,X10) != s(bool,X11) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f200,f272])).
fof(f5622,plain,(
  ( ! [X158,X156,X157,X155] : (s(bool,X155) != s(bool,X156) | s('type.sum.sum'(X157,bool),'const.sum.INR_1'(s(bool,'T_0'))) = s('type.sum.sum'(X157,bool),'const.sum.INR_1'(s(bool,X156))) | s(bool,X155) = s(bool,X158) | s(bool,'T_0') = s(bool,X158)) )),
  inference(superposition,[],[f209,f272])).
fof(f5578,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s(X3,X4) = s(X3,'const.bool.COND_3'(s(bool,X2),s(X3,X4),s(X3,X5))) | s(bool,X2) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f193,f272])).
fof(f5437,plain,(
  ( ! [X528,X524,X526,X520,X522,X525,X527,X521,X523] : (s(X526,happ(s(fun(bool,X526),X520),s(bool,X528))) != s(X526,happ(s(fun(bool,X526),X525),s(bool,X528))) | s(X526,'const.sum.sum_CASE_3'(s('type.sum.sum'(X527,bool),X522),s(fun(X527,X526),X524),s(fun(bool,X526),X525))) = s(X526,'const.sum.sum_CASE_3'(s('type.sum.sum'(X527,bool),X523),s(fun(X527,X526),X521),s(fun(bool,X526),X520))) | s('type.sum.sum'(X527,bool),X522) = s('type.sum.sum'(X527,bool),'const.sum.INL_1'(s(X527,sK17(X520,X521,X522,X523,X524,X525,X526,bool,X527)))) | s('type.sum.sum'(X527,bool),X522) != s('type.sum.sum'(X527,bool),X523) | s(bool,'T_0') = s(bool,X528) | s(bool,'T_0') = s(bool,sK18(X520,X521,X522,X523,X524,X525,X526,bool,X527))) )),
  inference(superposition,[],[f231,f272])).
fof(f5436,plain,(
  ( ! [X511,X516,X518,X512,X514,X517,X519,X513,X515] : (s(X517,happ(s(fun(bool,X517),X511),s(bool,X519))) != s(X517,happ(s(fun(bool,X517),X516),s(bool,X519))) | s(X517,'const.sum.sum_CASE_3'(s('type.sum.sum'(X518,bool),X513),s(fun(X518,X517),X515),s(fun(bool,X517),X516))) = s(X517,'const.sum.sum_CASE_3'(s('type.sum.sum'(X518,bool),X514),s(fun(X518,X517),X512),s(fun(bool,X517),X511))) | s(X517,happ(s(fun(X518,X517),X512),s(X518,sK17(X511,X512,X513,X514,X515,X516,X517,bool,X518)))) != s(X517,happ(s(fun(X518,X517),X515),s(X518,sK17(X511,X512,X513,X514,X515,X516,X517,bool,X518)))) | s('type.sum.sum'(X518,bool),X513) != s('type.sum.sum'(X518,bool),X514) | s(bool,'T_0') = s(bool,X519) | s(bool,'T_0') = s(bool,sK18(X511,X512,X513,X514,X515,X516,X517,bool,X518))) )),
  inference(superposition,[],[f232,f272])).
fof(f5435,plain,(
  ( ! [X509,X505,X507,X502,X508,X510,X504,X506,X503] : (s(X508,happ(s(fun(bool,X508),X503),s(bool,X510))) != s(X508,happ(s(fun(bool,X508),X506),s(bool,X510))) | s('type.sum.sum'(bool,X509),X504) = s('type.sum.sum'(bool,X509),'const.sum.INR_1'(s(X509,sK18(X502,X503,X504,X505,X506,X507,X508,X509,bool)))) | s(X508,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X509),X504),s(fun(bool,X508),X506),s(fun(X509,X508),X507))) = s(X508,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X509),X505),s(fun(bool,X508),X503),s(fun(X509,X508),X502))) | s('type.sum.sum'(bool,X509),X504) != s('type.sum.sum'(bool,X509),X505) | s(bool,'T_0') = s(bool,X510) | s(bool,'T_0') = s(bool,sK17(X502,X503,X504,X505,X506,X507,X508,X509,bool))) )),
  inference(superposition,[],[f230,f272])).
fof(f5434,plain,(
  ( ! [X499,X500,X496,X498,X495,X501,X497] : (~p(s(bool,happ(s(fun(bool,bool),X498),s(bool,X501)))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X500,bool),X495),s(fun(bool,bool),X497),s('type.sum.sum'(X500,bool),X499)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X500,bool),X496),s(fun(bool,bool),X498),s('type.sum.sum'(X500,bool),X499)))) | p(s(bool,happ(s(fun(X500,bool),X495),s(X500,sK14(X495,X496,X497,X498,X499,bool,X500))))) | s(bool,'T_0') = s(bool,X501) | s(bool,'T_0') = s(bool,sK15(X495,X496,X497,X498,X499,bool,X500))) )),
  inference(superposition,[],[f217,f272])).
fof(f5433,plain,(
  ( ! [X492,X494,X488,X490,X493,X489,X491] : (~p(s(bool,happ(s(fun(bool,bool),X491),s(bool,X494)))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X493,bool),X488),s(fun(bool,bool),X490),s('type.sum.sum'(X493,bool),X492)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X493,bool),X489),s(fun(bool,bool),X491),s('type.sum.sum'(X493,bool),X492)))) | ~p(s(bool,happ(s(fun(X493,bool),X489),s(X493,sK14(X488,X489,X490,X491,X492,bool,X493))))) | s(bool,'T_0') = s(bool,X494) | s(bool,'T_0') = s(bool,sK15(X488,X489,X490,X491,X492,bool,X493))) )),
  inference(superposition,[],[f218,f272])).
fof(f5432,plain,(
  ( ! [X484,X486,X482,X485,X487,X481,X483] : (~p(s(bool,happ(s(fun(bool,bool),X482),s(bool,X487)))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(bool,bool),X481),s(fun(X486,bool),X483),s('type.sum.sum'(bool,X486),X485)))) | p(s(bool,happ(s(fun(X486,bool),X483),s(X486,sK15(X481,X482,X483,X484,X485,X486,bool))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(bool,bool),X482),s(fun(X486,bool),X484),s('type.sum.sum'(bool,X486),X485)))) | s(bool,'T_0') = s(bool,X487) | s(bool,'T_0') = s(bool,sK14(X481,X482,X483,X484,X485,X486,bool))) )),
  inference(superposition,[],[f216,f272])).
fof(f5430,plain,(
  ( ! [X476,X474,X475,X473] : (~p(s(bool,happ(s(fun('type.sum.sum'(X474,bool),bool),X473),s('type.sum.sum'(X474,bool),'const.sum.INR_1'(s(bool,X475)))))) | p(s(bool,happ(s(fun('type.sum.sum'(X474,bool),bool),X473),s('type.sum.sum'(X474,bool),X476)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X474,bool),bool),X473),s('type.sum.sum'(X474,bool),'const.sum.INL_1'(s(X474,sK11(X473,bool,X474))))))) | s(bool,'T_0') = s(bool,X475) | s(bool,'T_0') = s(bool,sK12(X473,bool,X474))) )),
  inference(superposition,[],[f206,f272])).
fof(f5429,plain,(
  ( ! [X472,X471,X470] : (s('type.sum.sum'(X471,bool),X470) = s('type.sum.sum'(X471,bool),'const.sum.INR_1'(s(bool,X472))) | s('type.sum.sum'(X471,bool),X470) = s('type.sum.sum'(X471,bool),'const.sum.INL_1'(s(X471,sK9(X470,bool,X471)))) | s(bool,'T_0') = s(bool,X472) | s(bool,'T_0') = s(bool,sK10(X470,bool,X471))) )),
  inference(superposition,[],[f205,f272])).
fof(f5428,plain,(
  ( ! [X469,X467,X468] : (s('type.sum.sum'(bool,X468),X467) = s('type.sum.sum'(bool,X468),'const.sum.INL_1'(s(bool,X469))) | s('type.sum.sum'(bool,X468),X467) = s('type.sum.sum'(bool,X468),'const.sum.INR_1'(s(X468,sK10(X467,X468,bool)))) | s(bool,'T_0') = s(bool,X469) | s(bool,'T_0') = s(bool,sK9(X467,X468,bool))) )),
  inference(superposition,[],[f205,f272])).
fof(f5427,plain,(
  ( ! [X466,X465,X464] : (~p(s(bool,happ(s(fun(bool,bool),X465),s(bool,X466)))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X465)))) | s(bool,'T_0') = s(bool,X466) | s(bool,'T_0') = s(bool,sK8(X464,X465,bool))) )),
  inference(superposition,[],[f196,f272])).
fof(f5422,plain,(
  ( ! [X445,X441,X443,X438,X444,X440,X442,X437,X439] : (s(bool,X445) != s(bool,happ(s(fun(X437,bool),X443),s(X437,sK18(X438,X439,X440,X441,X442,X443,bool,X437,X444)))) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X444,X437),X440),s(fun(X444,bool),X442),s(fun(X437,bool),X443))) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X444,X437),X441),s(fun(X444,bool),X439),s(fun(X437,bool),X438))) | s('type.sum.sum'(X444,X437),X440) = s('type.sum.sum'(X444,X437),'const.sum.INL_1'(s(X444,sK17(X438,X439,X440,X441,X442,X443,bool,X437,X444)))) | s('type.sum.sum'(X444,X437),X440) != s('type.sum.sum'(X444,X437),X441) | s(bool,'T_0') = s(bool,X445) | s(bool,'T_0') = s(bool,happ(s(fun(X437,bool),X438),s(X437,sK18(X438,X439,X440,X441,X442,X443,bool,X437,X444))))) )),
  inference(superposition,[],[f231,f272])).
fof(f5421,plain,(
  ( ! [X428,X430,X436,X432,X434,X429,X431,X433,X435] : (s(bool,X436) != s(bool,happ(s(fun(X428,bool),X434),s(X428,sK18(X429,X430,X431,X432,X433,X434,bool,X428,X435)))) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X435,X428),X431),s(fun(X435,bool),X433),s(fun(X428,bool),X434))) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X435,X428),X432),s(fun(X435,bool),X430),s(fun(X428,bool),X429))) | s(bool,happ(s(fun(X435,bool),X430),s(X435,sK17(X429,X430,X431,X432,X433,X434,bool,X428,X435)))) != s(bool,happ(s(fun(X435,bool),X433),s(X435,sK17(X429,X430,X431,X432,X433,X434,bool,X428,X435)))) | s('type.sum.sum'(X435,X428),X431) != s('type.sum.sum'(X435,X428),X432) | s(bool,'T_0') = s(bool,X436) | s(bool,'T_0') = s(bool,happ(s(fun(X428,bool),X429),s(X428,sK18(X429,X430,X431,X432,X433,X434,bool,X428,X435))))) )),
  inference(superposition,[],[f232,f272])).
fof(f5420,plain,(
  ( ! [X420,X422,X424,X426,X421,X423,X419,X425,X427] : (s(bool,X427) != s(bool,happ(s(fun(X419,bool),X421),s(X419,sK18(X421,X422,X423,X424,X425,X420,bool,X419,X426)))) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X426,X419),X423),s(fun(X426,bool),X425),s(fun(X419,bool),X420))) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X426,X419),X424),s(fun(X426,bool),X422),s(fun(X419,bool),X421))) | s('type.sum.sum'(X426,X419),X423) = s('type.sum.sum'(X426,X419),'const.sum.INL_1'(s(X426,sK17(X421,X422,X423,X424,X425,X420,bool,X419,X426)))) | s('type.sum.sum'(X426,X419),X423) != s('type.sum.sum'(X426,X419),X424) | s(bool,'T_0') = s(bool,X427) | s(bool,'T_0') = s(bool,happ(s(fun(X419,bool),X420),s(X419,sK18(X421,X422,X423,X424,X425,X420,bool,X419,X426))))) )),
  inference(superposition,[],[f231,f272])).
fof(f5419,plain,(
  ( ! [X412,X414,X410,X416,X418,X413,X415,X411,X417] : (s(bool,X418) != s(bool,happ(s(fun(X410,bool),X412),s(X410,sK18(X412,X413,X414,X415,X416,X411,bool,X410,X417)))) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X417,X410),X414),s(fun(X417,bool),X416),s(fun(X410,bool),X411))) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X417,X410),X415),s(fun(X417,bool),X413),s(fun(X410,bool),X412))) | s(bool,happ(s(fun(X417,bool),X413),s(X417,sK17(X412,X413,X414,X415,X416,X411,bool,X410,X417)))) != s(bool,happ(s(fun(X417,bool),X416),s(X417,sK17(X412,X413,X414,X415,X416,X411,bool,X410,X417)))) | s('type.sum.sum'(X417,X410),X414) != s('type.sum.sum'(X417,X410),X415) | s(bool,'T_0') = s(bool,X418) | s(bool,'T_0') = s(bool,happ(s(fun(X410,bool),X411),s(X410,sK18(X412,X413,X414,X415,X416,X411,bool,X410,X417))))) )),
  inference(superposition,[],[f232,f272])).
fof(f5418,plain,(
  ( ! [X408,X405,X407,X401,X403,X409,X404,X406,X402] : (s(bool,X409) != s(bool,happ(s(fun(X401,bool),X406),s(X401,sK17(X403,X402,X404,X405,X406,X407,bool,X408,X401)))) | s('type.sum.sum'(X401,X408),X404) = s('type.sum.sum'(X401,X408),'const.sum.INR_1'(s(X408,sK18(X403,X402,X404,X405,X406,X407,bool,X408,X401)))) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X401,X408),X404),s(fun(X401,bool),X406),s(fun(X408,bool),X407))) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X401,X408),X405),s(fun(X401,bool),X402),s(fun(X408,bool),X403))) | s('type.sum.sum'(X401,X408),X404) != s('type.sum.sum'(X401,X408),X405) | s(bool,'T_0') = s(bool,X409) | s(bool,'T_0') = s(bool,happ(s(fun(X401,bool),X402),s(X401,sK17(X403,X402,X404,X405,X406,X407,bool,X408,X401))))) )),
  inference(superposition,[],[f230,f272])).
fof(f5417,plain,(
  ( ! [X397,X399,X393,X395,X396,X398,X392,X394,X400] : (s(bool,X400) != s(bool,happ(s(fun(X392,bool),X395),s(X392,sK17(X394,X395,X396,X397,X393,X398,bool,X399,X392)))) | s('type.sum.sum'(X392,X399),X396) = s('type.sum.sum'(X392,X399),'const.sum.INR_1'(s(X399,sK18(X394,X395,X396,X397,X393,X398,bool,X399,X392)))) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X392,X399),X396),s(fun(X392,bool),X393),s(fun(X399,bool),X398))) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X392,X399),X397),s(fun(X392,bool),X395),s(fun(X399,bool),X394))) | s('type.sum.sum'(X392,X399),X396) != s('type.sum.sum'(X392,X399),X397) | s(bool,'T_0') = s(bool,X400) | s(bool,'T_0') = s(bool,happ(s(fun(X392,bool),X393),s(X392,sK17(X394,X395,X396,X397,X393,X398,bool,X399,X392))))) )),
  inference(superposition,[],[f230,f272])).
fof(f5413,plain,(
  ( ! [X374,X376,X373,X375] : (s(bool,X376) != s(bool,f(s(X373,sK13(X374,X375,bool,X373)))) | s(fun(X373,bool),X374) = s(fun(X373,bool),X375) | s(bool,'T_0') = s(bool,X376) | s(bool,'T_0') = s(bool,g(s(X373,sK13(X374,X375,bool,X373))))) )),
  inference(superposition,[],[f207,f272])).
fof(f5403,plain,(
  ( ! [X325,X327,X323,X324,X326,X322] : (s('type.sum.sum'(X324,bool),'const.sum.INR_1'(s(bool,X325))) != s('type.sum.sum'(X324,bool),'const.bool.COND_3'(s(bool,X326),s('type.sum.sum'(X324,bool),'const.sum.INR_1'(s(bool,X323))),s('type.sum.sum'(X324,bool),'const.sum.INL_1'(s(X324,X327))))) | s(bool,X322) = s(bool,X325) | s(bool,'T_0') = s(bool,X323) | s(bool,'T_0') = s(bool,X322)) )),
  inference(superposition,[],[f213,f272])).
fof(f5402,plain,(
  ( ! [X321,X316,X318,X320,X317,X319] : (s('type.sum.sum'(X318,bool),'const.sum.INR_1'(s(bool,X317))) != s('type.sum.sum'(X318,bool),'const.bool.COND_3'(s(bool,X319),s('type.sum.sum'(X318,bool),'const.sum.INR_1'(s(bool,X320))),s('type.sum.sum'(X318,bool),'const.sum.INL_1'(s(X318,X321))))) | s(bool,X316) = s(bool,X320) | s(bool,'T_0') = s(bool,X317) | s(bool,'T_0') = s(bool,X316)) )),
  inference(superposition,[],[f213,f272])).
fof(f5398,plain,(
  ( ! [X300,X298,X301,X299] : (s(bool,X299) != s(bool,X300) | s('type.sum.sum'(X301,bool),'const.sum.INR_1'(s(bool,X298))) = s('type.sum.sum'(X301,bool),'const.sum.INR_1'(s(bool,X300))) | s(bool,'T_0') = s(bool,X299) | s(bool,'T_0') = s(bool,X298)) )),
  inference(superposition,[],[f209,f272])).
fof(f5396,plain,(
  ( ! [X292,X290,X293,X291] : (s('type.sum.sum'(X292,bool),'const.sum.INR_1'(s(bool,X291))) != s('type.sum.sum'(X292,bool),'const.sum.INR_1'(s(bool,X293))) | s(bool,X290) = s(bool,X293) | s(bool,'T_0') = s(bool,X291) | s(bool,'T_0') = s(bool,X290)) )),
  inference(superposition,[],[f208,f272])).
fof(f5266,plain,(
  ( ! [X80,X83,X81,X79,X84,X82] : (s(bool,'T_0') = s(bool,happ(s(fun(X80,bool),X81),s(X80,X83))) | s(bool,X84) = s(bool,happ(s(fun('type.sum.sum'(X79,X80),bool),sK16(X81,X82,X80,bool,X79)),s('type.sum.sum'(X79,X80),'const.sum.INR_1'(s(X80,X83))))) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f272,f226])).
fof(f5265,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (s(bool,'T_0') = s(bool,happ(s(fun(X73,bool),X76),s(X73,X77))) | s(bool,X78) = s(bool,happ(s(fun('type.sum.sum'(X73,X74),bool),sK16(X75,X76,X74,bool,X73)),s('type.sum.sum'(X73,X74),'const.sum.INL_1'(s(X73,X77))))) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f272,f225])).
fof(f5264,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'T_0') = s(bool,happ(s(fun(X67,bool),X70),s(X67,X69))) | s(bool,X72) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X67,X68),'const.sum.INL_1'(s(X67,X69))),s(fun(X67,bool),X70),s(fun(X68,bool),X71))) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f272,f228])).
fof(f5263,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(X62,bool),X65),s(X62,X63))) | s(bool,X66) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X63))),s(fun(X61,bool),X64),s(fun(X62,bool),X65))) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f272,f227])).
fof(f5262,plain,(
  ( ! [X59,X60,X58] : (s(bool,'T_0') = s(bool,X59) | s(bool,X60) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X58),s(bool,X59))) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f272,f194])).
fof(f5261,plain,(
  ( ! [X57,X56,X55] : (s(bool,'T_0') = s(bool,X55) | s(bool,X57) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X55),s(bool,X56))) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f272,f193])).
fof(f5260,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,happ(s(fun(X47,bool),X48),s(X47,X51))) | s(bool,X52) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X47,bool),X48),s(fun(X49,bool),X50),s('type.sum.sum'(X47,X49),'const.sum.INL_1'(s(X47,X51))))) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f272,f211])).
fof(f5259,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(X43,bool),X44),s(X43,X45))) | s(bool,X46) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X41,bool),X42),s(fun(X43,bool),X44),s('type.sum.sum'(X41,X43),'const.sum.INR_1'(s(X43,X45))))) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f272,f210])).
fof(f5235,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f272,f272])).
fof(f5201,plain,(
  ( ! [X92,X90,X88,X93,X91,X89] : (s(bool,X93) = s(bool,happ(s(fun(X89,bool),X90),s(X89,X92))) | s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,happ(s(fun('type.sum.sum'(X88,X89),bool),sK16(X90,X91,X89,bool,X88)),s('type.sum.sum'(X88,X89),'const.sum.INR_1'(s(X89,X92)))))) )),
  inference(superposition,[],[f272,f226])).
fof(f5200,plain,(
  ( ! [X87,X85,X83,X86,X84,X82] : (s(bool,X87) = s(bool,happ(s(fun(X82,bool),X85),s(X82,X86))) | s(bool,'T_0') = s(bool,X87) | s(bool,'T_0') = s(bool,happ(s(fun('type.sum.sum'(X82,X83),bool),sK16(X84,X85,X83,bool,X82)),s('type.sum.sum'(X82,X83),'const.sum.INL_1'(s(X82,X86)))))) )),
  inference(superposition,[],[f272,f225])).
fof(f5199,plain,(
  ( ! [X80,X78,X76,X81,X79,X77] : (s(bool,X81) = s(bool,happ(s(fun(X76,bool),X79),s(X76,X78))) | s(bool,'T_0') = s(bool,X81) | s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X76,X77),'const.sum.INL_1'(s(X76,X78))),s(fun(X76,bool),X79),s(fun(X77,bool),X80)))) )),
  inference(superposition,[],[f272,f228])).
fof(f5198,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (s(bool,X75) = s(bool,happ(s(fun(X71,bool),X74),s(X71,X72))) | s(bool,'T_0') = s(bool,X75) | s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X70,X71),'const.sum.INR_1'(s(X71,X72))),s(fun(X70,bool),X73),s(fun(X71,bool),X74)))) )),
  inference(superposition,[],[f272,f227])).
fof(f5197,plain,(
  ( ! [X68,X69,X67] : (s(bool,X68) = s(bool,X69) | s(bool,'T_0') = s(bool,X69) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X67),s(bool,X68)))) )),
  inference(superposition,[],[f272,f194])).
fof(f5196,plain,(
  ( ! [X66,X64,X65] : (s(bool,X64) = s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X64),s(bool,X65)))) )),
  inference(superposition,[],[f272,f193])).
fof(f5194,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (s(bool,X60) = s(bool,happ(s(fun(X55,bool),X56),s(X55,X59))) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,'const.sum.SUM_ALL_3'(s(fun(X55,bool),X56),s(fun(X57,bool),X58),s('type.sum.sum'(X55,X57),'const.sum.INL_1'(s(X55,X59)))))) )),
  inference(superposition,[],[f272,f211])).
fof(f5193,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (s(bool,X54) = s(bool,happ(s(fun(X51,bool),X52),s(X51,X53))) | s(bool,'T_0') = s(bool,X54) | s(bool,'T_0') = s(bool,'const.sum.SUM_ALL_3'(s(fun(X49,bool),X50),s(fun(X51,bool),X52),s('type.sum.sum'(X49,X51),'const.sum.INR_1'(s(X51,X53)))))) )),
  inference(superposition,[],[f272,f210])).
fof(f4170,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X1,X2),X3),s(fun(X1,X0),X4),s(fun(X2,X0),X5))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X1,X2),X6),s(fun(X1,X0),X7),s(fun(X2,X0),X5))) | s(X0,happ(s(fun(X1,X0),X4),s(X1,sK17(X5,X4,X6,X3,X7,X5,X0,X2,X1)))) != s(X0,happ(s(fun(X1,X0),X7),s(X1,sK17(X5,X4,X6,X3,X7,X5,X0,X2,X1)))) | s('type.sum.sum'(X1,X2),X3) != s('type.sum.sum'(X1,X2),X6)) )),
  inference(equality_resolution,[],[f232])).
fof(f4185,plain,(
  ( ! [X61,X59,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X68),s(fun(X71,X60),X70),s(fun(X59,X60),X66))) = s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X69),s(fun(X71,X60),X67),s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))))) | s(X60,happ(s(fun(X59,X60),X66),s(X59,sK18(happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X67,X68,X69,X70,X66,X60,X59,X71)))) != s(X60,happ(s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))),s(X59,sK18(happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X67,X68,X69,X70,X66,X60,X59,X71)))) | s(X60,happ(s(fun(X71,X60),X67),s(X71,sK17(happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X67,X68,X69,X70,X66,X60,X59,X71)))) != s(X60,happ(s(fun(X71,X60),X70),s(X71,sK17(happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X67,X68,X69,X70,X66,X60,X59,X71)))) | s('type.sum.sum'(X71,X59),X68) != s('type.sum.sum'(X71,X59),X69)) )),
  inference(forward_demodulation,[],[f4161,f226])).
fof(f4161,plain,(
  ( ! [X61,X59,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,happ(s(fun(X59,X60),X66),s(X59,sK18(happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X67,X68,X69,X70,X66,X60,X59,X71)))) != s(X60,happ(s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))),s(X59,sK18(happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X67,X68,X69,X70,X66,X60,X59,X71)))) | s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X68),s(fun(X71,X60),X70),s(fun(X59,X60),X66))) = s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X69),s(fun(X71,X60),X67),s(fun(X59,X60),happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65))))))) | s(X60,happ(s(fun(X71,X60),X67),s(X71,sK17(happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X67,X68,X69,X70,X66,X60,X59,X71)))) != s(X60,happ(s(fun(X71,X60),X70),s(X71,sK17(happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X67,X68,X69,X70,X66,X60,X59,X71)))) | s('type.sum.sum'(X71,X59),X68) != s('type.sum.sum'(X71,X59),X69)) )),
  inference(superposition,[],[f232,f226])).
fof(f4184,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X55),s(fun(X58,X47),X57),s(fun(X46,X47),X53))) = s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X56),s(fun(X58,X47),X54),s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))))) | s(X47,happ(s(fun(X46,X47),X53),s(X46,sK18(happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X54,X55,X56,X57,X53,X47,X46,X58)))) != s(X47,happ(s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))),s(X46,sK18(happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X54,X55,X56,X57,X53,X47,X46,X58)))) | s(X47,happ(s(fun(X58,X47),X54),s(X58,sK17(happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X54,X55,X56,X57,X53,X47,X46,X58)))) != s(X47,happ(s(fun(X58,X47),X57),s(X58,sK17(happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X54,X55,X56,X57,X53,X47,X46,X58)))) | s('type.sum.sum'(X58,X46),X55) != s('type.sum.sum'(X58,X46),X56)) )),
  inference(forward_demodulation,[],[f4160,f225])).
fof(f4160,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (s(X47,happ(s(fun(X46,X47),X53),s(X46,sK18(happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X54,X55,X56,X57,X53,X47,X46,X58)))) != s(X47,happ(s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))),s(X46,sK18(happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X54,X55,X56,X57,X53,X47,X46,X58)))) | s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X55),s(fun(X58,X47),X57),s(fun(X46,X47),X53))) = s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X56),s(fun(X58,X47),X54),s(fun(X46,X47),happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52))))))) | s(X47,happ(s(fun(X58,X47),X54),s(X58,sK17(happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X54,X55,X56,X57,X53,X47,X46,X58)))) != s(X47,happ(s(fun(X58,X47),X57),s(X58,sK17(happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X54,X55,X56,X57,X53,X47,X46,X58)))) | s('type.sum.sum'(X58,X46),X55) != s('type.sum.sum'(X58,X46),X56)) )),
  inference(superposition,[],[f232,f225])).
fof(f4183,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X44,X42,X40] : (s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X42),s(fun(X45,X34),X44),s(fun(X33,X34),X40))) = s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X43),s(fun(X45,X34),X41),s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))))) | s(X34,happ(s(fun(X33,X34),X40),s(X33,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X41,X42,X43,X44,X40,X34,X33,X45)))) != s(X34,happ(s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))),s(X33,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X41,X42,X43,X44,X40,X34,X33,X45)))) | s(X34,happ(s(fun(X45,X34),X41),s(X45,sK17('const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X41,X42,X43,X44,X40,X34,X33,X45)))) != s(X34,happ(s(fun(X45,X34),X44),s(X45,sK17('const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X41,X42,X43,X44,X40,X34,X33,X45)))) | s('type.sum.sum'(X45,X33),X42) != s('type.sum.sum'(X45,X33),X43)) )),
  inference(forward_demodulation,[],[f4159,f228])).
fof(f4159,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X44,X42,X40] : (s(X34,happ(s(fun(X33,X34),X40),s(X33,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X41,X42,X43,X44,X40,X34,X33,X45)))) != s(X34,happ(s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))),s(X33,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X41,X42,X43,X44,X40,X34,X33,X45)))) | s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X42),s(fun(X45,X34),X44),s(fun(X33,X34),X40))) = s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X43),s(fun(X45,X34),X41),s(fun(X33,X34),'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39))))) | s(X34,happ(s(fun(X45,X34),X41),s(X45,sK17('const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X41,X42,X43,X44,X40,X34,X33,X45)))) != s(X34,happ(s(fun(X45,X34),X44),s(X45,sK17('const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X41,X42,X43,X44,X40,X34,X33,X45)))) | s('type.sum.sum'(X45,X33),X42) != s('type.sum.sum'(X45,X33),X43)) )),
  inference(superposition,[],[f232,f228])).
fof(f4182,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X32,X22,X20] : (s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X29),s(fun(X32,X21),X31),s(fun(X20,X21),X27))) = s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X30),s(fun(X32,X21),X28),s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))))) | s(X21,happ(s(fun(X20,X21),X27),s(X20,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X28,X29,X30,X31,X27,X21,X20,X32)))) != s(X21,happ(s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))),s(X20,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X28,X29,X30,X31,X27,X21,X20,X32)))) | s(X21,happ(s(fun(X32,X21),X28),s(X32,sK17('const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X28,X29,X30,X31,X27,X21,X20,X32)))) != s(X21,happ(s(fun(X32,X21),X31),s(X32,sK17('const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X28,X29,X30,X31,X27,X21,X20,X32)))) | s('type.sum.sum'(X32,X20),X29) != s('type.sum.sum'(X32,X20),X30)) )),
  inference(forward_demodulation,[],[f4158,f227])).
fof(f4158,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X32,X22,X20] : (s(X21,happ(s(fun(X20,X21),X27),s(X20,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X28,X29,X30,X31,X27,X21,X20,X32)))) != s(X21,happ(s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))),s(X20,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X28,X29,X30,X31,X27,X21,X20,X32)))) | s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X29),s(fun(X32,X21),X31),s(fun(X20,X21),X27))) = s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X30),s(fun(X32,X21),X28),s(fun(X20,X21),'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26))))) | s(X21,happ(s(fun(X32,X21),X28),s(X32,sK17('const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X28,X29,X30,X31,X27,X21,X20,X32)))) != s(X21,happ(s(fun(X32,X21),X31),s(X32,sK17('const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X28,X29,X30,X31,X27,X21,X20,X32)))) | s('type.sum.sum'(X32,X20),X29) != s('type.sum.sum'(X32,X20),X30)) )),
  inference(superposition,[],[f232,f227])).
fof(f4181,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X16),s(fun(X19,X11),X18),s(fun(X10,X11),X14))) = s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X17),s(fun(X19,X11),X15),s(fun(X10,X11),X13))) | s(X11,happ(s(fun(X10,X11),X13),s(X10,sK18('const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X15,X16,X17,X18,X14,X11,X10,X19)))) != s(X11,happ(s(fun(X10,X11),X14),s(X10,sK18('const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X15,X16,X17,X18,X14,X11,X10,X19)))) | s(X11,happ(s(fun(X19,X11),X15),s(X19,sK17('const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X15,X16,X17,X18,X14,X11,X10,X19)))) != s(X11,happ(s(fun(X19,X11),X18),s(X19,sK17('const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X15,X16,X17,X18,X14,X11,X10,X19)))) | s('type.sum.sum'(X19,X10),X16) != s('type.sum.sum'(X19,X10),X17)) )),
  inference(forward_demodulation,[],[f4157,f194])).
fof(f4157,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (s(X11,happ(s(fun(X10,X11),X13),s(X10,sK18('const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X15,X16,X17,X18,X14,X11,X10,X19)))) != s(X11,happ(s(fun(X10,X11),X14),s(X10,sK18('const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X15,X16,X17,X18,X14,X11,X10,X19)))) | s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X16),s(fun(X19,X11),X18),s(fun(X10,X11),X14))) = s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X17),s(fun(X19,X11),X15),s(fun(X10,X11),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13))))) | s(X11,happ(s(fun(X19,X11),X15),s(X19,sK17('const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X15,X16,X17,X18,X14,X11,X10,X19)))) != s(X11,happ(s(fun(X19,X11),X18),s(X19,sK17('const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X15,X16,X17,X18,X14,X11,X10,X19)))) | s('type.sum.sum'(X19,X10),X16) != s('type.sum.sum'(X19,X10),X17)) )),
  inference(superposition,[],[f232,f194])).
fof(f4180,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X6),s(fun(X9,X1),X8),s(fun(X0,X1),X4))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X7),s(fun(X9,X1),X5),s(fun(X0,X1),X2))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,sK18('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X5,X6,X7,X8,X4,X1,X0,X9)))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,sK18('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X5,X6,X7,X8,X4,X1,X0,X9)))) | s(X1,happ(s(fun(X9,X1),X5),s(X9,sK17('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X5,X6,X7,X8,X4,X1,X0,X9)))) != s(X1,happ(s(fun(X9,X1),X8),s(X9,sK17('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X5,X6,X7,X8,X4,X1,X0,X9)))) | s('type.sum.sum'(X9,X0),X6) != s('type.sum.sum'(X9,X0),X7)) )),
  inference(forward_demodulation,[],[f4156,f193])).
fof(f4156,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,happ(s(fun(X0,X1),X2),s(X0,sK18('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X5,X6,X7,X8,X4,X1,X0,X9)))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,sK18('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X5,X6,X7,X8,X4,X1,X0,X9)))) | s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X6),s(fun(X9,X1),X8),s(fun(X0,X1),X4))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X7),s(fun(X9,X1),X5),s(fun(X0,X1),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3))))) | s(X1,happ(s(fun(X9,X1),X5),s(X9,sK17('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X5,X6,X7,X8,X4,X1,X0,X9)))) != s(X1,happ(s(fun(X9,X1),X8),s(X9,sK17('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X5,X6,X7,X8,X4,X1,X0,X9)))) | s('type.sum.sum'(X9,X0),X6) != s('type.sum.sum'(X9,X0),X7)) )),
  inference(superposition,[],[f232,f193])).
fof(f4176,plain,(
  ( ! [X61,X59,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X69),s(fun(X71,X60),X67),s(fun(X59,X60),X66))) = s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X68),s(fun(X71,X60),X70),s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))))) | s(X60,happ(s(fun(X59,X60),X66),s(X59,sK18(X66,X67,X68,X69,X70,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X60,X59,X71)))) != s(X60,happ(s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))),s(X59,sK18(X66,X67,X68,X69,X70,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X60,X59,X71)))) | s(X60,happ(s(fun(X71,X60),X67),s(X71,sK17(X66,X67,X68,X69,X70,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X60,X59,X71)))) != s(X60,happ(s(fun(X71,X60),X70),s(X71,sK17(X66,X67,X68,X69,X70,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X60,X59,X71)))) | s('type.sum.sum'(X71,X59),X68) != s('type.sum.sum'(X71,X59),X69)) )),
  inference(forward_demodulation,[],[f4151,f226])).
fof(f4151,plain,(
  ( ! [X61,X59,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,happ(s(fun(X59,X60),X66),s(X59,sK18(X66,X67,X68,X69,X70,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X60,X59,X71)))) != s(X60,happ(s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))),s(X59,sK18(X66,X67,X68,X69,X70,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X60,X59,X71)))) | s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X69),s(fun(X71,X60),X67),s(fun(X59,X60),X66))) = s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X68),s(fun(X71,X60),X70),s(fun(X59,X60),happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65))))))) | s(X60,happ(s(fun(X71,X60),X67),s(X71,sK17(X66,X67,X68,X69,X70,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X60,X59,X71)))) != s(X60,happ(s(fun(X71,X60),X70),s(X71,sK17(X66,X67,X68,X69,X70,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X60,X59,X71)))) | s('type.sum.sum'(X71,X59),X68) != s('type.sum.sum'(X71,X59),X69)) )),
  inference(superposition,[],[f232,f226])).
fof(f4175,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X56),s(fun(X58,X47),X54),s(fun(X46,X47),X53))) = s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X55),s(fun(X58,X47),X57),s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))))) | s(X47,happ(s(fun(X46,X47),X53),s(X46,sK18(X53,X54,X55,X56,X57,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X47,X46,X58)))) != s(X47,happ(s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))),s(X46,sK18(X53,X54,X55,X56,X57,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X47,X46,X58)))) | s(X47,happ(s(fun(X58,X47),X54),s(X58,sK17(X53,X54,X55,X56,X57,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X47,X46,X58)))) != s(X47,happ(s(fun(X58,X47),X57),s(X58,sK17(X53,X54,X55,X56,X57,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X47,X46,X58)))) | s('type.sum.sum'(X58,X46),X55) != s('type.sum.sum'(X58,X46),X56)) )),
  inference(forward_demodulation,[],[f4150,f225])).
fof(f4150,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (s(X47,happ(s(fun(X46,X47),X53),s(X46,sK18(X53,X54,X55,X56,X57,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X47,X46,X58)))) != s(X47,happ(s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))),s(X46,sK18(X53,X54,X55,X56,X57,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X47,X46,X58)))) | s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X56),s(fun(X58,X47),X54),s(fun(X46,X47),X53))) = s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X55),s(fun(X58,X47),X57),s(fun(X46,X47),happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52))))))) | s(X47,happ(s(fun(X58,X47),X54),s(X58,sK17(X53,X54,X55,X56,X57,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X47,X46,X58)))) != s(X47,happ(s(fun(X58,X47),X57),s(X58,sK17(X53,X54,X55,X56,X57,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X47,X46,X58)))) | s('type.sum.sum'(X58,X46),X55) != s('type.sum.sum'(X58,X46),X56)) )),
  inference(superposition,[],[f232,f225])).
fof(f4174,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X44,X42,X40] : (s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X43),s(fun(X45,X34),X41),s(fun(X33,X34),X40))) = s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X42),s(fun(X45,X34),X44),s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))))) | s(X34,happ(s(fun(X33,X34),X40),s(X33,sK18(X40,X41,X42,X43,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X34,X33,X45)))) != s(X34,happ(s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))),s(X33,sK18(X40,X41,X42,X43,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X34,X33,X45)))) | s(X34,happ(s(fun(X45,X34),X41),s(X45,sK17(X40,X41,X42,X43,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X34,X33,X45)))) != s(X34,happ(s(fun(X45,X34),X44),s(X45,sK17(X40,X41,X42,X43,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X34,X33,X45)))) | s('type.sum.sum'(X45,X33),X42) != s('type.sum.sum'(X45,X33),X43)) )),
  inference(forward_demodulation,[],[f4149,f228])).
fof(f4149,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X44,X42,X40] : (s(X34,happ(s(fun(X33,X34),X40),s(X33,sK18(X40,X41,X42,X43,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X34,X33,X45)))) != s(X34,happ(s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))),s(X33,sK18(X40,X41,X42,X43,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X34,X33,X45)))) | s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X43),s(fun(X45,X34),X41),s(fun(X33,X34),X40))) = s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X42),s(fun(X45,X34),X44),s(fun(X33,X34),'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39))))) | s(X34,happ(s(fun(X45,X34),X41),s(X45,sK17(X40,X41,X42,X43,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X34,X33,X45)))) != s(X34,happ(s(fun(X45,X34),X44),s(X45,sK17(X40,X41,X42,X43,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X34,X33,X45)))) | s('type.sum.sum'(X45,X33),X42) != s('type.sum.sum'(X45,X33),X43)) )),
  inference(superposition,[],[f232,f228])).
fof(f4173,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X32,X22,X20] : (s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X30),s(fun(X32,X21),X28),s(fun(X20,X21),X27))) = s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X29),s(fun(X32,X21),X31),s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))))) | s(X21,happ(s(fun(X20,X21),X27),s(X20,sK18(X27,X28,X29,X30,X31,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X21,X20,X32)))) != s(X21,happ(s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))),s(X20,sK18(X27,X28,X29,X30,X31,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X21,X20,X32)))) | s(X21,happ(s(fun(X32,X21),X28),s(X32,sK17(X27,X28,X29,X30,X31,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X21,X20,X32)))) != s(X21,happ(s(fun(X32,X21),X31),s(X32,sK17(X27,X28,X29,X30,X31,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X21,X20,X32)))) | s('type.sum.sum'(X32,X20),X29) != s('type.sum.sum'(X32,X20),X30)) )),
  inference(forward_demodulation,[],[f4148,f227])).
fof(f4148,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X32,X22,X20] : (s(X21,happ(s(fun(X20,X21),X27),s(X20,sK18(X27,X28,X29,X30,X31,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X21,X20,X32)))) != s(X21,happ(s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))),s(X20,sK18(X27,X28,X29,X30,X31,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X21,X20,X32)))) | s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X30),s(fun(X32,X21),X28),s(fun(X20,X21),X27))) = s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X29),s(fun(X32,X21),X31),s(fun(X20,X21),'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26))))) | s(X21,happ(s(fun(X32,X21),X28),s(X32,sK17(X27,X28,X29,X30,X31,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X21,X20,X32)))) != s(X21,happ(s(fun(X32,X21),X31),s(X32,sK17(X27,X28,X29,X30,X31,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X21,X20,X32)))) | s('type.sum.sum'(X32,X20),X29) != s('type.sum.sum'(X32,X20),X30)) )),
  inference(superposition,[],[f232,f227])).
fof(f4172,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X16),s(fun(X19,X11),X18),s(fun(X10,X11),X13))) = s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X17),s(fun(X19,X11),X15),s(fun(X10,X11),X14))) | s(X11,happ(s(fun(X10,X11),X13),s(X10,sK18(X14,X15,X16,X17,X18,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X11,X10,X19)))) != s(X11,happ(s(fun(X10,X11),X14),s(X10,sK18(X14,X15,X16,X17,X18,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X11,X10,X19)))) | s(X11,happ(s(fun(X19,X11),X15),s(X19,sK17(X14,X15,X16,X17,X18,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X11,X10,X19)))) != s(X11,happ(s(fun(X19,X11),X18),s(X19,sK17(X14,X15,X16,X17,X18,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X11,X10,X19)))) | s('type.sum.sum'(X19,X10),X16) != s('type.sum.sum'(X19,X10),X17)) )),
  inference(forward_demodulation,[],[f4147,f194])).
fof(f4147,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (s(X11,happ(s(fun(X10,X11),X13),s(X10,sK18(X14,X15,X16,X17,X18,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X11,X10,X19)))) != s(X11,happ(s(fun(X10,X11),X14),s(X10,sK18(X14,X15,X16,X17,X18,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X11,X10,X19)))) | s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X17),s(fun(X19,X11),X15),s(fun(X10,X11),X14))) = s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X16),s(fun(X19,X11),X18),s(fun(X10,X11),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13))))) | s(X11,happ(s(fun(X19,X11),X15),s(X19,sK17(X14,X15,X16,X17,X18,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X11,X10,X19)))) != s(X11,happ(s(fun(X19,X11),X18),s(X19,sK17(X14,X15,X16,X17,X18,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X11,X10,X19)))) | s('type.sum.sum'(X19,X10),X16) != s('type.sum.sum'(X19,X10),X17)) )),
  inference(superposition,[],[f232,f194])).
fof(f4171,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X6),s(fun(X9,X1),X8),s(fun(X0,X1),X2))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X7),s(fun(X9,X1),X5),s(fun(X0,X1),X4))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,sK18(X4,X5,X6,X7,X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X1,X0,X9)))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,sK18(X4,X5,X6,X7,X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X1,X0,X9)))) | s(X1,happ(s(fun(X9,X1),X5),s(X9,sK17(X4,X5,X6,X7,X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X1,X0,X9)))) != s(X1,happ(s(fun(X9,X1),X8),s(X9,sK17(X4,X5,X6,X7,X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X1,X0,X9)))) | s('type.sum.sum'(X9,X0),X6) != s('type.sum.sum'(X9,X0),X7)) )),
  inference(forward_demodulation,[],[f4146,f193])).
fof(f4146,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,happ(s(fun(X0,X1),X2),s(X0,sK18(X4,X5,X6,X7,X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X1,X0,X9)))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,sK18(X4,X5,X6,X7,X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X1,X0,X9)))) | s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X7),s(fun(X9,X1),X5),s(fun(X0,X1),X4))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X6),s(fun(X9,X1),X8),s(fun(X0,X1),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3))))) | s(X1,happ(s(fun(X9,X1),X5),s(X9,sK17(X4,X5,X6,X7,X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X1,X0,X9)))) != s(X1,happ(s(fun(X9,X1),X8),s(X9,sK17(X4,X5,X6,X7,X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X1,X0,X9)))) | s('type.sum.sum'(X9,X0),X6) != s('type.sum.sum'(X9,X0),X7)) )),
  inference(superposition,[],[f232,f193])).
fof(f4032,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X1,X2),X3),s(fun(X1,X0),X4),s(fun(X2,X0),X5))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X1,X2),X6),s(fun(X1,X0),X7),s(fun(X2,X0),X5))) | s('type.sum.sum'(X1,X2),X6) = s('type.sum.sum'(X1,X2),'const.sum.INL_1'(s(X1,sK17(X5,X4,X6,X3,X7,X5,X0,X2,X1)))) | s('type.sum.sum'(X1,X2),X3) != s('type.sum.sum'(X1,X2),X6)) )),
  inference(equality_resolution,[],[f231])).
fof(f4047,plain,(
  ( ! [X61,X59,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X68),s(fun(X71,X60),X70),s(fun(X59,X60),X66))) = s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X69),s(fun(X71,X60),X67),s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))))) | s(X60,happ(s(fun(X59,X60),X66),s(X59,sK18(happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X67,X68,X69,X70,X66,X60,X59,X71)))) != s(X60,happ(s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))),s(X59,sK18(happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X67,X68,X69,X70,X66,X60,X59,X71)))) | s('type.sum.sum'(X71,X59),X68) = s('type.sum.sum'(X71,X59),'const.sum.INL_1'(s(X71,sK17(happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X67,X68,X69,X70,X66,X60,X59,X71)))) | s('type.sum.sum'(X71,X59),X68) != s('type.sum.sum'(X71,X59),X69)) )),
  inference(forward_demodulation,[],[f4023,f226])).
fof(f4023,plain,(
  ( ! [X61,X59,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,happ(s(fun(X59,X60),X66),s(X59,sK18(happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X67,X68,X69,X70,X66,X60,X59,X71)))) != s(X60,happ(s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))),s(X59,sK18(happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X67,X68,X69,X70,X66,X60,X59,X71)))) | s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X68),s(fun(X71,X60),X70),s(fun(X59,X60),X66))) = s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X69),s(fun(X71,X60),X67),s(fun(X59,X60),happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65))))))) | s('type.sum.sum'(X71,X59),X68) = s('type.sum.sum'(X71,X59),'const.sum.INL_1'(s(X71,sK17(happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X67,X68,X69,X70,X66,X60,X59,X71)))) | s('type.sum.sum'(X71,X59),X68) != s('type.sum.sum'(X71,X59),X69)) )),
  inference(superposition,[],[f231,f226])).
fof(f4046,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X55),s(fun(X58,X47),X57),s(fun(X46,X47),X53))) = s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X56),s(fun(X58,X47),X54),s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))))) | s(X47,happ(s(fun(X46,X47),X53),s(X46,sK18(happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X54,X55,X56,X57,X53,X47,X46,X58)))) != s(X47,happ(s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))),s(X46,sK18(happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X54,X55,X56,X57,X53,X47,X46,X58)))) | s('type.sum.sum'(X58,X46),X55) = s('type.sum.sum'(X58,X46),'const.sum.INL_1'(s(X58,sK17(happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X54,X55,X56,X57,X53,X47,X46,X58)))) | s('type.sum.sum'(X58,X46),X55) != s('type.sum.sum'(X58,X46),X56)) )),
  inference(forward_demodulation,[],[f4022,f225])).
fof(f4022,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (s(X47,happ(s(fun(X46,X47),X53),s(X46,sK18(happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X54,X55,X56,X57,X53,X47,X46,X58)))) != s(X47,happ(s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))),s(X46,sK18(happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X54,X55,X56,X57,X53,X47,X46,X58)))) | s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X55),s(fun(X58,X47),X57),s(fun(X46,X47),X53))) = s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X56),s(fun(X58,X47),X54),s(fun(X46,X47),happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52))))))) | s('type.sum.sum'(X58,X46),X55) = s('type.sum.sum'(X58,X46),'const.sum.INL_1'(s(X58,sK17(happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X54,X55,X56,X57,X53,X47,X46,X58)))) | s('type.sum.sum'(X58,X46),X55) != s('type.sum.sum'(X58,X46),X56)) )),
  inference(superposition,[],[f231,f225])).
fof(f4045,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X44,X42,X40] : (s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X42),s(fun(X45,X34),X44),s(fun(X33,X34),X40))) = s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X43),s(fun(X45,X34),X41),s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))))) | s(X34,happ(s(fun(X33,X34),X40),s(X33,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X41,X42,X43,X44,X40,X34,X33,X45)))) != s(X34,happ(s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))),s(X33,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X41,X42,X43,X44,X40,X34,X33,X45)))) | s('type.sum.sum'(X45,X33),X42) = s('type.sum.sum'(X45,X33),'const.sum.INL_1'(s(X45,sK17('const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X41,X42,X43,X44,X40,X34,X33,X45)))) | s('type.sum.sum'(X45,X33),X42) != s('type.sum.sum'(X45,X33),X43)) )),
  inference(forward_demodulation,[],[f4021,f228])).
fof(f4021,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X44,X42,X40] : (s(X34,happ(s(fun(X33,X34),X40),s(X33,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X41,X42,X43,X44,X40,X34,X33,X45)))) != s(X34,happ(s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))),s(X33,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X41,X42,X43,X44,X40,X34,X33,X45)))) | s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X42),s(fun(X45,X34),X44),s(fun(X33,X34),X40))) = s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X43),s(fun(X45,X34),X41),s(fun(X33,X34),'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39))))) | s('type.sum.sum'(X45,X33),X42) = s('type.sum.sum'(X45,X33),'const.sum.INL_1'(s(X45,sK17('const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X41,X42,X43,X44,X40,X34,X33,X45)))) | s('type.sum.sum'(X45,X33),X42) != s('type.sum.sum'(X45,X33),X43)) )),
  inference(superposition,[],[f231,f228])).
fof(f4044,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X32,X22,X20] : (s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X29),s(fun(X32,X21),X31),s(fun(X20,X21),X27))) = s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X30),s(fun(X32,X21),X28),s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))))) | s(X21,happ(s(fun(X20,X21),X27),s(X20,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X28,X29,X30,X31,X27,X21,X20,X32)))) != s(X21,happ(s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))),s(X20,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X28,X29,X30,X31,X27,X21,X20,X32)))) | s('type.sum.sum'(X32,X20),X29) = s('type.sum.sum'(X32,X20),'const.sum.INL_1'(s(X32,sK17('const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X28,X29,X30,X31,X27,X21,X20,X32)))) | s('type.sum.sum'(X32,X20),X29) != s('type.sum.sum'(X32,X20),X30)) )),
  inference(forward_demodulation,[],[f4020,f227])).
fof(f4020,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X32,X22,X20] : (s(X21,happ(s(fun(X20,X21),X27),s(X20,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X28,X29,X30,X31,X27,X21,X20,X32)))) != s(X21,happ(s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))),s(X20,sK18('const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X28,X29,X30,X31,X27,X21,X20,X32)))) | s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X29),s(fun(X32,X21),X31),s(fun(X20,X21),X27))) = s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X30),s(fun(X32,X21),X28),s(fun(X20,X21),'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26))))) | s('type.sum.sum'(X32,X20),X29) = s('type.sum.sum'(X32,X20),'const.sum.INL_1'(s(X32,sK17('const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X28,X29,X30,X31,X27,X21,X20,X32)))) | s('type.sum.sum'(X32,X20),X29) != s('type.sum.sum'(X32,X20),X30)) )),
  inference(superposition,[],[f231,f227])).
fof(f4043,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X16),s(fun(X19,X11),X18),s(fun(X10,X11),X14))) = s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X17),s(fun(X19,X11),X15),s(fun(X10,X11),X13))) | s(X11,happ(s(fun(X10,X11),X13),s(X10,sK18('const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X15,X16,X17,X18,X14,X11,X10,X19)))) != s(X11,happ(s(fun(X10,X11),X14),s(X10,sK18('const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X15,X16,X17,X18,X14,X11,X10,X19)))) | s('type.sum.sum'(X19,X10),X16) = s('type.sum.sum'(X19,X10),'const.sum.INL_1'(s(X19,sK17('const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X15,X16,X17,X18,X14,X11,X10,X19)))) | s('type.sum.sum'(X19,X10),X16) != s('type.sum.sum'(X19,X10),X17)) )),
  inference(forward_demodulation,[],[f4019,f194])).
fof(f4019,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (s(X11,happ(s(fun(X10,X11),X13),s(X10,sK18('const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X15,X16,X17,X18,X14,X11,X10,X19)))) != s(X11,happ(s(fun(X10,X11),X14),s(X10,sK18('const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X15,X16,X17,X18,X14,X11,X10,X19)))) | s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X16),s(fun(X19,X11),X18),s(fun(X10,X11),X14))) = s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X17),s(fun(X19,X11),X15),s(fun(X10,X11),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13))))) | s('type.sum.sum'(X19,X10),X16) = s('type.sum.sum'(X19,X10),'const.sum.INL_1'(s(X19,sK17('const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X15,X16,X17,X18,X14,X11,X10,X19)))) | s('type.sum.sum'(X19,X10),X16) != s('type.sum.sum'(X19,X10),X17)) )),
  inference(superposition,[],[f231,f194])).
fof(f4042,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X6),s(fun(X9,X1),X8),s(fun(X0,X1),X4))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X7),s(fun(X9,X1),X5),s(fun(X0,X1),X2))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,sK18('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X5,X6,X7,X8,X4,X1,X0,X9)))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,sK18('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X5,X6,X7,X8,X4,X1,X0,X9)))) | s('type.sum.sum'(X9,X0),X6) = s('type.sum.sum'(X9,X0),'const.sum.INL_1'(s(X9,sK17('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X5,X6,X7,X8,X4,X1,X0,X9)))) | s('type.sum.sum'(X9,X0),X6) != s('type.sum.sum'(X9,X0),X7)) )),
  inference(forward_demodulation,[],[f4018,f193])).
fof(f4018,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,happ(s(fun(X0,X1),X2),s(X0,sK18('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X5,X6,X7,X8,X4,X1,X0,X9)))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,sK18('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X5,X6,X7,X8,X4,X1,X0,X9)))) | s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X6),s(fun(X9,X1),X8),s(fun(X0,X1),X4))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X7),s(fun(X9,X1),X5),s(fun(X0,X1),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3))))) | s('type.sum.sum'(X9,X0),X6) = s('type.sum.sum'(X9,X0),'const.sum.INL_1'(s(X9,sK17('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X5,X6,X7,X8,X4,X1,X0,X9)))) | s('type.sum.sum'(X9,X0),X6) != s('type.sum.sum'(X9,X0),X7)) )),
  inference(superposition,[],[f231,f193])).
fof(f4038,plain,(
  ( ! [X61,X59,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X69),s(fun(X71,X60),X67),s(fun(X59,X60),X66))) = s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X68),s(fun(X71,X60),X70),s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))))) | s(X60,happ(s(fun(X59,X60),X66),s(X59,sK18(X66,X67,X68,X69,X70,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X60,X59,X71)))) != s(X60,happ(s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))),s(X59,sK18(X66,X67,X68,X69,X70,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X60,X59,X71)))) | s('type.sum.sum'(X71,X59),X68) = s('type.sum.sum'(X71,X59),'const.sum.INL_1'(s(X71,sK17(X66,X67,X68,X69,X70,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X60,X59,X71)))) | s('type.sum.sum'(X71,X59),X68) != s('type.sum.sum'(X71,X59),X69)) )),
  inference(forward_demodulation,[],[f4013,f226])).
fof(f4013,plain,(
  ( ! [X61,X59,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,happ(s(fun(X59,X60),X66),s(X59,sK18(X66,X67,X68,X69,X70,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X60,X59,X71)))) != s(X60,happ(s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))),s(X59,sK18(X66,X67,X68,X69,X70,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X60,X59,X71)))) | s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X69),s(fun(X71,X60),X67),s(fun(X59,X60),X66))) = s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X59),X68),s(fun(X71,X60),X70),s(fun(X59,X60),happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65))))))) | s('type.sum.sum'(X71,X59),X68) = s('type.sum.sum'(X71,X59),'const.sum.INL_1'(s(X71,sK17(X66,X67,X68,X69,X70,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X60,X59,X71)))) | s('type.sum.sum'(X71,X59),X68) != s('type.sum.sum'(X71,X59),X69)) )),
  inference(superposition,[],[f231,f226])).
fof(f4037,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X56),s(fun(X58,X47),X54),s(fun(X46,X47),X53))) = s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X55),s(fun(X58,X47),X57),s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))))) | s(X47,happ(s(fun(X46,X47),X53),s(X46,sK18(X53,X54,X55,X56,X57,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X47,X46,X58)))) != s(X47,happ(s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))),s(X46,sK18(X53,X54,X55,X56,X57,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X47,X46,X58)))) | s('type.sum.sum'(X58,X46),X55) = s('type.sum.sum'(X58,X46),'const.sum.INL_1'(s(X58,sK17(X53,X54,X55,X56,X57,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X47,X46,X58)))) | s('type.sum.sum'(X58,X46),X55) != s('type.sum.sum'(X58,X46),X56)) )),
  inference(forward_demodulation,[],[f4012,f225])).
fof(f4012,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (s(X47,happ(s(fun(X46,X47),X53),s(X46,sK18(X53,X54,X55,X56,X57,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X47,X46,X58)))) != s(X47,happ(s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))),s(X46,sK18(X53,X54,X55,X56,X57,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X47,X46,X58)))) | s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X56),s(fun(X58,X47),X54),s(fun(X46,X47),X53))) = s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X58,X46),X55),s(fun(X58,X47),X57),s(fun(X46,X47),happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52))))))) | s('type.sum.sum'(X58,X46),X55) = s('type.sum.sum'(X58,X46),'const.sum.INL_1'(s(X58,sK17(X53,X54,X55,X56,X57,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X47,X46,X58)))) | s('type.sum.sum'(X58,X46),X55) != s('type.sum.sum'(X58,X46),X56)) )),
  inference(superposition,[],[f231,f225])).
fof(f4036,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X44,X42,X40] : (s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X43),s(fun(X45,X34),X41),s(fun(X33,X34),X40))) = s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X42),s(fun(X45,X34),X44),s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))))) | s(X34,happ(s(fun(X33,X34),X40),s(X33,sK18(X40,X41,X42,X43,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X34,X33,X45)))) != s(X34,happ(s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))),s(X33,sK18(X40,X41,X42,X43,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X34,X33,X45)))) | s('type.sum.sum'(X45,X33),X42) = s('type.sum.sum'(X45,X33),'const.sum.INL_1'(s(X45,sK17(X40,X41,X42,X43,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X34,X33,X45)))) | s('type.sum.sum'(X45,X33),X42) != s('type.sum.sum'(X45,X33),X43)) )),
  inference(forward_demodulation,[],[f4011,f228])).
fof(f4011,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X44,X42,X40] : (s(X34,happ(s(fun(X33,X34),X40),s(X33,sK18(X40,X41,X42,X43,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X34,X33,X45)))) != s(X34,happ(s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))),s(X33,sK18(X40,X41,X42,X43,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X34,X33,X45)))) | s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X43),s(fun(X45,X34),X41),s(fun(X33,X34),X40))) = s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X33),X42),s(fun(X45,X34),X44),s(fun(X33,X34),'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39))))) | s('type.sum.sum'(X45,X33),X42) = s('type.sum.sum'(X45,X33),'const.sum.INL_1'(s(X45,sK17(X40,X41,X42,X43,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X34,X33,X45)))) | s('type.sum.sum'(X45,X33),X42) != s('type.sum.sum'(X45,X33),X43)) )),
  inference(superposition,[],[f231,f228])).
fof(f4035,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X32,X22,X20] : (s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X30),s(fun(X32,X21),X28),s(fun(X20,X21),X27))) = s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X29),s(fun(X32,X21),X31),s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))))) | s(X21,happ(s(fun(X20,X21),X27),s(X20,sK18(X27,X28,X29,X30,X31,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X21,X20,X32)))) != s(X21,happ(s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))),s(X20,sK18(X27,X28,X29,X30,X31,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X21,X20,X32)))) | s('type.sum.sum'(X32,X20),X29) = s('type.sum.sum'(X32,X20),'const.sum.INL_1'(s(X32,sK17(X27,X28,X29,X30,X31,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X21,X20,X32)))) | s('type.sum.sum'(X32,X20),X29) != s('type.sum.sum'(X32,X20),X30)) )),
  inference(forward_demodulation,[],[f4010,f227])).
fof(f4010,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X32,X22,X20] : (s(X21,happ(s(fun(X20,X21),X27),s(X20,sK18(X27,X28,X29,X30,X31,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X21,X20,X32)))) != s(X21,happ(s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))),s(X20,sK18(X27,X28,X29,X30,X31,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X21,X20,X32)))) | s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X30),s(fun(X32,X21),X28),s(fun(X20,X21),X27))) = s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X20),X29),s(fun(X32,X21),X31),s(fun(X20,X21),'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26))))) | s('type.sum.sum'(X32,X20),X29) = s('type.sum.sum'(X32,X20),'const.sum.INL_1'(s(X32,sK17(X27,X28,X29,X30,X31,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X21,X20,X32)))) | s('type.sum.sum'(X32,X20),X29) != s('type.sum.sum'(X32,X20),X30)) )),
  inference(superposition,[],[f231,f227])).
fof(f4034,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X16),s(fun(X19,X11),X18),s(fun(X10,X11),X13))) = s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X17),s(fun(X19,X11),X15),s(fun(X10,X11),X14))) | s(X11,happ(s(fun(X10,X11),X13),s(X10,sK18(X14,X15,X16,X17,X18,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X11,X10,X19)))) != s(X11,happ(s(fun(X10,X11),X14),s(X10,sK18(X14,X15,X16,X17,X18,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X11,X10,X19)))) | s('type.sum.sum'(X19,X10),X16) = s('type.sum.sum'(X19,X10),'const.sum.INL_1'(s(X19,sK17(X14,X15,X16,X17,X18,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X11,X10,X19)))) | s('type.sum.sum'(X19,X10),X16) != s('type.sum.sum'(X19,X10),X17)) )),
  inference(forward_demodulation,[],[f4009,f194])).
fof(f4009,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (s(X11,happ(s(fun(X10,X11),X13),s(X10,sK18(X14,X15,X16,X17,X18,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X11,X10,X19)))) != s(X11,happ(s(fun(X10,X11),X14),s(X10,sK18(X14,X15,X16,X17,X18,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X11,X10,X19)))) | s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X17),s(fun(X19,X11),X15),s(fun(X10,X11),X14))) = s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X10),X16),s(fun(X19,X11),X18),s(fun(X10,X11),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13))))) | s('type.sum.sum'(X19,X10),X16) = s('type.sum.sum'(X19,X10),'const.sum.INL_1'(s(X19,sK17(X14,X15,X16,X17,X18,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X11,X10,X19)))) | s('type.sum.sum'(X19,X10),X16) != s('type.sum.sum'(X19,X10),X17)) )),
  inference(superposition,[],[f231,f194])).
fof(f4033,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X6),s(fun(X9,X1),X8),s(fun(X0,X1),X2))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X7),s(fun(X9,X1),X5),s(fun(X0,X1),X4))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,sK18(X4,X5,X6,X7,X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X1,X0,X9)))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,sK18(X4,X5,X6,X7,X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X1,X0,X9)))) | s('type.sum.sum'(X9,X0),X6) = s('type.sum.sum'(X9,X0),'const.sum.INL_1'(s(X9,sK17(X4,X5,X6,X7,X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X1,X0,X9)))) | s('type.sum.sum'(X9,X0),X6) != s('type.sum.sum'(X9,X0),X7)) )),
  inference(forward_demodulation,[],[f4008,f193])).
fof(f4008,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,happ(s(fun(X0,X1),X2),s(X0,sK18(X4,X5,X6,X7,X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X1,X0,X9)))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,sK18(X4,X5,X6,X7,X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X1,X0,X9)))) | s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X7),s(fun(X9,X1),X5),s(fun(X0,X1),X4))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X0),X6),s(fun(X9,X1),X8),s(fun(X0,X1),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3))))) | s('type.sum.sum'(X9,X0),X6) = s('type.sum.sum'(X9,X0),'const.sum.INL_1'(s(X9,sK17(X4,X5,X6,X7,X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X1,X0,X9)))) | s('type.sum.sum'(X9,X0),X6) != s('type.sum.sum'(X9,X0),X7)) )),
  inference(superposition,[],[f231,f193])).
fof(f3664,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK18(X3,X4,X2,X5,X4,X6,X7,X1,X0)))) | s(X7,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X2),s(fun(X0,X7),X4),s(fun(X1,X7),X6))) = s(X7,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X5),s(fun(X0,X7),X4),s(fun(X1,X7),X3))) | s('type.sum.sum'(X0,X1),X2) != s('type.sum.sum'(X0,X1),X5)) )),
  inference(equality_resolution,[],[f230])).
fof(f3678,plain,(
  ( ! [X61,X59,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X71),X68),s(fun(X59,X60),X66),s(fun(X71,X60),X70))) = s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X71),X69),s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))),s(fun(X71,X60),X67))) | s(X60,happ(s(fun(X59,X60),X66),s(X59,sK17(X67,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X68,X69,X66,X70,X60,X71,X59)))) != s(X60,happ(s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))),s(X59,sK17(X67,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X68,X69,X66,X70,X60,X71,X59)))) | s('type.sum.sum'(X59,X71),X68) = s('type.sum.sum'(X59,X71),'const.sum.INR_1'(s(X71,sK18(X67,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X68,X69,X66,X70,X60,X71,X59)))) | s('type.sum.sum'(X59,X71),X68) != s('type.sum.sum'(X59,X71),X69)) )),
  inference(forward_demodulation,[],[f3657,f226])).
fof(f3657,plain,(
  ( ! [X61,X59,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,happ(s(fun(X59,X60),X66),s(X59,sK17(X67,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X68,X69,X66,X70,X60,X71,X59)))) != s(X60,happ(s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))),s(X59,sK17(X67,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X68,X69,X66,X70,X60,X71,X59)))) | s('type.sum.sum'(X59,X71),X68) = s('type.sum.sum'(X59,X71),'const.sum.INR_1'(s(X71,sK18(X67,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X68,X69,X66,X70,X60,X71,X59)))) | s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X71),X68),s(fun(X59,X60),X66),s(fun(X71,X60),X70))) = s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X71),X69),s(fun(X59,X60),happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65))))),s(fun(X71,X60),X67))) | s('type.sum.sum'(X59,X71),X68) != s('type.sum.sum'(X59,X71),X69)) )),
  inference(superposition,[],[f230,f226])).
fof(f3677,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X58),X55),s(fun(X46,X47),X53),s(fun(X58,X47),X57))) = s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X58),X56),s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))),s(fun(X58,X47),X54))) | s(X47,happ(s(fun(X46,X47),X53),s(X46,sK17(X54,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X55,X56,X53,X57,X47,X58,X46)))) != s(X47,happ(s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))),s(X46,sK17(X54,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X55,X56,X53,X57,X47,X58,X46)))) | s('type.sum.sum'(X46,X58),X55) = s('type.sum.sum'(X46,X58),'const.sum.INR_1'(s(X58,sK18(X54,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X55,X56,X53,X57,X47,X58,X46)))) | s('type.sum.sum'(X46,X58),X55) != s('type.sum.sum'(X46,X58),X56)) )),
  inference(forward_demodulation,[],[f3656,f225])).
fof(f3656,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (s(X47,happ(s(fun(X46,X47),X53),s(X46,sK17(X54,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X55,X56,X53,X57,X47,X58,X46)))) != s(X47,happ(s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))),s(X46,sK17(X54,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X55,X56,X53,X57,X47,X58,X46)))) | s('type.sum.sum'(X46,X58),X55) = s('type.sum.sum'(X46,X58),'const.sum.INR_1'(s(X58,sK18(X54,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X55,X56,X53,X57,X47,X58,X46)))) | s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X58),X55),s(fun(X46,X47),X53),s(fun(X58,X47),X57))) = s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X58),X56),s(fun(X46,X47),happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52))))),s(fun(X58,X47),X54))) | s('type.sum.sum'(X46,X58),X55) != s('type.sum.sum'(X46,X58),X56)) )),
  inference(superposition,[],[f230,f225])).
fof(f3676,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X44,X42,X40] : (s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X33,X45),X42),s(fun(X33,X34),X40),s(fun(X45,X34),X44))) = s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X33,X45),X43),s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))),s(fun(X45,X34),X41))) | s(X34,happ(s(fun(X33,X34),X40),s(X33,sK17(X41,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X42,X43,X40,X44,X34,X45,X33)))) != s(X34,happ(s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))),s(X33,sK17(X41,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X42,X43,X40,X44,X34,X45,X33)))) | s('type.sum.sum'(X33,X45),X42) = s('type.sum.sum'(X33,X45),'const.sum.INR_1'(s(X45,sK18(X41,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X42,X43,X40,X44,X34,X45,X33)))) | s('type.sum.sum'(X33,X45),X42) != s('type.sum.sum'(X33,X45),X43)) )),
  inference(forward_demodulation,[],[f3655,f228])).
fof(f3655,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X44,X42,X40] : (s(X34,happ(s(fun(X33,X34),X40),s(X33,sK17(X41,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X42,X43,X40,X44,X34,X45,X33)))) != s(X34,happ(s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))),s(X33,sK17(X41,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X42,X43,X40,X44,X34,X45,X33)))) | s('type.sum.sum'(X33,X45),X42) = s('type.sum.sum'(X33,X45),'const.sum.INR_1'(s(X45,sK18(X41,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X42,X43,X40,X44,X34,X45,X33)))) | s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X33,X45),X42),s(fun(X33,X34),X40),s(fun(X45,X34),X44))) = s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X33,X45),X43),s(fun(X33,X34),'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39))),s(fun(X45,X34),X41))) | s('type.sum.sum'(X33,X45),X42) != s('type.sum.sum'(X33,X45),X43)) )),
  inference(superposition,[],[f230,f228])).
fof(f3675,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X32,X22,X20] : (s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X32),X29),s(fun(X20,X21),X27),s(fun(X32,X21),X31))) = s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X32),X30),s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))),s(fun(X32,X21),X28))) | s(X21,happ(s(fun(X20,X21),X27),s(X20,sK17(X28,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X29,X30,X27,X31,X21,X32,X20)))) != s(X21,happ(s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))),s(X20,sK17(X28,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X29,X30,X27,X31,X21,X32,X20)))) | s('type.sum.sum'(X20,X32),X29) = s('type.sum.sum'(X20,X32),'const.sum.INR_1'(s(X32,sK18(X28,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X29,X30,X27,X31,X21,X32,X20)))) | s('type.sum.sum'(X20,X32),X29) != s('type.sum.sum'(X20,X32),X30)) )),
  inference(forward_demodulation,[],[f3654,f227])).
fof(f3654,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X32,X22,X20] : (s(X21,happ(s(fun(X20,X21),X27),s(X20,sK17(X28,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X29,X30,X27,X31,X21,X32,X20)))) != s(X21,happ(s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))),s(X20,sK17(X28,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X29,X30,X27,X31,X21,X32,X20)))) | s('type.sum.sum'(X20,X32),X29) = s('type.sum.sum'(X20,X32),'const.sum.INR_1'(s(X32,sK18(X28,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X29,X30,X27,X31,X21,X32,X20)))) | s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X32),X29),s(fun(X20,X21),X27),s(fun(X32,X21),X31))) = s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X32),X30),s(fun(X20,X21),'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26))),s(fun(X32,X21),X28))) | s('type.sum.sum'(X20,X32),X29) != s('type.sum.sum'(X20,X32),X30)) )),
  inference(superposition,[],[f230,f227])).
fof(f3674,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X10,X19),X16),s(fun(X10,X11),X14),s(fun(X19,X11),X18))) = s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X10,X19),X17),s(fun(X10,X11),X13),s(fun(X19,X11),X15))) | s(X11,happ(s(fun(X10,X11),X13),s(X10,sK17(X15,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X16,X17,X14,X18,X11,X19,X10)))) != s(X11,happ(s(fun(X10,X11),X14),s(X10,sK17(X15,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X16,X17,X14,X18,X11,X19,X10)))) | s('type.sum.sum'(X10,X19),X16) = s('type.sum.sum'(X10,X19),'const.sum.INR_1'(s(X19,sK18(X15,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X16,X17,X14,X18,X11,X19,X10)))) | s('type.sum.sum'(X10,X19),X16) != s('type.sum.sum'(X10,X19),X17)) )),
  inference(forward_demodulation,[],[f3653,f194])).
fof(f3653,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (s(X11,happ(s(fun(X10,X11),X13),s(X10,sK17(X15,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X16,X17,X14,X18,X11,X19,X10)))) != s(X11,happ(s(fun(X10,X11),X14),s(X10,sK17(X15,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X16,X17,X14,X18,X11,X19,X10)))) | s('type.sum.sum'(X10,X19),X16) = s('type.sum.sum'(X10,X19),'const.sum.INR_1'(s(X19,sK18(X15,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X16,X17,X14,X18,X11,X19,X10)))) | s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X10,X19),X16),s(fun(X10,X11),X14),s(fun(X19,X11),X18))) = s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X10,X19),X17),s(fun(X10,X11),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13))),s(fun(X19,X11),X15))) | s('type.sum.sum'(X10,X19),X16) != s('type.sum.sum'(X10,X19),X17)) )),
  inference(superposition,[],[f230,f194])).
fof(f3673,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X9),X6),s(fun(X0,X1),X4),s(fun(X9,X1),X8))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X9),X7),s(fun(X0,X1),X2),s(fun(X9,X1),X5))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,sK17(X5,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X6,X7,X4,X8,X1,X9,X0)))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,sK17(X5,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X6,X7,X4,X8,X1,X9,X0)))) | s('type.sum.sum'(X0,X9),X6) = s('type.sum.sum'(X0,X9),'const.sum.INR_1'(s(X9,sK18(X5,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X6,X7,X4,X8,X1,X9,X0)))) | s('type.sum.sum'(X0,X9),X6) != s('type.sum.sum'(X0,X9),X7)) )),
  inference(forward_demodulation,[],[f3652,f193])).
fof(f3652,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,happ(s(fun(X0,X1),X2),s(X0,sK17(X5,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X6,X7,X4,X8,X1,X9,X0)))) != s(X1,happ(s(fun(X0,X1),X4),s(X0,sK17(X5,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X6,X7,X4,X8,X1,X9,X0)))) | s('type.sum.sum'(X0,X9),X6) = s('type.sum.sum'(X0,X9),'const.sum.INR_1'(s(X9,sK18(X5,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X6,X7,X4,X8,X1,X9,X0)))) | s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X9),X6),s(fun(X0,X1),X4),s(fun(X9,X1),X8))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X9),X7),s(fun(X0,X1),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3))),s(fun(X9,X1),X5))) | s('type.sum.sum'(X0,X9),X6) != s('type.sum.sum'(X0,X9),X7)) )),
  inference(superposition,[],[f230,f193])).
fof(f3670,plain,(
  ( ! [X61,X59,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X71),X69),s(fun(X59,X60),X67),s(fun(X71,X60),X66))) = s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X71),X68),s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))),s(fun(X71,X60),X70))) | s(X60,happ(s(fun(X59,X60),X67),s(X59,sK17(X66,X67,X68,X69,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X70,X60,X71,X59)))) != s(X60,happ(s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))),s(X59,sK17(X66,X67,X68,X69,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X70,X60,X71,X59)))) | s('type.sum.sum'(X59,X71),X68) = s('type.sum.sum'(X59,X71),'const.sum.INR_1'(s(X71,sK18(X66,X67,X68,X69,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X70,X60,X71,X59)))) | s('type.sum.sum'(X59,X71),X68) != s('type.sum.sum'(X59,X71),X69)) )),
  inference(forward_demodulation,[],[f3648,f226])).
fof(f3648,plain,(
  ( ! [X61,X59,X70,X68,X66,X64,X62,X60,X71,X69,X67,X65,X63] : (s(X60,happ(s(fun(X59,X60),X67),s(X59,sK17(X66,X67,X68,X69,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X70,X60,X71,X59)))) != s(X60,happ(s(fun(X59,X60),happ(s(fun(X62,fun(X59,X60)),X63),s(X62,X65))),s(X59,sK17(X66,X67,X68,X69,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X70,X60,X71,X59)))) | s('type.sum.sum'(X59,X71),X68) = s('type.sum.sum'(X59,X71),'const.sum.INR_1'(s(X71,sK18(X66,X67,X68,X69,happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65)))),X70,X60,X71,X59)))) | s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X71),X69),s(fun(X59,X60),X67),s(fun(X71,X60),X66))) = s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X71),X68),s(fun(X59,X60),happ(s(fun('type.sum.sum'(X61,X62),fun(X59,X60)),sK16(X63,X64,X62,fun(X59,X60),X61)),s('type.sum.sum'(X61,X62),'const.sum.INR_1'(s(X62,X65))))),s(fun(X71,X60),X70))) | s('type.sum.sum'(X59,X71),X68) != s('type.sum.sum'(X59,X71),X69)) )),
  inference(superposition,[],[f230,f226])).
fof(f3669,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X58),X56),s(fun(X46,X47),X54),s(fun(X58,X47),X53))) = s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X58),X55),s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))),s(fun(X58,X47),X57))) | s(X47,happ(s(fun(X46,X47),X54),s(X46,sK17(X53,X54,X55,X56,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X57,X47,X58,X46)))) != s(X47,happ(s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))),s(X46,sK17(X53,X54,X55,X56,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X57,X47,X58,X46)))) | s('type.sum.sum'(X46,X58),X55) = s('type.sum.sum'(X46,X58),'const.sum.INR_1'(s(X58,sK18(X53,X54,X55,X56,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X57,X47,X58,X46)))) | s('type.sum.sum'(X46,X58),X55) != s('type.sum.sum'(X46,X58),X56)) )),
  inference(forward_demodulation,[],[f3647,f225])).
fof(f3647,plain,(
  ( ! [X57,X47,X54,X52,X50,X48,X58,X56,X46,X55,X53,X51,X49] : (s(X47,happ(s(fun(X46,X47),X54),s(X46,sK17(X53,X54,X55,X56,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X57,X47,X58,X46)))) != s(X47,happ(s(fun(X46,X47),happ(s(fun(X48,fun(X46,X47)),X51),s(X48,X52))),s(X46,sK17(X53,X54,X55,X56,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X57,X47,X58,X46)))) | s('type.sum.sum'(X46,X58),X55) = s('type.sum.sum'(X46,X58),'const.sum.INR_1'(s(X58,sK18(X53,X54,X55,X56,happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52)))),X57,X47,X58,X46)))) | s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X58),X56),s(fun(X46,X47),X54),s(fun(X58,X47),X53))) = s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X58),X55),s(fun(X46,X47),happ(s(fun('type.sum.sum'(X48,X49),fun(X46,X47)),sK16(X50,X51,X49,fun(X46,X47),X48)),s('type.sum.sum'(X48,X49),'const.sum.INL_1'(s(X48,X52))))),s(fun(X58,X47),X57))) | s('type.sum.sum'(X46,X58),X55) != s('type.sum.sum'(X46,X58),X56)) )),
  inference(superposition,[],[f230,f225])).
fof(f3668,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X44,X42,X40] : (s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X33,X45),X43),s(fun(X33,X34),X41),s(fun(X45,X34),X40))) = s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X33,X45),X42),s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))),s(fun(X45,X34),X44))) | s(X34,happ(s(fun(X33,X34),X41),s(X33,sK17(X40,X41,X42,X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X44,X34,X45,X33)))) != s(X34,happ(s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))),s(X33,sK17(X40,X41,X42,X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X44,X34,X45,X33)))) | s('type.sum.sum'(X33,X45),X42) = s('type.sum.sum'(X33,X45),'const.sum.INR_1'(s(X45,sK18(X40,X41,X42,X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X44,X34,X45,X33)))) | s('type.sum.sum'(X33,X45),X42) != s('type.sum.sum'(X33,X45),X43)) )),
  inference(forward_demodulation,[],[f3646,f228])).
fof(f3646,plain,(
  ( ! [X39,X37,X35,X33,X45,X43,X41,X38,X36,X34,X44,X42,X40] : (s(X34,happ(s(fun(X33,X34),X41),s(X33,sK17(X40,X41,X42,X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X44,X34,X45,X33)))) != s(X34,happ(s(fun(X33,X34),happ(s(fun(X35,fun(X33,X34)),X38),s(X35,X37))),s(X33,sK17(X40,X41,X42,X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X44,X34,X45,X33)))) | s('type.sum.sum'(X33,X45),X42) = s('type.sum.sum'(X33,X45),'const.sum.INR_1'(s(X45,sK18(X40,X41,X42,X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39)),X44,X34,X45,X33)))) | s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X33,X45),X43),s(fun(X33,X34),X41),s(fun(X45,X34),X40))) = s(X34,'const.sum.sum_CASE_3'(s('type.sum.sum'(X33,X45),X42),s(fun(X33,X34),'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,X37))),s(fun(X35,fun(X33,X34)),X38),s(fun(X36,fun(X33,X34)),X39))),s(fun(X45,X34),X44))) | s('type.sum.sum'(X33,X45),X42) != s('type.sum.sum'(X33,X45),X43)) )),
  inference(superposition,[],[f230,f228])).
fof(f3667,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X32,X22,X20] : (s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X32),X30),s(fun(X20,X21),X28),s(fun(X32,X21),X27))) = s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X32),X29),s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))),s(fun(X32,X21),X31))) | s(X21,happ(s(fun(X20,X21),X28),s(X20,sK17(X27,X28,X29,X30,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X31,X21,X32,X20)))) != s(X21,happ(s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))),s(X20,sK17(X27,X28,X29,X30,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X31,X21,X32,X20)))) | s('type.sum.sum'(X20,X32),X29) = s('type.sum.sum'(X20,X32),'const.sum.INR_1'(s(X32,sK18(X27,X28,X29,X30,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X31,X21,X32,X20)))) | s('type.sum.sum'(X20,X32),X29) != s('type.sum.sum'(X20,X32),X30)) )),
  inference(forward_demodulation,[],[f3645,f227])).
fof(f3645,plain,(
  ( ! [X30,X28,X26,X24,X23,X21,X31,X29,X27,X25,X32,X22,X20] : (s(X21,happ(s(fun(X20,X21),X28),s(X20,sK17(X27,X28,X29,X30,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X31,X21,X32,X20)))) != s(X21,happ(s(fun(X20,X21),happ(s(fun(X23,fun(X20,X21)),X26),s(X23,X24))),s(X20,sK17(X27,X28,X29,X30,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X31,X21,X32,X20)))) | s('type.sum.sum'(X20,X32),X29) = s('type.sum.sum'(X20,X32),'const.sum.INR_1'(s(X32,sK18(X27,X28,X29,X30,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26)),X31,X21,X32,X20)))) | s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X32),X30),s(fun(X20,X21),X28),s(fun(X32,X21),X27))) = s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X32),X29),s(fun(X20,X21),'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),'const.sum.INR_1'(s(X23,X24))),s(fun(X22,fun(X20,X21)),X25),s(fun(X23,fun(X20,X21)),X26))),s(fun(X32,X21),X31))) | s('type.sum.sum'(X20,X32),X29) != s('type.sum.sum'(X20,X32),X30)) )),
  inference(superposition,[],[f230,f227])).
fof(f3666,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X10,X19),X16),s(fun(X10,X11),X13),s(fun(X19,X11),X18))) = s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X10,X19),X17),s(fun(X10,X11),X15),s(fun(X19,X11),X14))) | s(X11,happ(s(fun(X10,X11),X13),s(X10,sK17(X14,X15,X16,X17,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X18,X11,X19,X10)))) != s(X11,happ(s(fun(X10,X11),X15),s(X10,sK17(X14,X15,X16,X17,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X18,X11,X19,X10)))) | s('type.sum.sum'(X10,X19),X16) = s('type.sum.sum'(X10,X19),'const.sum.INR_1'(s(X19,sK18(X14,X15,X16,X17,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X18,X11,X19,X10)))) | s('type.sum.sum'(X10,X19),X16) != s('type.sum.sum'(X10,X19),X17)) )),
  inference(forward_demodulation,[],[f3644,f194])).
fof(f3644,plain,(
  ( ! [X14,X12,X10,X19,X17,X15,X13,X11,X18,X16] : (s(X11,happ(s(fun(X10,X11),X13),s(X10,sK17(X14,X15,X16,X17,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X18,X11,X19,X10)))) != s(X11,happ(s(fun(X10,X11),X15),s(X10,sK17(X14,X15,X16,X17,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X18,X11,X19,X10)))) | s('type.sum.sum'(X10,X19),X16) = s('type.sum.sum'(X10,X19),'const.sum.INR_1'(s(X19,sK18(X14,X15,X16,X17,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13)),X18,X11,X19,X10)))) | s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X10,X19),X17),s(fun(X10,X11),X15),s(fun(X19,X11),X14))) = s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(X10,X19),X16),s(fun(X10,X11),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X10,X11),X12),s(fun(X10,X11),X13))),s(fun(X19,X11),X18))) | s('type.sum.sum'(X10,X19),X16) != s('type.sum.sum'(X10,X19),X17)) )),
  inference(superposition,[],[f230,f194])).
fof(f3665,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X9),X6),s(fun(X0,X1),X2),s(fun(X9,X1),X8))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X9),X7),s(fun(X0,X1),X5),s(fun(X9,X1),X4))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,sK17(X4,X5,X6,X7,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X8,X1,X9,X0)))) != s(X1,happ(s(fun(X0,X1),X5),s(X0,sK17(X4,X5,X6,X7,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X8,X1,X9,X0)))) | s('type.sum.sum'(X0,X9),X6) = s('type.sum.sum'(X0,X9),'const.sum.INR_1'(s(X9,sK18(X4,X5,X6,X7,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X8,X1,X9,X0)))) | s('type.sum.sum'(X0,X9),X6) != s('type.sum.sum'(X0,X9),X7)) )),
  inference(forward_demodulation,[],[f3643,f193])).
fof(f3643,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (s(X1,happ(s(fun(X0,X1),X2),s(X0,sK17(X4,X5,X6,X7,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X8,X1,X9,X0)))) != s(X1,happ(s(fun(X0,X1),X5),s(X0,sK17(X4,X5,X6,X7,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X8,X1,X9,X0)))) | s('type.sum.sum'(X0,X9),X6) = s('type.sum.sum'(X0,X9),'const.sum.INR_1'(s(X9,sK18(X4,X5,X6,X7,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3)),X8,X1,X9,X0)))) | s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X9),X7),s(fun(X0,X1),X5),s(fun(X9,X1),X4))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X9),X6),s(fun(X0,X1),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,X1),X2),s(fun(X0,X1),X3))),s(fun(X9,X1),X8))) | s('type.sum.sum'(X0,X9),X6) != s('type.sum.sum'(X0,X9),X7)) )),
  inference(superposition,[],[f230,f193])).
fof(f3480,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK18(X3,X4,X2,X2,X5,X6,X7,X1,X0)))) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK17(X3,X4,X2,X2,X5,X6,X7,X1,X0)))) | s(X7,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X2),s(fun(X0,X7),X4),s(fun(X1,X7),X3))) = s(X7,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X2),s(fun(X0,X7),X5),s(fun(X1,X7),X6)))) )),
  inference(equality_resolution,[],[f229])).
fof(f3508,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X86,X84] : (s(X95,'const.sum.sum_CASE_3'(s('type.sum.sum'(X83,X84),X90),s(fun(X83,X95),X92),s(fun(X84,X95),X91))) = s(X95,'const.sum.sum_CASE_3'(s('type.sum.sum'(X83,X84),happ(s(fun(X86,'type.sum.sum'(X83,X84)),X87),s(X86,X89))),s(fun(X83,X95),X93),s(fun(X84,X95),X94))) | s('type.sum.sum'(X83,X84),happ(s(fun(X86,'type.sum.sum'(X83,X84)),X87),s(X86,X89))) = s('type.sum.sum'(X83,X84),'const.sum.INL_1'(s(X83,sK17(X91,X92,happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89)))),X90,X93,X94,X95,X84,X83)))) | s('type.sum.sum'(X83,X84),happ(s(fun(X86,'type.sum.sum'(X83,X84)),X87),s(X86,X89))) = s('type.sum.sum'(X83,X84),'const.sum.INR_1'(s(X84,sK18(X91,X92,happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89)))),X90,X93,X94,X95,X84,X83)))) | s('type.sum.sum'(X83,X84),X90) != s('type.sum.sum'(X83,X84),happ(s(fun(X86,'type.sum.sum'(X83,X84)),X87),s(X86,X89)))) )),
  inference(forward_demodulation,[],[f3507,f226])).
fof(f3507,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X86,X84] : (s('type.sum.sum'(X83,X84),happ(s(fun(X86,'type.sum.sum'(X83,X84)),X87),s(X86,X89))) = s('type.sum.sum'(X83,X84),'const.sum.INL_1'(s(X83,sK17(X91,X92,happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89)))),X90,X93,X94,X95,X84,X83)))) | s('type.sum.sum'(X83,X84),happ(s(fun(X86,'type.sum.sum'(X83,X84)),X87),s(X86,X89))) = s('type.sum.sum'(X83,X84),'const.sum.INR_1'(s(X84,sK18(X91,X92,happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89)))),X90,X93,X94,X95,X84,X83)))) | s('type.sum.sum'(X83,X84),X90) != s('type.sum.sum'(X83,X84),happ(s(fun(X86,'type.sum.sum'(X83,X84)),X87),s(X86,X89))) | s(X95,'const.sum.sum_CASE_3'(s('type.sum.sum'(X83,X84),X90),s(fun(X83,X95),X92),s(fun(X84,X95),X91))) = s(X95,'const.sum.sum_CASE_3'(s('type.sum.sum'(X83,X84),happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89))))),s(fun(X83,X95),X93),s(fun(X84,X95),X94)))) )),
  inference(forward_demodulation,[],[f3506,f226])).
fof(f3506,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X86,X84] : (s('type.sum.sum'(X83,X84),happ(s(fun(X86,'type.sum.sum'(X83,X84)),X87),s(X86,X89))) = s('type.sum.sum'(X83,X84),'const.sum.INR_1'(s(X84,sK18(X91,X92,happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89)))),X90,X93,X94,X95,X84,X83)))) | s('type.sum.sum'(X83,X84),X90) != s('type.sum.sum'(X83,X84),happ(s(fun(X86,'type.sum.sum'(X83,X84)),X87),s(X86,X89))) | s('type.sum.sum'(X83,X84),happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89))))) = s('type.sum.sum'(X83,X84),'const.sum.INL_1'(s(X83,sK17(X91,X92,happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89)))),X90,X93,X94,X95,X84,X83)))) | s(X95,'const.sum.sum_CASE_3'(s('type.sum.sum'(X83,X84),X90),s(fun(X83,X95),X92),s(fun(X84,X95),X91))) = s(X95,'const.sum.sum_CASE_3'(s('type.sum.sum'(X83,X84),happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89))))),s(fun(X83,X95),X93),s(fun(X84,X95),X94)))) )),
  inference(forward_demodulation,[],[f3479,f226])).
fof(f3479,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X86,X84] : (s('type.sum.sum'(X83,X84),X90) != s('type.sum.sum'(X83,X84),happ(s(fun(X86,'type.sum.sum'(X83,X84)),X87),s(X86,X89))) | s('type.sum.sum'(X83,X84),happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89))))) = s('type.sum.sum'(X83,X84),'const.sum.INR_1'(s(X84,sK18(X91,X92,happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89)))),X90,X93,X94,X95,X84,X83)))) | s('type.sum.sum'(X83,X84),happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89))))) = s('type.sum.sum'(X83,X84),'const.sum.INL_1'(s(X83,sK17(X91,X92,happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89)))),X90,X93,X94,X95,X84,X83)))) | s(X95,'const.sum.sum_CASE_3'(s('type.sum.sum'(X83,X84),X90),s(fun(X83,X95),X92),s(fun(X84,X95),X91))) = s(X95,'const.sum.sum_CASE_3'(s('type.sum.sum'(X83,X84),happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89))))),s(fun(X83,X95),X93),s(fun(X84,X95),X94)))) )),
  inference(superposition,[],[f229,f226])).
fof(f3505,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X79,X77,X75,X73,X82] : (s(X82,'const.sum.sum_CASE_3'(s('type.sum.sum'(X70,X71),X77),s(fun(X70,X82),X79),s(fun(X71,X82),X78))) = s(X82,'const.sum.sum_CASE_3'(s('type.sum.sum'(X70,X71),happ(s(fun(X72,'type.sum.sum'(X70,X71)),X75),s(X72,X76))),s(fun(X70,X82),X80),s(fun(X71,X82),X81))) | s('type.sum.sum'(X70,X71),happ(s(fun(X72,'type.sum.sum'(X70,X71)),X75),s(X72,X76))) = s('type.sum.sum'(X70,X71),'const.sum.INL_1'(s(X70,sK17(X78,X79,happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76)))),X77,X80,X81,X82,X71,X70)))) | s('type.sum.sum'(X70,X71),happ(s(fun(X72,'type.sum.sum'(X70,X71)),X75),s(X72,X76))) = s('type.sum.sum'(X70,X71),'const.sum.INR_1'(s(X71,sK18(X78,X79,happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76)))),X77,X80,X81,X82,X71,X70)))) | s('type.sum.sum'(X70,X71),X77) != s('type.sum.sum'(X70,X71),happ(s(fun(X72,'type.sum.sum'(X70,X71)),X75),s(X72,X76)))) )),
  inference(forward_demodulation,[],[f3504,f225])).
fof(f3504,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X79,X77,X75,X73,X82] : (s('type.sum.sum'(X70,X71),happ(s(fun(X72,'type.sum.sum'(X70,X71)),X75),s(X72,X76))) = s('type.sum.sum'(X70,X71),'const.sum.INL_1'(s(X70,sK17(X78,X79,happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76)))),X77,X80,X81,X82,X71,X70)))) | s('type.sum.sum'(X70,X71),happ(s(fun(X72,'type.sum.sum'(X70,X71)),X75),s(X72,X76))) = s('type.sum.sum'(X70,X71),'const.sum.INR_1'(s(X71,sK18(X78,X79,happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76)))),X77,X80,X81,X82,X71,X70)))) | s('type.sum.sum'(X70,X71),X77) != s('type.sum.sum'(X70,X71),happ(s(fun(X72,'type.sum.sum'(X70,X71)),X75),s(X72,X76))) | s(X82,'const.sum.sum_CASE_3'(s('type.sum.sum'(X70,X71),X77),s(fun(X70,X82),X79),s(fun(X71,X82),X78))) = s(X82,'const.sum.sum_CASE_3'(s('type.sum.sum'(X70,X71),happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76))))),s(fun(X70,X82),X80),s(fun(X71,X82),X81)))) )),
  inference(forward_demodulation,[],[f3503,f225])).
fof(f3503,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X79,X77,X75,X73,X82] : (s('type.sum.sum'(X70,X71),happ(s(fun(X72,'type.sum.sum'(X70,X71)),X75),s(X72,X76))) = s('type.sum.sum'(X70,X71),'const.sum.INR_1'(s(X71,sK18(X78,X79,happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76)))),X77,X80,X81,X82,X71,X70)))) | s('type.sum.sum'(X70,X71),X77) != s('type.sum.sum'(X70,X71),happ(s(fun(X72,'type.sum.sum'(X70,X71)),X75),s(X72,X76))) | s('type.sum.sum'(X70,X71),happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76))))) = s('type.sum.sum'(X70,X71),'const.sum.INL_1'(s(X70,sK17(X78,X79,happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76)))),X77,X80,X81,X82,X71,X70)))) | s(X82,'const.sum.sum_CASE_3'(s('type.sum.sum'(X70,X71),X77),s(fun(X70,X82),X79),s(fun(X71,X82),X78))) = s(X82,'const.sum.sum_CASE_3'(s('type.sum.sum'(X70,X71),happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76))))),s(fun(X70,X82),X80),s(fun(X71,X82),X81)))) )),
  inference(forward_demodulation,[],[f3478,f225])).
fof(f3478,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X79,X77,X75,X73,X82] : (s('type.sum.sum'(X70,X71),X77) != s('type.sum.sum'(X70,X71),happ(s(fun(X72,'type.sum.sum'(X70,X71)),X75),s(X72,X76))) | s('type.sum.sum'(X70,X71),happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76))))) = s('type.sum.sum'(X70,X71),'const.sum.INR_1'(s(X71,sK18(X78,X79,happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76)))),X77,X80,X81,X82,X71,X70)))) | s('type.sum.sum'(X70,X71),happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76))))) = s('type.sum.sum'(X70,X71),'const.sum.INL_1'(s(X70,sK17(X78,X79,happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76)))),X77,X80,X81,X82,X71,X70)))) | s(X82,'const.sum.sum_CASE_3'(s('type.sum.sum'(X70,X71),X77),s(fun(X70,X82),X79),s(fun(X71,X82),X78))) = s(X82,'const.sum.sum_CASE_3'(s('type.sum.sum'(X70,X71),happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76))))),s(fun(X70,X82),X80),s(fun(X71,X82),X81)))) )),
  inference(superposition,[],[f229,f225])).
fof(f3502,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X69,X67,X65,X63] : (s(X69,'const.sum.sum_CASE_3'(s('type.sum.sum'(X57,X58),X64),s(fun(X57,X69),X66),s(fun(X58,X69),X65))) = s(X69,'const.sum.sum_CASE_3'(s('type.sum.sum'(X57,X58),happ(s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(X59,X61))),s(fun(X57,X69),X67),s(fun(X58,X69),X68))) | s('type.sum.sum'(X57,X58),happ(s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(X59,X61))) = s('type.sum.sum'(X57,X58),'const.sum.INL_1'(s(X57,sK17(X65,X66,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63)),X64,X67,X68,X69,X58,X57)))) | s('type.sum.sum'(X57,X58),happ(s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(X59,X61))) = s('type.sum.sum'(X57,X58),'const.sum.INR_1'(s(X58,sK18(X65,X66,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63)),X64,X67,X68,X69,X58,X57)))) | s('type.sum.sum'(X57,X58),X64) != s('type.sum.sum'(X57,X58),happ(s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(X59,X61)))) )),
  inference(forward_demodulation,[],[f3501,f228])).
fof(f3501,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X69,X67,X65,X63] : (s('type.sum.sum'(X57,X58),happ(s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(X59,X61))) = s('type.sum.sum'(X57,X58),'const.sum.INL_1'(s(X57,sK17(X65,X66,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63)),X64,X67,X68,X69,X58,X57)))) | s('type.sum.sum'(X57,X58),happ(s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(X59,X61))) = s('type.sum.sum'(X57,X58),'const.sum.INR_1'(s(X58,sK18(X65,X66,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63)),X64,X67,X68,X69,X58,X57)))) | s('type.sum.sum'(X57,X58),X64) != s('type.sum.sum'(X57,X58),happ(s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(X59,X61))) | s(X69,'const.sum.sum_CASE_3'(s('type.sum.sum'(X57,X58),X64),s(fun(X57,X69),X66),s(fun(X58,X69),X65))) = s(X69,'const.sum.sum_CASE_3'(s('type.sum.sum'(X57,X58),'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63))),s(fun(X57,X69),X67),s(fun(X58,X69),X68)))) )),
  inference(forward_demodulation,[],[f3500,f228])).
fof(f3500,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X69,X67,X65,X63] : (s('type.sum.sum'(X57,X58),happ(s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(X59,X61))) = s('type.sum.sum'(X57,X58),'const.sum.INR_1'(s(X58,sK18(X65,X66,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63)),X64,X67,X68,X69,X58,X57)))) | s('type.sum.sum'(X57,X58),X64) != s('type.sum.sum'(X57,X58),happ(s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(X59,X61))) | s('type.sum.sum'(X57,X58),'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63))) = s('type.sum.sum'(X57,X58),'const.sum.INL_1'(s(X57,sK17(X65,X66,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63)),X64,X67,X68,X69,X58,X57)))) | s(X69,'const.sum.sum_CASE_3'(s('type.sum.sum'(X57,X58),X64),s(fun(X57,X69),X66),s(fun(X58,X69),X65))) = s(X69,'const.sum.sum_CASE_3'(s('type.sum.sum'(X57,X58),'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63))),s(fun(X57,X69),X67),s(fun(X58,X69),X68)))) )),
  inference(forward_demodulation,[],[f3477,f228])).
fof(f3477,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X69,X67,X65,X63] : (s('type.sum.sum'(X57,X58),X64) != s('type.sum.sum'(X57,X58),happ(s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(X59,X61))) | s('type.sum.sum'(X57,X58),'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63))) = s('type.sum.sum'(X57,X58),'const.sum.INR_1'(s(X58,sK18(X65,X66,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63)),X64,X67,X68,X69,X58,X57)))) | s('type.sum.sum'(X57,X58),'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63))) = s('type.sum.sum'(X57,X58),'const.sum.INL_1'(s(X57,sK17(X65,X66,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63)),X64,X67,X68,X69,X58,X57)))) | s(X69,'const.sum.sum_CASE_3'(s('type.sum.sum'(X57,X58),X64),s(fun(X57,X69),X66),s(fun(X58,X69),X65))) = s(X69,'const.sum.sum_CASE_3'(s('type.sum.sum'(X57,X58),'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63))),s(fun(X57,X69),X67),s(fun(X58,X69),X68)))) )),
  inference(superposition,[],[f229,f228])).
fof(f3499,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X56,X46,X44,X55,X53,X51,X49] : (s(X56,'const.sum.sum_CASE_3'(s('type.sum.sum'(X44,X45),X51),s(fun(X44,X56),X53),s(fun(X45,X56),X52))) = s(X56,'const.sum.sum_CASE_3'(s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48))),s(fun(X44,X56),X54),s(fun(X45,X56),X55))) | s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48))) = s('type.sum.sum'(X44,X45),'const.sum.INL_1'(s(X44,sK17(X52,X53,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X51,X54,X55,X56,X45,X44)))) | s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48))) = s('type.sum.sum'(X44,X45),'const.sum.INR_1'(s(X45,sK18(X52,X53,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X51,X54,X55,X56,X45,X44)))) | s('type.sum.sum'(X44,X45),X51) != s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48)))) )),
  inference(forward_demodulation,[],[f3498,f227])).
fof(f3498,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X56,X46,X44,X55,X53,X51,X49] : (s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48))) = s('type.sum.sum'(X44,X45),'const.sum.INL_1'(s(X44,sK17(X52,X53,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X51,X54,X55,X56,X45,X44)))) | s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48))) = s('type.sum.sum'(X44,X45),'const.sum.INR_1'(s(X45,sK18(X52,X53,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X51,X54,X55,X56,X45,X44)))) | s('type.sum.sum'(X44,X45),X51) != s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48))) | s(X56,'const.sum.sum_CASE_3'(s('type.sum.sum'(X44,X45),X51),s(fun(X44,X56),X53),s(fun(X45,X56),X52))) = s(X56,'const.sum.sum_CASE_3'(s('type.sum.sum'(X44,X45),'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50))),s(fun(X44,X56),X54),s(fun(X45,X56),X55)))) )),
  inference(forward_demodulation,[],[f3497,f227])).
fof(f3497,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X56,X46,X44,X55,X53,X51,X49] : (s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48))) = s('type.sum.sum'(X44,X45),'const.sum.INR_1'(s(X45,sK18(X52,X53,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X51,X54,X55,X56,X45,X44)))) | s('type.sum.sum'(X44,X45),X51) != s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48))) | s('type.sum.sum'(X44,X45),'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50))) = s('type.sum.sum'(X44,X45),'const.sum.INL_1'(s(X44,sK17(X52,X53,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X51,X54,X55,X56,X45,X44)))) | s(X56,'const.sum.sum_CASE_3'(s('type.sum.sum'(X44,X45),X51),s(fun(X44,X56),X53),s(fun(X45,X56),X52))) = s(X56,'const.sum.sum_CASE_3'(s('type.sum.sum'(X44,X45),'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50))),s(fun(X44,X56),X54),s(fun(X45,X56),X55)))) )),
  inference(forward_demodulation,[],[f3476,f227])).
fof(f3476,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X56,X46,X44,X55,X53,X51,X49] : (s('type.sum.sum'(X44,X45),X51) != s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48))) | s('type.sum.sum'(X44,X45),'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50))) = s('type.sum.sum'(X44,X45),'const.sum.INR_1'(s(X45,sK18(X52,X53,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X51,X54,X55,X56,X45,X44)))) | s('type.sum.sum'(X44,X45),'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50))) = s('type.sum.sum'(X44,X45),'const.sum.INL_1'(s(X44,sK17(X52,X53,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X51,X54,X55,X56,X45,X44)))) | s(X56,'const.sum.sum_CASE_3'(s('type.sum.sum'(X44,X45),X51),s(fun(X44,X56),X53),s(fun(X45,X56),X52))) = s(X56,'const.sum.sum_CASE_3'(s('type.sum.sum'(X44,X45),'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50))),s(fun(X44,X56),X54),s(fun(X45,X56),X55)))) )),
  inference(superposition,[],[f229,f227])).
fof(f3496,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (s(X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X34,X35),X37),s(fun(X34,X43),X41),s(fun(X35,X43),X42))) = s(X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X34,X35),X38),s(fun(X34,X43),X40),s(fun(X35,X43),X39))) | s('type.sum.sum'(X34,X35),X37) = s('type.sum.sum'(X34,X35),'const.sum.INL_1'(s(X34,sK17(X39,X40,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37)),X38,X41,X42,X43,X35,X34)))) | s('type.sum.sum'(X34,X35),X37) = s('type.sum.sum'(X34,X35),'const.sum.INR_1'(s(X35,sK18(X39,X40,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37)),X38,X41,X42,X43,X35,X34)))) | s('type.sum.sum'(X34,X35),X37) != s('type.sum.sum'(X34,X35),X38)) )),
  inference(forward_demodulation,[],[f3495,f194])).
fof(f3495,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (s('type.sum.sum'(X34,X35),X37) = s('type.sum.sum'(X34,X35),'const.sum.INL_1'(s(X34,sK17(X39,X40,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37)),X38,X41,X42,X43,X35,X34)))) | s('type.sum.sum'(X34,X35),X37) = s('type.sum.sum'(X34,X35),'const.sum.INR_1'(s(X35,sK18(X39,X40,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37)),X38,X41,X42,X43,X35,X34)))) | s('type.sum.sum'(X34,X35),X37) != s('type.sum.sum'(X34,X35),X38) | s(X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X34,X35),X38),s(fun(X34,X43),X40),s(fun(X35,X43),X39))) = s(X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X34,X35),'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37))),s(fun(X34,X43),X41),s(fun(X35,X43),X42)))) )),
  inference(forward_demodulation,[],[f3494,f194])).
fof(f3494,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (s('type.sum.sum'(X34,X35),X37) = s('type.sum.sum'(X34,X35),'const.sum.INR_1'(s(X35,sK18(X39,X40,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37)),X38,X41,X42,X43,X35,X34)))) | s('type.sum.sum'(X34,X35),X37) != s('type.sum.sum'(X34,X35),X38) | s('type.sum.sum'(X34,X35),'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37))) = s('type.sum.sum'(X34,X35),'const.sum.INL_1'(s(X34,sK17(X39,X40,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37)),X38,X41,X42,X43,X35,X34)))) | s(X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X34,X35),X38),s(fun(X34,X43),X40),s(fun(X35,X43),X39))) = s(X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X34,X35),'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37))),s(fun(X34,X43),X41),s(fun(X35,X43),X42)))) )),
  inference(forward_demodulation,[],[f3475,f194])).
fof(f3475,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (s('type.sum.sum'(X34,X35),X37) != s('type.sum.sum'(X34,X35),X38) | s('type.sum.sum'(X34,X35),'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37))) = s('type.sum.sum'(X34,X35),'const.sum.INR_1'(s(X35,sK18(X39,X40,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37)),X38,X41,X42,X43,X35,X34)))) | s('type.sum.sum'(X34,X35),'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37))) = s('type.sum.sum'(X34,X35),'const.sum.INL_1'(s(X34,sK17(X39,X40,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37)),X38,X41,X42,X43,X35,X34)))) | s(X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X34,X35),X38),s(fun(X34,X43),X40),s(fun(X35,X43),X39))) = s(X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X34,X35),'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37))),s(fun(X34,X43),X41),s(fun(X35,X43),X42)))) )),
  inference(superposition,[],[f229,f194])).
fof(f3493,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X24,X25),X26),s(fun(X24,X33),X31),s(fun(X25,X33),X32))) = s(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X24,X25),X28),s(fun(X24,X33),X30),s(fun(X25,X33),X29))) | s('type.sum.sum'(X24,X25),X26) = s('type.sum.sum'(X24,X25),'const.sum.INL_1'(s(X24,sK17(X29,X30,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27)),X28,X31,X32,X33,X25,X24)))) | s('type.sum.sum'(X24,X25),X26) = s('type.sum.sum'(X24,X25),'const.sum.INR_1'(s(X25,sK18(X29,X30,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27)),X28,X31,X32,X33,X25,X24)))) | s('type.sum.sum'(X24,X25),X26) != s('type.sum.sum'(X24,X25),X28)) )),
  inference(forward_demodulation,[],[f3492,f193])).
fof(f3492,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s('type.sum.sum'(X24,X25),X26) = s('type.sum.sum'(X24,X25),'const.sum.INL_1'(s(X24,sK17(X29,X30,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27)),X28,X31,X32,X33,X25,X24)))) | s('type.sum.sum'(X24,X25),X26) = s('type.sum.sum'(X24,X25),'const.sum.INR_1'(s(X25,sK18(X29,X30,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27)),X28,X31,X32,X33,X25,X24)))) | s('type.sum.sum'(X24,X25),X26) != s('type.sum.sum'(X24,X25),X28) | s(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X24,X25),X28),s(fun(X24,X33),X30),s(fun(X25,X33),X29))) = s(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X24,X25),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27))),s(fun(X24,X33),X31),s(fun(X25,X33),X32)))) )),
  inference(forward_demodulation,[],[f3491,f193])).
fof(f3491,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s('type.sum.sum'(X24,X25),X26) = s('type.sum.sum'(X24,X25),'const.sum.INR_1'(s(X25,sK18(X29,X30,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27)),X28,X31,X32,X33,X25,X24)))) | s('type.sum.sum'(X24,X25),X26) != s('type.sum.sum'(X24,X25),X28) | s('type.sum.sum'(X24,X25),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27))) = s('type.sum.sum'(X24,X25),'const.sum.INL_1'(s(X24,sK17(X29,X30,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27)),X28,X31,X32,X33,X25,X24)))) | s(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X24,X25),X28),s(fun(X24,X33),X30),s(fun(X25,X33),X29))) = s(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X24,X25),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27))),s(fun(X24,X33),X31),s(fun(X25,X33),X32)))) )),
  inference(forward_demodulation,[],[f3474,f193])).
fof(f3474,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s('type.sum.sum'(X24,X25),X26) != s('type.sum.sum'(X24,X25),X28) | s('type.sum.sum'(X24,X25),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27))) = s('type.sum.sum'(X24,X25),'const.sum.INR_1'(s(X25,sK18(X29,X30,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27)),X28,X31,X32,X33,X25,X24)))) | s('type.sum.sum'(X24,X25),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27))) = s('type.sum.sum'(X24,X25),'const.sum.INL_1'(s(X24,sK17(X29,X30,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27)),X28,X31,X32,X33,X25,X24)))) | s(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X24,X25),X28),s(fun(X24,X33),X30),s(fun(X25,X33),X29))) = s(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X24,X25),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27))),s(fun(X24,X33),X31),s(fun(X25,X33),X32)))) )),
  inference(superposition,[],[f229,f193])).
fof(f3473,plain,(
  ( ! [X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s('type.sum.sum'(X15,X16),X17) != s('type.sum.sum'(X15,X16),X18) | s('type.sum.sum'(X15,X16),'const.sum.INL_1'(s(X15,sK9(X17,X16,X15)))) = s('type.sum.sum'(X15,X16),'const.sum.INR_1'(s(X16,sK18(X19,X20,'const.sum.INL_1'(s(X15,sK9(X17,X16,X15))),X18,X21,X22,X23,X16,X15)))) | s('type.sum.sum'(X15,X16),'const.sum.INL_1'(s(X15,sK9(X17,X16,X15)))) = s('type.sum.sum'(X15,X16),'const.sum.INL_1'(s(X15,sK17(X19,X20,'const.sum.INL_1'(s(X15,sK9(X17,X16,X15))),X18,X21,X22,X23,X16,X15)))) | s(X23,'const.sum.sum_CASE_3'(s('type.sum.sum'(X15,X16),X18),s(fun(X15,X23),X20),s(fun(X16,X23),X19))) = s(X23,'const.sum.sum_CASE_3'(s('type.sum.sum'(X15,X16),'const.sum.INL_1'(s(X15,sK9(X17,X16,X15)))),s(fun(X15,X23),X21),s(fun(X16,X23),X22))) | s('type.sum.sum'(X15,X16),X17) = s('type.sum.sum'(X15,X16),'const.sum.INR_1'(s(X16,sK10(X17,X16,X15))))) )),
  inference(superposition,[],[f229,f205])).
fof(f3472,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X13,X11,X9] : (s('type.sum.sum'(X6,X7),X8) != s('type.sum.sum'(X6,X7),X9) | s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,sK10(X8,X7,X6)))) = s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,sK18(X10,X11,'const.sum.INR_1'(s(X7,sK10(X8,X7,X6))),X9,X12,X13,X14,X7,X6)))) | s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,sK10(X8,X7,X6)))) = s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,sK17(X10,X11,'const.sum.INR_1'(s(X7,sK10(X8,X7,X6))),X9,X12,X13,X14,X7,X6)))) | s(X14,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X7),X9),s(fun(X6,X14),X11),s(fun(X7,X14),X10))) = s(X14,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,sK10(X8,X7,X6)))),s(fun(X6,X14),X12),s(fun(X7,X14),X13))) | s('type.sum.sum'(X6,X7),X8) = s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,sK9(X8,X7,X6))))) )),
  inference(superposition,[],[f229,f205])).
fof(f3487,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X86,X84] : (s(X95,'const.sum.sum_CASE_3'(s('type.sum.sum'(X83,X84),X90),s(fun(X83,X95),X93),s(fun(X84,X95),X94))) = s(X95,'const.sum.sum_CASE_3'(s('type.sum.sum'(X83,X84),happ(s(fun(X86,'type.sum.sum'(X83,X84)),X87),s(X86,X89))),s(fun(X83,X95),X92),s(fun(X84,X95),X91))) | s('type.sum.sum'(X83,X84),X90) != s('type.sum.sum'(X83,X84),happ(s(fun(X86,'type.sum.sum'(X83,X84)),X87),s(X86,X89))) | s('type.sum.sum'(X83,X84),X90) = s('type.sum.sum'(X83,X84),'const.sum.INR_1'(s(X84,sK18(X91,X92,X90,happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89)))),X93,X94,X95,X84,X83)))) | s('type.sum.sum'(X83,X84),X90) = s('type.sum.sum'(X83,X84),'const.sum.INL_1'(s(X83,sK17(X91,X92,X90,happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89)))),X93,X94,X95,X84,X83))))) )),
  inference(forward_demodulation,[],[f3470,f226])).
fof(f3470,plain,(
  ( ! [X94,X92,X90,X88,X87,X85,X83,X95,X93,X91,X89,X86,X84] : (s('type.sum.sum'(X83,X84),X90) != s('type.sum.sum'(X83,X84),happ(s(fun(X86,'type.sum.sum'(X83,X84)),X87),s(X86,X89))) | s('type.sum.sum'(X83,X84),X90) = s('type.sum.sum'(X83,X84),'const.sum.INR_1'(s(X84,sK18(X91,X92,X90,happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89)))),X93,X94,X95,X84,X83)))) | s('type.sum.sum'(X83,X84),X90) = s('type.sum.sum'(X83,X84),'const.sum.INL_1'(s(X83,sK17(X91,X92,X90,happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89)))),X93,X94,X95,X84,X83)))) | s(X95,'const.sum.sum_CASE_3'(s('type.sum.sum'(X83,X84),X90),s(fun(X83,X95),X93),s(fun(X84,X95),X94))) = s(X95,'const.sum.sum_CASE_3'(s('type.sum.sum'(X83,X84),happ(s(fun('type.sum.sum'(X85,X86),'type.sum.sum'(X83,X84)),sK16(X87,X88,X86,'type.sum.sum'(X83,X84),X85)),s('type.sum.sum'(X85,X86),'const.sum.INR_1'(s(X86,X89))))),s(fun(X83,X95),X92),s(fun(X84,X95),X91)))) )),
  inference(superposition,[],[f229,f226])).
fof(f3486,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X79,X77,X75,X73,X82] : (s(X82,'const.sum.sum_CASE_3'(s('type.sum.sum'(X70,X71),X77),s(fun(X70,X82),X80),s(fun(X71,X82),X81))) = s(X82,'const.sum.sum_CASE_3'(s('type.sum.sum'(X70,X71),happ(s(fun(X72,'type.sum.sum'(X70,X71)),X75),s(X72,X76))),s(fun(X70,X82),X79),s(fun(X71,X82),X78))) | s('type.sum.sum'(X70,X71),X77) != s('type.sum.sum'(X70,X71),happ(s(fun(X72,'type.sum.sum'(X70,X71)),X75),s(X72,X76))) | s('type.sum.sum'(X70,X71),X77) = s('type.sum.sum'(X70,X71),'const.sum.INR_1'(s(X71,sK18(X78,X79,X77,happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76)))),X80,X81,X82,X71,X70)))) | s('type.sum.sum'(X70,X71),X77) = s('type.sum.sum'(X70,X71),'const.sum.INL_1'(s(X70,sK17(X78,X79,X77,happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76)))),X80,X81,X82,X71,X70))))) )),
  inference(forward_demodulation,[],[f3469,f225])).
fof(f3469,plain,(
  ( ! [X80,X70,X78,X76,X74,X72,X81,X71,X79,X77,X75,X73,X82] : (s('type.sum.sum'(X70,X71),X77) != s('type.sum.sum'(X70,X71),happ(s(fun(X72,'type.sum.sum'(X70,X71)),X75),s(X72,X76))) | s('type.sum.sum'(X70,X71),X77) = s('type.sum.sum'(X70,X71),'const.sum.INR_1'(s(X71,sK18(X78,X79,X77,happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76)))),X80,X81,X82,X71,X70)))) | s('type.sum.sum'(X70,X71),X77) = s('type.sum.sum'(X70,X71),'const.sum.INL_1'(s(X70,sK17(X78,X79,X77,happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76)))),X80,X81,X82,X71,X70)))) | s(X82,'const.sum.sum_CASE_3'(s('type.sum.sum'(X70,X71),X77),s(fun(X70,X82),X80),s(fun(X71,X82),X81))) = s(X82,'const.sum.sum_CASE_3'(s('type.sum.sum'(X70,X71),happ(s(fun('type.sum.sum'(X72,X73),'type.sum.sum'(X70,X71)),sK16(X74,X75,X73,'type.sum.sum'(X70,X71),X72)),s('type.sum.sum'(X72,X73),'const.sum.INL_1'(s(X72,X76))))),s(fun(X70,X82),X79),s(fun(X71,X82),X78)))) )),
  inference(superposition,[],[f229,f225])).
fof(f3485,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X69,X67,X65,X63] : (s(X69,'const.sum.sum_CASE_3'(s('type.sum.sum'(X57,X58),X64),s(fun(X57,X69),X67),s(fun(X58,X69),X68))) = s(X69,'const.sum.sum_CASE_3'(s('type.sum.sum'(X57,X58),happ(s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(X59,X61))),s(fun(X57,X69),X66),s(fun(X58,X69),X65))) | s('type.sum.sum'(X57,X58),X64) != s('type.sum.sum'(X57,X58),happ(s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(X59,X61))) | s('type.sum.sum'(X57,X58),X64) = s('type.sum.sum'(X57,X58),'const.sum.INR_1'(s(X58,sK18(X65,X66,X64,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63)),X67,X68,X69,X58,X57)))) | s('type.sum.sum'(X57,X58),X64) = s('type.sum.sum'(X57,X58),'const.sum.INL_1'(s(X57,sK17(X65,X66,X64,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63)),X67,X68,X69,X58,X57))))) )),
  inference(forward_demodulation,[],[f3468,f228])).
fof(f3468,plain,(
  ( ! [X61,X59,X57,X68,X66,X64,X62,X60,X58,X69,X67,X65,X63] : (s('type.sum.sum'(X57,X58),X64) != s('type.sum.sum'(X57,X58),happ(s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(X59,X61))) | s('type.sum.sum'(X57,X58),X64) = s('type.sum.sum'(X57,X58),'const.sum.INR_1'(s(X58,sK18(X65,X66,X64,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63)),X67,X68,X69,X58,X57)))) | s('type.sum.sum'(X57,X58),X64) = s('type.sum.sum'(X57,X58),'const.sum.INL_1'(s(X57,sK17(X65,X66,X64,'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63)),X67,X68,X69,X58,X57)))) | s(X69,'const.sum.sum_CASE_3'(s('type.sum.sum'(X57,X58),X64),s(fun(X57,X69),X67),s(fun(X58,X69),X68))) = s(X69,'const.sum.sum_CASE_3'(s('type.sum.sum'(X57,X58),'const.sum.sum_CASE_3'(s('type.sum.sum'(X59,X60),'const.sum.INL_1'(s(X59,X61))),s(fun(X59,'type.sum.sum'(X57,X58)),X62),s(fun(X60,'type.sum.sum'(X57,X58)),X63))),s(fun(X57,X69),X66),s(fun(X58,X69),X65)))) )),
  inference(superposition,[],[f229,f228])).
fof(f3484,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X56,X46,X44,X55,X53,X51,X49] : (s(X56,'const.sum.sum_CASE_3'(s('type.sum.sum'(X44,X45),X51),s(fun(X44,X56),X54),s(fun(X45,X56),X55))) = s(X56,'const.sum.sum_CASE_3'(s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48))),s(fun(X44,X56),X53),s(fun(X45,X56),X52))) | s('type.sum.sum'(X44,X45),X51) != s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48))) | s('type.sum.sum'(X44,X45),X51) = s('type.sum.sum'(X44,X45),'const.sum.INR_1'(s(X45,sK18(X52,X53,X51,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X54,X55,X56,X45,X44)))) | s('type.sum.sum'(X44,X45),X51) = s('type.sum.sum'(X44,X45),'const.sum.INL_1'(s(X44,sK17(X52,X53,X51,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X54,X55,X56,X45,X44))))) )),
  inference(forward_demodulation,[],[f3467,f227])).
fof(f3467,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X56,X46,X44,X55,X53,X51,X49] : (s('type.sum.sum'(X44,X45),X51) != s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48))) | s('type.sum.sum'(X44,X45),X51) = s('type.sum.sum'(X44,X45),'const.sum.INR_1'(s(X45,sK18(X52,X53,X51,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X54,X55,X56,X45,X44)))) | s('type.sum.sum'(X44,X45),X51) = s('type.sum.sum'(X44,X45),'const.sum.INL_1'(s(X44,sK17(X52,X53,X51,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X54,X55,X56,X45,X44)))) | s(X56,'const.sum.sum_CASE_3'(s('type.sum.sum'(X44,X45),X51),s(fun(X44,X56),X54),s(fun(X45,X56),X55))) = s(X56,'const.sum.sum_CASE_3'(s('type.sum.sum'(X44,X45),'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50))),s(fun(X44,X56),X53),s(fun(X45,X56),X52)))) )),
  inference(superposition,[],[f229,f227])).
fof(f3483,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (s(X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X34,X35),X37),s(fun(X34,X43),X40),s(fun(X35,X43),X39))) = s(X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X34,X35),X38),s(fun(X34,X43),X41),s(fun(X35,X43),X42))) | s('type.sum.sum'(X34,X35),X37) != s('type.sum.sum'(X34,X35),X38) | s('type.sum.sum'(X34,X35),X38) = s('type.sum.sum'(X34,X35),'const.sum.INR_1'(s(X35,sK18(X39,X40,X38,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37)),X41,X42,X43,X35,X34)))) | s('type.sum.sum'(X34,X35),X38) = s('type.sum.sum'(X34,X35),'const.sum.INL_1'(s(X34,sK17(X39,X40,X38,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37)),X41,X42,X43,X35,X34))))) )),
  inference(forward_demodulation,[],[f3466,f194])).
fof(f3466,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (s('type.sum.sum'(X34,X35),X37) != s('type.sum.sum'(X34,X35),X38) | s('type.sum.sum'(X34,X35),X38) = s('type.sum.sum'(X34,X35),'const.sum.INR_1'(s(X35,sK18(X39,X40,X38,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37)),X41,X42,X43,X35,X34)))) | s('type.sum.sum'(X34,X35),X38) = s('type.sum.sum'(X34,X35),'const.sum.INL_1'(s(X34,sK17(X39,X40,X38,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37)),X41,X42,X43,X35,X34)))) | s(X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X34,X35),X38),s(fun(X34,X43),X41),s(fun(X35,X43),X42))) = s(X43,'const.sum.sum_CASE_3'(s('type.sum.sum'(X34,X35),'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X34,X35),X36),s('type.sum.sum'(X34,X35),X37))),s(fun(X34,X43),X40),s(fun(X35,X43),X39)))) )),
  inference(superposition,[],[f229,f194])).
fof(f3482,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X24,X25),X26),s(fun(X24,X33),X30),s(fun(X25,X33),X29))) = s(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X24,X25),X28),s(fun(X24,X33),X31),s(fun(X25,X33),X32))) | s('type.sum.sum'(X24,X25),X26) != s('type.sum.sum'(X24,X25),X28) | s('type.sum.sum'(X24,X25),X28) = s('type.sum.sum'(X24,X25),'const.sum.INR_1'(s(X25,sK18(X29,X30,X28,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27)),X31,X32,X33,X25,X24)))) | s('type.sum.sum'(X24,X25),X28) = s('type.sum.sum'(X24,X25),'const.sum.INL_1'(s(X24,sK17(X29,X30,X28,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27)),X31,X32,X33,X25,X24))))) )),
  inference(forward_demodulation,[],[f3465,f193])).
fof(f3465,plain,(
  ( ! [X30,X28,X26,X24,X33,X31,X29,X27,X25,X32] : (s('type.sum.sum'(X24,X25),X26) != s('type.sum.sum'(X24,X25),X28) | s('type.sum.sum'(X24,X25),X28) = s('type.sum.sum'(X24,X25),'const.sum.INR_1'(s(X25,sK18(X29,X30,X28,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27)),X31,X32,X33,X25,X24)))) | s('type.sum.sum'(X24,X25),X28) = s('type.sum.sum'(X24,X25),'const.sum.INL_1'(s(X24,sK17(X29,X30,X28,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27)),X31,X32,X33,X25,X24)))) | s(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X24,X25),X28),s(fun(X24,X33),X31),s(fun(X25,X33),X32))) = s(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X24,X25),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X24,X25),X26),s('type.sum.sum'(X24,X25),X27))),s(fun(X24,X33),X30),s(fun(X25,X33),X29)))) )),
  inference(superposition,[],[f229,f193])).
fof(f3464,plain,(
  ( ! [X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s('type.sum.sum'(X15,X16),X17) != s('type.sum.sum'(X15,X16),X18) | s('type.sum.sum'(X15,X16),X18) = s('type.sum.sum'(X15,X16),'const.sum.INR_1'(s(X16,sK18(X19,X20,X18,'const.sum.INL_1'(s(X15,sK9(X17,X16,X15))),X21,X22,X23,X16,X15)))) | s('type.sum.sum'(X15,X16),X18) = s('type.sum.sum'(X15,X16),'const.sum.INL_1'(s(X15,sK17(X19,X20,X18,'const.sum.INL_1'(s(X15,sK9(X17,X16,X15))),X21,X22,X23,X16,X15)))) | s(X23,'const.sum.sum_CASE_3'(s('type.sum.sum'(X15,X16),X18),s(fun(X15,X23),X21),s(fun(X16,X23),X22))) = s(X23,'const.sum.sum_CASE_3'(s('type.sum.sum'(X15,X16),'const.sum.INL_1'(s(X15,sK9(X17,X16,X15)))),s(fun(X15,X23),X20),s(fun(X16,X23),X19))) | s('type.sum.sum'(X15,X16),X17) = s('type.sum.sum'(X15,X16),'const.sum.INR_1'(s(X16,sK10(X17,X16,X15))))) )),
  inference(superposition,[],[f229,f205])).
fof(f3463,plain,(
  ( ! [X6,X14,X12,X10,X8,X7,X13,X11,X9] : (s('type.sum.sum'(X6,X7),X8) != s('type.sum.sum'(X6,X7),X9) | s('type.sum.sum'(X6,X7),X9) = s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,sK18(X10,X11,X9,'const.sum.INR_1'(s(X7,sK10(X8,X7,X6))),X12,X13,X14,X7,X6)))) | s('type.sum.sum'(X6,X7),X9) = s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,sK17(X10,X11,X9,'const.sum.INR_1'(s(X7,sK10(X8,X7,X6))),X12,X13,X14,X7,X6)))) | s(X14,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X7),X9),s(fun(X6,X14),X12),s(fun(X7,X14),X13))) = s(X14,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,sK10(X8,X7,X6)))),s(fun(X6,X14),X11),s(fun(X7,X14),X10))) | s('type.sum.sum'(X6,X7),X8) = s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,sK9(X8,X7,X6))))) )),
  inference(superposition,[],[f229,f205])).
fof(f3392,plain,(
  ( ! [X59,X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X59,bool),X56),s(fun(X49,bool),happ(s(fun(X51,fun(X49,bool)),X52),s(X51,X54))),s('type.sum.sum'(X59,X49),X58)))) | ~p(s(bool,happ(s(fun(X49,bool),happ(s(fun(X51,fun(X49,bool)),X52),s(X51,X54))),s(X49,sK15(X55,X56,X57,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X58,X49,X59))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X59,bool),X55),s(fun(X49,bool),X57),s('type.sum.sum'(X59,X49),X58)))) | ~p(s(bool,happ(s(fun(X59,bool),X56),s(X59,sK14(X55,X56,X57,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X58,X49,X59)))))) )),
  inference(forward_demodulation,[],[f3380,f226])).
fof(f3380,plain,(
  ( ! [X59,X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun(X49,bool),happ(s(fun(X51,fun(X49,bool)),X52),s(X51,X54))),s(X49,sK15(X55,X56,X57,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X58,X49,X59))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X59,bool),X55),s(fun(X49,bool),X57),s('type.sum.sum'(X59,X49),X58)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X59,bool),X56),s(fun(X49,bool),happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54))))),s('type.sum.sum'(X59,X49),X58)))) | ~p(s(bool,happ(s(fun(X59,bool),X56),s(X59,sK14(X55,X56,X57,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X58,X49,X59)))))) )),
  inference(superposition,[],[f218,f226])).
fof(f3391,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X38,X46,X44,X42,X40] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X48,bool),X45),s(fun(X38,bool),happ(s(fun(X39,fun(X38,bool)),X42),s(X39,X43))),s('type.sum.sum'(X48,X38),X47)))) | ~p(s(bool,happ(s(fun(X38,bool),happ(s(fun(X39,fun(X38,bool)),X42),s(X39,X43))),s(X38,sK15(X44,X45,X46,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X47,X38,X48))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X48,bool),X44),s(fun(X38,bool),X46),s('type.sum.sum'(X48,X38),X47)))) | ~p(s(bool,happ(s(fun(X48,bool),X45),s(X48,sK14(X44,X45,X46,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X47,X38,X48)))))) )),
  inference(forward_demodulation,[],[f3379,f225])).
fof(f3379,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X38,X46,X44,X42,X40] : (~p(s(bool,happ(s(fun(X38,bool),happ(s(fun(X39,fun(X38,bool)),X42),s(X39,X43))),s(X38,sK15(X44,X45,X46,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X47,X38,X48))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X48,bool),X44),s(fun(X38,bool),X46),s('type.sum.sum'(X48,X38),X47)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X48,bool),X45),s(fun(X38,bool),happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43))))),s('type.sum.sum'(X48,X38),X47)))) | ~p(s(bool,happ(s(fun(X48,bool),X45),s(X48,sK14(X44,X45,X46,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X47,X38,X48)))))) )),
  inference(superposition,[],[f218,f225])).
fof(f3390,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X37,bool),X34),s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X31),s(X28,X30))),s('type.sum.sum'(X37,X27),X36)))) | ~p(s(bool,happ(s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X31),s(X28,X30))),s(X27,sK15(X33,X34,X35,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X36,X27,X37))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X37,bool),X33),s(fun(X27,bool),X35),s('type.sum.sum'(X37,X27),X36)))) | ~p(s(bool,happ(s(fun(X37,bool),X34),s(X37,sK14(X33,X34,X35,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X36,X27,X37)))))) )),
  inference(forward_demodulation,[],[f3378,f228])).
fof(f3378,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (~p(s(bool,happ(s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X31),s(X28,X30))),s(X27,sK15(X33,X34,X35,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X36,X27,X37))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X37,bool),X33),s(fun(X27,bool),X35),s('type.sum.sum'(X37,X27),X36)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X37,bool),X34),s(fun(X27,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32))),s('type.sum.sum'(X37,X27),X36)))) | ~p(s(bool,happ(s(fun(X37,bool),X34),s(X37,sK14(X33,X34,X35,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X36,X27,X37)))))) )),
  inference(superposition,[],[f218,f228])).
fof(f3389,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X26,bool),X23),s(fun(X16,bool),happ(s(fun(X18,fun(X16,bool)),X21),s(X18,X19))),s('type.sum.sum'(X26,X16),X25)))) | ~p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X18,fun(X16,bool)),X21),s(X18,X19))),s(X16,sK15(X22,X23,X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X25,X16,X26))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X26,bool),X22),s(fun(X16,bool),X24),s('type.sum.sum'(X26,X16),X25)))) | ~p(s(bool,happ(s(fun(X26,bool),X23),s(X26,sK14(X22,X23,X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X25,X16,X26)))))) )),
  inference(forward_demodulation,[],[f3377,f227])).
fof(f3377,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X18,fun(X16,bool)),X21),s(X18,X19))),s(X16,sK15(X22,X23,X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X25,X16,X26))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X26,bool),X22),s(fun(X16,bool),X24),s('type.sum.sum'(X26,X16),X25)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X26,bool),X23),s(fun(X16,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21))),s('type.sum.sum'(X26,X16),X25)))) | ~p(s(bool,happ(s(fun(X26,bool),X23),s(X26,sK14(X22,X23,X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X25,X16,X26)))))) )),
  inference(superposition,[],[f218,f227])).
fof(f3388,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X15,bool),X12),s(fun(X8,bool),X10),s('type.sum.sum'(X15,X8),X14)))) | ~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK15(X11,X12,X13,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X14,X8,X15))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X15,bool),X11),s(fun(X8,bool),X13),s('type.sum.sum'(X15,X8),X14)))) | ~p(s(bool,happ(s(fun(X15,bool),X12),s(X15,sK14(X11,X12,X13,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X14,X8,X15)))))) )),
  inference(forward_demodulation,[],[f3376,f194])).
fof(f3376,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK15(X11,X12,X13,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X14,X8,X15))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X15,bool),X11),s(fun(X8,bool),X13),s('type.sum.sum'(X15,X8),X14)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X15,bool),X12),s(fun(X8,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10))),s('type.sum.sum'(X15,X8),X14)))) | ~p(s(bool,happ(s(fun(X15,bool),X12),s(X15,sK14(X11,X12,X13,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X14,X8,X15)))))) )),
  inference(superposition,[],[f218,f194])).
fof(f3387,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X7,bool),X4),s(fun(X0,bool),X1),s('type.sum.sum'(X7,X0),X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X3,X4,X5,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X6,X0,X7))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X7,bool),X3),s(fun(X0,bool),X5),s('type.sum.sum'(X7,X0),X6)))) | ~p(s(bool,happ(s(fun(X7,bool),X4),s(X7,sK14(X3,X4,X5,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X6,X0,X7)))))) )),
  inference(forward_demodulation,[],[f3375,f193])).
fof(f3375,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X3,X4,X5,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X6,X0,X7))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X7,bool),X3),s(fun(X0,bool),X5),s('type.sum.sum'(X7,X0),X6)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X7,bool),X4),s(fun(X0,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s('type.sum.sum'(X7,X0),X6)))) | ~p(s(bool,happ(s(fun(X7,bool),X4),s(X7,sK14(X3,X4,X5,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X6,X0,X7)))))) )),
  inference(superposition,[],[f218,f193])).
fof(f3271,plain,(
  ( ! [X59,X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X59,bool),X56),s(fun(X49,bool),happ(s(fun(X51,fun(X49,bool)),X52),s(X51,X54))),s('type.sum.sum'(X59,X49),X58)))) | ~p(s(bool,happ(s(fun(X49,bool),happ(s(fun(X51,fun(X49,bool)),X52),s(X51,X54))),s(X49,sK15(X55,X56,X57,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X58,X49,X59))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X59,bool),X55),s(fun(X49,bool),X57),s('type.sum.sum'(X59,X49),X58)))) | p(s(bool,happ(s(fun(X59,bool),X55),s(X59,sK14(X55,X56,X57,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X58,X49,X59)))))) )),
  inference(forward_demodulation,[],[f3259,f226])).
fof(f3259,plain,(
  ( ! [X59,X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun(X49,bool),happ(s(fun(X51,fun(X49,bool)),X52),s(X51,X54))),s(X49,sK15(X55,X56,X57,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X58,X49,X59))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X59,bool),X55),s(fun(X49,bool),X57),s('type.sum.sum'(X59,X49),X58)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X59,bool),X56),s(fun(X49,bool),happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54))))),s('type.sum.sum'(X59,X49),X58)))) | p(s(bool,happ(s(fun(X59,bool),X55),s(X59,sK14(X55,X56,X57,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X58,X49,X59)))))) )),
  inference(superposition,[],[f217,f226])).
fof(f3270,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X38,X46,X44,X42,X40] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X48,bool),X45),s(fun(X38,bool),happ(s(fun(X39,fun(X38,bool)),X42),s(X39,X43))),s('type.sum.sum'(X48,X38),X47)))) | ~p(s(bool,happ(s(fun(X38,bool),happ(s(fun(X39,fun(X38,bool)),X42),s(X39,X43))),s(X38,sK15(X44,X45,X46,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X47,X38,X48))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X48,bool),X44),s(fun(X38,bool),X46),s('type.sum.sum'(X48,X38),X47)))) | p(s(bool,happ(s(fun(X48,bool),X44),s(X48,sK14(X44,X45,X46,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X47,X38,X48)))))) )),
  inference(forward_demodulation,[],[f3258,f225])).
fof(f3258,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X38,X46,X44,X42,X40] : (~p(s(bool,happ(s(fun(X38,bool),happ(s(fun(X39,fun(X38,bool)),X42),s(X39,X43))),s(X38,sK15(X44,X45,X46,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X47,X38,X48))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X48,bool),X44),s(fun(X38,bool),X46),s('type.sum.sum'(X48,X38),X47)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X48,bool),X45),s(fun(X38,bool),happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43))))),s('type.sum.sum'(X48,X38),X47)))) | p(s(bool,happ(s(fun(X48,bool),X44),s(X48,sK14(X44,X45,X46,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X47,X38,X48)))))) )),
  inference(superposition,[],[f217,f225])).
fof(f3269,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X37,bool),X34),s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X31),s(X28,X30))),s('type.sum.sum'(X37,X27),X36)))) | ~p(s(bool,happ(s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X31),s(X28,X30))),s(X27,sK15(X33,X34,X35,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X36,X27,X37))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X37,bool),X33),s(fun(X27,bool),X35),s('type.sum.sum'(X37,X27),X36)))) | p(s(bool,happ(s(fun(X37,bool),X33),s(X37,sK14(X33,X34,X35,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X36,X27,X37)))))) )),
  inference(forward_demodulation,[],[f3257,f228])).
fof(f3257,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (~p(s(bool,happ(s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X31),s(X28,X30))),s(X27,sK15(X33,X34,X35,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X36,X27,X37))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X37,bool),X33),s(fun(X27,bool),X35),s('type.sum.sum'(X37,X27),X36)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X37,bool),X34),s(fun(X27,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32))),s('type.sum.sum'(X37,X27),X36)))) | p(s(bool,happ(s(fun(X37,bool),X33),s(X37,sK14(X33,X34,X35,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X36,X27,X37)))))) )),
  inference(superposition,[],[f217,f228])).
fof(f3268,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X26,bool),X23),s(fun(X16,bool),happ(s(fun(X18,fun(X16,bool)),X21),s(X18,X19))),s('type.sum.sum'(X26,X16),X25)))) | ~p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X18,fun(X16,bool)),X21),s(X18,X19))),s(X16,sK15(X22,X23,X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X25,X16,X26))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X26,bool),X22),s(fun(X16,bool),X24),s('type.sum.sum'(X26,X16),X25)))) | p(s(bool,happ(s(fun(X26,bool),X22),s(X26,sK14(X22,X23,X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X25,X16,X26)))))) )),
  inference(forward_demodulation,[],[f3256,f227])).
fof(f3256,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X18,fun(X16,bool)),X21),s(X18,X19))),s(X16,sK15(X22,X23,X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X25,X16,X26))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X26,bool),X22),s(fun(X16,bool),X24),s('type.sum.sum'(X26,X16),X25)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X26,bool),X23),s(fun(X16,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21))),s('type.sum.sum'(X26,X16),X25)))) | p(s(bool,happ(s(fun(X26,bool),X22),s(X26,sK14(X22,X23,X24,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X25,X16,X26)))))) )),
  inference(superposition,[],[f217,f227])).
fof(f3267,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X15,bool),X12),s(fun(X8,bool),X10),s('type.sum.sum'(X15,X8),X14)))) | ~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK15(X11,X12,X13,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X14,X8,X15))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X15,bool),X11),s(fun(X8,bool),X13),s('type.sum.sum'(X15,X8),X14)))) | p(s(bool,happ(s(fun(X15,bool),X11),s(X15,sK14(X11,X12,X13,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X14,X8,X15)))))) )),
  inference(forward_demodulation,[],[f3255,f194])).
fof(f3255,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK15(X11,X12,X13,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X14,X8,X15))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X15,bool),X11),s(fun(X8,bool),X13),s('type.sum.sum'(X15,X8),X14)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X15,bool),X12),s(fun(X8,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10))),s('type.sum.sum'(X15,X8),X14)))) | p(s(bool,happ(s(fun(X15,bool),X11),s(X15,sK14(X11,X12,X13,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X14,X8,X15)))))) )),
  inference(superposition,[],[f217,f194])).
fof(f3266,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X7,bool),X4),s(fun(X0,bool),X1),s('type.sum.sum'(X7,X0),X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X3,X4,X5,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X6,X0,X7))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X7,bool),X3),s(fun(X0,bool),X5),s('type.sum.sum'(X7,X0),X6)))) | p(s(bool,happ(s(fun(X7,bool),X3),s(X7,sK14(X3,X4,X5,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X6,X0,X7)))))) )),
  inference(forward_demodulation,[],[f3254,f193])).
fof(f3254,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X3,X4,X5,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X6,X0,X7))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X7,bool),X3),s(fun(X0,bool),X5),s('type.sum.sum'(X7,X0),X6)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X7,bool),X4),s(fun(X0,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s('type.sum.sum'(X7,X0),X6)))) | p(s(bool,happ(s(fun(X7,bool),X3),s(X7,sK14(X3,X4,X5,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X6,X0,X7)))))) )),
  inference(superposition,[],[f217,f193])).
fof(f3128,plain,(
  ( ! [X59,X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X49,bool),happ(s(fun(X51,fun(X49,bool)),X52),s(X51,X54))),s(fun(X59,bool),X57),s('type.sum.sum'(X49,X59),X58)))) | ~p(s(bool,happ(s(fun(X49,bool),happ(s(fun(X51,fun(X49,bool)),X52),s(X51,X54))),s(X49,sK14(X55,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X56,X57,X58,X59,X49))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X49,bool),X55),s(fun(X59,bool),X56),s('type.sum.sum'(X49,X59),X58)))) | p(s(bool,happ(s(fun(X59,bool),X56),s(X59,sK15(X55,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X56,X57,X58,X59,X49)))))) )),
  inference(forward_demodulation,[],[f3116,f226])).
fof(f3116,plain,(
  ( ! [X59,X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,happ(s(fun(X49,bool),happ(s(fun(X51,fun(X49,bool)),X52),s(X51,X54))),s(X49,sK14(X55,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X56,X57,X58,X59,X49))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X49,bool),X55),s(fun(X59,bool),X56),s('type.sum.sum'(X49,X59),X58)))) | p(s(bool,happ(s(fun(X59,bool),X56),s(X59,sK15(X55,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X56,X57,X58,X59,X49))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X49,bool),happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54))))),s(fun(X59,bool),X57),s('type.sum.sum'(X49,X59),X58))))) )),
  inference(superposition,[],[f216,f226])).
fof(f3127,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X38,X46,X44,X42,X40] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X38,bool),happ(s(fun(X39,fun(X38,bool)),X42),s(X39,X43))),s(fun(X48,bool),X46),s('type.sum.sum'(X38,X48),X47)))) | ~p(s(bool,happ(s(fun(X38,bool),happ(s(fun(X39,fun(X38,bool)),X42),s(X39,X43))),s(X38,sK14(X44,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X45,X46,X47,X48,X38))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X38,bool),X44),s(fun(X48,bool),X45),s('type.sum.sum'(X38,X48),X47)))) | p(s(bool,happ(s(fun(X48,bool),X45),s(X48,sK15(X44,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X45,X46,X47,X48,X38)))))) )),
  inference(forward_demodulation,[],[f3115,f225])).
fof(f3115,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X38,X46,X44,X42,X40] : (~p(s(bool,happ(s(fun(X38,bool),happ(s(fun(X39,fun(X38,bool)),X42),s(X39,X43))),s(X38,sK14(X44,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X45,X46,X47,X48,X38))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X38,bool),X44),s(fun(X48,bool),X45),s('type.sum.sum'(X38,X48),X47)))) | p(s(bool,happ(s(fun(X48,bool),X45),s(X48,sK15(X44,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X45,X46,X47,X48,X38))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X38,bool),happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43))))),s(fun(X48,bool),X46),s('type.sum.sum'(X38,X48),X47))))) )),
  inference(superposition,[],[f216,f225])).
fof(f3126,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X31),s(X28,X30))),s(fun(X37,bool),X35),s('type.sum.sum'(X27,X37),X36)))) | ~p(s(bool,happ(s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X31),s(X28,X30))),s(X27,sK14(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X34,X35,X36,X37,X27))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X27,bool),X33),s(fun(X37,bool),X34),s('type.sum.sum'(X27,X37),X36)))) | p(s(bool,happ(s(fun(X37,bool),X34),s(X37,sK15(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X34,X35,X36,X37,X27)))))) )),
  inference(forward_demodulation,[],[f3114,f228])).
fof(f3114,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (~p(s(bool,happ(s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X31),s(X28,X30))),s(X27,sK14(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X34,X35,X36,X37,X27))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X27,bool),X33),s(fun(X37,bool),X34),s('type.sum.sum'(X27,X37),X36)))) | p(s(bool,happ(s(fun(X37,bool),X34),s(X37,sK15(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X34,X35,X36,X37,X27))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X27,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32))),s(fun(X37,bool),X35),s('type.sum.sum'(X27,X37),X36))))) )),
  inference(superposition,[],[f216,f228])).
fof(f3125,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X16,bool),happ(s(fun(X18,fun(X16,bool)),X21),s(X18,X19))),s(fun(X26,bool),X24),s('type.sum.sum'(X16,X26),X25)))) | ~p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X18,fun(X16,bool)),X21),s(X18,X19))),s(X16,sK14(X22,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X23,X24,X25,X26,X16))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X16,bool),X22),s(fun(X26,bool),X23),s('type.sum.sum'(X16,X26),X25)))) | p(s(bool,happ(s(fun(X26,bool),X23),s(X26,sK15(X22,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X23,X24,X25,X26,X16)))))) )),
  inference(forward_demodulation,[],[f3113,f227])).
fof(f3113,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X18,fun(X16,bool)),X21),s(X18,X19))),s(X16,sK14(X22,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X23,X24,X25,X26,X16))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X16,bool),X22),s(fun(X26,bool),X23),s('type.sum.sum'(X16,X26),X25)))) | p(s(bool,happ(s(fun(X26,bool),X23),s(X26,sK15(X22,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X23,X24,X25,X26,X16))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X16,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21))),s(fun(X26,bool),X24),s('type.sum.sum'(X16,X26),X25))))) )),
  inference(superposition,[],[f216,f227])).
fof(f3124,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X8,bool),X10),s(fun(X15,bool),X13),s('type.sum.sum'(X8,X15),X14)))) | ~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK14(X11,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X12,X13,X14,X15,X8))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X8,bool),X11),s(fun(X15,bool),X12),s('type.sum.sum'(X8,X15),X14)))) | p(s(bool,happ(s(fun(X15,bool),X12),s(X15,sK15(X11,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X12,X13,X14,X15,X8)))))) )),
  inference(forward_demodulation,[],[f3112,f194])).
fof(f3112,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK14(X11,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X12,X13,X14,X15,X8))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X8,bool),X11),s(fun(X15,bool),X12),s('type.sum.sum'(X8,X15),X14)))) | p(s(bool,happ(s(fun(X15,bool),X12),s(X15,sK15(X11,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X12,X13,X14,X15,X8))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X8,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10))),s(fun(X15,bool),X13),s('type.sum.sum'(X8,X15),X14))))) )),
  inference(superposition,[],[f216,f194])).
fof(f3123,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X1),s(fun(X7,bool),X5),s('type.sum.sum'(X0,X7),X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X3,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X4,X5,X6,X7,X0))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X3),s(fun(X7,bool),X4),s('type.sum.sum'(X0,X7),X6)))) | p(s(bool,happ(s(fun(X7,bool),X4),s(X7,sK15(X3,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X4,X5,X6,X7,X0)))))) )),
  inference(forward_demodulation,[],[f3111,f193])).
fof(f3111,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X3,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X4,X5,X6,X7,X0))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X3),s(fun(X7,bool),X4),s('type.sum.sum'(X0,X7),X6)))) | p(s(bool,happ(s(fun(X7,bool),X4),s(X7,sK15(X3,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X4,X5,X6,X7,X0))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(fun(X7,bool),X5),s('type.sum.sum'(X0,X7),X6))))) )),
  inference(superposition,[],[f216,f193])).
fof(f2854,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (p(s(bool,happ(s(fun(X14,bool),X19),s(X14,X18)))) | ~p(s(bool,happ(s(fun(X14,bool),X15),s(X14,X18)))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK15(X15,X19,X17,X20,'const.sum.INL_1'(s(X14,X18)),X16,X14))))) | p(s(bool,happ(s(fun(X14,bool),X15),s(X14,sK14(X15,X19,X17,X20,'const.sum.INL_1'(s(X14,X18)),X16,X14)))))) )),
  inference(forward_demodulation,[],[f2828,f211])).
fof(f2828,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (~p(s(bool,happ(s(fun(X14,bool),X15),s(X14,X18)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X14,bool),X19),s(fun(X16,bool),X20),s('type.sum.sum'(X14,X16),'const.sum.INL_1'(s(X14,X18)))))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK15(X15,X19,X17,X20,'const.sum.INL_1'(s(X14,X18)),X16,X14))))) | p(s(bool,happ(s(fun(X14,bool),X15),s(X14,sK14(X15,X19,X17,X20,'const.sum.INL_1'(s(X14,X18)),X16,X14)))))) )),
  inference(superposition,[],[f215,f211])).
fof(f2853,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,happ(s(fun(X9,bool),X13),s(X9,X11)))) | ~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X11)))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK15(X8,X12,X10,X13,'const.sum.INR_1'(s(X9,X11)),X9,X7))))) | p(s(bool,happ(s(fun(X7,bool),X8),s(X7,sK14(X8,X12,X10,X13,'const.sum.INR_1'(s(X9,X11)),X9,X7)))))) )),
  inference(forward_demodulation,[],[f2827,f210])).
fof(f2827,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,X11)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X7,bool),X12),s(fun(X9,bool),X13),s('type.sum.sum'(X7,X9),'const.sum.INR_1'(s(X9,X11)))))) | p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK15(X8,X12,X10,X13,'const.sum.INR_1'(s(X9,X11)),X9,X7))))) | p(s(bool,happ(s(fun(X7,bool),X8),s(X7,sK14(X8,X12,X10,X13,'const.sum.INR_1'(s(X9,X11)),X9,X7)))))) )),
  inference(superposition,[],[f215,f210])).
fof(f2851,plain,(
  ( ! [X70,X68,X76,X74,X72,X71,X69,X67,X77,X75,X73] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X67,bool),X76),s(fun(X68,bool),X77),s('type.sum.sum'(X67,X68),happ(s(fun(X70,'type.sum.sum'(X67,X68)),X71),s(X70,X73)))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X67,bool),X74),s(fun(X68,bool),X75),s('type.sum.sum'(X67,X68),happ(s(fun(X70,'type.sum.sum'(X67,X68)),X71),s(X70,X73)))))) | p(s(bool,happ(s(fun(X68,bool),X75),s(X68,sK15(X74,X76,X75,X77,happ(s(fun('type.sum.sum'(X69,X70),'type.sum.sum'(X67,X68)),sK16(X71,X72,X70,'type.sum.sum'(X67,X68),X69)),s('type.sum.sum'(X69,X70),'const.sum.INR_1'(s(X70,X73)))),X68,X67))))) | p(s(bool,happ(s(fun(X67,bool),X74),s(X67,sK14(X74,X76,X75,X77,happ(s(fun('type.sum.sum'(X69,X70),'type.sum.sum'(X67,X68)),sK16(X71,X72,X70,'type.sum.sum'(X67,X68),X69)),s('type.sum.sum'(X69,X70),'const.sum.INR_1'(s(X70,X73)))),X68,X67)))))) )),
  inference(forward_demodulation,[],[f2823,f226])).
fof(f2823,plain,(
  ( ! [X70,X68,X76,X74,X72,X71,X69,X67,X77,X75,X73] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X67,bool),X74),s(fun(X68,bool),X75),s('type.sum.sum'(X67,X68),happ(s(fun(X70,'type.sum.sum'(X67,X68)),X71),s(X70,X73)))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X67,bool),X76),s(fun(X68,bool),X77),s('type.sum.sum'(X67,X68),happ(s(fun('type.sum.sum'(X69,X70),'type.sum.sum'(X67,X68)),sK16(X71,X72,X70,'type.sum.sum'(X67,X68),X69)),s('type.sum.sum'(X69,X70),'const.sum.INR_1'(s(X70,X73)))))))) | p(s(bool,happ(s(fun(X68,bool),X75),s(X68,sK15(X74,X76,X75,X77,happ(s(fun('type.sum.sum'(X69,X70),'type.sum.sum'(X67,X68)),sK16(X71,X72,X70,'type.sum.sum'(X67,X68),X69)),s('type.sum.sum'(X69,X70),'const.sum.INR_1'(s(X70,X73)))),X68,X67))))) | p(s(bool,happ(s(fun(X67,bool),X74),s(X67,sK14(X74,X76,X75,X77,happ(s(fun('type.sum.sum'(X69,X70),'type.sum.sum'(X67,X68)),sK16(X71,X72,X70,'type.sum.sum'(X67,X68),X69)),s('type.sum.sum'(X69,X70),'const.sum.INR_1'(s(X70,X73)))),X68,X67)))))) )),
  inference(superposition,[],[f215,f226])).
fof(f2850,plain,(
  ( ! [X61,X59,X57,X66,X64,X62,X60,X58,X56,X65,X63] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X56,bool),X65),s(fun(X57,bool),X66),s('type.sum.sum'(X56,X57),happ(s(fun(X58,'type.sum.sum'(X56,X57)),X61),s(X58,X62)))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X56,bool),X63),s(fun(X57,bool),X64),s('type.sum.sum'(X56,X57),happ(s(fun(X58,'type.sum.sum'(X56,X57)),X61),s(X58,X62)))))) | p(s(bool,happ(s(fun(X57,bool),X64),s(X57,sK15(X63,X65,X64,X66,happ(s(fun('type.sum.sum'(X58,X59),'type.sum.sum'(X56,X57)),sK16(X60,X61,X59,'type.sum.sum'(X56,X57),X58)),s('type.sum.sum'(X58,X59),'const.sum.INL_1'(s(X58,X62)))),X57,X56))))) | p(s(bool,happ(s(fun(X56,bool),X63),s(X56,sK14(X63,X65,X64,X66,happ(s(fun('type.sum.sum'(X58,X59),'type.sum.sum'(X56,X57)),sK16(X60,X61,X59,'type.sum.sum'(X56,X57),X58)),s('type.sum.sum'(X58,X59),'const.sum.INL_1'(s(X58,X62)))),X57,X56)))))) )),
  inference(forward_demodulation,[],[f2822,f225])).
fof(f2822,plain,(
  ( ! [X61,X59,X57,X66,X64,X62,X60,X58,X56,X65,X63] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X56,bool),X63),s(fun(X57,bool),X64),s('type.sum.sum'(X56,X57),happ(s(fun(X58,'type.sum.sum'(X56,X57)),X61),s(X58,X62)))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X56,bool),X65),s(fun(X57,bool),X66),s('type.sum.sum'(X56,X57),happ(s(fun('type.sum.sum'(X58,X59),'type.sum.sum'(X56,X57)),sK16(X60,X61,X59,'type.sum.sum'(X56,X57),X58)),s('type.sum.sum'(X58,X59),'const.sum.INL_1'(s(X58,X62)))))))) | p(s(bool,happ(s(fun(X57,bool),X64),s(X57,sK15(X63,X65,X64,X66,happ(s(fun('type.sum.sum'(X58,X59),'type.sum.sum'(X56,X57)),sK16(X60,X61,X59,'type.sum.sum'(X56,X57),X58)),s('type.sum.sum'(X58,X59),'const.sum.INL_1'(s(X58,X62)))),X57,X56))))) | p(s(bool,happ(s(fun(X56,bool),X63),s(X56,sK14(X63,X65,X64,X66,happ(s(fun('type.sum.sum'(X58,X59),'type.sum.sum'(X56,X57)),sK16(X60,X61,X59,'type.sum.sum'(X56,X57),X58)),s('type.sum.sum'(X58,X59),'const.sum.INL_1'(s(X58,X62)))),X57,X56)))))) )),
  inference(superposition,[],[f215,f225])).
fof(f2849,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X46,X55,X53,X51,X49] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X45,bool),X54),s(fun(X46,bool),X55),s('type.sum.sum'(X45,X46),happ(s(fun(X47,'type.sum.sum'(X45,X46)),X50),s(X47,X49)))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X45,bool),X52),s(fun(X46,bool),X53),s('type.sum.sum'(X45,X46),happ(s(fun(X47,'type.sum.sum'(X45,X46)),X50),s(X47,X49)))))) | p(s(bool,happ(s(fun(X46,bool),X53),s(X46,sK15(X52,X54,X53,X55,'const.sum.sum_CASE_3'(s('type.sum.sum'(X47,X48),'const.sum.INL_1'(s(X47,X49))),s(fun(X47,'type.sum.sum'(X45,X46)),X50),s(fun(X48,'type.sum.sum'(X45,X46)),X51)),X46,X45))))) | p(s(bool,happ(s(fun(X45,bool),X52),s(X45,sK14(X52,X54,X53,X55,'const.sum.sum_CASE_3'(s('type.sum.sum'(X47,X48),'const.sum.INL_1'(s(X47,X49))),s(fun(X47,'type.sum.sum'(X45,X46)),X50),s(fun(X48,'type.sum.sum'(X45,X46)),X51)),X46,X45)))))) )),
  inference(forward_demodulation,[],[f2821,f228])).
fof(f2821,plain,(
  ( ! [X47,X45,X54,X52,X50,X48,X46,X55,X53,X51,X49] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X45,bool),X52),s(fun(X46,bool),X53),s('type.sum.sum'(X45,X46),happ(s(fun(X47,'type.sum.sum'(X45,X46)),X50),s(X47,X49)))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X45,bool),X54),s(fun(X46,bool),X55),s('type.sum.sum'(X45,X46),'const.sum.sum_CASE_3'(s('type.sum.sum'(X47,X48),'const.sum.INL_1'(s(X47,X49))),s(fun(X47,'type.sum.sum'(X45,X46)),X50),s(fun(X48,'type.sum.sum'(X45,X46)),X51)))))) | p(s(bool,happ(s(fun(X46,bool),X53),s(X46,sK15(X52,X54,X53,X55,'const.sum.sum_CASE_3'(s('type.sum.sum'(X47,X48),'const.sum.INL_1'(s(X47,X49))),s(fun(X47,'type.sum.sum'(X45,X46)),X50),s(fun(X48,'type.sum.sum'(X45,X46)),X51)),X46,X45))))) | p(s(bool,happ(s(fun(X45,bool),X52),s(X45,sK14(X52,X54,X53,X55,'const.sum.sum_CASE_3'(s('type.sum.sum'(X47,X48),'const.sum.INL_1'(s(X47,X49))),s(fun(X47,'type.sum.sum'(X45,X46)),X50),s(fun(X48,'type.sum.sum'(X45,X46)),X51)),X46,X45)))))) )),
  inference(superposition,[],[f215,f228])).
fof(f2848,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X44,X42,X40] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X34,bool),X43),s(fun(X35,bool),X44),s('type.sum.sum'(X34,X35),happ(s(fun(X37,'type.sum.sum'(X34,X35)),X40),s(X37,X38)))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X34,bool),X41),s(fun(X35,bool),X42),s('type.sum.sum'(X34,X35),happ(s(fun(X37,'type.sum.sum'(X34,X35)),X40),s(X37,X38)))))) | p(s(bool,happ(s(fun(X35,bool),X42),s(X35,sK15(X41,X43,X42,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X36,X37),'const.sum.INR_1'(s(X37,X38))),s(fun(X36,'type.sum.sum'(X34,X35)),X39),s(fun(X37,'type.sum.sum'(X34,X35)),X40)),X35,X34))))) | p(s(bool,happ(s(fun(X34,bool),X41),s(X34,sK14(X41,X43,X42,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X36,X37),'const.sum.INR_1'(s(X37,X38))),s(fun(X36,'type.sum.sum'(X34,X35)),X39),s(fun(X37,'type.sum.sum'(X34,X35)),X40)),X35,X34)))))) )),
  inference(forward_demodulation,[],[f2820,f227])).
fof(f2820,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X44,X42,X40] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X34,bool),X41),s(fun(X35,bool),X42),s('type.sum.sum'(X34,X35),happ(s(fun(X37,'type.sum.sum'(X34,X35)),X40),s(X37,X38)))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X34,bool),X43),s(fun(X35,bool),X44),s('type.sum.sum'(X34,X35),'const.sum.sum_CASE_3'(s('type.sum.sum'(X36,X37),'const.sum.INR_1'(s(X37,X38))),s(fun(X36,'type.sum.sum'(X34,X35)),X39),s(fun(X37,'type.sum.sum'(X34,X35)),X40)))))) | p(s(bool,happ(s(fun(X35,bool),X42),s(X35,sK15(X41,X43,X42,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X36,X37),'const.sum.INR_1'(s(X37,X38))),s(fun(X36,'type.sum.sum'(X34,X35)),X39),s(fun(X37,'type.sum.sum'(X34,X35)),X40)),X35,X34))))) | p(s(bool,happ(s(fun(X34,bool),X41),s(X34,sK14(X41,X43,X42,X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X36,X37),'const.sum.INR_1'(s(X37,X38))),s(fun(X36,'type.sum.sum'(X34,X35)),X39),s(fun(X37,'type.sum.sum'(X34,X35)),X40)),X35,X34)))))) )),
  inference(superposition,[],[f215,f227])).
fof(f2847,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X32] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X26,bool),X32),s(fun(X27,bool),X33),s('type.sum.sum'(X26,X27),X29)))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X26,bool),X30),s(fun(X27,bool),X31),s('type.sum.sum'(X26,X27),X29)))) | p(s(bool,happ(s(fun(X27,bool),X31),s(X27,sK15(X30,X32,X31,X33,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X26,X27),X28),s('type.sum.sum'(X26,X27),X29)),X27,X26))))) | p(s(bool,happ(s(fun(X26,bool),X30),s(X26,sK14(X30,X32,X31,X33,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X26,X27),X28),s('type.sum.sum'(X26,X27),X29)),X27,X26)))))) )),
  inference(forward_demodulation,[],[f2819,f194])).
fof(f2819,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X32] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X26,bool),X30),s(fun(X27,bool),X31),s('type.sum.sum'(X26,X27),X29)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X26,bool),X32),s(fun(X27,bool),X33),s('type.sum.sum'(X26,X27),'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X26,X27),X28),s('type.sum.sum'(X26,X27),X29)))))) | p(s(bool,happ(s(fun(X27,bool),X31),s(X27,sK15(X30,X32,X31,X33,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X26,X27),X28),s('type.sum.sum'(X26,X27),X29)),X27,X26))))) | p(s(bool,happ(s(fun(X26,bool),X30),s(X26,sK14(X30,X32,X31,X33,'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X26,X27),X28),s('type.sum.sum'(X26,X27),X29)),X27,X26)))))) )),
  inference(superposition,[],[f215,f194])).
fof(f2846,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X18,bool),X24),s(fun(X19,bool),X25),s('type.sum.sum'(X18,X19),X20)))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X18,bool),X22),s(fun(X19,bool),X23),s('type.sum.sum'(X18,X19),X20)))) | p(s(bool,happ(s(fun(X19,bool),X23),s(X19,sK15(X22,X24,X23,X25,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X18,X19),X20),s('type.sum.sum'(X18,X19),X21)),X19,X18))))) | p(s(bool,happ(s(fun(X18,bool),X22),s(X18,sK14(X22,X24,X23,X25,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X18,X19),X20),s('type.sum.sum'(X18,X19),X21)),X19,X18)))))) )),
  inference(forward_demodulation,[],[f2818,f193])).
fof(f2818,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X18,bool),X22),s(fun(X19,bool),X23),s('type.sum.sum'(X18,X19),X20)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X18,bool),X24),s(fun(X19,bool),X25),s('type.sum.sum'(X18,X19),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X18,X19),X20),s('type.sum.sum'(X18,X19),X21)))))) | p(s(bool,happ(s(fun(X19,bool),X23),s(X19,sK15(X22,X24,X23,X25,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X18,X19),X20),s('type.sum.sum'(X18,X19),X21)),X19,X18))))) | p(s(bool,happ(s(fun(X18,bool),X22),s(X18,sK14(X22,X24,X23,X25,'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X18,X19),X20),s('type.sum.sum'(X18,X19),X21)),X19,X18)))))) )),
  inference(superposition,[],[f215,f193])).
fof(f2845,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X16] : (p(s(bool,happ(s(fun(X11,bool),X16),s(X11,sK9(X13,X12,X11))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X11,bool),X14),s(fun(X12,bool),X15),s('type.sum.sum'(X11,X12),X13)))) | p(s(bool,happ(s(fun(X12,bool),X15),s(X12,sK15(X14,X16,X15,X17,'const.sum.INL_1'(s(X11,sK9(X13,X12,X11))),X12,X11))))) | p(s(bool,happ(s(fun(X11,bool),X14),s(X11,sK14(X14,X16,X15,X17,'const.sum.INL_1'(s(X11,sK9(X13,X12,X11))),X12,X11))))) | s('type.sum.sum'(X11,X12),X13) = s('type.sum.sum'(X11,X12),'const.sum.INR_1'(s(X12,sK10(X13,X12,X11))))) )),
  inference(forward_demodulation,[],[f2817,f211])).
fof(f2817,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X16] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X11,bool),X14),s(fun(X12,bool),X15),s('type.sum.sum'(X11,X12),X13)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X11,bool),X16),s(fun(X12,bool),X17),s('type.sum.sum'(X11,X12),'const.sum.INL_1'(s(X11,sK9(X13,X12,X11))))))) | p(s(bool,happ(s(fun(X12,bool),X15),s(X12,sK15(X14,X16,X15,X17,'const.sum.INL_1'(s(X11,sK9(X13,X12,X11))),X12,X11))))) | p(s(bool,happ(s(fun(X11,bool),X14),s(X11,sK14(X14,X16,X15,X17,'const.sum.INL_1'(s(X11,sK9(X13,X12,X11))),X12,X11))))) | s('type.sum.sum'(X11,X12),X13) = s('type.sum.sum'(X11,X12),'const.sum.INR_1'(s(X12,sK10(X13,X12,X11))))) )),
  inference(superposition,[],[f215,f205])).
fof(f2844,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(X5,bool),X10),s(X5,sK10(X6,X5,X4))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X4,bool),X7),s(fun(X5,bool),X8),s('type.sum.sum'(X4,X5),X6)))) | p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK15(X7,X9,X8,X10,'const.sum.INR_1'(s(X5,sK10(X6,X5,X4))),X5,X4))))) | p(s(bool,happ(s(fun(X4,bool),X7),s(X4,sK14(X7,X9,X8,X10,'const.sum.INR_1'(s(X5,sK10(X6,X5,X4))),X5,X4))))) | s('type.sum.sum'(X4,X5),X6) = s('type.sum.sum'(X4,X5),'const.sum.INL_1'(s(X4,sK9(X6,X5,X4))))) )),
  inference(forward_demodulation,[],[f2816,f210])).
fof(f2816,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X4,bool),X7),s(fun(X5,bool),X8),s('type.sum.sum'(X4,X5),X6)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X4,bool),X9),s(fun(X5,bool),X10),s('type.sum.sum'(X4,X5),'const.sum.INR_1'(s(X5,sK10(X6,X5,X4))))))) | p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK15(X7,X9,X8,X10,'const.sum.INR_1'(s(X5,sK10(X6,X5,X4))),X5,X4))))) | p(s(bool,happ(s(fun(X4,bool),X7),s(X4,sK14(X7,X9,X8,X10,'const.sum.INR_1'(s(X5,sK10(X6,X5,X4))),X5,X4))))) | s('type.sum.sum'(X4,X5),X6) = s('type.sum.sum'(X4,X5),'const.sum.INL_1'(s(X4,sK9(X6,X5,X4))))) )),
  inference(superposition,[],[f215,f205])).
fof(f2842,plain,(
  ( ! [X59,X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (p(s(bool,happ(s(fun(X49,bool),happ(s(fun(X51,fun(X49,bool)),X52),s(X51,X54))),s(X49,sK15(X56,X58,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X59,X57,X49,X55))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X55,bool),X56),s(fun(X49,bool),happ(s(fun(X51,fun(X49,bool)),X52),s(X51,X54))),s('type.sum.sum'(X55,X49),X57)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X55,bool),X58),s(fun(X49,bool),X59),s('type.sum.sum'(X55,X49),X57)))) | p(s(bool,happ(s(fun(X55,bool),X56),s(X55,sK14(X56,X58,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X59,X57,X49,X55)))))) )),
  inference(forward_demodulation,[],[f2814,f226])).
fof(f2814,plain,(
  ( ! [X59,X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X55,bool),X56),s(fun(X49,bool),happ(s(fun(X51,fun(X49,bool)),X52),s(X51,X54))),s('type.sum.sum'(X55,X49),X57)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X55,bool),X58),s(fun(X49,bool),X59),s('type.sum.sum'(X55,X49),X57)))) | p(s(bool,happ(s(fun(X49,bool),happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54))))),s(X49,sK15(X56,X58,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X59,X57,X49,X55))))) | p(s(bool,happ(s(fun(X55,bool),X56),s(X55,sK14(X56,X58,happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X59,X57,X49,X55)))))) )),
  inference(superposition,[],[f215,f226])).
fof(f2841,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X38,X46,X44,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),happ(s(fun(X39,fun(X38,bool)),X42),s(X39,X43))),s(X38,sK15(X45,X47,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X48,X46,X38,X44))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X44,bool),X45),s(fun(X38,bool),happ(s(fun(X39,fun(X38,bool)),X42),s(X39,X43))),s('type.sum.sum'(X44,X38),X46)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X44,bool),X47),s(fun(X38,bool),X48),s('type.sum.sum'(X44,X38),X46)))) | p(s(bool,happ(s(fun(X44,bool),X45),s(X44,sK14(X45,X47,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X48,X46,X38,X44)))))) )),
  inference(forward_demodulation,[],[f2813,f225])).
fof(f2813,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X38,X46,X44,X42,X40] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X44,bool),X45),s(fun(X38,bool),happ(s(fun(X39,fun(X38,bool)),X42),s(X39,X43))),s('type.sum.sum'(X44,X38),X46)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X44,bool),X47),s(fun(X38,bool),X48),s('type.sum.sum'(X44,X38),X46)))) | p(s(bool,happ(s(fun(X38,bool),happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43))))),s(X38,sK15(X45,X47,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X48,X46,X38,X44))))) | p(s(bool,happ(s(fun(X44,bool),X45),s(X44,sK14(X45,X47,happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X48,X46,X38,X44)))))) )),
  inference(superposition,[],[f215,f225])).
fof(f2840,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (p(s(bool,happ(s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X31),s(X28,X30))),s(X27,sK15(X34,X36,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X37,X35,X27,X33))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X33,bool),X34),s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X31),s(X28,X30))),s('type.sum.sum'(X33,X27),X35)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X33,bool),X36),s(fun(X27,bool),X37),s('type.sum.sum'(X33,X27),X35)))) | p(s(bool,happ(s(fun(X33,bool),X34),s(X33,sK14(X34,X36,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X37,X35,X27,X33)))))) )),
  inference(forward_demodulation,[],[f2812,f228])).
fof(f2812,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X33,bool),X34),s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X31),s(X28,X30))),s('type.sum.sum'(X33,X27),X35)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X33,bool),X36),s(fun(X27,bool),X37),s('type.sum.sum'(X33,X27),X35)))) | p(s(bool,happ(s(fun(X27,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32))),s(X27,sK15(X34,X36,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X37,X35,X27,X33))))) | p(s(bool,happ(s(fun(X33,bool),X34),s(X33,sK14(X34,X36,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X37,X35,X27,X33)))))) )),
  inference(superposition,[],[f215,f228])).
fof(f2839,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X18,fun(X16,bool)),X21),s(X18,X19))),s(X16,sK15(X23,X25,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X26,X24,X16,X22))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X22,bool),X23),s(fun(X16,bool),happ(s(fun(X18,fun(X16,bool)),X21),s(X18,X19))),s('type.sum.sum'(X22,X16),X24)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X22,bool),X25),s(fun(X16,bool),X26),s('type.sum.sum'(X22,X16),X24)))) | p(s(bool,happ(s(fun(X22,bool),X23),s(X22,sK14(X23,X25,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X26,X24,X16,X22)))))) )),
  inference(forward_demodulation,[],[f2811,f227])).
fof(f2811,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X22,bool),X23),s(fun(X16,bool),happ(s(fun(X18,fun(X16,bool)),X21),s(X18,X19))),s('type.sum.sum'(X22,X16),X24)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X22,bool),X25),s(fun(X16,bool),X26),s('type.sum.sum'(X22,X16),X24)))) | p(s(bool,happ(s(fun(X16,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21))),s(X16,sK15(X23,X25,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X26,X24,X16,X22))))) | p(s(bool,happ(s(fun(X22,bool),X23),s(X22,sK14(X23,X25,'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X26,X24,X16,X22)))))) )),
  inference(superposition,[],[f215,f227])).
fof(f2838,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK15(X12,X14,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X15,X13,X8,X11))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X11,bool),X12),s(fun(X8,bool),X10),s('type.sum.sum'(X11,X8),X13)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X11,bool),X14),s(fun(X8,bool),X15),s('type.sum.sum'(X11,X8),X13)))) | p(s(bool,happ(s(fun(X11,bool),X12),s(X11,sK14(X12,X14,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X15,X13,X8,X11)))))) )),
  inference(forward_demodulation,[],[f2810,f194])).
fof(f2810,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X11,bool),X12),s(fun(X8,bool),X10),s('type.sum.sum'(X11,X8),X13)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X11,bool),X14),s(fun(X8,bool),X15),s('type.sum.sum'(X11,X8),X13)))) | p(s(bool,happ(s(fun(X8,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10))),s(X8,sK15(X12,X14,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X15,X13,X8,X11))))) | p(s(bool,happ(s(fun(X11,bool),X12),s(X11,sK14(X12,X14,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X15,X13,X8,X11)))))) )),
  inference(superposition,[],[f215,f194])).
fof(f2837,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X4,X6,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X7,X5,X0,X3))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X3,bool),X4),s(fun(X0,bool),X1),s('type.sum.sum'(X3,X0),X5)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X3,bool),X6),s(fun(X0,bool),X7),s('type.sum.sum'(X3,X0),X5)))) | p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK14(X4,X6,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X7,X5,X0,X3)))))) )),
  inference(forward_demodulation,[],[f2809,f193])).
fof(f2809,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X3,bool),X4),s(fun(X0,bool),X1),s('type.sum.sum'(X3,X0),X5)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X3,bool),X6),s(fun(X0,bool),X7),s('type.sum.sum'(X3,X0),X5)))) | p(s(bool,happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(X0,sK15(X4,X6,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X7,X5,X0,X3))))) | p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK14(X4,X6,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X7,X5,X0,X3)))))) )),
  inference(superposition,[],[f215,f193])).
fof(f2836,plain,(
  ( ! [X59,X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (p(s(bool,happ(s(fun(X49,bool),happ(s(fun(X51,fun(X49,bool)),X52),s(X51,X54))),s(X49,sK14(happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X58,X56,X59,X57,X55,X49))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X49,bool),happ(s(fun(X51,fun(X49,bool)),X52),s(X51,X54))),s(fun(X55,bool),X56),s('type.sum.sum'(X49,X55),X57)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X49,bool),X58),s(fun(X55,bool),X59),s('type.sum.sum'(X49,X55),X57)))) | p(s(bool,happ(s(fun(X55,bool),X56),s(X55,sK15(happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X58,X56,X59,X57,X55,X49)))))) )),
  inference(forward_demodulation,[],[f2808,f226])).
fof(f2808,plain,(
  ( ! [X59,X57,X54,X52,X50,X58,X56,X55,X53,X51,X49] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X49,bool),happ(s(fun(X51,fun(X49,bool)),X52),s(X51,X54))),s(fun(X55,bool),X56),s('type.sum.sum'(X49,X55),X57)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X49,bool),X58),s(fun(X55,bool),X59),s('type.sum.sum'(X49,X55),X57)))) | p(s(bool,happ(s(fun(X55,bool),X56),s(X55,sK15(happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X58,X56,X59,X57,X55,X49))))) | p(s(bool,happ(s(fun(X49,bool),happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54))))),s(X49,sK14(happ(s(fun('type.sum.sum'(X50,X51),fun(X49,bool)),sK16(X52,X53,X51,fun(X49,bool),X50)),s('type.sum.sum'(X50,X51),'const.sum.INR_1'(s(X51,X54)))),X58,X56,X59,X57,X55,X49)))))) )),
  inference(superposition,[],[f215,f226])).
fof(f2835,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X38,X46,X44,X42,X40] : (p(s(bool,happ(s(fun(X38,bool),happ(s(fun(X39,fun(X38,bool)),X42),s(X39,X43))),s(X38,sK14(happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X47,X45,X48,X46,X44,X38))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X38,bool),happ(s(fun(X39,fun(X38,bool)),X42),s(X39,X43))),s(fun(X44,bool),X45),s('type.sum.sum'(X38,X44),X46)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X38,bool),X47),s(fun(X44,bool),X48),s('type.sum.sum'(X38,X44),X46)))) | p(s(bool,happ(s(fun(X44,bool),X45),s(X44,sK15(happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X47,X45,X48,X46,X44,X38)))))) )),
  inference(forward_demodulation,[],[f2807,f225])).
fof(f2807,plain,(
  ( ! [X39,X47,X45,X43,X41,X48,X38,X46,X44,X42,X40] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X38,bool),happ(s(fun(X39,fun(X38,bool)),X42),s(X39,X43))),s(fun(X44,bool),X45),s('type.sum.sum'(X38,X44),X46)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X38,bool),X47),s(fun(X44,bool),X48),s('type.sum.sum'(X38,X44),X46)))) | p(s(bool,happ(s(fun(X44,bool),X45),s(X44,sK15(happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X47,X45,X48,X46,X44,X38))))) | p(s(bool,happ(s(fun(X38,bool),happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43))))),s(X38,sK14(happ(s(fun('type.sum.sum'(X39,X40),fun(X38,bool)),sK16(X41,X42,X40,fun(X38,bool),X39)),s('type.sum.sum'(X39,X40),'const.sum.INL_1'(s(X39,X43)))),X47,X45,X48,X46,X44,X38)))))) )),
  inference(superposition,[],[f215,f225])).
fof(f2834,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (p(s(bool,happ(s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X31),s(X28,X30))),s(X27,sK14('const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X36,X34,X37,X35,X33,X27))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X31),s(X28,X30))),s(fun(X33,bool),X34),s('type.sum.sum'(X27,X33),X35)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X27,bool),X36),s(fun(X33,bool),X37),s('type.sum.sum'(X27,X33),X35)))) | p(s(bool,happ(s(fun(X33,bool),X34),s(X33,sK15('const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X36,X34,X37,X35,X33,X27)))))) )),
  inference(forward_demodulation,[],[f2806,f228])).
fof(f2806,plain,(
  ( ! [X30,X28,X37,X35,X33,X31,X29,X27,X36,X34,X32] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X27,bool),happ(s(fun(X28,fun(X27,bool)),X31),s(X28,X30))),s(fun(X33,bool),X34),s('type.sum.sum'(X27,X33),X35)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X27,bool),X36),s(fun(X33,bool),X37),s('type.sum.sum'(X27,X33),X35)))) | p(s(bool,happ(s(fun(X33,bool),X34),s(X33,sK15('const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X36,X34,X37,X35,X33,X27))))) | p(s(bool,happ(s(fun(X27,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32))),s(X27,sK14('const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X30))),s(fun(X28,fun(X27,bool)),X31),s(fun(X29,fun(X27,bool)),X32)),X36,X34,X37,X35,X33,X27)))))) )),
  inference(superposition,[],[f215,f228])).
fof(f2833,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (p(s(bool,happ(s(fun(X16,bool),happ(s(fun(X18,fun(X16,bool)),X21),s(X18,X19))),s(X16,sK14('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X25,X23,X26,X24,X22,X16))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X16,bool),happ(s(fun(X18,fun(X16,bool)),X21),s(X18,X19))),s(fun(X22,bool),X23),s('type.sum.sum'(X16,X22),X24)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X16,bool),X25),s(fun(X22,bool),X26),s('type.sum.sum'(X16,X22),X24)))) | p(s(bool,happ(s(fun(X22,bool),X23),s(X22,sK15('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X25,X23,X26,X24,X22,X16)))))) )),
  inference(forward_demodulation,[],[f2805,f227])).
fof(f2805,plain,(
  ( ! [X26,X24,X23,X21,X19,X17,X25,X22,X20,X18,X16] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X16,bool),happ(s(fun(X18,fun(X16,bool)),X21),s(X18,X19))),s(fun(X22,bool),X23),s('type.sum.sum'(X16,X22),X24)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X16,bool),X25),s(fun(X22,bool),X26),s('type.sum.sum'(X16,X22),X24)))) | p(s(bool,happ(s(fun(X22,bool),X23),s(X22,sK15('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X25,X23,X26,X24,X22,X16))))) | p(s(bool,happ(s(fun(X16,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21))),s(X16,sK14('const.sum.sum_CASE_3'(s('type.sum.sum'(X17,X18),'const.sum.INR_1'(s(X18,X19))),s(fun(X17,fun(X16,bool)),X20),s(fun(X18,fun(X16,bool)),X21)),X25,X23,X26,X24,X22,X16)))))) )),
  inference(superposition,[],[f215,f227])).
fof(f2832,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (p(s(bool,happ(s(fun(X8,bool),X10),s(X8,sK14('const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X14,X12,X15,X13,X11,X8))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X8,bool),X10),s(fun(X11,bool),X12),s('type.sum.sum'(X8,X11),X13)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X8,bool),X14),s(fun(X11,bool),X15),s('type.sum.sum'(X8,X11),X13)))) | p(s(bool,happ(s(fun(X11,bool),X12),s(X11,sK15('const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X14,X12,X15,X13,X11,X8)))))) )),
  inference(forward_demodulation,[],[f2804,f194])).
fof(f2804,plain,(
  ( ! [X14,X12,X10,X8,X15,X13,X11,X9] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X8,bool),X10),s(fun(X11,bool),X12),s('type.sum.sum'(X8,X11),X13)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X8,bool),X14),s(fun(X11,bool),X15),s('type.sum.sum'(X8,X11),X13)))) | p(s(bool,happ(s(fun(X11,bool),X12),s(X11,sK15('const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X14,X12,X15,X13,X11,X8))))) | p(s(bool,happ(s(fun(X8,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10))),s(X8,sK14('const.bool.COND_3'(s(bool,'F_0'),s(fun(X8,bool),X9),s(fun(X8,bool),X10)),X14,X12,X15,X13,X11,X8)))))) )),
  inference(superposition,[],[f215,f194])).
fof(f2831,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X6,X4,X7,X5,X3,X0))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X1),s(fun(X3,bool),X4),s('type.sum.sum'(X0,X3),X5)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X6),s(fun(X3,bool),X7),s('type.sum.sum'(X0,X3),X5)))) | p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK15('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X6,X4,X7,X5,X3,X0)))))) )),
  inference(forward_demodulation,[],[f2803,f193])).
fof(f2803,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X1),s(fun(X3,bool),X4),s('type.sum.sum'(X0,X3),X5)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X6),s(fun(X3,bool),X7),s('type.sum.sum'(X0,X3),X5)))) | p(s(bool,happ(s(fun(X3,bool),X4),s(X3,sK15('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X6,X4,X7,X5,X3,X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2))),s(X0,sK14('const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X6,X4,X7,X5,X3,X0)))))) )),
  inference(superposition,[],[f215,f193])).
fof(f2703,plain,(
  ( ! [X12,X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X8,bool),X11),s(X8,sK11(sK16(X10,X11,X9,bool,X8),X9,X8))))) | ~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK12(sK16(X10,X11,X9,bool,X8),X9,X8))))) | p(s(bool,happ(s(fun('type.sum.sum'(X8,X9),bool),sK16(X10,X11,X9,bool,X8)),s('type.sum.sum'(X8,X9),X12))))) )),
  inference(forward_demodulation,[],[f2688,f225])).
fof(f2688,plain,(
  ( ! [X12,X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(X9,bool),X10),s(X9,sK12(sK16(X10,X11,X9,bool,X8),X9,X8))))) | p(s(bool,happ(s(fun('type.sum.sum'(X8,X9),bool),sK16(X10,X11,X9,bool,X8)),s('type.sum.sum'(X8,X9),X12)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X8,X9),bool),sK16(X10,X11,X9,bool,X8)),s('type.sum.sum'(X8,X9),'const.sum.INL_1'(s(X8,sK11(sK16(X10,X11,X9,bool,X8),X9,X8)))))))) )),
  inference(superposition,[],[f206,f226])).
fof(f2700,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun('type.sum.sum'(X34,X35),bool),happ(s(fun(X37,fun('type.sum.sum'(X34,X35),bool)),X38),s(X37,X40))),s('type.sum.sum'(X34,X35),'const.sum.INL_1'(s(X34,sK11(happ(s(fun('type.sum.sum'(X36,X37),fun('type.sum.sum'(X34,X35),bool)),sK16(X38,X39,X37,fun('type.sum.sum'(X34,X35),bool),X36)),s('type.sum.sum'(X36,X37),'const.sum.INR_1'(s(X37,X40)))),X35,X34))))))) | p(s(bool,happ(s(fun('type.sum.sum'(X34,X35),bool),happ(s(fun(X37,fun('type.sum.sum'(X34,X35),bool)),X38),s(X37,X40))),s('type.sum.sum'(X34,X35),X41)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X34,X35),bool),happ(s(fun(X37,fun('type.sum.sum'(X34,X35),bool)),X38),s(X37,X40))),s('type.sum.sum'(X34,X35),'const.sum.INR_1'(s(X35,sK12(happ(s(fun('type.sum.sum'(X36,X37),fun('type.sum.sum'(X34,X35),bool)),sK16(X38,X39,X37,fun('type.sum.sum'(X34,X35),bool),X36)),s('type.sum.sum'(X36,X37),'const.sum.INR_1'(s(X37,X40)))),X35,X34)))))))) )),
  inference(forward_demodulation,[],[f2699,f226])).
fof(f2699,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (p(s(bool,happ(s(fun('type.sum.sum'(X34,X35),bool),happ(s(fun(X37,fun('type.sum.sum'(X34,X35),bool)),X38),s(X37,X40))),s('type.sum.sum'(X34,X35),X41)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X34,X35),bool),happ(s(fun(X37,fun('type.sum.sum'(X34,X35),bool)),X38),s(X37,X40))),s('type.sum.sum'(X34,X35),'const.sum.INR_1'(s(X35,sK12(happ(s(fun('type.sum.sum'(X36,X37),fun('type.sum.sum'(X34,X35),bool)),sK16(X38,X39,X37,fun('type.sum.sum'(X34,X35),bool),X36)),s('type.sum.sum'(X36,X37),'const.sum.INR_1'(s(X37,X40)))),X35,X34))))))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X34,X35),bool),happ(s(fun('type.sum.sum'(X36,X37),fun('type.sum.sum'(X34,X35),bool)),sK16(X38,X39,X37,fun('type.sum.sum'(X34,X35),bool),X36)),s('type.sum.sum'(X36,X37),'const.sum.INR_1'(s(X37,X40))))),s('type.sum.sum'(X34,X35),'const.sum.INL_1'(s(X34,sK11(happ(s(fun('type.sum.sum'(X36,X37),fun('type.sum.sum'(X34,X35),bool)),sK16(X38,X39,X37,fun('type.sum.sum'(X34,X35),bool),X36)),s('type.sum.sum'(X36,X37),'const.sum.INR_1'(s(X37,X40)))),X35,X34)))))))) )),
  inference(forward_demodulation,[],[f2683,f226])).
fof(f2683,plain,(
  ( ! [X39,X37,X35,X41,X38,X36,X34,X40] : (~p(s(bool,happ(s(fun('type.sum.sum'(X34,X35),bool),happ(s(fun(X37,fun('type.sum.sum'(X34,X35),bool)),X38),s(X37,X40))),s('type.sum.sum'(X34,X35),'const.sum.INR_1'(s(X35,sK12(happ(s(fun('type.sum.sum'(X36,X37),fun('type.sum.sum'(X34,X35),bool)),sK16(X38,X39,X37,fun('type.sum.sum'(X34,X35),bool),X36)),s('type.sum.sum'(X36,X37),'const.sum.INR_1'(s(X37,X40)))),X35,X34))))))) | p(s(bool,happ(s(fun('type.sum.sum'(X34,X35),bool),happ(s(fun('type.sum.sum'(X36,X37),fun('type.sum.sum'(X34,X35),bool)),sK16(X38,X39,X37,fun('type.sum.sum'(X34,X35),bool),X36)),s('type.sum.sum'(X36,X37),'const.sum.INR_1'(s(X37,X40))))),s('type.sum.sum'(X34,X35),X41)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X34,X35),bool),happ(s(fun('type.sum.sum'(X36,X37),fun('type.sum.sum'(X34,X35),bool)),sK16(X38,X39,X37,fun('type.sum.sum'(X34,X35),bool),X36)),s('type.sum.sum'(X36,X37),'const.sum.INR_1'(s(X37,X40))))),s('type.sum.sum'(X34,X35),'const.sum.INL_1'(s(X34,sK11(happ(s(fun('type.sum.sum'(X36,X37),fun('type.sum.sum'(X34,X35),bool)),sK16(X38,X39,X37,fun('type.sum.sum'(X34,X35),bool),X36)),s('type.sum.sum'(X36,X37),'const.sum.INR_1'(s(X37,X40)))),X35,X34)))))))) )),
  inference(superposition,[],[f206,f226])).
fof(f2698,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X32] : (~p(s(bool,happ(s(fun('type.sum.sum'(X26,X27),bool),happ(s(fun(X28,fun('type.sum.sum'(X26,X27),bool)),X31),s(X28,X32))),s('type.sum.sum'(X26,X27),'const.sum.INL_1'(s(X26,sK11(happ(s(fun('type.sum.sum'(X28,X29),fun('type.sum.sum'(X26,X27),bool)),sK16(X30,X31,X29,fun('type.sum.sum'(X26,X27),bool),X28)),s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X32)))),X27,X26))))))) | p(s(bool,happ(s(fun('type.sum.sum'(X26,X27),bool),happ(s(fun(X28,fun('type.sum.sum'(X26,X27),bool)),X31),s(X28,X32))),s('type.sum.sum'(X26,X27),X33)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X26,X27),bool),happ(s(fun(X28,fun('type.sum.sum'(X26,X27),bool)),X31),s(X28,X32))),s('type.sum.sum'(X26,X27),'const.sum.INR_1'(s(X27,sK12(happ(s(fun('type.sum.sum'(X28,X29),fun('type.sum.sum'(X26,X27),bool)),sK16(X30,X31,X29,fun('type.sum.sum'(X26,X27),bool),X28)),s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X32)))),X27,X26)))))))) )),
  inference(forward_demodulation,[],[f2697,f225])).
fof(f2697,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X32] : (p(s(bool,happ(s(fun('type.sum.sum'(X26,X27),bool),happ(s(fun(X28,fun('type.sum.sum'(X26,X27),bool)),X31),s(X28,X32))),s('type.sum.sum'(X26,X27),X33)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X26,X27),bool),happ(s(fun(X28,fun('type.sum.sum'(X26,X27),bool)),X31),s(X28,X32))),s('type.sum.sum'(X26,X27),'const.sum.INR_1'(s(X27,sK12(happ(s(fun('type.sum.sum'(X28,X29),fun('type.sum.sum'(X26,X27),bool)),sK16(X30,X31,X29,fun('type.sum.sum'(X26,X27),bool),X28)),s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X32)))),X27,X26))))))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X26,X27),bool),happ(s(fun('type.sum.sum'(X28,X29),fun('type.sum.sum'(X26,X27),bool)),sK16(X30,X31,X29,fun('type.sum.sum'(X26,X27),bool),X28)),s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X32))))),s('type.sum.sum'(X26,X27),'const.sum.INL_1'(s(X26,sK11(happ(s(fun('type.sum.sum'(X28,X29),fun('type.sum.sum'(X26,X27),bool)),sK16(X30,X31,X29,fun('type.sum.sum'(X26,X27),bool),X28)),s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X32)))),X27,X26)))))))) )),
  inference(forward_demodulation,[],[f2682,f225])).
fof(f2682,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X32] : (~p(s(bool,happ(s(fun('type.sum.sum'(X26,X27),bool),happ(s(fun(X28,fun('type.sum.sum'(X26,X27),bool)),X31),s(X28,X32))),s('type.sum.sum'(X26,X27),'const.sum.INR_1'(s(X27,sK12(happ(s(fun('type.sum.sum'(X28,X29),fun('type.sum.sum'(X26,X27),bool)),sK16(X30,X31,X29,fun('type.sum.sum'(X26,X27),bool),X28)),s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X32)))),X27,X26))))))) | p(s(bool,happ(s(fun('type.sum.sum'(X26,X27),bool),happ(s(fun('type.sum.sum'(X28,X29),fun('type.sum.sum'(X26,X27),bool)),sK16(X30,X31,X29,fun('type.sum.sum'(X26,X27),bool),X28)),s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X32))))),s('type.sum.sum'(X26,X27),X33)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X26,X27),bool),happ(s(fun('type.sum.sum'(X28,X29),fun('type.sum.sum'(X26,X27),bool)),sK16(X30,X31,X29,fun('type.sum.sum'(X26,X27),bool),X28)),s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X32))))),s('type.sum.sum'(X26,X27),'const.sum.INL_1'(s(X26,sK11(happ(s(fun('type.sum.sum'(X28,X29),fun('type.sum.sum'(X26,X27),bool)),sK16(X30,X31,X29,fun('type.sum.sum'(X26,X27),bool),X28)),s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,X32)))),X27,X26)))))))) )),
  inference(superposition,[],[f206,f225])).
fof(f2696,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun('type.sum.sum'(X18,X19),bool),happ(s(fun(X20,fun('type.sum.sum'(X18,X19),bool)),X23),s(X20,X22))),s('type.sum.sum'(X18,X19),'const.sum.INL_1'(s(X18,sK11('const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X21),'const.sum.INL_1'(s(X20,X22))),s(fun(X20,fun('type.sum.sum'(X18,X19),bool)),X23),s(fun(X21,fun('type.sum.sum'(X18,X19),bool)),X24)),X19,X18))))))) | p(s(bool,happ(s(fun('type.sum.sum'(X18,X19),bool),happ(s(fun(X20,fun('type.sum.sum'(X18,X19),bool)),X23),s(X20,X22))),s('type.sum.sum'(X18,X19),X25)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X18,X19),bool),happ(s(fun(X20,fun('type.sum.sum'(X18,X19),bool)),X23),s(X20,X22))),s('type.sum.sum'(X18,X19),'const.sum.INR_1'(s(X19,sK12('const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X21),'const.sum.INL_1'(s(X20,X22))),s(fun(X20,fun('type.sum.sum'(X18,X19),bool)),X23),s(fun(X21,fun('type.sum.sum'(X18,X19),bool)),X24)),X19,X18)))))))) )),
  inference(forward_demodulation,[],[f2695,f228])).
fof(f2695,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (p(s(bool,happ(s(fun('type.sum.sum'(X18,X19),bool),happ(s(fun(X20,fun('type.sum.sum'(X18,X19),bool)),X23),s(X20,X22))),s('type.sum.sum'(X18,X19),X25)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X18,X19),bool),happ(s(fun(X20,fun('type.sum.sum'(X18,X19),bool)),X23),s(X20,X22))),s('type.sum.sum'(X18,X19),'const.sum.INR_1'(s(X19,sK12('const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X21),'const.sum.INL_1'(s(X20,X22))),s(fun(X20,fun('type.sum.sum'(X18,X19),bool)),X23),s(fun(X21,fun('type.sum.sum'(X18,X19),bool)),X24)),X19,X18))))))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X18,X19),bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X21),'const.sum.INL_1'(s(X20,X22))),s(fun(X20,fun('type.sum.sum'(X18,X19),bool)),X23),s(fun(X21,fun('type.sum.sum'(X18,X19),bool)),X24))),s('type.sum.sum'(X18,X19),'const.sum.INL_1'(s(X18,sK11('const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X21),'const.sum.INL_1'(s(X20,X22))),s(fun(X20,fun('type.sum.sum'(X18,X19),bool)),X23),s(fun(X21,fun('type.sum.sum'(X18,X19),bool)),X24)),X19,X18)))))))) )),
  inference(forward_demodulation,[],[f2681,f228])).
fof(f2681,plain,(
  ( ! [X24,X23,X21,X19,X25,X22,X20,X18] : (~p(s(bool,happ(s(fun('type.sum.sum'(X18,X19),bool),happ(s(fun(X20,fun('type.sum.sum'(X18,X19),bool)),X23),s(X20,X22))),s('type.sum.sum'(X18,X19),'const.sum.INR_1'(s(X19,sK12('const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X21),'const.sum.INL_1'(s(X20,X22))),s(fun(X20,fun('type.sum.sum'(X18,X19),bool)),X23),s(fun(X21,fun('type.sum.sum'(X18,X19),bool)),X24)),X19,X18))))))) | p(s(bool,happ(s(fun('type.sum.sum'(X18,X19),bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X21),'const.sum.INL_1'(s(X20,X22))),s(fun(X20,fun('type.sum.sum'(X18,X19),bool)),X23),s(fun(X21,fun('type.sum.sum'(X18,X19),bool)),X24))),s('type.sum.sum'(X18,X19),X25)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X18,X19),bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X21),'const.sum.INL_1'(s(X20,X22))),s(fun(X20,fun('type.sum.sum'(X18,X19),bool)),X23),s(fun(X21,fun('type.sum.sum'(X18,X19),bool)),X24))),s('type.sum.sum'(X18,X19),'const.sum.INL_1'(s(X18,sK11('const.sum.sum_CASE_3'(s('type.sum.sum'(X20,X21),'const.sum.INL_1'(s(X20,X22))),s(fun(X20,fun('type.sum.sum'(X18,X19),bool)),X23),s(fun(X21,fun('type.sum.sum'(X18,X19),bool)),X24)),X19,X18)))))))) )),
  inference(superposition,[],[f206,f228])).
fof(f2694,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun('type.sum.sum'(X10,X11),bool),happ(s(fun(X13,fun('type.sum.sum'(X10,X11),bool)),X16),s(X13,X14))),s('type.sum.sum'(X10,X11),'const.sum.INL_1'(s(X10,sK11('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,fun('type.sum.sum'(X10,X11),bool)),X15),s(fun(X13,fun('type.sum.sum'(X10,X11),bool)),X16)),X11,X10))))))) | p(s(bool,happ(s(fun('type.sum.sum'(X10,X11),bool),happ(s(fun(X13,fun('type.sum.sum'(X10,X11),bool)),X16),s(X13,X14))),s('type.sum.sum'(X10,X11),X17)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X10,X11),bool),happ(s(fun(X13,fun('type.sum.sum'(X10,X11),bool)),X16),s(X13,X14))),s('type.sum.sum'(X10,X11),'const.sum.INR_1'(s(X11,sK12('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,fun('type.sum.sum'(X10,X11),bool)),X15),s(fun(X13,fun('type.sum.sum'(X10,X11),bool)),X16)),X11,X10)))))))) )),
  inference(forward_demodulation,[],[f2693,f227])).
fof(f2693,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (p(s(bool,happ(s(fun('type.sum.sum'(X10,X11),bool),happ(s(fun(X13,fun('type.sum.sum'(X10,X11),bool)),X16),s(X13,X14))),s('type.sum.sum'(X10,X11),X17)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X10,X11),bool),happ(s(fun(X13,fun('type.sum.sum'(X10,X11),bool)),X16),s(X13,X14))),s('type.sum.sum'(X10,X11),'const.sum.INR_1'(s(X11,sK12('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,fun('type.sum.sum'(X10,X11),bool)),X15),s(fun(X13,fun('type.sum.sum'(X10,X11),bool)),X16)),X11,X10))))))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X10,X11),bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,fun('type.sum.sum'(X10,X11),bool)),X15),s(fun(X13,fun('type.sum.sum'(X10,X11),bool)),X16))),s('type.sum.sum'(X10,X11),'const.sum.INL_1'(s(X10,sK11('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,fun('type.sum.sum'(X10,X11),bool)),X15),s(fun(X13,fun('type.sum.sum'(X10,X11),bool)),X16)),X11,X10)))))))) )),
  inference(forward_demodulation,[],[f2680,f227])).
fof(f2680,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (~p(s(bool,happ(s(fun('type.sum.sum'(X10,X11),bool),happ(s(fun(X13,fun('type.sum.sum'(X10,X11),bool)),X16),s(X13,X14))),s('type.sum.sum'(X10,X11),'const.sum.INR_1'(s(X11,sK12('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,fun('type.sum.sum'(X10,X11),bool)),X15),s(fun(X13,fun('type.sum.sum'(X10,X11),bool)),X16)),X11,X10))))))) | p(s(bool,happ(s(fun('type.sum.sum'(X10,X11),bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,fun('type.sum.sum'(X10,X11),bool)),X15),s(fun(X13,fun('type.sum.sum'(X10,X11),bool)),X16))),s('type.sum.sum'(X10,X11),X17)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X10,X11),bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,fun('type.sum.sum'(X10,X11),bool)),X15),s(fun(X13,fun('type.sum.sum'(X10,X11),bool)),X16))),s('type.sum.sum'(X10,X11),'const.sum.INL_1'(s(X10,sK11('const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,fun('type.sum.sum'(X10,X11),bool)),X15),s(fun(X13,fun('type.sum.sum'(X10,X11),bool)),X16)),X11,X10)))))))) )),
  inference(superposition,[],[f206,f227])).
fof(f2692,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun('type.sum.sum'(X5,X6),bool),X8),s('type.sum.sum'(X5,X6),'const.sum.INL_1'(s(X5,sK11('const.bool.COND_3'(s(bool,'F_0'),s(fun('type.sum.sum'(X5,X6),bool),X7),s(fun('type.sum.sum'(X5,X6),bool),X8)),X6,X5))))))) | p(s(bool,happ(s(fun('type.sum.sum'(X5,X6),bool),X8),s('type.sum.sum'(X5,X6),X9)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X5,X6),bool),X8),s('type.sum.sum'(X5,X6),'const.sum.INR_1'(s(X6,sK12('const.bool.COND_3'(s(bool,'F_0'),s(fun('type.sum.sum'(X5,X6),bool),X7),s(fun('type.sum.sum'(X5,X6),bool),X8)),X6,X5)))))))) )),
  inference(forward_demodulation,[],[f2691,f194])).
fof(f2691,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,happ(s(fun('type.sum.sum'(X5,X6),bool),X8),s('type.sum.sum'(X5,X6),X9)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X5,X6),bool),X8),s('type.sum.sum'(X5,X6),'const.sum.INR_1'(s(X6,sK12('const.bool.COND_3'(s(bool,'F_0'),s(fun('type.sum.sum'(X5,X6),bool),X7),s(fun('type.sum.sum'(X5,X6),bool),X8)),X6,X5))))))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X5,X6),bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun('type.sum.sum'(X5,X6),bool),X7),s(fun('type.sum.sum'(X5,X6),bool),X8))),s('type.sum.sum'(X5,X6),'const.sum.INL_1'(s(X5,sK11('const.bool.COND_3'(s(bool,'F_0'),s(fun('type.sum.sum'(X5,X6),bool),X7),s(fun('type.sum.sum'(X5,X6),bool),X8)),X6,X5)))))))) )),
  inference(forward_demodulation,[],[f2679,f194])).
fof(f2679,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun('type.sum.sum'(X5,X6),bool),X8),s('type.sum.sum'(X5,X6),'const.sum.INR_1'(s(X6,sK12('const.bool.COND_3'(s(bool,'F_0'),s(fun('type.sum.sum'(X5,X6),bool),X7),s(fun('type.sum.sum'(X5,X6),bool),X8)),X6,X5))))))) | p(s(bool,happ(s(fun('type.sum.sum'(X5,X6),bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun('type.sum.sum'(X5,X6),bool),X7),s(fun('type.sum.sum'(X5,X6),bool),X8))),s('type.sum.sum'(X5,X6),X9)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X5,X6),bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun('type.sum.sum'(X5,X6),bool),X7),s(fun('type.sum.sum'(X5,X6),bool),X8))),s('type.sum.sum'(X5,X6),'const.sum.INL_1'(s(X5,sK11('const.bool.COND_3'(s(bool,'F_0'),s(fun('type.sum.sum'(X5,X6),bool),X7),s(fun('type.sum.sum'(X5,X6),bool),X8)),X6,X5)))))))) )),
  inference(superposition,[],[f206,f194])).
fof(f2690,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK11('const.bool.COND_3'(s(bool,'T_0'),s(fun('type.sum.sum'(X0,X1),bool),X2),s(fun('type.sum.sum'(X0,X1),bool),X3)),X1,X0))))))) | p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),X4)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK12('const.bool.COND_3'(s(bool,'T_0'),s(fun('type.sum.sum'(X0,X1),bool),X2),s(fun('type.sum.sum'(X0,X1),bool),X3)),X1,X0)))))))) )),
  inference(forward_demodulation,[],[f2689,f193])).
fof(f2689,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),X4)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK12('const.bool.COND_3'(s(bool,'T_0'),s(fun('type.sum.sum'(X0,X1),bool),X2),s(fun('type.sum.sum'(X0,X1),bool),X3)),X1,X0))))))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun('type.sum.sum'(X0,X1),bool),X2),s(fun('type.sum.sum'(X0,X1),bool),X3))),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK11('const.bool.COND_3'(s(bool,'T_0'),s(fun('type.sum.sum'(X0,X1),bool),X2),s(fun('type.sum.sum'(X0,X1),bool),X3)),X1,X0)))))))) )),
  inference(forward_demodulation,[],[f2678,f193])).
fof(f2678,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK12('const.bool.COND_3'(s(bool,'T_0'),s(fun('type.sum.sum'(X0,X1),bool),X2),s(fun('type.sum.sum'(X0,X1),bool),X3)),X1,X0))))))) | p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun('type.sum.sum'(X0,X1),bool),X2),s(fun('type.sum.sum'(X0,X1),bool),X3))),s('type.sum.sum'(X0,X1),X4)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun('type.sum.sum'(X0,X1),bool),X2),s(fun('type.sum.sum'(X0,X1),bool),X3))),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK11('const.bool.COND_3'(s(bool,'T_0'),s(fun('type.sum.sum'(X0,X1),bool),X2),s(fun('type.sum.sum'(X0,X1),bool),X3)),X1,X0)))))))) )),
  inference(superposition,[],[f206,f193])).
fof(f2051,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s('type.sum.sum'(X5,X6),'const.sum.INL_1'(s(X5,X8))) != s('type.sum.sum'(X5,X6),'const.sum.INR_1'(s(X6,X9))) | s(X6,X7) = s(X6,X9)) )),
  inference(superposition,[],[f213,f194])).
fof(f2049,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) != s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X4),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))),s('type.sum.sum'(X0,X1),X2))) | s(X1,X3) = s(X1,X5) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK10(X2,X1,X0))))) )),
  inference(superposition,[],[f213,f205])).
fof(f2028,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) != s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X4),s('type.sum.sum'(X0,X1),X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X5))))) | s(X1,X3) = s(X1,sK10(X2,X1,X0)) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK9(X2,X1,X0))))) )),
  inference(superposition,[],[f213,f205])).
fof(f2027,plain,(
  ( ! [X127,X125,X123,X132,X130,X128,X126,X124,X131,X129] : (s('type.sum.sum'(X129,X123),'const.sum.INR_1'(s(X123,X130))) != s('type.sum.sum'(X129,X123),'const.bool.COND_3'(s(bool,X131),s('type.sum.sum'(X129,X123),'const.sum.INR_1'(s(X123,happ(s(fun(X125,X123),X126),s(X125,X128))))),s('type.sum.sum'(X129,X123),'const.sum.INL_1'(s(X129,X132))))) | s(X123,X130) = s(X123,happ(s(fun('type.sum.sum'(X124,X125),X123),sK16(X126,X127,X125,X123,X124)),s('type.sum.sum'(X124,X125),'const.sum.INR_1'(s(X125,X128)))))) )),
  inference(superposition,[],[f213,f226])).
fof(f2026,plain,(
  ( ! [X121,X118,X116,X114,X122,X120,X119,X117,X115,X113] : (s('type.sum.sum'(X119,X113),'const.sum.INR_1'(s(X113,X120))) != s('type.sum.sum'(X119,X113),'const.bool.COND_3'(s(bool,X121),s('type.sum.sum'(X119,X113),'const.sum.INR_1'(s(X113,happ(s(fun(X114,X113),X117),s(X114,X118))))),s('type.sum.sum'(X119,X113),'const.sum.INL_1'(s(X119,X122))))) | s(X113,X120) = s(X113,happ(s(fun('type.sum.sum'(X114,X115),X113),sK16(X116,X117,X115,X113,X114)),s('type.sum.sum'(X114,X115),'const.sum.INL_1'(s(X114,X118)))))) )),
  inference(superposition,[],[f213,f225])).
fof(f2025,plain,(
  ( ! [X103,X111,X109,X107,X105,X112,X110,X108,X106,X104] : (s('type.sum.sum'(X109,X103),'const.sum.INR_1'(s(X103,X110))) != s('type.sum.sum'(X109,X103),'const.bool.COND_3'(s(bool,X111),s('type.sum.sum'(X109,X103),'const.sum.INR_1'(s(X103,happ(s(fun(X104,X103),X107),s(X104,X106))))),s('type.sum.sum'(X109,X103),'const.sum.INL_1'(s(X109,X112))))) | s(X103,X110) = s(X103,'const.sum.sum_CASE_3'(s('type.sum.sum'(X104,X105),'const.sum.INL_1'(s(X104,X106))),s(fun(X104,X103),X107),s(fun(X105,X103),X108)))) )),
  inference(superposition,[],[f213,f228])).
fof(f2024,plain,(
  ( ! [X94,X101,X99,X97,X95,X93,X102,X100,X98,X96] : (s('type.sum.sum'(X99,X93),'const.sum.INR_1'(s(X93,X100))) != s('type.sum.sum'(X99,X93),'const.bool.COND_3'(s(bool,X101),s('type.sum.sum'(X99,X93),'const.sum.INR_1'(s(X93,happ(s(fun(X95,X93),X98),s(X95,X96))))),s('type.sum.sum'(X99,X93),'const.sum.INL_1'(s(X99,X102))))) | s(X93,X100) = s(X93,'const.sum.sum_CASE_3'(s('type.sum.sum'(X94,X95),'const.sum.INR_1'(s(X95,X96))),s(fun(X94,X93),X97),s(fun(X95,X93),X98)))) )),
  inference(superposition,[],[f213,f227])).
fof(f2023,plain,(
  ( ! [X92,X90,X88,X87,X91,X89,X86] : (s('type.sum.sum'(X89,X86),'const.sum.INR_1'(s(X86,X90))) != s('type.sum.sum'(X89,X86),'const.bool.COND_3'(s(bool,X91),s('type.sum.sum'(X89,X86),'const.sum.INR_1'(s(X86,X88))),s('type.sum.sum'(X89,X86),'const.sum.INL_1'(s(X89,X92))))) | s(X86,X90) = s(X86,'const.bool.COND_3'(s(bool,'F_0'),s(X86,X87),s(X86,X88)))) )),
  inference(superposition,[],[f213,f194])).
fof(f2022,plain,(
  ( ! [X80,X85,X83,X81,X79,X84,X82] : (s('type.sum.sum'(X82,X79),'const.sum.INR_1'(s(X79,X83))) != s('type.sum.sum'(X82,X79),'const.bool.COND_3'(s(bool,X84),s('type.sum.sum'(X82,X79),'const.sum.INR_1'(s(X79,X80))),s('type.sum.sum'(X82,X79),'const.sum.INL_1'(s(X82,X85))))) | s(X79,X83) = s(X79,'const.bool.COND_3'(s(bool,'T_0'),s(X79,X80),s(X79,X81)))) )),
  inference(superposition,[],[f213,f193])).
fof(f2020,plain,(
  ( ! [X70,X68,X66,X64,X72,X71,X69,X67,X65] : (s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,X70))) != s('type.sum.sum'(X69,bool),'const.bool.COND_3'(s(bool,X71),s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X64,bool),X65),s(X64,X68))))),s('type.sum.sum'(X69,bool),'const.sum.INL_1'(s(X69,X72))))) | s(bool,X70) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X64,bool),X65),s(fun(X66,bool),X67),s('type.sum.sum'(X64,X66),'const.sum.INL_1'(s(X64,X68)))))) )),
  inference(superposition,[],[f213,f211])).
fof(f2019,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X55,X63] : (s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,X61))) != s('type.sum.sum'(X60,bool),'const.bool.COND_3'(s(bool,X62),s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X57,bool),X58),s(X57,X59))))),s('type.sum.sum'(X60,bool),'const.sum.INL_1'(s(X60,X63))))) | s(bool,X61) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X55,bool),X56),s(fun(X57,bool),X58),s('type.sum.sum'(X55,X57),'const.sum.INR_1'(s(X57,X59)))))) )),
  inference(superposition,[],[f213,f210])).
fof(f2013,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s('type.sum.sum'(X29,bool),'const.sum.INR_1'(s(bool,X30))) != s('type.sum.sum'(X29,bool),'const.bool.COND_3'(s(bool,X31),s('type.sum.sum'(X29,bool),'const.sum.INR_1'(s(bool,X28))),s('type.sum.sum'(X29,bool),'const.sum.INL_1'(s(X29,X32))))) | s(bool,'F_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f213,f163])).
fof(f2010,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X16] : (s('type.sum.sum'(X14,'type.sum.sum'(X11,X12)),'const.sum.INR_1'(s('type.sum.sum'(X11,X12),X15))) != s('type.sum.sum'(X14,'type.sum.sum'(X11,X12)),'const.bool.COND_3'(s(bool,X16),s('type.sum.sum'(X14,'type.sum.sum'(X11,X12)),'const.sum.INR_1'(s('type.sum.sum'(X11,X12),X13))),s('type.sum.sum'(X14,'type.sum.sum'(X11,X12)),'const.sum.INL_1'(s(X14,X17))))) | s('type.sum.sum'(X11,X12),X15) = s('type.sum.sum'(X11,X12),'const.sum.INL_1'(s(X11,sK9(X13,X12,X11)))) | s('type.sum.sum'(X11,X12),X13) = s('type.sum.sum'(X11,X12),'const.sum.INR_1'(s(X12,sK10(X13,X12,X11))))) )),
  inference(superposition,[],[f213,f205])).
fof(f2009,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (s('type.sum.sum'(X7,'type.sum.sum'(X4,X5)),'const.sum.INR_1'(s('type.sum.sum'(X4,X5),X8))) != s('type.sum.sum'(X7,'type.sum.sum'(X4,X5)),'const.bool.COND_3'(s(bool,X9),s('type.sum.sum'(X7,'type.sum.sum'(X4,X5)),'const.sum.INR_1'(s('type.sum.sum'(X4,X5),X6))),s('type.sum.sum'(X7,'type.sum.sum'(X4,X5)),'const.sum.INL_1'(s(X7,X10))))) | s('type.sum.sum'(X4,X5),X8) = s('type.sum.sum'(X4,X5),'const.sum.INR_1'(s(X5,sK10(X6,X5,X4)))) | s('type.sum.sum'(X4,X5),X6) = s('type.sum.sum'(X4,X5),'const.sum.INL_1'(s(X4,sK9(X6,X5,X4))))) )),
  inference(superposition,[],[f213,f205])).
fof(f1990,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(X0,X1),X2) != s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X3),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X4))),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X5))))) | s(X1,X4) = s(X1,sK10(X2,X1,X0)) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK9(X2,X1,X0))))) )),
  inference(superposition,[],[f213,f205])).
fof(f1989,plain,(
  ( ! [X127,X125,X123,X132,X130,X128,X126,X124,X131,X129] : (s('type.sum.sum'(X129,X123),'const.sum.INR_1'(s(X123,happ(s(fun(X125,X123),X126),s(X125,X128))))) != s('type.sum.sum'(X129,X123),'const.bool.COND_3'(s(bool,X130),s('type.sum.sum'(X129,X123),'const.sum.INR_1'(s(X123,X131))),s('type.sum.sum'(X129,X123),'const.sum.INL_1'(s(X129,X132))))) | s(X123,X131) = s(X123,happ(s(fun('type.sum.sum'(X124,X125),X123),sK16(X126,X127,X125,X123,X124)),s('type.sum.sum'(X124,X125),'const.sum.INR_1'(s(X125,X128)))))) )),
  inference(superposition,[],[f213,f226])).
fof(f1988,plain,(
  ( ! [X121,X118,X116,X114,X122,X120,X119,X117,X115,X113] : (s('type.sum.sum'(X119,X113),'const.sum.INR_1'(s(X113,happ(s(fun(X114,X113),X117),s(X114,X118))))) != s('type.sum.sum'(X119,X113),'const.bool.COND_3'(s(bool,X120),s('type.sum.sum'(X119,X113),'const.sum.INR_1'(s(X113,X121))),s('type.sum.sum'(X119,X113),'const.sum.INL_1'(s(X119,X122))))) | s(X113,X121) = s(X113,happ(s(fun('type.sum.sum'(X114,X115),X113),sK16(X116,X117,X115,X113,X114)),s('type.sum.sum'(X114,X115),'const.sum.INL_1'(s(X114,X118)))))) )),
  inference(superposition,[],[f213,f225])).
fof(f1987,plain,(
  ( ! [X103,X111,X109,X107,X105,X112,X110,X108,X106,X104] : (s('type.sum.sum'(X109,X103),'const.sum.INR_1'(s(X103,happ(s(fun(X104,X103),X107),s(X104,X106))))) != s('type.sum.sum'(X109,X103),'const.bool.COND_3'(s(bool,X110),s('type.sum.sum'(X109,X103),'const.sum.INR_1'(s(X103,X111))),s('type.sum.sum'(X109,X103),'const.sum.INL_1'(s(X109,X112))))) | s(X103,X111) = s(X103,'const.sum.sum_CASE_3'(s('type.sum.sum'(X104,X105),'const.sum.INL_1'(s(X104,X106))),s(fun(X104,X103),X107),s(fun(X105,X103),X108)))) )),
  inference(superposition,[],[f213,f228])).
fof(f1986,plain,(
  ( ! [X94,X101,X99,X97,X95,X93,X102,X100,X98,X96] : (s('type.sum.sum'(X99,X93),'const.sum.INR_1'(s(X93,happ(s(fun(X95,X93),X98),s(X95,X96))))) != s('type.sum.sum'(X99,X93),'const.bool.COND_3'(s(bool,X100),s('type.sum.sum'(X99,X93),'const.sum.INR_1'(s(X93,X101))),s('type.sum.sum'(X99,X93),'const.sum.INL_1'(s(X99,X102))))) | s(X93,X101) = s(X93,'const.sum.sum_CASE_3'(s('type.sum.sum'(X94,X95),'const.sum.INR_1'(s(X95,X96))),s(fun(X94,X93),X97),s(fun(X95,X93),X98)))) )),
  inference(superposition,[],[f213,f227])).
fof(f1985,plain,(
  ( ! [X92,X90,X88,X87,X91,X89,X86] : (s('type.sum.sum'(X89,X86),'const.sum.INR_1'(s(X86,X88))) != s('type.sum.sum'(X89,X86),'const.bool.COND_3'(s(bool,X90),s('type.sum.sum'(X89,X86),'const.sum.INR_1'(s(X86,X91))),s('type.sum.sum'(X89,X86),'const.sum.INL_1'(s(X89,X92))))) | s(X86,X91) = s(X86,'const.bool.COND_3'(s(bool,'F_0'),s(X86,X87),s(X86,X88)))) )),
  inference(superposition,[],[f213,f194])).
fof(f1984,plain,(
  ( ! [X80,X85,X83,X81,X79,X84,X82] : (s('type.sum.sum'(X82,X79),'const.sum.INR_1'(s(X79,X80))) != s('type.sum.sum'(X82,X79),'const.bool.COND_3'(s(bool,X83),s('type.sum.sum'(X82,X79),'const.sum.INR_1'(s(X79,X84))),s('type.sum.sum'(X82,X79),'const.sum.INL_1'(s(X82,X85))))) | s(X79,X84) = s(X79,'const.bool.COND_3'(s(bool,'T_0'),s(X79,X80),s(X79,X81)))) )),
  inference(superposition,[],[f213,f193])).
fof(f1982,plain,(
  ( ! [X70,X68,X66,X64,X72,X71,X69,X67,X65] : (s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X64,bool),X65),s(X64,X68))))) != s('type.sum.sum'(X69,bool),'const.bool.COND_3'(s(bool,X70),s('type.sum.sum'(X69,bool),'const.sum.INR_1'(s(bool,X71))),s('type.sum.sum'(X69,bool),'const.sum.INL_1'(s(X69,X72))))) | s(bool,X71) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X64,bool),X65),s(fun(X66,bool),X67),s('type.sum.sum'(X64,X66),'const.sum.INL_1'(s(X64,X68)))))) )),
  inference(superposition,[],[f213,f211])).
fof(f1981,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56,X55,X63] : (s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X57,bool),X58),s(X57,X59))))) != s('type.sum.sum'(X60,bool),'const.bool.COND_3'(s(bool,X61),s('type.sum.sum'(X60,bool),'const.sum.INR_1'(s(bool,X62))),s('type.sum.sum'(X60,bool),'const.sum.INL_1'(s(X60,X63))))) | s(bool,X62) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X55,bool),X56),s(fun(X57,bool),X58),s('type.sum.sum'(X55,X57),'const.sum.INR_1'(s(X57,X59)))))) )),
  inference(superposition,[],[f213,f210])).
fof(f1977,plain,(
  ( ! [X14,X12,X17,X15,X13,X11,X16] : (s('type.sum.sum'(X14,'type.sum.sum'(X11,X12)),'const.sum.INR_1'(s('type.sum.sum'(X11,X12),X13))) != s('type.sum.sum'(X14,'type.sum.sum'(X11,X12)),'const.bool.COND_3'(s(bool,X15),s('type.sum.sum'(X14,'type.sum.sum'(X11,X12)),'const.sum.INR_1'(s('type.sum.sum'(X11,X12),X16))),s('type.sum.sum'(X14,'type.sum.sum'(X11,X12)),'const.sum.INL_1'(s(X14,X17))))) | s('type.sum.sum'(X11,X12),X16) = s('type.sum.sum'(X11,X12),'const.sum.INL_1'(s(X11,sK9(X13,X12,X11)))) | s('type.sum.sum'(X11,X12),X13) = s('type.sum.sum'(X11,X12),'const.sum.INR_1'(s(X12,sK10(X13,X12,X11))))) )),
  inference(superposition,[],[f213,f205])).
fof(f1976,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X9] : (s('type.sum.sum'(X7,'type.sum.sum'(X4,X5)),'const.sum.INR_1'(s('type.sum.sum'(X4,X5),X6))) != s('type.sum.sum'(X7,'type.sum.sum'(X4,X5)),'const.bool.COND_3'(s(bool,X8),s('type.sum.sum'(X7,'type.sum.sum'(X4,X5)),'const.sum.INR_1'(s('type.sum.sum'(X4,X5),X9))),s('type.sum.sum'(X7,'type.sum.sum'(X4,X5)),'const.sum.INL_1'(s(X7,X10))))) | s('type.sum.sum'(X4,X5),X9) = s('type.sum.sum'(X4,X5),'const.sum.INR_1'(s(X5,sK10(X6,X5,X4)))) | s('type.sum.sum'(X4,X5),X6) = s('type.sum.sum'(X4,X5),'const.sum.INL_1'(s(X4,sK9(X6,X5,X4))))) )),
  inference(superposition,[],[f213,f205])).
fof(f1857,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) != s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,'F_0'),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4)))))) )),
  inference(unit_resulting_resolution,[],[f262,f212])).
fof(f1680,plain,(
  ( ! [X292,X294,X290,X296,X293,X295,X291] : (p(s(bool,'|dollar|exists_1'(s(fun(X290,bool),happ(s(fun(X292,fun(X290,bool)),X293),s(X292,X295)))))) | s(X290,X296) = s(X290,sK8(X296,happ(s(fun('type.sum.sum'(X291,X292),fun(X290,bool)),sK16(X293,X294,X292,fun(X290,bool),X291)),s('type.sum.sum'(X291,X292),'const.sum.INR_1'(s(X292,X295)))),X290))) )),
  inference(superposition,[],[f195,f226])).
fof(f1709,plain,(
  ( ! [X288,X285,X287,X283,X289,X284,X286] : (p(s(bool,'|dollar|exists_1'(s(fun(X283,bool),happ(s(fun(X285,fun(X283,bool)),X286),s(X285,X288)))))) | ~p(s(bool,happ(s(fun(X283,bool),happ(s(fun(X285,fun(X283,bool)),X286),s(X285,X288))),s(X283,sK8(X289,happ(s(fun('type.sum.sum'(X284,X285),fun(X283,bool)),sK16(X286,X287,X285,fun(X283,bool),X284)),s('type.sum.sum'(X284,X285),'const.sum.INR_1'(s(X285,X288)))),X283)))))) )),
  inference(forward_demodulation,[],[f1679,f226])).
fof(f1679,plain,(
  ( ! [X288,X285,X287,X283,X289,X284,X286] : (~p(s(bool,happ(s(fun(X283,bool),happ(s(fun(X285,fun(X283,bool)),X286),s(X285,X288))),s(X283,sK8(X289,happ(s(fun('type.sum.sum'(X284,X285),fun(X283,bool)),sK16(X286,X287,X285,fun(X283,bool),X284)),s('type.sum.sum'(X284,X285),'const.sum.INR_1'(s(X285,X288)))),X283))))) | p(s(bool,'|dollar|exists_1'(s(fun(X283,bool),happ(s(fun('type.sum.sum'(X284,X285),fun(X283,bool)),sK16(X286,X287,X285,fun(X283,bool),X284)),s('type.sum.sum'(X284,X285),'const.sum.INR_1'(s(X285,X288))))))))) )),
  inference(superposition,[],[f196,f226])).
fof(f1641,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (s('type.sum.sum'(X34,X28),'const.sum.INR_1'(s(X28,X35))) != s('type.sum.sum'(X34,X28),'const.sum.INR_1'(s(X28,happ(s(fun(X30,X28),X31),s(X30,X33))))) | s(X28,X35) = s(X28,happ(s(fun('type.sum.sum'(X29,X30),X28),sK16(X31,X32,X30,X28,X29)),s('type.sum.sum'(X29,X30),'const.sum.INR_1'(s(X30,X33)))))) )),
  inference(superposition,[],[f208,f226])).
fof(f1634,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X3,happ(s(fun(X1,X3),X4),s(X1,sK10(X2,X1,X0)))) = s(X3,happ(s(fun('type.sum.sum'(X0,X1),X3),sK16(X4,X5,X1,X3,X0)),s('type.sum.sum'(X0,X1),X2))) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK9(X2,X1,X0))))) )),
  inference(superposition,[],[f226,f205])).
fof(f1515,plain,(
  ( ! [X269,X271,X273,X268,X270,X272,X274] : (p(s(bool,'|dollar|exists_1'(s(fun(X268,bool),happ(s(fun(X269,fun(X268,bool)),X272),s(X269,X273)))))) | s(X268,X274) = s(X268,sK8(X274,happ(s(fun('type.sum.sum'(X269,X270),fun(X268,bool)),sK16(X271,X272,X270,fun(X268,bool),X269)),s('type.sum.sum'(X269,X270),'const.sum.INL_1'(s(X269,X273)))),X268))) )),
  inference(superposition,[],[f195,f225])).
fof(f1542,plain,(
  ( ! [X261,X263,X265,X267,X262,X264,X266] : (p(s(bool,'|dollar|exists_1'(s(fun(X261,bool),happ(s(fun(X262,fun(X261,bool)),X265),s(X262,X266)))))) | ~p(s(bool,happ(s(fun(X261,bool),happ(s(fun(X262,fun(X261,bool)),X265),s(X262,X266))),s(X261,sK8(X267,happ(s(fun('type.sum.sum'(X262,X263),fun(X261,bool)),sK16(X264,X265,X263,fun(X261,bool),X262)),s('type.sum.sum'(X262,X263),'const.sum.INL_1'(s(X262,X266)))),X261)))))) )),
  inference(forward_demodulation,[],[f1514,f225])).
fof(f1514,plain,(
  ( ! [X261,X263,X265,X267,X262,X264,X266] : (~p(s(bool,happ(s(fun(X261,bool),happ(s(fun(X262,fun(X261,bool)),X265),s(X262,X266))),s(X261,sK8(X267,happ(s(fun('type.sum.sum'(X262,X263),fun(X261,bool)),sK16(X264,X265,X263,fun(X261,bool),X262)),s('type.sum.sum'(X262,X263),'const.sum.INL_1'(s(X262,X266)))),X261))))) | p(s(bool,'|dollar|exists_1'(s(fun(X261,bool),happ(s(fun('type.sum.sum'(X262,X263),fun(X261,bool)),sK16(X264,X265,X263,fun(X261,bool),X262)),s('type.sum.sum'(X262,X263),'const.sum.INL_1'(s(X262,X266))))))))) )),
  inference(superposition,[],[f196,f225])).
fof(f1479,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (s('type.sum.sum'(X34,X28),'const.sum.INR_1'(s(X28,X35))) != s('type.sum.sum'(X34,X28),'const.sum.INR_1'(s(X28,happ(s(fun(X29,X28),X32),s(X29,X33))))) | s(X28,X35) = s(X28,happ(s(fun('type.sum.sum'(X29,X30),X28),sK16(X31,X32,X30,X28,X29)),s('type.sum.sum'(X29,X30),'const.sum.INL_1'(s(X29,X33)))))) )),
  inference(superposition,[],[f208,f225])).
fof(f1472,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X3,happ(s(fun(X0,X3),X4),s(X0,sK9(X2,X1,X0)))) = s(X3,happ(s(fun('type.sum.sum'(X0,X1),X3),sK16(X5,X4,X1,X3,X0)),s('type.sum.sum'(X0,X1),X2))) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK10(X2,X1,X0))))) )),
  inference(superposition,[],[f225,f205])).
fof(f1339,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s('type.sum.sum'(X38,X39),X40) != s('type.sum.sum'(X38,X39),X41) | s('type.sum.sum'(X42,'type.sum.sum'(X38,X39)),'const.sum.INR_1'(s('type.sum.sum'(X38,X39),X41))) = s('type.sum.sum'(X42,'type.sum.sum'(X38,X39)),'const.sum.INR_1'(s('type.sum.sum'(X38,X39),'const.sum.INL_1'(s(X38,sK9(X40,X39,X38)))))) | s('type.sum.sum'(X38,X39),X40) = s('type.sum.sum'(X38,X39),'const.sum.INR_1'(s(X39,sK10(X40,X39,X38))))) )),
  inference(superposition,[],[f209,f205])).
fof(f1337,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s('type.sum.sum'(X31,'type.sum.sum'(X28,X29)),'const.sum.INR_1'(s('type.sum.sum'(X28,X29),X30))) != s('type.sum.sum'(X31,'type.sum.sum'(X28,X29)),'const.sum.INR_1'(s('type.sum.sum'(X28,X29),X32))) | s('type.sum.sum'(X28,X29),X32) = s('type.sum.sum'(X28,X29),'const.sum.INL_1'(s(X28,sK9(X30,X29,X28)))) | s('type.sum.sum'(X28,X29),X30) = s('type.sum.sum'(X28,X29),'const.sum.INR_1'(s(X29,sK10(X30,X29,X28))))) )),
  inference(superposition,[],[f208,f205])).
fof(f1332,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(bool,happ(s(fun(X7,bool),X10),s(X7,sK9(X9,X8,X7)))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X7,bool),X10),s(fun(X8,bool),X11),s('type.sum.sum'(X7,X8),X9))) | s('type.sum.sum'(X7,X8),X9) = s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,sK10(X9,X8,X7))))) )),
  inference(superposition,[],[f211,f205])).
fof(f1331,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X4,happ(s(fun(X1,X4),X5),s(X1,sK9(X3,X2,X1)))) = s(X4,'const.sum.sum_CASE_3'(s('type.sum.sum'(X1,X2),X3),s(fun(X1,X4),X5),s(fun(X2,X4),X6))) | s('type.sum.sum'(X1,X2),X3) = s('type.sum.sum'(X1,X2),'const.sum.INR_1'(s(X2,sK10(X3,X2,X1))))) )),
  inference(superposition,[],[f228,f205])).
fof(f1330,plain,(
  ( ! [X0] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X0) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(X0,sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(X0,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),X0) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(X0,sK1,sK0))))) )),
  inference(superposition,[],[f259,f205])).
fof(f1324,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s('type.sum.sum'(X46,X47),X48) != s('type.sum.sum'(X46,X47),X49) | s('type.sum.sum'(X50,'type.sum.sum'(X46,X47)),'const.sum.INR_1'(s('type.sum.sum'(X46,X47),X49))) = s('type.sum.sum'(X50,'type.sum.sum'(X46,X47)),'const.sum.INR_1'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,sK10(X48,X47,X46)))))) | s('type.sum.sum'(X46,X47),X48) = s('type.sum.sum'(X46,X47),'const.sum.INL_1'(s(X46,sK9(X48,X47,X46))))) )),
  inference(superposition,[],[f209,f205])).
fof(f1322,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s('type.sum.sum'(X39,'type.sum.sum'(X36,X37)),'const.sum.INR_1'(s('type.sum.sum'(X36,X37),X38))) != s('type.sum.sum'(X39,'type.sum.sum'(X36,X37)),'const.sum.INR_1'(s('type.sum.sum'(X36,X37),X40))) | s('type.sum.sum'(X36,X37),X40) = s('type.sum.sum'(X36,X37),'const.sum.INR_1'(s(X37,sK10(X38,X37,X36)))) | s('type.sum.sum'(X36,X37),X38) = s('type.sum.sum'(X36,X37),'const.sum.INL_1'(s(X36,sK9(X38,X37,X36))))) )),
  inference(superposition,[],[f208,f205])).
fof(f1316,plain,(
  ( ! [X14,X12,X15,X13] : (s('type.sum.sum'(X12,X13),X14) != s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X15))) | s(X13,X15) = s(X13,sK10(X14,X13,X12)) | s('type.sum.sum'(X12,X13),X14) = s('type.sum.sum'(X12,X13),'const.sum.INL_1'(s(X12,sK9(X14,X13,X12))))) )),
  inference(superposition,[],[f208,f205])).
fof(f1315,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(bool,happ(s(fun(X8,bool),X10),s(X8,sK10(X9,X8,X7)))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X7,bool),X11),s(fun(X8,bool),X10),s('type.sum.sum'(X7,X8),X9))) | s('type.sum.sum'(X7,X8),X9) = s('type.sum.sum'(X7,X8),'const.sum.INL_1'(s(X7,sK9(X9,X8,X7))))) )),
  inference(superposition,[],[f210,f205])).
fof(f1314,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X4,happ(s(fun(X2,X4),X5),s(X2,sK10(X3,X2,X1)))) = s(X4,'const.sum.sum_CASE_3'(s('type.sum.sum'(X1,X2),X3),s(fun(X1,X4),X6),s(fun(X2,X4),X5))) | s('type.sum.sum'(X1,X2),X3) = s('type.sum.sum'(X1,X2),'const.sum.INL_1'(s(X1,sK9(X3,X2,X1))))) )),
  inference(superposition,[],[f227,f205])).
fof(f1313,plain,(
  ( ! [X0] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X0) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(X0,sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(X0,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),X0) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(X0,sK1,sK0))))) )),
  inference(superposition,[],[f260,f205])).
fof(f1219,plain,(
  ( ! [X231,X229,X233,X230,X228,X234,X232] : (p(s(bool,'|dollar|exists_1'(s(fun(X228,bool),happ(s(fun(X229,fun(X228,bool)),X232),s(X229,X231)))))) | s(X228,X234) = s(X228,sK8(X234,'const.sum.sum_CASE_3'(s('type.sum.sum'(X229,X230),'const.sum.INL_1'(s(X229,X231))),s(fun(X229,fun(X228,bool)),X232),s(fun(X230,fun(X228,bool)),X233)),X228))) )),
  inference(superposition,[],[f195,f228])).
fof(f1241,plain,(
  ( ! [X222,X227,X225,X223,X221,X226,X224] : (p(s(bool,'|dollar|exists_1'(s(fun(X221,bool),happ(s(fun(X222,fun(X221,bool)),X225),s(X222,X224)))))) | ~p(s(bool,happ(s(fun(X221,bool),happ(s(fun(X222,fun(X221,bool)),X225),s(X222,X224))),s(X221,sK8(X227,'const.sum.sum_CASE_3'(s('type.sum.sum'(X222,X223),'const.sum.INL_1'(s(X222,X224))),s(fun(X222,fun(X221,bool)),X225),s(fun(X223,fun(X221,bool)),X226)),X221)))))) )),
  inference(forward_demodulation,[],[f1218,f228])).
fof(f1218,plain,(
  ( ! [X222,X227,X225,X223,X221,X226,X224] : (~p(s(bool,happ(s(fun(X221,bool),happ(s(fun(X222,fun(X221,bool)),X225),s(X222,X224))),s(X221,sK8(X227,'const.sum.sum_CASE_3'(s('type.sum.sum'(X222,X223),'const.sum.INL_1'(s(X222,X224))),s(fun(X222,fun(X221,bool)),X225),s(fun(X223,fun(X221,bool)),X226)),X221))))) | p(s(bool,'|dollar|exists_1'(s(fun(X221,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X222,X223),'const.sum.INL_1'(s(X222,X224))),s(fun(X222,fun(X221,bool)),X225),s(fun(X223,fun(X221,bool)),X226))))))) )),
  inference(superposition,[],[f196,f228])).
fof(f1189,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (s('type.sum.sum'(X34,X28),'const.sum.INR_1'(s(X28,X35))) != s('type.sum.sum'(X34,X28),'const.sum.INR_1'(s(X28,happ(s(fun(X29,X28),X32),s(X29,X31))))) | s(X28,X35) = s(X28,'const.sum.sum_CASE_3'(s('type.sum.sum'(X29,X30),'const.sum.INL_1'(s(X29,X31))),s(fun(X29,X28),X32),s(fun(X30,X28),X33)))) )),
  inference(superposition,[],[f208,f228])).
fof(f1060,plain,(
  ( ! [X212,X210,X208,X213,X211,X209,X207] : (p(s(bool,'|dollar|exists_1'(s(fun(X207,bool),happ(s(fun(X209,fun(X207,bool)),X212),s(X209,X210)))))) | s(X207,X213) = s(X207,sK8(X213,'const.sum.sum_CASE_3'(s('type.sum.sum'(X208,X209),'const.sum.INR_1'(s(X209,X210))),s(fun(X208,fun(X207,bool)),X211),s(fun(X209,fun(X207,bool)),X212)),X207))) )),
  inference(superposition,[],[f195,f227])).
fof(f1079,plain,(
  ( ! [X206,X204,X202,X200,X205,X203,X201] : (p(s(bool,'|dollar|exists_1'(s(fun(X200,bool),happ(s(fun(X202,fun(X200,bool)),X205),s(X202,X203)))))) | ~p(s(bool,happ(s(fun(X200,bool),happ(s(fun(X202,fun(X200,bool)),X205),s(X202,X203))),s(X200,sK8(X206,'const.sum.sum_CASE_3'(s('type.sum.sum'(X201,X202),'const.sum.INR_1'(s(X202,X203))),s(fun(X201,fun(X200,bool)),X204),s(fun(X202,fun(X200,bool)),X205)),X200)))))) )),
  inference(forward_demodulation,[],[f1059,f227])).
fof(f1059,plain,(
  ( ! [X206,X204,X202,X200,X205,X203,X201] : (~p(s(bool,happ(s(fun(X200,bool),happ(s(fun(X202,fun(X200,bool)),X205),s(X202,X203))),s(X200,sK8(X206,'const.sum.sum_CASE_3'(s('type.sum.sum'(X201,X202),'const.sum.INR_1'(s(X202,X203))),s(fun(X201,fun(X200,bool)),X204),s(fun(X202,fun(X200,bool)),X205)),X200))))) | p(s(bool,'|dollar|exists_1'(s(fun(X200,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X201,X202),'const.sum.INR_1'(s(X202,X203))),s(fun(X201,fun(X200,bool)),X204),s(fun(X202,fun(X200,bool)),X205))))))) )),
  inference(superposition,[],[f196,f227])).
fof(f1033,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (s('type.sum.sum'(X34,X28),'const.sum.INR_1'(s(X28,X35))) != s('type.sum.sum'(X34,X28),'const.sum.INR_1'(s(X28,happ(s(fun(X30,X28),X33),s(X30,X31))))) | s(X28,X35) = s(X28,'const.sum.sum_CASE_3'(s('type.sum.sum'(X29,X30),'const.sum.INR_1'(s(X30,X31))),s(fun(X29,X28),X32),s(fun(X30,X28),X33)))) )),
  inference(superposition,[],[f208,f227])).
fof(f890,plain,(
  ( ! [X103,X101,X105,X102,X100,X106,X104] : (s('type.sum.sum'(X105,bool),'const.sum.INR_1'(s(bool,X106))) != s('type.sum.sum'(X105,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X100,bool),X101),s(X100,X104))))) | s(bool,X106) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X100,bool),X101),s(fun(X102,bool),X103),s('type.sum.sum'(X100,X102),'const.sum.INL_1'(s(X100,X104)))))) )),
  inference(superposition,[],[f208,f211])).
fof(f769,plain,(
  ( ! [X101,X99,X97,X95,X100,X98,X96] : (s('type.sum.sum'(X100,bool),'const.sum.INR_1'(s(bool,X101))) != s('type.sum.sum'(X100,bool),'const.sum.INR_1'(s(bool,happ(s(fun(X97,bool),X98),s(X97,X99))))) | s(bool,X101) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X95,bool),X96),s(fun(X97,bool),X98),s('type.sum.sum'(X95,X97),'const.sum.INR_1'(s(X97,X99)))))) )),
  inference(superposition,[],[f208,f210])).
fof(f637,plain,(
  ( ! [X10,X8,X9] : (s(bool,'F_0') != s(bool,X9) | s('type.sum.sum'(X10,bool),'const.sum.INR_1'(s(bool,X8))) = s('type.sum.sum'(X10,bool),'const.sum.INR_1'(s(bool,X9))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f209,f163])).
fof(f635,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,'F_0'))) = s('type.sum.sum'(X4,bool),'const.sum.INR_1'(s(bool,X3))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f209,f163])).
fof(f527,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s('type.sum.sum'(X20,X17),'const.sum.INR_1'(s(X17,X19))) != s('type.sum.sum'(X20,X17),'const.sum.INR_1'(s(X17,X21))) | s(X17,X21) = s(X17,'const.bool.COND_3'(s(bool,'F_0'),s(X17,X18),s(X17,X19)))) )),
  inference(superposition,[],[f208,f194])).
fof(f526,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.sum.sum'(X15,X12),'const.sum.INR_1'(s(X12,X13))) != s('type.sum.sum'(X15,X12),'const.sum.INR_1'(s(X12,X16))) | s(X12,X16) = s(X12,'const.bool.COND_3'(s(bool,'T_0'),s(X12,X13),s(X12,X14)))) )),
  inference(superposition,[],[f208,f193])).
fof(f523,plain,(
  ( ! [X6,X7,X5] : (s('type.sum.sum'(X6,bool),'const.sum.INR_1'(s(bool,'F_0'))) != s('type.sum.sum'(X6,bool),'const.sum.INR_1'(s(bool,X7))) | s(bool,X5) = s(bool,X7) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f208,f163])).
fof(f522,plain,(
  ( ! [X4,X2,X3] : (s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,X2))) != s('type.sum.sum'(X3,bool),'const.sum.INR_1'(s(bool,X4))) | s(bool,'F_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f208,f163])).
fof(f458,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,sK8(X0,X1,bool))) )),
  inference(superposition,[],[f196,f163])).
fof(f461,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X6)))) | ~p(s(bool,happ(s(fun(X4,bool),X6),s(X4,sK8(X7,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X4,bool),X5),s(fun(X4,bool),X6)),X4)))))) )),
  inference(forward_demodulation,[],[f457,f194])).
fof(f457,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,happ(s(fun(X4,bool),X6),s(X4,sK8(X7,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X4,bool),X5),s(fun(X4,bool),X6)),X4))))) | p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X4,bool),X5),s(fun(X4,bool),X6))))))) )),
  inference(superposition,[],[f196,f194])).
fof(f460,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X3,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X0)))))) )),
  inference(forward_demodulation,[],[f456,f193])).
fof(f456,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X3,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(superposition,[],[f196,f193])).
fof(f425,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,sK8(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(subsumption_resolution,[],[f424,f262])).
fof(f424,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | s(X0,X2) = s(X0,sK8(X2,X1,X0)) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f195,f163])).
fof(f423,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'|dollar|exists_1'(s(fun(X4,bool),X6)))) | s(X4,X7) = s(X4,sK8(X7,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X4,bool),X5),s(fun(X4,bool),X6)),X4))) )),
  inference(superposition,[],[f195,f194])).
fof(f422,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X3) = s(X0,sK8(X3,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X0,bool),X1),s(fun(X0,bool),X2)),X0))) )),
  inference(superposition,[],[f195,f193])).
fof(f363,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f194,f163])).
fof(f351,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f194,f163])).
fof(f325,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f193,f163])).
fof(f306,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f303,f262])).
fof(f303,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f201,f163])).
fof(f274,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f169,f163])).
fof(f272,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f163,f163])).
fof(f263,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f262,f165])).
fof(f229,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s('type.sum.sum'(X0,X1),X5) != s('type.sum.sum'(X0,X1),X6) | s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK18(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK17(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X5),s(fun(X0,X2),X7),s(fun(X1,X2),X8))) = s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X6),s(fun(X0,X2),X4),s(fun(X1,X2),X3)))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.sum.sum'(X0,X1),X5) != s('type.sum.sum'(X0,X1),X6) | (s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK17(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) & s(X2,happ(s(fun(X0,X2),X4),s(X0,sK17(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X2,happ(s(fun(X0,X2),X7),s(X0,sK17(X8,X7,X6,X5,X4,X3,X2,X1,X0))))) | (s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK18(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) & s(X2,happ(s(fun(X1,X2),X3),s(X1,sK18(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X2,happ(s(fun(X1,X2),X8),s(X1,sK18(X8,X7,X6,X5,X4,X3,X2,X1,X0))))) | s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X5),s(fun(X0,X2),X7),s(fun(X1,X2),X8))) = s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X6),s(fun(X0,X2),X4),s(fun(X1,X2),X3))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17,sK18])],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.sum.sum'(X0,X1),X5) != s('type.sum.sum'(X0,X1),X6) | ? [X10] : (s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X10))) & s(X2,happ(s(fun(X0,X2),X4),s(X0,X10))) != s(X2,happ(s(fun(X0,X2),X7),s(X0,X10)))) | ? [X9] : (s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X9))) & s(X2,happ(s(fun(X1,X2),X3),s(X1,X9))) != s(X2,happ(s(fun(X1,X2),X8),s(X1,X9)))) | s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X5),s(fun(X0,X2),X7),s(fun(X1,X2),X8))) = s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X6),s(fun(X0,X2),X4),s(fun(X1,X2),X3))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.sum.sum'(X0,X1),X5) != s('type.sum.sum'(X0,X1),X6) | ? [X10] : (s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X10))) & s(X2,happ(s(fun(X0,X2),X4),s(X0,X10))) != s(X2,happ(s(fun(X0,X2),X7),s(X0,X10)))) | ? [X9] : (s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X9))) & s(X2,happ(s(fun(X1,X2),X3),s(X1,X9))) != s(X2,happ(s(fun(X1,X2),X8),s(X1,X9))))) | s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X5),s(fun(X0,X2),X7),s(fun(X1,X2),X8))) = s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X6),s(fun(X0,X2),X4),s(fun(X1,X2),X3))))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.sum.sum'(X0,X1),X5) = s('type.sum.sum'(X0,X1),X6) & ! [X10] : (s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X10))) => s(X2,happ(s(fun(X0,X2),X4),s(X0,X10))) = s(X2,happ(s(fun(X0,X2),X7),s(X0,X10)))) & ! [X9] : (s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X9))) => s(X2,happ(s(fun(X1,X2),X3),s(X1,X9))) = s(X2,happ(s(fun(X1,X2),X8),s(X1,X9))))) => s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X5),s(fun(X0,X2),X7),s(fun(X1,X2),X8))) = s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X6),s(fun(X0,X2),X4),s(fun(X1,X2),X3))))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5,X6,X7,X8] : ((s('type.sum.sum'(X0,X1),X5) = s('type.sum.sum'(X0,X1),X6) & ! [X10] : (s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X10))) => s(X2,happ(s(fun(X0,X2),X4),s(X0,X10))) = s(X2,happ(s(fun(X0,X2),X7),s(X0,X10)))) & ! [X9] : (s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X9))) => s(X2,happ(s(fun(X1,X2),X3),s(X1,X9))) = s(X2,happ(s(fun(X1,X2),X8),s(X1,X9))))) => s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X5),s(fun(X0,X2),X7),s(fun(X1,X2),X8))) = s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X6),s(fun(X0,X2),X4),s(fun(X1,X2),X3))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6,X7,X16] : ! [X24,X25,X26,X27,X17,X28] : ((s('type.sum.sum'(X6,X7),X26) = s('type.sum.sum'(X6,X7),X27) & ! [X20] : (s('type.sum.sum'(X6,X7),X27) = s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,X20))) => s(X16,happ(s(fun(X6,X16),X17),s(X6,X20))) = s(X16,happ(s(fun(X6,X16),X25),s(X6,X20)))) & ! [X21] : (s('type.sum.sum'(X6,X7),X27) = s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X21))) => s(X16,happ(s(fun(X7,X16),X24),s(X7,X21))) = s(X16,happ(s(fun(X7,X16),X28),s(X7,X21))))) => s(X16,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X7),X26),s(fun(X6,X16),X17),s(fun(X7,X16),X28))) = s(X16,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X7),X27),s(fun(X6,X16),X25),s(fun(X7,X16),X24))))),
  file('thm.sum.SUM_ALL_CONG','thm.sum.sum_case_cong')).
fof(f230,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X2,happ(s(fun(X0,X2),X4),s(X0,sK17(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X2,happ(s(fun(X0,X2),X7),s(X0,sK17(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK18(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X5),s(fun(X0,X2),X7),s(fun(X1,X2),X8))) = s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X6),s(fun(X0,X2),X4),s(fun(X1,X2),X3))) | s('type.sum.sum'(X0,X1),X5) != s('type.sum.sum'(X0,X1),X6)) )),
  inference(cnf_transformation,[],[f150])).
fof(f231,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X2,happ(s(fun(X1,X2),X3),s(X1,sK18(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X2,happ(s(fun(X1,X2),X8),s(X1,sK18(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X5),s(fun(X0,X2),X7),s(fun(X1,X2),X8))) = s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X6),s(fun(X0,X2),X4),s(fun(X1,X2),X3))) | s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK17(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.sum.sum'(X0,X1),X5) != s('type.sum.sum'(X0,X1),X6)) )),
  inference(cnf_transformation,[],[f150])).
fof(f232,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X2,happ(s(fun(X1,X2),X3),s(X1,sK18(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X2,happ(s(fun(X1,X2),X8),s(X1,sK18(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X5),s(fun(X0,X2),X7),s(fun(X1,X2),X8))) = s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X6),s(fun(X0,X2),X4),s(fun(X1,X2),X3))) | s(X2,happ(s(fun(X0,X2),X4),s(X0,sK17(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) != s(X2,happ(s(fun(X0,X2),X7),s(X0,sK17(X8,X7,X6,X5,X4,X3,X2,X1,X0)))) | s('type.sum.sum'(X0,X1),X5) != s('type.sum.sum'(X0,X1),X6)) )),
  inference(cnf_transformation,[],[f150])).
fof(f228,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X2,X1),X4),s(X2,X3))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X2,X0),'const.sum.INL_1'(s(X2,X3))),s(fun(X2,X1),X4),s(fun(X0,X1),X5)))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X4),s(X2,X3))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X2,X0),'const.sum.INL_1'(s(X2,X3))),s(fun(X2,X1),X4),s(fun(X0,X1),X5)))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X4),s(X2,X3))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X2,X0),'const.sum.INL_1'(s(X2,X3))),s(fun(X2,X1),X4),s(fun(X0,X1),X5)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X7,X16,X6] : ! [X20,X17,X28] : s(X16,happ(s(fun(X6,X16),X17),s(X6,X20))) = s(X16,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,X20))),s(fun(X6,X16),X17),s(fun(X7,X16),X28)))),
  file('thm.sum.SUM_ALL_CONG','thm.sum.sum_case_def|split|0')).
fof(f227,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X2,X1),X5),s(X2,X3))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X2),'const.sum.INR_1'(s(X2,X3))),s(fun(X0,X1),X4),s(fun(X2,X1),X5)))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X5),s(X2,X3))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X2),'const.sum.INR_1'(s(X2,X3))),s(fun(X0,X1),X4),s(fun(X2,X1),X5)))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X5),s(X2,X3))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X2),'const.sum.INR_1'(s(X2,X3))),s(fun(X0,X1),X4),s(fun(X2,X1),X5)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6,X16,X7] : ! [X21,X17,X28] : s(X16,happ(s(fun(X7,X16),X28),s(X7,X21))) = s(X16,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X21))),s(fun(X6,X16),X17),s(fun(X7,X16),X28)))),
  file('thm.sum.SUM_ALL_CONG','thm.sum.sum_case_def|split|1')).
fof(f225,plain,(
  ( ! [X4,X2,X0,X7,X3,X1] : (s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))) = s(X1,happ(s(fun('type.sum.sum'(X0,X2),X1),sK16(X4,X3,X2,X1,X0)),s('type.sum.sum'(X0,X2),'const.sum.INL_1'(s(X0,X7)))))) )),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2,X3,X4] : (! [X7] : s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))) = s(X1,happ(s(fun('type.sum.sum'(X0,X2),X1),sK16(X4,X3,X2,X1,X0)),s('type.sum.sum'(X0,X2),'const.sum.INL_1'(s(X0,X7))))) & ! [X6] : s(X1,happ(s(fun(X2,X1),X4),s(X2,X6))) = s(X1,happ(s(fun('type.sum.sum'(X0,X2),X1),sK16(X4,X3,X2,X1,X0)),s('type.sum.sum'(X0,X2),'const.sum.INR_1'(s(X2,X6))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2,X3,X4] : ? [X5] : (! [X7] : s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))) = s(X1,happ(s(fun('type.sum.sum'(X0,X2),X1),X5),s('type.sum.sum'(X0,X2),'const.sum.INL_1'(s(X0,X7))))) & ! [X6] : s(X1,happ(s(fun(X2,X1),X4),s(X2,X6))) = s(X1,happ(s(fun('type.sum.sum'(X0,X2),X1),X5),s('type.sum.sum'(X0,X2),'const.sum.INR_1'(s(X2,X6))))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2] : ! [X3,X4] : ? [X5] : (! [X7] : s(X1,happ(s(fun(X0,X1),X3),s(X0,X7))) = s(X1,happ(s(fun('type.sum.sum'(X0,X2),X1),X5),s('type.sum.sum'(X0,X2),'const.sum.INL_1'(s(X0,X7))))) & ! [X6] : s(X1,happ(s(fun(X2,X1),X4),s(X2,X6))) = s(X1,happ(s(fun('type.sum.sum'(X0,X2),X1),X5),s('type.sum.sum'(X0,X2),'const.sum.INR_1'(s(X2,X6))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X16,X7] : ! [X17,X18] : ? [X19] : (! [X20] : s(X16,happ(s(fun(X6,X16),X17),s(X6,X20))) = s(X16,happ(s(fun('type.sum.sum'(X6,X7),X16),X19),s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,X20))))) & ! [X21] : s(X16,happ(s(fun(X7,X16),X18),s(X7,X21))) = s(X16,happ(s(fun('type.sum.sum'(X6,X7),X16),X19),s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X21))))))),
  file('thm.sum.SUM_ALL_CONG','thm.sum.sum_Axiom')).
fof(f226,plain,(
  ( ! [X6,X4,X2,X0,X3,X1] : (s(X1,happ(s(fun(X2,X1),X4),s(X2,X6))) = s(X1,happ(s(fun('type.sum.sum'(X0,X2),X1),sK16(X4,X3,X2,X1,X0)),s('type.sum.sum'(X0,X2),'const.sum.INR_1'(s(X2,X6)))))) )),
  inference(cnf_transformation,[],[f149])).
fof(f215,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X6),s(fun(X1,bool),X4),s('type.sum.sum'(X0,X1),X2)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X3),s('type.sum.sum'(X0,X1),X2)))) | p(s(bool,happ(s(fun(X1,bool),X4),s(X1,sK15(X6,X5,X4,X3,X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK14(X6,X5,X4,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK14(X6,X5,X4,X3,X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK14(X6,X5,X4,X3,X2,X1,X0)))))) | (p(s(bool,happ(s(fun(X1,bool),X4),s(X1,sK15(X6,X5,X4,X3,X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,sK15(X6,X5,X4,X3,X2,X1,X0)))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X6),s(fun(X1,bool),X4),s('type.sum.sum'(X0,X1),X2)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X3),s('type.sum.sum'(X0,X1),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (? [X8] : (p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X5),s(X0,X8))))) | ? [X7] : (p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X7)))) & ~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X7))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X6),s(fun(X1,bool),X4),s('type.sum.sum'(X0,X1),X2)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X3),s('type.sum.sum'(X0,X1),X2)))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((? [X8] : (p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X5),s(X0,X8))))) | ? [X7] : (p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X7)))) & ~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X7)))))) | (~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X6),s(fun(X1,bool),X4),s('type.sum.sum'(X0,X1),X2)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X3),s('type.sum.sum'(X0,X1),X2))))))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((! [X8] : (p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X8)))) => p(s(bool,happ(s(fun(X0,bool),X5),s(X0,X8))))) & ! [X7] : (p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X7)))) => p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X7)))))) => (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X6),s(fun(X1,bool),X4),s('type.sum.sum'(X0,X1),X2)))) => p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X3),s('type.sum.sum'(X0,X1),X2))))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6] : ((! [X8] : (p(s(bool,happ(s(fun(X0,bool),X6),s(X0,X8)))) => p(s(bool,happ(s(fun(X0,bool),X5),s(X0,X8))))) & ! [X7] : (p(s(bool,happ(s(fun(X1,bool),X4),s(X1,X7)))) => p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X7)))))) => (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X6),s(fun(X1,bool),X4),s('type.sum.sum'(X0,X1),X2)))) => p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X3),s('type.sum.sum'(X0,X1),X2))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6,X7] : ! [X8,X13,X12,X11,X10] : ((! [X20] : (p(s(bool,happ(s(fun(X6,bool),X10),s(X6,X20)))) => p(s(bool,happ(s(fun(X6,bool),X11),s(X6,X20))))) & ! [X21] : (p(s(bool,happ(s(fun(X7,bool),X12),s(X7,X21)))) => p(s(bool,happ(s(fun(X7,bool),X13),s(X7,X21)))))) => (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X6,bool),X10),s(fun(X7,bool),X12),s('type.sum.sum'(X6,X7),X8)))) => p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X6,bool),X11),s(fun(X7,bool),X13),s('type.sum.sum'(X6,X7),X8))))))),
  file('thm.sum.SUM_ALL_CONG','thm.sum.SUM_ALL_MONO')).
fof(f216,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK14(X6,X5,X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X6),s(fun(X1,bool),X4),s('type.sum.sum'(X0,X1),X2)))) | p(s(bool,happ(s(fun(X1,bool),X4),s(X1,sK15(X6,X5,X4,X3,X2,X1,X0))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X3),s('type.sum.sum'(X0,X1),X2))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f217,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,sK15(X6,X5,X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X6),s(fun(X1,bool),X4),s('type.sum.sum'(X0,X1),X2)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X3),s('type.sum.sum'(X0,X1),X2)))) | p(s(bool,happ(s(fun(X0,bool),X6),s(X0,sK14(X6,X5,X4,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f218,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,sK15(X6,X5,X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X6),s(fun(X1,bool),X4),s('type.sum.sum'(X0,X1),X2)))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X3),s('type.sum.sum'(X0,X1),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X5),s(X0,sK14(X6,X5,X4,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f146])).
fof(f212,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,X2)) | s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) != s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4)))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) != s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4))))) | (p(s(bool,X2)) & s(X1,X3) = s(X1,X5))) & (~p(s(bool,X2)) | s(X1,X3) != s(X1,X5) | s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) = s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4)))))))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) != s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4))))) | (p(s(bool,X2)) & s(X1,X3) = s(X1,X5))) & ((~p(s(bool,X2)) | s(X1,X3) != s(X1,X5)) | s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) = s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4)))))))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) = s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4))))) <=> (p(s(bool,X2)) & s(X1,X3) = s(X1,X5)))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) = s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4))))) <=> (p(s(bool,X2)) & s(X1,X3) = s(X1,X5)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X7,X6] : ! [X10,X20,X21,X29] : (s('type.sum.sum'(X7,X6),'const.sum.INR_1'(s(X6,X29))) = s('type.sum.sum'(X7,X6),'const.bool.COND_3'(s(bool,X10),s('type.sum.sum'(X7,X6),'const.sum.INR_1'(s(X6,X20))),s('type.sum.sum'(X7,X6),'const.sum.INL_1'(s(X7,X21))))) <=> (p(s(bool,X10)) & s(X6,X20) = s(X6,X29)))),
  file('thm.sum.SUM_ALL_CONG','thm.sum.cond_sum_expand|split|0')).
fof(f213,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) != s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4))))) | s(X1,X3) = s(X1,X5)) )),
  inference(cnf_transformation,[],[f145])).
fof(f214,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,X2)) | s(X1,X3) != s(X1,X5) | s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X5))) = s('type.sum.sum'(X0,X1),'const.bool.COND_3'(s(bool,X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4)))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f211,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,happ(s(fun(X1,bool),X2),s(X1,X4))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X1,bool),X2),s(fun(X0,bool),X3),s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X4)))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,happ(s(fun(X1,bool),X2),s(X1,X4))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X1,bool),X2),s(fun(X0,bool),X3),s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X4)))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(bool,happ(s(fun(X1,bool),X2),s(X1,X4))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X1,bool),X2),s(fun(X0,bool),X3),s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X4)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7,X6] : ! [X10,X12,X20] : s(bool,happ(s(fun(X6,bool),X10),s(X6,X20))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X6,bool),X10),s(fun(X7,bool),X12),s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,X20)))))),
  file('thm.sum.SUM_ALL_CONG','thm.sum.SUM_ALL_def|split|0')).
fof(f210,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,happ(s(fun(X1,bool),X3),s(X1,X4))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X2),s(fun(X1,bool),X3),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X4)))))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,happ(s(fun(X1,bool),X3),s(X1,X4))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X2),s(fun(X1,bool),X3),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X4)))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(bool,happ(s(fun(X1,bool),X3),s(X1,X4))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X2),s(fun(X1,bool),X3),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X4)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7] : ! [X10,X12,X21] : s(bool,happ(s(fun(X7,bool),X12),s(X7,X21))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X6,bool),X10),s(fun(X7,bool),X12),s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X21)))))),
  file('thm.sum.SUM_ALL_CONG','thm.sum.SUM_ALL_def|split|1')).
fof(f208,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) | s(X1,X2) = s(X1,X3)) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3] : ((s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) | s(X1,X2) = s(X1,X3)) & (s(X1,X2) != s(X1,X3) | s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3)))))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : ! [X2,X3] : (s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X7] : ! [X21,X20] : (s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X20))) = s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X21))) <=> s(X7,X20) = s(X7,X21))),
  file('thm.sum.SUM_ALL_CONG','thm.sum.INR_|49||49|')).
fof(f209,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) != s(X1,X3) | s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3)))) )),
  inference(cnf_transformation,[],[f143])).
fof(f207,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK13(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK13(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK13(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK13(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.sum.SUM_ALL_CONG',aHL_EXT)).
fof(f206,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK12(X2,X1,X0))))))) | p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),X3)))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK11(X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK11(X2,X1,X0))))))) | ~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK12(X2,X1,X0))))))) | ! [X3] : p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),X3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11,sK12])],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : (? [X5] : ~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X5)))))) | ? [X4] : ~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X4)))))) | ! [X3] : p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),X3)))))),
  inference(rectify,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : (? [X4] : ~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4)))))) | ? [X3] : ~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3)))))) | ! [X5] : p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),X5)))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : ((? [X4] : ~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4)))))) | ? [X3] : ~p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))))))) | ! [X5] : p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),X5)))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : ((! [X4] : p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4)))))) & ! [X3] : p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))))))) => ! [X5] : p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),X5)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ! [X2] : ((! [X4] : p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4)))))) & ! [X3] : p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))))))) => ! [X5] : p(s(bool,happ(s(fun('type.sum.sum'(X0,X1),bool),X2),s('type.sum.sum'(X0,X1),X5)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : ! [X10] : ((! [X20] : p(s(bool,happ(s(fun('type.sum.sum'(X6,X7),bool),X10),s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,X20)))))) & ! [X21] : p(s(bool,happ(s(fun('type.sum.sum'(X6,X7),bool),X10),s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X21))))))) => ! [X8] : p(s(bool,happ(s(fun('type.sum.sum'(X6,X7),bool),X10),s('type.sum.sum'(X6,X7),X8)))))),
  file('thm.sum.SUM_ALL_CONG','thm.sum.sum_INDUCT')).
fof(f205,plain,(
  ( ! [X2,X0,X1] : (s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK10(X2,X1,X0)))) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK9(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2] : (s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK9(X2,X1,X0)))) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK10(X2,X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9,sK10])],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2] : (? [X4] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4))) | ? [X3] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ! [X2] : (? [X4] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4))) | ? [X3] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7] : ! [X22] : (? [X20] : s('type.sum.sum'(X6,X7),X22) = s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,X20))) | ? [X21] : s('type.sum.sum'(X6,X7),X22) = s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X21))))),
  file('thm.sum.SUM_ALL_CONG','thm.sum.sum_CASES')).
fof(f200,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & (((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0)))))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) => p(s(bool,X1))) & (p(s(bool,X1)) => p(s(bool,X0)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X32,X33] : (s(bool,X32) = s(bool,X33) <=> ((p(s(bool,X32)) => p(s(bool,X33))) & (p(s(bool,X33)) => p(s(bool,X32)))))),
  file('thm.sum.SUM_ALL_CONG','thm.bool.EQ_IMP_THM')).
fof(f201,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f138])).
fof(f204,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f138])).
fof(f195,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,sK8(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK8(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8])],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X10,X23] : (! [X20] : (s(X6,X20) = s(X6,X23) => p(s(bool,happ(s(fun(X6,bool),X10),s(X6,X20))))) => p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X10)))))),
  file('thm.sum.SUM_ALL_CONG','thm.bool.JRH_INDUCT_UTIL')).
fof(f196,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK8(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f194,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X32,X33] : s(X6,X33) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X32),s(X6,X33)))),
  file('thm.sum.SUM_ALL_CONG','thm.bool.COND_CLAUSES|split|1')).
fof(f193,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X32,X33] : s(X6,X32) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X32),s(X6,X33)))),
  file('thm.sum.SUM_ALL_CONG','thm.bool.COND_CLAUSES|split|0')).
fof(f169,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X23] : (s(bool,'F_0') = s(bool,X23) <=> ~p(s(bool,X23)))),
  file('thm.sum.SUM_ALL_CONG','thm.bool.EQ_CLAUSES|split|3')).
fof(f170,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f120])).
fof(f165,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X23] : (s(bool,'T_0') = s(bool,X23) <=> p(s(bool,X23)))),
  file('thm.sum.SUM_ALL_CONG','thm.bool.EQ_CLAUSES|split|0')).
fof(f166,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f118])).
fof(f163,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X23] : (s(bool,'T_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X23))),
  file('thm.sum.SUM_ALL_CONG','thm.bool.BOOL_CASES_AX')).
fof(f262,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f159,f157])).
fof(f159,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f27])).
fof(f27,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.sum.SUM_ALL_CONG','thm.bool.NOT_CLAUSES|split|2')).
fof(f157,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.sum.SUM_ALL_CONG',aHL_TRUTH)).
fof(f260,plain,(
  ( ! [X8] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X8))) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,X8))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,X8)))) )),
  inference(backward_demodulation,[],[f153,f155])).
fof(f259,plain,(
  ( ! [X9] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,X9))) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,X9))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,X9)))) )),
  inference(backward_demodulation,[],[f153,f154])).
fof(f153,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),sK3)),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),sK3) & ! [X9] : (s('type.sum.sum'(sK0,sK1),sK3) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,X9))) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,X9))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,X9)))) & ! [X8] : (s('type.sum.sum'(sK0,sK1),sK3) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X8))) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,X8))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,X8)))) & s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK4),s(fun(sK1,bool),sK6),s('type.sum.sum'(sK0,sK1),sK2))) != s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7])],[f96])).
fof(f96,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : (s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),X3) & ! [X9] : (s('type.sum.sum'(X0,X1),X3) != s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X9))) | s(bool,happ(s(fun(X0,bool),X4),s(X0,X9))) = s(bool,happ(s(fun(X0,bool),X5),s(X0,X9)))) & ! [X8] : (s('type.sum.sum'(X0,X1),X3) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X8))) | s(bool,happ(s(fun(X1,bool),X6),s(X1,X8))) = s(bool,happ(s(fun(X1,bool),X7),s(X1,X8)))) & s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X4),s(fun(X1,bool),X6),s('type.sum.sum'(X0,X1),X2))) != s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X7),s('type.sum.sum'(X0,X1),X3))))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7] : ((s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),X3) & ! [X9] : (s('type.sum.sum'(X0,X1),X3) != s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X9))) | s(bool,happ(s(fun(X0,bool),X4),s(X0,X9))) = s(bool,happ(s(fun(X0,bool),X5),s(X0,X9)))) & ! [X8] : (s('type.sum.sum'(X0,X1),X3) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X8))) | s(bool,happ(s(fun(X1,bool),X6),s(X1,X8))) = s(bool,happ(s(fun(X1,bool),X7),s(X1,X8))))) & s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X4),s(fun(X1,bool),X6),s('type.sum.sum'(X0,X1),X2))) != s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X7),s('type.sum.sum'(X0,X1),X3))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7] : ((s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),X3) & ! [X9] : (s('type.sum.sum'(X0,X1),X3) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X9))) => s(bool,happ(s(fun(X0,bool),X4),s(X0,X9))) = s(bool,happ(s(fun(X0,bool),X5),s(X0,X9)))) & ! [X8] : (s('type.sum.sum'(X0,X1),X3) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X8))) => s(bool,happ(s(fun(X1,bool),X6),s(X1,X8))) = s(bool,happ(s(fun(X1,bool),X7),s(X1,X8))))) => s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X4),s(fun(X1,bool),X6),s('type.sum.sum'(X0,X1),X2))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X7),s('type.sum.sum'(X0,X1),X3))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : ((s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),X3) & ! [X9] : (s('type.sum.sum'(X0,X1),X3) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X9))) => s(bool,happ(s(fun(X0,bool),X4),s(X0,X9))) = s(bool,happ(s(fun(X0,bool),X5),s(X0,X9)))) & ! [X8] : (s('type.sum.sum'(X0,X1),X3) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X8))) => s(bool,happ(s(fun(X1,bool),X6),s(X1,X8))) = s(bool,happ(s(fun(X1,bool),X7),s(X1,X8))))) => s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X4),s(fun(X1,bool),X6),s('type.sum.sum'(X0,X1),X2))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X0,bool),X5),s(fun(X1,bool),X7),s('type.sum.sum'(X0,X1),X3))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11,X12,X13] : ((s('type.sum.sum'(X6,X7),X8) = s('type.sum.sum'(X6,X7),X9) & ! [X14] : (s('type.sum.sum'(X6,X7),X9) = s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,X14))) => s(bool,happ(s(fun(X6,bool),X10),s(X6,X14))) = s(bool,happ(s(fun(X6,bool),X11),s(X6,X14)))) & ! [X15] : (s('type.sum.sum'(X6,X7),X9) = s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X15))) => s(bool,happ(s(fun(X7,bool),X12),s(X7,X15))) = s(bool,happ(s(fun(X7,bool),X13),s(X7,X15))))) => s(bool,'const.sum.SUM_ALL_3'(s(fun(X6,bool),X10),s(fun(X7,bool),X12),s('type.sum.sum'(X6,X7),X8))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X6,bool),X11),s(fun(X7,bool),X13),s('type.sum.sum'(X6,X7),X9))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11,X12,X13] : ((s('type.sum.sum'(X6,X7),X8) = s('type.sum.sum'(X6,X7),X9) & ! [X14] : (s('type.sum.sum'(X6,X7),X9) = s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,X14))) => s(bool,happ(s(fun(X6,bool),X10),s(X6,X14))) = s(bool,happ(s(fun(X6,bool),X11),s(X6,X14)))) & ! [X15] : (s('type.sum.sum'(X6,X7),X9) = s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X15))) => s(bool,happ(s(fun(X7,bool),X12),s(X7,X15))) = s(bool,happ(s(fun(X7,bool),X13),s(X7,X15))))) => s(bool,'const.sum.SUM_ALL_3'(s(fun(X6,bool),X10),s(fun(X7,bool),X12),s('type.sum.sum'(X6,X7),X8))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(X6,bool),X11),s(fun(X7,bool),X13),s('type.sum.sum'(X6,X7),X9))))),
  file('thm.sum.SUM_ALL_CONG','thm.sum.SUM_ALL_CONG')).
fof(f154,plain,(
  ( ! [X9] : (s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,X9))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,X9))) | s('type.sum.sum'(sK0,sK1),sK3) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,X9)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f155,plain,(
  ( ! [X8] : (s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,X8))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,X8))) | s('type.sum.sum'(sK0,sK1),sK3) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X8)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f156,plain,(
  s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK4),s(fun(sK1,bool),sK6),s('type.sum.sum'(sK0,sK1),sK2))) != s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK3)))),
  inference(cnf_transformation,[],[f115])).
fof(f9853,plain,(
  ( ! [X0] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(sK2,sK1,sK0)))))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(sK2,sK1,sK0))))))) ) | $spl28),
  inference(forward_demodulation,[],[f9846,f9812])).
fof(f9812,plain,(
  ( ! [X12,X11] : (s(bool,happ(s(fun(sK0,bool),X11),s(sK0,sK9(sK2,sK1,sK0)))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X11),s(fun(sK1,bool),X12),s('type.sum.sum'(sK0,sK1),sK2)))) ) | $spl28),
  inference(superposition,[],[f211,f9178])).
fof(f9846,plain,(
  ( ! [X0] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(sK2,sK1,sK0)))))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2)))))) ) | $spl28),
  inference(backward_demodulation,[],[f9812,f9627])).
fof(f9769,plain,(
  $false | ($spl30 | $spl50)),
  inference(trivial_inequality_removal,[],[f9768])).
fof(f9768,plain,(
  ( ! [X0] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK2,sK1,sK0)))))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK2,sK1,sK0))))))) ) | ($spl30 | $spl50)),
  inference(forward_demodulation,[],[f9767,f9180])).
fof(f9767,plain,(
  ( ! [X0] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK2,sK1,sK0)))))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(sK2,sK1,sK0))))))) ) | $spl50),
  inference(forward_demodulation,[],[f9760,f9722])).
fof(f9722,plain,(
  ( ! [X2,X3] : (s(bool,happ(s(fun(sK1,bool),X2),s(sK1,sK10(sK2,sK1,sK0)))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X3),s(fun(sK1,bool),X2),s('type.sum.sum'(sK0,sK1),sK2)))) ) | $spl50),
  inference(superposition,[],[f210,f9597])).
fof(f9597,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0)))) | $spl50),
  inference(unit_resulting_resolution,[],[f9338,f205])).
fof(f9760,plain,(
  ( ! [X0] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK2,sK1,sK0)))))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2)))))) ) | $spl50),
  inference(backward_demodulation,[],[f9722,f9627])).
fof(f9698,plain,(
  $false | ($spl29 | $spl30)),
  inference(trivial_inequality_removal,[],[f9697])).
fof(f9697,plain,(
  ( ! [X0] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK2,sK1,sK0)))))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK2,sK1,sK0))))))) ) | ($spl29 | $spl30)),
  inference(forward_demodulation,[],[f9696,f9180])).
fof(f9696,plain,(
  ( ! [X0] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK2,sK1,sK0)))))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(sK2,sK1,sK0))))))) ) | $spl29),
  inference(forward_demodulation,[],[f9689,f9651])).
fof(f9651,plain,(
  ( ! [X2,X3] : (s(bool,happ(s(fun(sK1,bool),X2),s(sK1,sK10(sK2,sK1,sK0)))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X3),s(fun(sK1,bool),X2),s('type.sum.sum'(sK0,sK1),sK2)))) ) | $spl29),
  inference(superposition,[],[f210,f9416])).
fof(f9416,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0)))) | $spl29),
  inference(global_subsumption,[],[f156,f155,f154,f153,f259,f260,f261,f157,f262,f163,f166,f165,f170,f169,f193,f194,f196,f195,f204,f201,f200,f205,f206,f207,f209,f208,f210,f211,f214,f213,f212,f218,f217,f216,f215,f226,f225,f227,f228,f232,f231,f230,f229,f263,f272,f274,f306,f325,f351,f363,f417,f422,f423,f425,f460,f461,f458,f522,f523,f526,f527,f635,f637,f769,f890,f1033,f1079,f1060,f1189,f1241,f1219,f1313,f1314,f1315,f1316,f1322,f1324,f1330,f1331,f1332,f1337,f1339,f1472,f1479,f1542,f1515,f1634,f1641,f1709,f1680,f1857,f1976,f1977,f1981,f1982,f1984,f1985,f1986,f1987,f1988,f1989,f1990,f2009,f2010,f2013,f2019,f2020,f2022,f2023,f2024,f2025,f2026,f2027,f2028,f2049,f2051,f2690,f2692,f2694,f2696,f2698,f2700,f2703,f2831,f2832,f2833,f2834,f2835,f2836,f2837,f2838,f2839,f2840,f2841,f2842,f2844,f2845,f2846,f2847,f2848,f2849,f2850,f2851,f2853,f2854,f3123,f3124,f3125,f3126,f3127,f3128,f3266,f3267,f3268,f3269,f3270,f3271,f3387,f3388,f3389,f3390,f3391,f3392,f3481,f3463,f3464,f3482,f3483,f3484,f3485,f3486,f3487,f3490,f3472,f3473,f3493,f3496,f3499,f3502,f3505,f3508,f3480,f3665,f3666,f3667,f3668,f3669,f3670,f3673,f3674,f3675,f3676,f3677,f3678,f3664,f4033,f4034,f4035,f4036,f4037,f4038,f4042,f4043,f4044,f4045,f4046,f4047,f4032,f4171,f4172,f4173,f4174,f4175,f4176,f4180,f4181,f4182,f4183,f4184,f4185,f4170,f5193,f5194,f5196,f5197,f5198,f5199,f5200,f5201,f5235,f5259,f5260,f5261,f5262,f5263,f5264,f5265,f5266,f5396,f5398,f5402,f5403,f5413,f5417,f5418,f5419,f5420,f5421,f5422,f5427,f5428,f5429,f5430,f5432,f5433,f5434,f5435,f5436,f5437,f5578,f5622,f6373,f6374,f6375,f6376,f6378,f6379,f6380,f6381,f6382,f6383,f5729,f5731,f5735,f5736,f5746,f5750,f5751,f5752,f5753,f5754,f5755,f5760,f5761,f5762,f5763,f5764,f5765,f5766,f5767,f5768,f5769,f5770,f6775,f6776,f6777,f6784,f6785,f6786,f6787,f6788,f6789,f6790,f6962,f6803,f6805,f6809,f6810,f6821,f6822,f6823,f6824,f6825,f6826,f6828,f6829,f6831,f6836,f6837,f6860,f6861,f6866,f6867,f6868,f6869,f6870,f6871,f6876,f6877,f6878,f6879,f6880,f6881,f6882,f6883,f6884,f6885,f6886,f6966,f6967,f6968,f6969,f6930,f6932,f6936,f6937,f7117,f7148,f7149,f7153,f7154,f7155,f7156,f7157,f7158,f7159,f7299,f7301,f7303,f7304,f7250,f7309,f7256,f7257,f7482,f7483,f7487,f7488,f7489,f7490,f7491,f7492,f7493,f7494,f7578,f7579,f7580,f7593,f7595,f7599,f7600,f7612,f7613,f7614,f7615,f7616,f7617,f7619,f7620,f7622,f7627,f7628,f7629,f7656,f7657,f7662,f7663,f7664,f7665,f7666,f7667,f7668,f7673,f7674,f7675,f7676,f7677,f7678,f7682,f7683,f7684,f7685,f7686,f7992,f8048,f8049,f8055,f8056,f8057,f8058,f8059,f8060,f8061,f8062,f8178,f8123,f8129,f8130,f8463,f8492,f8493,f8500,f8501,f8502,f8503,f8504,f8505,f8506,f8507,f8508,f9166,f9154,f9155,f9171,f9172,f9173,f9174,f9175,f9176,f9162,f9163,f9164,f9177,f9250,f9249,f9248,f9246,f9245,f9244,f9266,f9254,f9255,f9271,f9272,f9273,f9274,f9275,f9276,f9262,f9263,f9324,f9329,f9285,f9287,f9336,f9288,f9289,f9290,f9291,f9341,f9342,f2843,f9264])).
fof(f9342,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(X29,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK2),s(fun(sK0,X29),X26),s(fun(sK1,X29),X25))) = s(X29,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X24),s(fun(sK0,X29),X27),s(fun(sK1,X29),X28))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X24) | s('type.sum.sum'(sK0,sK1),X24) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X25,X26,X24,'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),X27,X28,X29,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),X24) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X25,X26,X24,'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),X27,X28,X29,sK1,sK0))))) ) | $spl29),
  inference(forward_demodulation,[],[f9293,f9244])).
fof(f9293,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X24) | s('type.sum.sum'(sK0,sK1),X24) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X25,X26,X24,'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),X27,X28,X29,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),X24) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X25,X26,X24,'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),X27,X28,X29,sK1,sK0)))) | s(X29,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X24),s(fun(sK0,X29),X27),s(fun(sK1,X29),X28))) = s(X29,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0)))),s(fun(sK0,X29),X26),s(fun(sK1,X29),X25)))) ) | $spl29),
  inference(superposition,[],[f229,f9244])).
fof(f9341,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(X23,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK2),s(fun(sK0,X23),X21),s(fun(sK1,X23),X22))) = s(X23,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X18),s(fun(sK0,X23),X20),s(fun(sK1,X23),X19))) | s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X19,X20,'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),X18,X21,X22,X23,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X19,X20,'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),X18,X21,X22,X23,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X18)) ) | $spl29),
  inference(forward_demodulation,[],[f9340,f9244])).
fof(f9340,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X19,X20,'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),X18,X21,X22,X23,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X19,X20,'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),X18,X21,X22,X23,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X18) | s(X23,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X18),s(fun(sK0,X23),X20),s(fun(sK1,X23),X19))) = s(X23,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0)))),s(fun(sK0,X23),X21),s(fun(sK1,X23),X22)))) ) | $spl29),
  inference(forward_demodulation,[],[f9339,f9244])).
fof(f9339,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X19,X20,'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),X18,X21,X22,X23,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X18) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X19,X20,'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),X18,X21,X22,X23,sK1,sK0)))) | s(X23,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X18),s(fun(sK0,X23),X20),s(fun(sK1,X23),X19))) = s(X23,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0)))),s(fun(sK0,X23),X21),s(fun(sK1,X23),X22)))) ) | $spl29),
  inference(forward_demodulation,[],[f9292,f9244])).
fof(f9292,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),X18) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK18(X19,X20,'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),X18,X21,X22,X23,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK17(X19,X20,'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),X18,X21,X22,X23,sK1,sK0)))) | s(X23,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X18),s(fun(sK0,X23),X20),s(fun(sK1,X23),X19))) = s(X23,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0)))),s(fun(sK0,X23),X21),s(fun(sK1,X23),X22)))) ) | $spl29),
  inference(superposition,[],[f229,f9244])).
fof(f9290,plain,(
  ( ! [X14,X15,X13] : (s(X13,happ(s(fun(sK1,X13),X14),s(sK1,sK10(sK2,sK1,sK0)))) = s(X13,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK2),s(fun(sK0,X13),X15),s(fun(sK1,X13),X14)))) ) | $spl29),
  inference(superposition,[],[f227,f9244])).
fof(f9289,plain,(
  ( ! [X12,X10,X11] : (s(X10,happ(s(fun(sK1,X10),X11),s(sK1,sK10(sK2,sK1,sK0)))) = s(X10,happ(s(fun('type.sum.sum'(sK0,sK1),X10),sK16(X11,X12,sK1,X10,sK0)),s('type.sum.sum'(sK0,sK1),sK2)))) ) | $spl29),
  inference(superposition,[],[f226,f9244])).
fof(f9288,plain,(
  ( ! [X6,X4,X5] : (s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X4))) != s('type.sum.sum'(sK0,sK1),'const.bool.COND_3'(s(bool,X5),s('type.sum.sum'(sK0,sK1),sK2),s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,X6))))) | s(sK1,X4) = s(sK1,sK10(sK2,sK1,sK0))) ) | $spl29),
  inference(superposition,[],[f213,f9244])).
fof(f9336,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),X3),s(sK1,sK10(sK2,sK1,sK0))))) | ~p(s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK10(sK2,sK1,sK0))))) | p(s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK15(X0,X2,X1,X3,sK3,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK14(X0,X2,X1,X3,sK3,sK1,sK0)))))) ) | $spl29),
  inference(forward_demodulation,[],[f9335,f9287])).
fof(f9335,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK10(sK2,sK1,sK0))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X2),s(fun(sK1,bool),X3),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,happ(s(fun(sK1,bool),X1),s(sK1,sK15(X0,X2,X1,X3,sK3,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),X0),s(sK0,sK14(X0,X2,X1,X3,sK3,sK1,sK0)))))) ) | $spl29),
  inference(backward_demodulation,[],[f9287,f2843])).
fof(f9287,plain,(
  ( ! [X2,X3] : (s(bool,happ(s(fun(sK1,bool),X2),s(sK1,sK10(sK2,sK1,sK0)))) = s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X3),s(fun(sK1,bool),X2),s('type.sum.sum'(sK0,sK1),sK2)))) ) | $spl29),
  inference(superposition,[],[f210,f9244])).
fof(f9285,plain,(
  ( ! [X0] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X0))) | s(sK1,X0) = s(sK1,sK10(sK2,sK1,sK0))) ) | $spl29),
  inference(superposition,[],[f208,f9244])).
fof(f9244,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0)))) | $spl29),
  inference(unit_resulting_resolution,[],[f9177,f205])).
fof(f9245,plain,(
  ( ! [X0,X1] : (s('type.sum.sum'(X0,'type.sum.sum'(sK0,sK1)),'const.sum.INL_1'(s(X0,X1))) != s('type.sum.sum'(X0,'type.sum.sum'(sK0,sK1)),'const.sum.INR_1'(s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))))))) ) | $spl29),
  inference(unit_resulting_resolution,[],[f9177,f2051])).
fof(f9246,plain,(
  ( ! [X2,X0,X1] : (s('type.sum.sum'(X0,'type.sum.sum'(sK0,sK1)),'const.sum.INR_1'(s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))))) != s('type.sum.sum'(X0,'type.sum.sum'(sK0,sK1)),'const.bool.COND_3'(s(bool,X1),s('type.sum.sum'(X0,'type.sum.sum'(sK0,sK1)),'const.sum.INR_1'(s('type.sum.sum'(sK0,sK1),sK2))),s('type.sum.sum'(X0,'type.sum.sum'(sK0,sK1)),'const.sum.INL_1'(s(X0,X2)))))) ) | $spl29),
  inference(unit_resulting_resolution,[],[f9177,f213])).
fof(f9248,plain,(
  ( ! [X0,X1] : (s('type.sum.sum'(X0,'type.sum.sum'(sK0,sK1)),'const.sum.INL_1'(s(X0,X1))) != s('type.sum.sum'(X0,'type.sum.sum'(sK0,sK1)),'const.sum.INR_1'(s('type.sum.sum'(sK0,sK1),sK2)))) ) | $spl29),
  inference(unit_resulting_resolution,[],[f9177,f2051])).
fof(f9249,plain,(
  ( ! [X2,X0,X1] : (s('type.sum.sum'(X0,'type.sum.sum'(sK0,sK1)),'const.sum.INR_1'(s('type.sum.sum'(sK0,sK1),sK2))) != s('type.sum.sum'(X0,'type.sum.sum'(sK0,sK1)),'const.bool.COND_3'(s(bool,X1),s('type.sum.sum'(X0,'type.sum.sum'(sK0,sK1)),'const.sum.INR_1'(s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))))),s('type.sum.sum'(X0,'type.sum.sum'(sK0,sK1)),'const.sum.INL_1'(s(X0,X2)))))) ) | $spl29),
  inference(unit_resulting_resolution,[],[f9177,f213])).
fof(f9250,plain,(
  ( ! [X0] : (s('type.sum.sum'(X0,'type.sum.sum'(sK0,sK1)),'const.sum.INR_1'(s('type.sum.sum'(sK0,sK1),sK2))) != s('type.sum.sum'(X0,'type.sum.sum'(sK0,sK1)),'const.sum.INR_1'(s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))))))) ) | $spl29),
  inference(unit_resulting_resolution,[],[f9177,f208])).
fof(f9177,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))) | $spl29),
  inference(cnf_transformation,[],[f9177_D])).
fof(f9177_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f9689,plain,(
  ( ! [X0] : (s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK2,sK1,sK0)))))) != s('type.sum.sum'(X0,bool),'const.sum.INR_1'(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2)))))) ) | $spl29),
  inference(backward_demodulation,[],[f9651,f9627])).
fof(f9582,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X0),s(fun(sK1,bool),X1),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK14(sK5,X0,sK7,X1,sK2,sK1,sK0)))))) ) | $spl78),
  inference(cnf_transformation,[],[f9582_D])).
fof(f9582_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X0),s(fun(sK1,bool),X1),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK14(sK5,X0,sK7,X1,sK2,sK1,sK0)))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f9456,plain,(
  p(s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK2,sK1,sK0))))) | $spl58),
  inference(cnf_transformation,[],[f9456_D])).
fof(f9456_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK2,sK1,sK0))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f9581,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK2,sK1,sK0))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X0),s(fun(sK1,bool),X1),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK14(sK5,X0,sK7,X1,sK2,sK1,sK0)))))) ) | ($spl0 | $spl30 | $spl48)),
  inference(forward_demodulation,[],[f9580,f9180])).
fof(f9580,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(sK2,sK1,sK0))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X0),s(fun(sK1,bool),X1),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK14(sK5,X0,sK7,X1,sK2,sK1,sK0)))))) ) | ($spl0 | $spl48)),
  inference(forward_demodulation,[],[f9579,f9337])).
fof(f9579,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X0),s(fun(sK1,bool),X1),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK15(sK5,X0,sK7,X1,sK2,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK14(sK5,X0,sK7,X1,sK2,sK1,sK0)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f9498,f157])).
fof(f9498,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X0),s(fun(sK1,bool),X1),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK15(sK5,X0,sK7,X1,sK2,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK14(sK5,X0,sK7,X1,sK2,sK1,sK0)))))) ) | $spl0),
  inference(superposition,[],[f215,f419])).
fof(f2569,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s('type.sum.sum'(X28,X25),'const.sum.INR_1'(s(X25,X27))) = s('type.sum.sum'(X28,X25),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X28,X25),'const.sum.INR_1'(s(X25,X26))),s('type.sum.sum'(X28,X25),'const.sum.INL_1'(s(X28,X29))))) | s(X25,X26) != s(X25,X27)) ) | $spl14),
  inference(cnf_transformation,[],[f2569_D])).
fof(f2569_D,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s('type.sum.sum'(X28,X25),'const.sum.INR_1'(s(X25,X27))) = s('type.sum.sum'(X28,X25),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X28,X25),'const.sum.INR_1'(s(X25,X26))),s('type.sum.sum'(X28,X25),'const.sum.INL_1'(s(X28,X29))))) | s(X25,X26) != s(X25,X27)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f9470,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl76),
  inference(cnf_transformation,[],[f9470_D])).
fof(f9470_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f9474,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3) | s('type.sum.sum'(X4,X1),'const.sum.INR_1'(s(X1,X3))) = s('type.sum.sum'(X4,X1),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X4,X1),'const.sum.INR_1'(s(X1,X2))),s('type.sum.sum'(X4,X1),'const.sum.INL_1'(s(X4,X5)))))) )),
  inference(condensation,[],[f9473])).
fof(f9473,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (s(X88,X89) != s(X88,X90) | s('type.sum.sum'(X91,X88),'const.sum.INR_1'(s(X88,X90))) = s('type.sum.sum'(X91,X88),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X91,X88),'const.sum.INR_1'(s(X88,X89))),s('type.sum.sum'(X91,X88),'const.sum.INL_1'(s(X91,X92))))) | s(bool,X87) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f5597,f306])).
fof(f5597,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,X87)) | s(X88,X89) != s(X88,X90) | s('type.sum.sum'(X91,X88),'const.sum.INR_1'(s(X88,X90))) = s('type.sum.sum'(X91,X88),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X91,X88),'const.sum.INR_1'(s(X88,X89))),s('type.sum.sum'(X91,X88),'const.sum.INL_1'(s(X91,X92))))) | s(bool,X87) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f214,f272])).
fof(f9472,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(X1,X2) != s(X1,X3) | s('type.sum.sum'(X4,X1),'const.sum.INR_1'(s(X1,X3))) = s('type.sum.sum'(X4,X1),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X4,X1),'const.sum.INR_1'(s(X1,X2))),s('type.sum.sum'(X4,X1),'const.sum.INL_1'(s(X4,X5)))))) )),
  inference(condensation,[],[f9471])).
fof(f9471,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (s(X88,X89) != s(X88,X90) | s('type.sum.sum'(X91,X88),'const.sum.INR_1'(s(X88,X90))) = s('type.sum.sum'(X91,X88),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X91,X88),'const.sum.INR_1'(s(X88,X89))),s('type.sum.sum'(X91,X88),'const.sum.INL_1'(s(X91,X92))))) | s(bool,X87) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f5791,f306])).
fof(f5791,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,X87)) | s(X88,X89) != s(X88,X90) | s('type.sum.sum'(X91,X88),'const.sum.INR_1'(s(X88,X90))) = s('type.sum.sum'(X91,X88),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X91,X88),'const.sum.INR_1'(s(X88,X89))),s('type.sum.sum'(X91,X88),'const.sum.INL_1'(s(X91,X92))))) | s(bool,X87) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f214,f272])).
fof(f9466,plain,(
  ( ! [X103,X102] : (s('type.sum.sum'(X102,bool),'const.sum.INR_1'(s(bool,'F_0'))) != s('type.sum.sum'(X102,bool),'const.sum.INL_1'(s(X102,X103)))) ) | $spl72),
  inference(cnf_transformation,[],[f9466_D])).
fof(f9466_D,plain,(
  ( ! [X103,X102] : (s('type.sum.sum'(X102,bool),'const.sum.INR_1'(s(bool,'F_0'))) != s('type.sum.sum'(X102,bool),'const.sum.INL_1'(s(X102,X103)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f8524,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0'))) != s('type.sum.sum'(X1,bool),'const.sum.INL_1'(s(X1,X2)))) )),
  inference(condensation,[],[f8467])).
fof(f8467,plain,(
  ( ! [X17,X15,X18,X16] : (s('type.sum.sum'(X16,bool),'const.sum.INR_1'(s(bool,'F_0'))) != s('type.sum.sum'(X16,bool),'const.sum.INL_1'(s(X16,X17))) | s(bool,X15) = s(bool,X18) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f2051,f163])).
fof(f8536,plain,(
  ( ! [X2,X0,X1] : (s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0'))) != s('type.sum.sum'(X1,bool),'const.sum.INL_1'(s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(forward_demodulation,[],[f8514,f713])).
fof(f713,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl8),
  inference(cnf_transformation,[],[f713_D])).
fof(f713_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f8514,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.sum.sum'(X1,bool),'const.sum.INL_1'(s(X1,X2))) != s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,f(s(bool,'F_0')))))) ) | $spl8),
  inference(condensation,[],[f8473])).
fof(f8473,plain,(
  ( ! [X45,X43,X46,X44] : (s('type.sum.sum'(X44,bool),'const.sum.INL_1'(s(X44,X45))) != s('type.sum.sum'(X44,bool),'const.sum.INR_1'(s(bool,f(s(bool,'F_0'))))) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X43)) ) | $spl8),
  inference(superposition,[],[f2051,f2638])).
fof(f2638,plain,(
  ( ! [X20] : (s(bool,X20) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X20)) ) | $spl8),
  inference(subsumption_resolution,[],[f2607,f262])).
fof(f2607,plain,(
  ( ! [X20] : (p(s(bool,'F_0')) | s(bool,X20) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X20)) ) | $spl8),
  inference(superposition,[],[f306,f713])).
fof(f8537,plain,(
  ( ! [X2,X0,X1] : (s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'F_0'))) != s('type.sum.sum'(X1,bool),'const.sum.INL_1'(s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl19)),
  inference(forward_demodulation,[],[f8511,f3424])).
fof(f3424,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl8 | $spl19)),
  inference(forward_demodulation,[],[f3408,f713])).
fof(f3408,plain,(
  s(bool,f(s(bool,'T_0'))) = s(bool,f(s(bool,'F_0'))) | ($spl8 | $spl19)),
  inference(unit_resulting_resolution,[],[f3371,f2638])).
fof(f8511,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s('type.sum.sum'(X1,bool),'const.sum.INL_1'(s(X1,X2))) != s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,f(s(bool,'T_0')))))) ) | ($spl8 | $spl19)),
  inference(condensation,[],[f8475])).
fof(f8475,plain,(
  ( ! [X54,X52,X55,X53] : (s('type.sum.sum'(X53,bool),'const.sum.INL_1'(s(X53,X54))) != s('type.sum.sum'(X53,bool),'const.sum.INR_1'(s(bool,f(s(bool,'T_0'))))) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X52)) ) | ($spl8 | $spl19)),
  inference(superposition,[],[f2051,f3633])).
fof(f3633,plain,(
  ( ! [X20] : (s(bool,X20) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl8 | $spl19)),
  inference(subsumption_resolution,[],[f3597,f262])).
fof(f3597,plain,(
  ( ! [X20] : (p(s(bool,'F_0')) | s(bool,X20) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl8 | $spl19)),
  inference(superposition,[],[f306,f3424])).
fof(f9469,plain,(
  ( ! [X78,X75] : (s(bool,X78) = s(bool,f(s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) ) | $spl74),
  inference(cnf_transformation,[],[f9469_D])).
fof(f9469_D,plain,(
  ( ! [X78,X75] : (s(bool,X78) = s(bool,f(s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f8480,plain,(
  ( ! [X78,X76,X77,X75] : (s('type.sum.sum'(X76,bool),'const.sum.INR_1'(s(bool,'F_0'))) != s('type.sum.sum'(X76,bool),'const.sum.INL_1'(s(X76,X77))) | s(bool,X78) = s(bool,f(s(bool,X75))) | s(bool,'T_0') = s(bool,X75)) ) | $spl8),
  inference(superposition,[],[f2051,f2766])).
fof(f2766,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl8),
  inference(subsumption_resolution,[],[f2730,f166])).
fof(f2730,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,'F_0') = s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | $spl8),
  inference(superposition,[],[f170,f2593])).
fof(f2593,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f713,f163])).
fof(f9465,plain,(
  ( ! [X101,X104] : (s(bool,X104) = s(bool,f(s(bool,f(s(bool,X101))))) | s(bool,'T_0') = s(bool,X101)) ) | $spl70),
  inference(cnf_transformation,[],[f9465_D])).
fof(f9465_D,plain,(
  ( ! [X101,X104] : (s(bool,X104) = s(bool,f(s(bool,f(s(bool,X101))))) | s(bool,'T_0') = s(bool,X101)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f8487,plain,(
  ( ! [X103,X101,X102,X104] : (s('type.sum.sum'(X102,bool),'const.sum.INR_1'(s(bool,'F_0'))) != s('type.sum.sum'(X102,bool),'const.sum.INL_1'(s(X102,X103))) | s(bool,X104) = s(bool,f(s(bool,f(s(bool,X101))))) | s(bool,'T_0') = s(bool,X101)) ) | $spl8),
  inference(superposition,[],[f2051,f4460])).
fof(f4460,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(condensation,[],[f4459])).
fof(f4459,plain,(
  ( ! [X64,X63] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X64))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f4293])).
fof(f4293,plain,(
  ( ! [X64,X63] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X64))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64)) ) | $spl8),
  inference(superposition,[],[f2766,f2769])).
fof(f2769,plain,(
  ( ! [X33,X34] : (s(bool,X34) = s(bool,f(s(bool,X33))) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X33)) ) | $spl8),
  inference(subsumption_resolution,[],[f2739,f166])).
fof(f2739,plain,(
  ( ! [X33,X34] : (p(s(bool,X33)) | s(bool,X34) = s(bool,f(s(bool,X33))) | s(bool,'T_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X33)) ) | $spl8),
  inference(superposition,[],[f306,f2593])).
fof(f9463,plain,(
  ( ! [X134,X138,X135,X139] : (s(bool,X138) = s(bool,'|dollar|exists_1'(s(fun(X134,bool),X135))) | s(X134,X139) = s(X134,sK8(X139,X135,X134))) ) | $spl68),
  inference(cnf_transformation,[],[f9463_D])).
fof(f9463_D,plain,(
  ( ! [X134,X138,X135,X139] : (s(bool,X138) = s(bool,'|dollar|exists_1'(s(fun(X134,bool),X135))) | s(X134,X139) = s(X134,sK8(X139,X135,X134))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f9460,plain,(
  ( ! [X161,X160] : (s('type.sum.sum'(X160,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X160,bool),'const.sum.INL_1'(s(X160,X161)))) ) | $spl62),
  inference(cnf_transformation,[],[f9460_D])).
fof(f9460_D,plain,(
  ( ! [X161,X160] : (s('type.sum.sum'(X160,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X160,bool),'const.sum.INL_1'(s(X160,X161)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f8494,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (s('type.sum.sum'(X136,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X136,bool),'const.sum.INL_1'(s(X136,X137))) | s(bool,X138) = s(bool,'|dollar|exists_1'(s(fun(X134,bool),X135))) | s(X134,X139) = s(X134,sK8(X139,X135,X134))) )),
  inference(superposition,[],[f2051,f425])).
fof(f9462,plain,(
  ( ! [X144,X140,X141] : (s(bool,X144) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X140),s(bool,X141))) | s(bool,'F_0') = s(bool,X140)) ) | $spl66),
  inference(cnf_transformation,[],[f9462_D])).
fof(f9462_D,plain,(
  ( ! [X144,X140,X141] : (s(bool,X144) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X140),s(bool,X141))) | s(bool,'F_0') = s(bool,X140)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f8495,plain,(
  ( ! [X144,X142,X140,X143,X141] : (s('type.sum.sum'(X142,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X142,bool),'const.sum.INL_1'(s(X142,X143))) | s(bool,X144) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X140),s(bool,X141))) | s(bool,'F_0') = s(bool,X140)) )),
  inference(superposition,[],[f2051,f325])).
fof(f9461,plain,(
  ( ! [X146,X149,X145] : (s(bool,X149) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X145),s(bool,X146))) | s(bool,'F_0') = s(bool,X146)) ) | $spl64),
  inference(cnf_transformation,[],[f9461_D])).
fof(f9461_D,plain,(
  ( ! [X146,X149,X145] : (s(bool,X149) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X145),s(bool,X146))) | s(bool,'F_0') = s(bool,X146)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f8496,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s('type.sum.sum'(X147,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X147,bool),'const.sum.INL_1'(s(X147,X148))) | s(bool,X149) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X145),s(bool,X146))) | s(bool,'F_0') = s(bool,X146)) )),
  inference(superposition,[],[f2051,f363])).
fof(f9459,plain,(
  ( ! [X158,X159,X162] : (s(bool,X162) = s(bool,sK13(X158,X159,bool,bool)) | s(fun(bool,bool),X158) = s(fun(bool,bool),X159)) ) | $spl60),
  inference(cnf_transformation,[],[f9459_D])).
fof(f9459_D,plain,(
  ( ! [X158,X159,X162] : (s(bool,X162) = s(bool,sK13(X158,X159,bool,bool)) | s(fun(bool,bool),X158) = s(fun(bool,bool),X159)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f8499,plain,(
  ( ! [X158,X161,X159,X162,X160] : (s('type.sum.sum'(X160,bool),'const.sum.INR_1'(s(bool,'T_0'))) != s('type.sum.sum'(X160,bool),'const.sum.INL_1'(s(X160,X161))) | s(bool,X162) = s(bool,sK13(X158,X159,bool,bool)) | s(fun(bool,bool),X158) = s(fun(bool,bool),X159)) ) | $spl6),
  inference(superposition,[],[f2051,f712])).
fof(f9454,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X5),s(fun(sK1,bool),X6),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK14(sK4,X5,sK6,X6,sK2,sK1,sK0)))))) ) | $spl56),
  inference(cnf_transformation,[],[f9454_D])).
fof(f9454_D,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X5),s(fun(sK1,bool),X6),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK14(sK4,X5,sK6,X6,sK2,sK1,sK0)))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f9441,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK2,sK1,sK0))))) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X5),s(fun(sK1,bool),X6),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK14(sK4,X5,sK6,X6,sK2,sK1,sK0)))))) ) | ($spl3 | $spl48)),
  inference(backward_demodulation,[],[f9337,f2852])).
fof(f2852,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X5),s(fun(sK1,bool),X6),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK15(sK4,X5,sK6,X6,sK2,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK14(sK4,X5,sK6,X6,sK2,sK1,sK0)))))) ) | $spl3),
  inference(subsumption_resolution,[],[f2826,f157])).
fof(f2826,plain,(
  ( ! [X6,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),X5),s(fun(sK1,bool),X6),s('type.sum.sum'(sK0,sK1),sK2)))) | p(s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK15(sK4,X5,sK6,X6,sK2,sK1,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK14(sK4,X5,sK6,X6,sK2,sK1,sK0)))))) ) | $spl3),
  inference(superposition,[],[f215,f430])).
fof(f430,plain,(
  s(bool,'T_0') = s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK4),s(fun(sK1,bool),sK6),s('type.sum.sum'(sK0,sK1),sK2))) | $spl3),
  inference(unit_resulting_resolution,[],[f262,f421,f306])).
fof(f9398,plain,(
  s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | $spl52),
  inference(cnf_transformation,[],[f9398_D])).
fof(f9398_D,plain,(
  s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f9400,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | $spl54),
  inference(cnf_transformation,[],[f9400_D])).
fof(f9400_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f9396,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | $spl28),
  inference(forward_demodulation,[],[f9392,f9178])).
fof(f9392,plain,(
  s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | $spl28),
  inference(trivial_inequality_removal,[],[f9357])).
fof(f9357,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),sK2) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | $spl28),
  inference(superposition,[],[f1330,f9178])).
fof(f9337,plain,(
  ( ! [X17] : (s(sK1,X17) = s(sK1,sK10(sK2,sK1,sK0))) ) | $spl48),
  inference(cnf_transformation,[],[f9337_D])).
fof(f9337_D,plain,(
  ( ! [X17] : (s(sK1,X17) = s(sK1,sK10(sK2,sK1,sK0))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f9338,plain,(
  ( ! [X16] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,X16)))) ) | $spl50),
  inference(cnf_transformation,[],[f9338_D])).
fof(f9338_D,plain,(
  ( ! [X16] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,X16)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f9291,plain,(
  ( ! [X17,X16] : (s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,X16))) | s(sK1,X17) = s(sK1,sK10(sK2,sK1,sK0))) ) | $spl29),
  inference(superposition,[],[f2051,f9244])).
fof(f9331,plain,(
  s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | $spl44),
  inference(cnf_transformation,[],[f9331_D])).
fof(f9331_D,plain,(
  s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f9333,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | $spl46),
  inference(cnf_transformation,[],[f9333_D])).
fof(f9333_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f9329,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | $spl29),
  inference(forward_demodulation,[],[f9319,f9244])).
fof(f9319,plain,(
  s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | $spl29),
  inference(trivial_inequality_removal,[],[f9284])).
fof(f9284,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),sK2) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | $spl29),
  inference(superposition,[],[f1313,f9244])).
fof(f9326,plain,(
  s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | $spl40),
  inference(cnf_transformation,[],[f9326_D])).
fof(f9326_D,plain,(
  s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f9328,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | $spl42),
  inference(cnf_transformation,[],[f9328_D])).
fof(f9328_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f9324,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | $spl29),
  inference(forward_demodulation,[],[f9320,f9244])).
fof(f9320,plain,(
  s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | $spl29),
  inference(trivial_inequality_removal,[],[f9283])).
fof(f9283,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),sK2) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK2,sK1,sK0))),sK1,sK0)))) | $spl29),
  inference(superposition,[],[f1330,f9244])).
fof(f9268,plain,(
  s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(sK3,sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(sK3,sK1,sK0)))) | $spl36),
  inference(cnf_transformation,[],[f9268_D])).
fof(f9268_D,plain,(
  s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(sK3,sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(sK3,sK1,sK0)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f9270,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK3,sK1,sK0)))) | $spl38),
  inference(cnf_transformation,[],[f9270_D])).
fof(f9270_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK3,sK1,sK0)))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f9266,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK3,sK1,sK0)))) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(sK3,sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(sK3,sK1,sK0))))),
  inference(forward_demodulation,[],[f9265,f153])).
fof(f9265,plain,(
  s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(sK3,sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(sK3,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK3) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK3,sK1,sK0))))),
  inference(trivial_inequality_removal,[],[f9253])).
fof(f9253,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),sK2) | s(bool,happ(s(fun(sK0,bool),sK4),s(sK0,sK9(sK3,sK1,sK0)))) = s(bool,happ(s(fun(sK0,bool),sK5),s(sK0,sK9(sK3,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK3) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK3,sK1,sK0))))),
  inference(superposition,[],[f1330,f153])).
fof(f9226,plain,(
  s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | $spl32),
  inference(cnf_transformation,[],[f9226_D])).
fof(f9226_D,plain,(
  s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f9228,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | $spl34),
  inference(cnf_transformation,[],[f9228_D])).
fof(f9228_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f9224,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | $spl28),
  inference(forward_demodulation,[],[f9220,f9178])).
fof(f9220,plain,(
  s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | $spl28),
  inference(trivial_inequality_removal,[],[f9187])).
fof(f9187,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),sK2) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))),sK1,sK0)))) | $spl28),
  inference(superposition,[],[f1313,f9178])).
fof(f9178,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))) | $spl28),
  inference(cnf_transformation,[],[f9178_D])).
fof(f9178_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f9180,plain,(
  s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK2,sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(sK2,sK1,sK0)))) | $spl30),
  inference(cnf_transformation,[],[f9180_D])).
fof(f9180_D,plain,(
  s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK2,sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(sK2,sK1,sK0)))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f9164,plain,(
  s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK2,sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(sK2,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK2,sK1,sK0))))),
  inference(equality_resolution,[],[f1313])).
fof(f9168,plain,(
  s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK3,sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(sK3,sK1,sK0)))) | $spl24),
  inference(cnf_transformation,[],[f9168_D])).
fof(f9168_D,plain,(
  s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK3,sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(sK3,sK1,sK0)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f9170,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK3,sK1,sK0)))) | $spl26),
  inference(cnf_transformation,[],[f9170_D])).
fof(f9170_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK3,sK1,sK0)))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f9166,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK3,sK1,sK0)))) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK3,sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(sK3,sK1,sK0))))),
  inference(forward_demodulation,[],[f9165,f153])).
fof(f9165,plain,(
  s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK3,sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(sK3,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK3) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK3,sK1,sK0))))),
  inference(trivial_inequality_removal,[],[f9153])).
fof(f9153,plain,(
  s('type.sum.sum'(sK0,sK1),sK2) != s('type.sum.sum'(sK0,sK1),sK2) | s(bool,happ(s(fun(sK1,bool),sK6),s(sK1,sK10(sK3,sK1,sK0)))) = s(bool,happ(s(fun(sK1,bool),sK7),s(sK1,sK10(sK3,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK3) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK3,sK1,sK0))))),
  inference(superposition,[],[f1313,f153])).
fof(f3372,plain,(
  ( ! [X68,X69] : (s(fun(bool,bool),X68) = s(fun(bool,bool),X69)) ) | $spl20),
  inference(cnf_transformation,[],[f3372_D])).
fof(f3372_D,plain,(
  ( ! [X68,X69] : (s(fun(bool,bool),X68) = s(fun(bool,bool),X69)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3374,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | $spl23),
  inference(cnf_transformation,[],[f3374_D])).
fof(f3374_D,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f3364,plain,(
  ( ! [X68,X69] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X68) = s(fun(bool,bool),X69)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f3346])).
fof(f3346,plain,(
  ( ! [X68,X69] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X68) = s(fun(bool,bool),X69) | s(fun(bool,bool),X68) = s(fun(bool,bool),X69)) ) | $spl6),
  inference(superposition,[],[f207,f712])).
fof(f3369,plain,(
  ( ! [X59,X60] : (p(s(bool,sK13(X59,X60,bool,bool))) | s(fun(bool,bool),X59) = s(fun(bool,bool),X60)) ) | $spl16),
  inference(cnf_transformation,[],[f3369_D])).
fof(f3369_D,plain,(
  ( ! [X59,X60] : (p(s(bool,sK13(X59,X60,bool,bool))) | s(fun(bool,bool),X59) = s(fun(bool,bool),X60)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3371,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | $spl19),
  inference(cnf_transformation,[],[f3371_D])).
fof(f3371_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f3342,plain,(
  ( ! [X59,X60] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | p(s(bool,sK13(X59,X60,bool,bool))) | s(fun(bool,bool),X59) = s(fun(bool,bool),X60)) ) | ($spl6 | $spl8)),
  inference(superposition,[],[f2764,f712])).
fof(f2764,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(bool,X1))) | p(s(bool,X1))) ) | $spl8),
  inference(subsumption_resolution,[],[f2727,f165])).
fof(f2727,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X1)) ) | $spl8),
  inference(superposition,[],[f165,f2593])).
fof(f2568,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,f(s(bool,X24)))) ) | $spl12),
  inference(cnf_transformation,[],[f2568_D])).
fof(f2568_D,plain,(
  ( ! [X24] : (s(bool,'T_0') = s(bool,f(s(bool,X24)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2567,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s('type.sum.sum'(X28,X25),'const.sum.INR_1'(s(X25,X27))) = s('type.sum.sum'(X28,X25),'const.bool.COND_3'(s(bool,'T_0'),s('type.sum.sum'(X28,X25),'const.sum.INR_1'(s(X25,X26))),s('type.sum.sum'(X28,X25),'const.sum.INL_1'(s(X28,X29))))) | s(X25,X26) != s(X25,X27) | s(bool,'T_0') = s(bool,f(s(bool,X24)))) ) | $spl9),
  inference(forward_demodulation,[],[f2566,f720])).
fof(f720,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl9),
  inference(unit_resulting_resolution,[],[f262,f714,f306])).
fof(f2566,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (s(X25,X26) != s(X25,X27) | s('type.sum.sum'(X28,X25),'const.sum.INR_1'(s(X25,X27))) = s('type.sum.sum'(X28,X25),'const.bool.COND_3'(s(bool,f(s(bool,'F_0'))),s('type.sum.sum'(X28,X25),'const.sum.INR_1'(s(X25,X26))),s('type.sum.sum'(X28,X25),'const.sum.INL_1'(s(X28,X29))))) | s(bool,'T_0') = s(bool,f(s(bool,X24)))) ) | $spl9),
  inference(subsumption_resolution,[],[f2522,f2413])).
fof(f2413,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,'T_0') = s(bool,f(s(bool,X13)))) ) | $spl9),
  inference(subsumption_resolution,[],[f2198,f720])).
fof(f2198,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X13)))) ) | $spl9),
  inference(superposition,[],[f165,f1955])).
fof(f1955,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X6)))) ) | $spl9),
  inference(subsumption_resolution,[],[f1937,f157])).
fof(f1937,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | s(bool,X6) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X6)))) ) | $spl9),
  inference(superposition,[],[f1809,f720])).
fof(f1809,plain,(
  ( ! [X21,X19] : (~p(s(bool,X21)) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,f(s(bool,X19)))) ) | $spl9),
  inference(forward_demodulation,[],[f1808,f193])).
fof(f1808,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,X21)) | s(bool,X21) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X19),s(bool,X20))) | s(bool,'T_0') = s(bool,f(s(bool,X19)))) ) | $spl9),
  inference(subsumption_resolution,[],[f1757,f157])).
fof(f1757,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | s(bool,X21) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X19),s(bool,X20))) | s(bool,'T_0') = s(bool,f(s(bool,X19)))) ) | $spl9),
  inference(superposition,[],[f204,f922])).
fof(f922,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X15),s(bool,X16))) | s(bool,'T_0') = s(bool,f(s(bool,X15)))) ) | $spl9),
  inference(superposition,[],[f793,f193])).
fof(f793,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(superposition,[],[f720,f163])).
fof(f2522,plain,(
  ( ! [X28,X26,X24,X29,X27,X25] : (~p(s(bool,X24)) | s(X25,X26) != s(X25,X27) | s('type.sum.sum'(X28,X25),'const.sum.INR_1'(s(X25,X27))) = s('type.sum.sum'(X28,X25),'const.bool.COND_3'(s(bool,f(s(bool,'F_0'))),s('type.sum.sum'(X28,X25),'const.sum.INR_1'(s(X25,X26))),s('type.sum.sum'(X28,X25),'const.sum.INL_1'(s(X28,X29))))) | s(bool,'T_0') = s(bool,f(s(bool,X24)))) ) | $spl9),
  inference(superposition,[],[f214,f1955])).
fof(f2396,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,g(s(bool,'F_0'))))) | $spl10),
  inference(cnf_transformation,[],[f2396_D])).
fof(f2396_D,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,g(s(bool,'F_0'))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f712,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK13(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl6),
  inference(cnf_transformation,[],[f712_D])).
fof(f712_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK13(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2305,plain,(
  ( ! [X159,X160] : (s(fun(bool,bool),X159) = s(fun(bool,bool),X160) | s(bool,'T_0') = s(bool,sK13(X159,X160,bool,bool)) | s(bool,'T_0') = s(bool,f(s(bool,g(s(bool,'F_0')))))) ) | ($spl3 | $spl9)),
  inference(trivial_inequality_removal,[],[f2178])).
fof(f2178,plain,(
  ( ! [X159,X160] : (s(bool,f(s(bool,'F_0'))) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X159) = s(fun(bool,bool),X160) | s(bool,'T_0') = s(bool,sK13(X159,X160,bool,bool)) | s(bool,'T_0') = s(bool,f(s(bool,g(s(bool,'F_0')))))) ) | ($spl3 | $spl9)),
  inference(superposition,[],[f705,f1955])).
fof(f705,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK13(X0,X1,X2,bool))) ) | $spl3),
  inference(forward_demodulation,[],[f701,f445])).
fof(f445,plain,(
  s(bool,'F_0') = s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2))) | $spl3),
  inference(unit_resulting_resolution,[],[f436,f163])).
fof(f436,plain,(
  s(bool,'T_0') != s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2))) | $spl3),
  inference(backward_demodulation,[],[f430,f261])).
fof(f701,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2))))) != s(X2,g(s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2))))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK13(X0,X1,X2,bool))) ) | $spl3),
  inference(superposition,[],[f207,f490])).
fof(f490,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2))) | s(bool,'T_0') = s(bool,X5)) ) | $spl3),
  inference(subsumption_resolution,[],[f476,f262])).
fof(f476,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2))) | s(bool,'T_0') = s(bool,X5)) ) | $spl3),
  inference(superposition,[],[f306,f445])).
fof(f711,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl4),
  inference(cnf_transformation,[],[f711_D])).
fof(f711_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f714,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl9),
  inference(cnf_transformation,[],[f714_D])).
fof(f714_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f709,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK13(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl3),
  inference(superposition,[],[f705,f163])).
fof(f419,plain,(
  s(bool,'T_0') = s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2))) | $spl0),
  inference(cnf_transformation,[],[f419_D])).
fof(f419_D,plain,(
  s(bool,'T_0') = s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f421,plain,(
  s(bool,'F_0') != s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK4),s(fun(sK1,bool),sK6),s('type.sum.sum'(sK0,sK1),sK2))) | $spl3),
  inference(cnf_transformation,[],[f421_D])).
fof(f421_D,plain,(
  s(bool,'F_0') != s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK4),s(fun(sK1,bool),sK6),s('type.sum.sum'(sK0,sK1),sK2))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f417,plain,(
  s(bool,'F_0') != s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK4),s(fun(sK1,bool),sK6),s('type.sum.sum'(sK0,sK1),sK2))) | s(bool,'T_0') = s(bool,'const.sum.SUM_ALL_3'(s(fun(sK0,bool),sK5),s(fun(sK1,bool),sK7),s('type.sum.sum'(sK0,sK1),sK2)))),
  inference(superposition,[],[f261,f163])).
% SZS output end Proof for thm.sum
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 129
% Passive clauses: 1087
% Generated clauses: 9743
% Final active clauses: 77
% Final passive clauses: 648
% Input formulas: 37
% Initial clauses: 88
% 
% Duplicate literals: 372
% Trivial inequalities: 103
% Fw subsumption resolutions: 404
% Fw demodulations: 1802
% Bw demodulations: 77
% Condensations: 522
% Global subsumptions: 6
% 
% Simple tautologies: 141
% Equational tautologies: 64
% Forward subsumptions: 5120
% Fw demodulations to eq. taut.: 92
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 589
% Factoring: 1
% Forward superposition: 2419
% Backward superposition: 3035
% Self superposition: 76
% Equality factoring: 37
% Equality resolution: 19
% 
% Split clauses: 27
% Split components: 55
% SAT solver clauses: 1152
% SAT solver unit clauses: 37
% SAT solver binary clauses: 238
% 
% Sat splits: 27
% Sat splitting refutations: 3
% 
% Memory used [KB]: 9722
% Time elapsed: 0.452 s
% ------------------------------
% ------------------------------
% Success in time 1.754 s
