Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.real
% (87668)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 718
% Passive clauses: 7994
% Generated clauses: 20970
% Final active clauses: 711
% Final passive clauses: 6824
% Input formulas: 54
% Initial clauses: 208
% 
% Duplicate literals: 471
% Trivial inequalities: 76
% Fw subsumption resolutions: 2018
% Fw demodulations: 845
% Bw demodulations: 7
% 
% Simple tautologies: 283
% Equational tautologies: 57
% Forward subsumptions: 8522
% 
% Binary resolution: 9028
% Factoring: 1259
% Forward superposition: 3780
% Backward superposition: 2981
% Self superposition: 38
% Equality factoring: 28
% Equality resolution: 23
% 
% Split clauses: 561
% Split components: 1131
% SAT solver clauses: 968
% SAT solver unit clauses: 27
% SAT solver binary clauses: 165
% 
% Sat splits: 561
% 
% TWLsolver clauses: 968
% TWLsolver calls for satisfiability: 395
% 
% Memory used [KB]: 8571
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real
% (87695)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 569
% Passive clauses: 3041
% Generated clauses: 20840
% Final active clauses: 344
% Final passive clauses: 1428
% Input formulas: 54
% Initial clauses: 208
% 
% Duplicate literals: 419
% Trivial inequalities: 45
% Fw subsumption resolutions: 1766
% Fw demodulations: 1686
% Bw demodulations: 1112
% Condensations: 535
% Global subsumptions: 22
% 
% Simple tautologies: 324
% Equational tautologies: 44
% Forward subsumptions: 9015
% Bw demodulations to eq. taut.: 18
% 
% Binary resolution: 4073
% Factoring: 99
% Forward superposition: 1860
% Backward superposition: 4943
% Self superposition: 80
% Equality factoring: 46
% Equality resolution: 7
% 
% Split clauses: 307
% Split components: 623
% SAT solver clauses: 3158
% SAT solver unit clauses: 215
% SAT solver binary clauses: 98
% SAT solver learnt clauses: 45
% SAT solver learnt literals: 13
% 
% Sat splits: 307
% Sat splitting refutations: 100
% 
% TWLsolver clauses: 3069
% TWLsolver calls for satisfiability: 1796
% 
% Memory used [KB]: 12920
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f13464,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f626,f635,f633,f1645,f2981,f2985,f2988,f2991,f2994,f2997,f3000,f3003,f3006,f3009,f3036,f2982,f3038,f2986,f3040,f2989,f3042,f2992,f3044,f2995,f3046,f2998,f3048,f3001,f3050,f3004,f3052,f3007,f3054,f3010,f3162,f3163,f3288,f3289,f3290,f3291,f3292,f3293,f3294,f3295,f3296,f3297,f3298,f3299,f3300,f3301,f3302,f3303,f3304,f3305,f3306,f3307,f3308,f3309,f3310,f3311,f3312,f3313,f3314,f3315,f3316,f3317,f3318,f3319,f3320,f3321,f3322,f3323,f3324,f3325,f3326,f3327,f3328,f3329,f3330,f3331,f3332,f3333,f3334,f3335,f3336,f3337,f3338,f3339,f3340,f3341,f3342,f3343,f3344,f3345,f4460,f4463,f4461,f4489,f4492,f4490,f4517,f4520,f4518,f4551,f4554,f4552,f4588,f4591,f4589,f4618,f4621,f4619,f4653,f4656,f4654,f4720,f4723,f4721,f5320,f5323,f5321,f5358,f5361,f5359,f5402,f5405,f5403,f5445,f5448,f5446,f5489,f5492,f5490,f5534,f5537,f5535,f5583,f5586,f5584,f5641,f5644,f2983,f5642,f5672,f5673,f5680,f5677,f5674,f5675,f5683,f5676,f5685,f5678,f5682,f5761,f5762,f5774,f5772,f11435,f11934,f11938,f11939,f11940,f11941,f11942,f11943,f11944,f11945,f11946,f11947,f11948,f11949,f11950,f11951,f11952,f11953,f11954,f11955,f11956,f11959,f11960,f11961,f11976,f11977,f12001,f12006,f12025,f12027,f3957,f1649,f1647,f12510,f12512,f12513,f12514,f12515,f12516,f12517,f12518,f12519,f12520,f12521,f12522,f12523,f12524,f12525,f12526,f12527,f12528,f12529,f12530,f12531,f12532,f12533,f12534,f12535,f12536,f12545,f12548,f12999,f13003,f13001,f13391,f13392,f13393,f13394,f13395,f13397])).
fof(f13397,plain,(
  $false | ($spl2 | $spl81)),
  inference(subsumption_resolution,[],[f13396,f265])).
fof(f265,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.real.REAL_POS_LE_ZERO',aHL_TRUTH)).
fof(f13396,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl81)),
  inference(forward_demodulation,[],[f13370,f12846])).
fof(f12846,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | $spl2),
  inference(subsumption_resolution,[],[f12728,f484])).
fof(f484,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f268,f265])).
fof(f268,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f31])).
fof(f31,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.real.REAL_POS_LE_ZERO','thm.bool.NOT_CLAUSES|split|1')).
fof(f12728,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | $spl2),
  inference(superposition,[],[f1354,f634])).
fof(f634,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(cnf_transformation,[],[f634_D])).
fof(f634_D,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1354,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(subsumption_resolution,[],[f1340,f484])).
fof(f1340,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(superposition,[],[f348,f274])).
fof(f274,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X7))),
  file('thm.real.REAL_POS_LE_ZERO','thm.bool.BOOL_CASES_AX')).
fof(f348,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6,X12] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X12)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X6)))))),
  file('thm.real.REAL_POS_LE_ZERO','thm.real.REAL_LE_TOTAL')).
fof(f13370,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)))) | ($spl2 | $spl81)),
  inference(backward_demodulation,[],[f13352,f13003])).
fof(f13352,plain,(
  s('type.realax.real',sK12) = s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))) | $spl2),
  inference(forward_demodulation,[],[f13225,f344])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X11] : ! [X8,X9] : s(X11,X8) = s(X11,'const.bool.COND_3'(s(bool,'T_0'),s(X11,X8),s(X11,X9)))),
  file('thm.real.REAL_POS_LE_ZERO','thm.bool.COND_CLAUSES|split|0')).
fof(f13225,plain,(
  s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(superposition,[],[f272,f12846])).
fof(f272,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.pos_1'(s('type.realax.real',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : s('type.realax.real','const.real.pos_1'(s('type.realax.real',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X6] : s('type.realax.real','const.real.pos_1'(s('type.realax.real',X6))) = s('type.realax.real','const.bool.COND_3'(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6))),s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.real.REAL_POS_LE_ZERO','thm.real.pos_def')).
fof(f13395,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f13357,f12572])).
fof(f12572,plain,(
  ~sP5('T_0','T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl0),
  inference(unit_resulting_resolution,[],[f265,f12551,f419])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f122,f158,f157,f156,f155])).
fof(f155,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f156,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f157,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f158,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f122,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X15,X16,X17] : ((p(s(bool,X17)) <=> s(bool,X15) = s(bool,X16)) <=> ((p(s(bool,X17)) | p(s(bool,X16)) | p(s(bool,X15))) & (p(s(bool,X17)) | ~p(s(bool,X15)) | ~p(s(bool,X16))) & (p(s(bool,X16)) | ~p(s(bool,X15)) | ~p(s(bool,X17))) & (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)))))),
  file('thm.real.REAL_POS_LE_ZERO','thm.sat.sat_2|49|_prove_ltb')).
fof(f12551,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl0),
  inference(forward_demodulation,[],[f264,f633])).
fof(f264,plain,(
  s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f129])).
fof(f129,plain,(
  ? [X0] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ~! [X0] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',X6))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',X6))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.real.REAL_POS_LE_ZERO','thm.real.REAL_POS_LE_ZERO')).
fof(f13357,plain,(
  sP5('T_0','T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f13352,f12869])).
fof(f12869,plain,(
  sP5('T_0','T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl0),
  inference(unit_resulting_resolution,[],[f265,f633,f418])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f240])).
fof(f13394,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f13356,f12573])).
fof(f12573,plain,(
  sP5('F_0','T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl0),
  inference(unit_resulting_resolution,[],[f484,f12551,f417])).
fof(f417,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f240])).
fof(f13356,plain,(
  ~sP5('F_0','T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f13352,f12868])).
fof(f12868,plain,(
  ~sP5('F_0','T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl0),
  inference(unit_resulting_resolution,[],[f484,f633,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f240])).
fof(f13393,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f13355,f12567])).
fof(f12567,plain,(
  ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f265,f12551,f419])).
fof(f13355,plain,(
  sP5('T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f13352,f12867])).
fof(f12867,plain,(
  sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f265,f633,f418])).
fof(f13392,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f13354,f12568])).
fof(f12568,plain,(
  sP5('F_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f484,f12551,f417])).
fof(f13354,plain,(
  ~sP5('F_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f13352,f12866])).
fof(f12866,plain,(
  ~sP5('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f484,f633,f420])).
fof(f13391,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f13353,f12551])).
fof(f13353,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f13352,f633])).
fof(f13001,plain,(
  s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl78),
  inference(cnf_transformation,[],[f13001_D])).
fof(f13001_D,plain,(
  s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f13003,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))))) | $spl81),
  inference(cnf_transformation,[],[f13003_D])).
fof(f13003_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f12999,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))))) | s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(subsumption_resolution,[],[f12874,f265])).
fof(f12874,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))))) | s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(superposition,[],[f353,f633])).
fof(f353,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & (s('type.realax.real',X0) != s('type.realax.real',X1) | (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0,X1] : (((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & (s('type.realax.real',X0) != s('type.realax.real',X1) | (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))))),
  inference(nnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) <=> s('type.realax.real',X0) = s('type.realax.real',X1))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X6,X12] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X12)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X6))))) <=> s('type.realax.real',X6) = s('type.realax.real',X12))),
  file('thm.real.REAL_POS_LE_ZERO','thm.real.REAL_LE_ANTISYM')).
fof(f12548,plain,(
  $false | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f12547,f506])).
fof(f506,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f484,f285])).
fof(f285,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X7] : (s(bool,'T_0') = s(bool,X7) <=> p(s(bool,X7)))),
  file('thm.real.REAL_POS_LE_ZERO','thm.bool.EQ_CLAUSES|split|0')).
fof(f12547,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f12546,f1169])).
fof(f1169,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f265,f1152,f352])).
fof(f352,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.real.REAL_POS_LE_ZERO','thm.bool.IMP_ANTISYM_AX')).
fof(f1152,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) )),
  inference(forward_demodulation,[],[f1146,f345])).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X11] : ! [X8,X9] : s(X11,X9) = s(X11,'const.bool.COND_3'(s(bool,'F_0'),s(X11,X8),s(X11,X9)))),
  file('thm.real.REAL_POS_LE_ZERO','thm.bool.COND_CLAUSES|split|1')).
fof(f1146,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X1),s('type.realax.real',X0))))))) )),
  inference(unit_resulting_resolution,[],[f345,f355])).
fof(f355,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real',X1) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f211])).
fof(f12546,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12))) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f634,f1647])).
fof(f12545,plain,(
  $false | $spl4),
  inference(trivial_inequality_removal,[],[f12544])).
fof(f12544,plain,(
  s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12))) != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12))) | $spl4),
  inference(forward_demodulation,[],[f12543,f12295])).
fof(f12295,plain,(
  s('type.realax.real',sK12) = s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))) | $spl4),
  inference(backward_demodulation,[],[f1647,f9273])).
fof(f9273,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pos_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f9145,f344])).
fof(f9145,plain,(
  s('type.realax.real','const.real.pos_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f272,f1169])).
fof(f12543,plain,(
  s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12))) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12))) | $spl4),
  inference(forward_demodulation,[],[f264,f1647])).
fof(f12536,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12487,f1191])).
fof(f1191,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) )),
  inference(unit_resulting_resolution,[],[f1152,f406])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f233])).
fof(f233,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f232])).
fof(f232,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f157])).
fof(f12487,plain,(
  ( ! [X0] : (~sP4('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)))) ) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12256])).
fof(f12256,plain,(
  ( ! [X0] : (~sP4('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)))) ) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f7311])).
fof(f7311,plain,(
  ( ! [X0] : (~sP4('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f265,f1201,f1200,f1215,f402])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f230])).
fof(f230,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f229])).
fof(f229,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f158])).
fof(f1215,plain,(
  ( ! [X0] : (~sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f642,f1152,f419])).
fof(f642,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(backward_demodulation,[],[f636,f264])).
fof(f636,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f635,f274])).
fof(f1200,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) )),
  inference(unit_resulting_resolution,[],[f1152,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f236])).
fof(f236,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f235])).
fof(f235,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f156])).
fof(f1201,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) )),
  inference(unit_resulting_resolution,[],[f1152,f414])).
fof(f414,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f155])).
fof(f12535,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12486,f1188])).
fof(f1188,plain,(
  ( ! [X0,X1] : (~sP4('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1152,f484,f1152,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f234])).
fof(f12486,plain,(
  ( ! [X0] : (sP4('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'F_0')) ) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12255])).
fof(f12255,plain,(
  ( ! [X0] : (sP4('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'F_0')) ) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f7236])).
fof(f7236,plain,(
  ( ! [X0] : (sP4('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f1209,f400])).
fof(f400,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f231])).
fof(f1209,plain,(
  ( ! [X0] : (sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f675,f1152,f418])).
fof(f675,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f642,f274])).
fof(f12534,plain,(
  $false | ($spl3 | $spl4 | $spl12)),
  inference(subsumption_resolution,[],[f12439,f1173])).
fof(f1173,plain,(
  ( ! [X0] : (~sP1('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f265,f1152,f376])).
fof(f376,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f221])).
fof(f221,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f12439,plain,(
  sP1('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'T_0') | ($spl3 | $spl4 | $spl12)),
  inference(backward_demodulation,[],[f12295,f12187])).
fof(f12187,plain,(
  sP1('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'T_0') | ($spl3 | $spl4 | $spl12)),
  inference(backward_demodulation,[],[f1647,f4144])).
fof(f4144,plain,(
  sP1('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | ($spl3 | $spl12)),
  inference(subsumption_resolution,[],[f4129,f484])).
fof(f4129,plain,(
  p(s(bool,'F_0')) | sP1('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | ($spl3 | $spl12)),
  inference(superposition,[],[f2986,f675])).
fof(f12533,plain,(
  $false | ($spl3 | $spl4 | $spl8)),
  inference(subsumption_resolution,[],[f12438,f1175])).
fof(f1175,plain,(
  ( ! [X0] : (~sP1('T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f265,f1152,f376])).
fof(f12438,plain,(
  sP1('T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12))) | ($spl3 | $spl4 | $spl8)),
  inference(backward_demodulation,[],[f12295,f12186])).
fof(f12186,plain,(
  sP1('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12))) | ($spl3 | $spl4 | $spl8)),
  inference(backward_demodulation,[],[f1647,f4143])).
fof(f4143,plain,(
  sP1('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl3 | $spl8)),
  inference(subsumption_resolution,[],[f4128,f484])).
fof(f4128,plain,(
  p(s(bool,'F_0')) | sP1('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl3 | $spl8)),
  inference(superposition,[],[f2982,f675])).
fof(f12532,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12429,f1202])).
fof(f1202,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f1152,f415])).
fof(f415,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f239])).
fof(f12429,plain,(
  ~sP2('F_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12166])).
fof(f12166,plain,(
  ~sP2('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f1605])).
fof(f1605,plain,(
  ~sP2('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f948,f492,f489,f484,f404])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f231])).
fof(f489,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f484,f407])).
fof(f407,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f234])).
fof(f492,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f484,f411])).
fof(f411,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f237])).
fof(f948,plain,(
  ~sP5('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f484,f675,f420])).
fof(f12531,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12428,f1203])).
fof(f1203,plain,(
  ( ! [X2,X0,X1] : (sP2('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f1152,f416])).
fof(f416,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f239])).
fof(f12428,plain,(
  ~sP2('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'F_0','F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12165])).
fof(f12165,plain,(
  ~sP2('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'F_0','F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f1575])).
fof(f1575,plain,(
  ~sP2('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f950,f493,f490,f484,f403])).
fof(f403,plain,(
  ( ! [X2,X0,X1] : (~sP2(X2,X1,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f231])).
fof(f490,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f484,f408])).
fof(f408,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f234])).
fof(f493,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f484,f412])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f237])).
fof(f950,plain,(
  ~sP5('F_0','F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f484,f675,f420])).
fof(f12530,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12427,f1191])).
fof(f12427,plain,(
  ~sP4('T_0','T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12164])).
fof(f12164,plain,(
  ~sP4('T_0','T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f1545])).
fof(f1545,plain,(
  ~sP4('T_0','T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f735,f491,f494,f265,f402])).
fof(f494,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f265,f414])).
fof(f491,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f265,f410])).
fof(f735,plain,(
  ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f265,f642,f419])).
fof(f12529,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12381,f6553])).
fof(f6553,plain,(
  ( ! [X0,X1] : (sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f265,f496,f1191,f1200,f402])).
fof(f496,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f265,f416])).
fof(f12381,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'T_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12117])).
fof(f12117,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'T_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f1409])).
fof(f1409,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f642,f1353,f419])).
fof(f1353,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))))) | $spl3),
  inference(subsumption_resolution,[],[f1339,f484])).
fof(f1339,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))))) | $spl3),
  inference(superposition,[],[f348,f675])).
fof(f12528,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12379,f12077])).
fof(f12077,plain,(
  ( ! [X0] : (sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)))) ) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f1208])).
fof(f1208,plain,(
  ( ! [X0] : (sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f636,f1152,f418])).
fof(f12379,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12115])).
fof(f12115,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f1407])).
fof(f1407,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f642,f1353,f419])).
fof(f12527,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12371,f1182])).
fof(f1182,plain,(
  ( ! [X0,X1] : (~sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1152,f484,f1152,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f231])).
fof(f12371,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12107])).
fof(f12107,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f1399])).
fof(f1399,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f675,f1353,f418])).
fof(f12526,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12348,f6553])).
fof(f12348,plain,(
  ( ! [X0] : (~sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'T_0')) ) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12080])).
fof(f12080,plain,(
  ( ! [X0] : (~sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'T_0')) ) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f1215])).
fof(f12525,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12347,f12077])).
fof(f12347,plain,(
  ( ! [X0] : (~sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)))) ) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12079])).
fof(f12079,plain,(
  ( ! [X0] : (~sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)))) ) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f1213])).
fof(f1213,plain,(
  ( ! [X0] : (~sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f642,f1152,f419])).
fof(f12524,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12345,f1182])).
fof(f12345,plain,(
  ( ! [X0] : (sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'F_0')) ) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12076])).
fof(f12076,plain,(
  ( ! [X0] : (sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'F_0')) ) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f1205])).
fof(f1205,plain,(
  ( ! [X0] : (sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f675,f1152,f418])).
fof(f12523,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12341,f1187])).
fof(f1187,plain,(
  ( ! [X0] : (~sP4('T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f265,f484,f1152,f405])).
fof(f12341,plain,(
  sP4('T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12064])).
fof(f12064,plain,(
  sP4('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f1027])).
fof(f1027,plain,(
  sP4('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f951,f400])).
fof(f951,plain,(
  sP5('T_0','F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f265,f675,f418])).
fof(f12522,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12340,f6177])).
fof(f6177,plain,(
  ( ! [X0] : (~sP5('T_0','F_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f1187,f400])).
fof(f12340,plain,(
  sP5('T_0','F_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12062])).
fof(f12062,plain,(
  sP5('T_0','F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f951])).
fof(f12521,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12339,f6667])).
fof(f6667,plain,(
  ( ! [X0] : (sP5('F_0','F_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f484,f493,f490,f1203,f403])).
fof(f12339,plain,(
  ~sP5('F_0','F_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12061])).
fof(f12061,plain,(
  ~sP5('F_0','F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f950])).
fof(f12520,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12338,f1181])).
fof(f1181,plain,(
  ( ! [X0] : (~sP5('T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f265,f484,f1152,f401])).
fof(f12338,plain,(
  sP5('T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12060])).
fof(f12060,plain,(
  sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f949])).
fof(f949,plain,(
  sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f265,f675,f418])).
fof(f12519,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12337,f6630])).
fof(f6630,plain,(
  ( ! [X0] : (sP5('F_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f484,f492,f490,f1202,f404])).
fof(f12337,plain,(
  ~sP5('F_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12059])).
fof(f12059,plain,(
  ~sP5('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f948])).
fof(f12518,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12336,f1196])).
fof(f1196,plain,(
  ( ! [X0] : (~sP3('T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f265,f484,f1152,f409])).
fof(f409,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f237])).
fof(f12336,plain,(
  sP3('T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12054])).
fof(f12054,plain,(
  sP3('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f851])).
fof(f851,plain,(
  sP3('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f716,f399])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f231])).
fof(f716,plain,(
  sP5('F_0','T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f484,f642,f417])).
fof(f12517,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12335,f1192])).
fof(f1192,plain,(
  ( ! [X0] : (~sP3('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f265,f484,f1152,f409])).
fof(f12335,plain,(
  sP3('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'T_0','F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12053])).
fof(f12053,plain,(
  sP3('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'T_0','F_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f798])).
fof(f798,plain,(
  sP3('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f714,f399])).
fof(f714,plain,(
  sP5('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f484,f642,f417])).
fof(f12516,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12334,f9133])).
fof(f9133,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f265,f1169,f418])).
fof(f12334,plain,(
  ~sP5('T_0','T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12050])).
fof(f12050,plain,(
  ~sP5('T_0','T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f738])).
fof(f738,plain,(
  ~sP5('T_0','T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f265,f642,f419])).
fof(f12515,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12333,f6256])).
fof(f6256,plain,(
  ( ! [X0] : (sP5('T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f265,f491,f496,f1191,f402])).
fof(f12333,plain,(
  ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'T_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12049])).
fof(f12049,plain,(
  ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'T_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f735])).
fof(f12514,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12332,f6478])).
fof(f6478,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f1196,f399])).
fof(f12332,plain,(
  sP5('F_0','T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12045])).
fof(f12045,plain,(
  sP5('F_0','T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f716])).
fof(f12513,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12331,f6290])).
fof(f6290,plain,(
  ( ! [X0] : (~sP5('F_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1192,f399])).
fof(f12331,plain,(
  sP5('F_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12)),'T_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12044])).
fof(f12044,plain,(
  sP5('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'T_0') | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f714])).
fof(f12512,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12511,f506])).
fof(f12511,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl4)),
  inference(forward_demodulation,[],[f12330,f1169])).
fof(f12330,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12041])).
fof(f12041,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f675])).
fof(f12510,plain,(
  $false | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f12329,f1169])).
fof(f12329,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f12295,f12030])).
fof(f12030,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12))) | ($spl3 | $spl4)),
  inference(backward_demodulation,[],[f1647,f642])).
fof(f1647,plain,(
  s('type.realax.real',sK12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl4),
  inference(cnf_transformation,[],[f1647_D])).
fof(f1647_D,plain,(
  s('type.realax.real',sK12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1649,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)))) | $spl7),
  inference(cnf_transformation,[],[f1649_D])).
fof(f1649_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f3957,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)))) | s('type.realax.real',sK12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(subsumption_resolution,[],[f3943,f265])).
fof(f3943,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)))) | s('type.realax.real',sK12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(superposition,[],[f353,f636])).
fof(f12027,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f12026,f1191])).
fof(f12026,plain,(
  ~sP4('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | ($spl3 | $spl7)),
  inference(forward_demodulation,[],[f11423,f11332])).
fof(f11332,plain,(
  s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl7),
  inference(superposition,[],[f5955,f345])).
fof(f5955,plain,(
  s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl7),
  inference(unit_resulting_resolution,[],[f1653,f1837])).
fof(f1837,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.pos_1'(s('type.realax.real',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) )),
  inference(superposition,[],[f272,f274])).
fof(f1653,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | $spl7),
  inference(unit_resulting_resolution,[],[f1649,f285])).
fof(f11423,plain,(
  ~sP4('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl3 | $spl7)),
  inference(superposition,[],[f7311,f5955])).
fof(f12025,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f12024,f1188])).
fof(f12024,plain,(
  sP4('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | ($spl3 | $spl7)),
  inference(forward_demodulation,[],[f11422,f11332])).
fof(f11422,plain,(
  sP4('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | ($spl3 | $spl7)),
  inference(superposition,[],[f7236,f5955])).
fof(f12006,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f12005,f6553])).
fof(f12005,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0') | ($spl3 | $spl7)),
  inference(forward_demodulation,[],[f11369,f11332])).
fof(f11369,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | ($spl3 | $spl7)),
  inference(superposition,[],[f1215,f5955])).
fof(f12001,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f12000,f1182])).
fof(f12000,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | ($spl3 | $spl7)),
  inference(forward_demodulation,[],[f11363,f11332])).
fof(f11363,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | ($spl3 | $spl7)),
  inference(superposition,[],[f1205,f5955])).
fof(f11977,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11783,f1191])).
fof(f11783,plain,(
  ( ! [X0] : (~sP4('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))))) ) | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f7311])).
fof(f11976,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11782,f1188])).
fof(f11782,plain,(
  ( ! [X0] : (sP4('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0')) ) | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f7236])).
fof(f11961,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11714,f11609])).
fof(f11609,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f1707])).
fof(f1707,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl3 | $spl7)),
  inference(unit_resulting_resolution,[],[f675,f1649,f420])).
fof(f11714,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real',sK12)),'F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f4700])).
fof(f4700,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | ($spl3 | $spl7)),
  inference(unit_resulting_resolution,[],[f484,f493,f1674,f1397,f403])).
fof(f1397,plain,(
  ( ! [X0,X1] : (sP2('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),X0,X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f1353,f416])).
fof(f1674,plain,(
  ( ! [X0,X1] : (sP4('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1649,f408])).
fof(f11960,plain,(
  $false | ($spl3 | $spl7 | $spl12)),
  inference(subsumption_resolution,[],[f11694,f1173])).
fof(f11694,plain,(
  sP1('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0') | ($spl3 | $spl7 | $spl12)),
  inference(backward_demodulation,[],[f11332,f4144])).
fof(f11959,plain,(
  $false | ($spl3 | $spl7 | $spl8)),
  inference(subsumption_resolution,[],[f11693,f1175])).
fof(f11693,plain,(
  sP1('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | ($spl3 | $spl7 | $spl8)),
  inference(backward_demodulation,[],[f11332,f4143])).
fof(f11956,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11576,f1202])).
fof(f11576,plain,(
  ~sP2('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f1605])).
fof(f11955,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11575,f1203])).
fof(f11575,plain,(
  ~sP2('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0','F_0') | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f1575])).
fof(f11954,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11574,f1191])).
fof(f11574,plain,(
  ~sP4('T_0','T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f1545])).
fof(f11953,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11527,f6553])).
fof(f11527,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0') | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f1409])).
fof(f11952,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11517,f1182])).
fof(f11517,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f1399])).
fof(f11951,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11490,f6553])).
fof(f11490,plain,(
  ( ! [X0] : (~sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0')) ) | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f1215])).
fof(f11950,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11486,f1182])).
fof(f11486,plain,(
  ( ! [X0] : (sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0')) ) | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f1205])).
fof(f11949,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11474,f1187])).
fof(f11474,plain,(
  sP4('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f1027])).
fof(f11948,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11472,f6177])).
fof(f11472,plain,(
  sP5('T_0','F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f951])).
fof(f11947,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11471,f6667])).
fof(f11471,plain,(
  ~sP5('F_0','F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f950])).
fof(f11946,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11470,f1181])).
fof(f11470,plain,(
  sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f949])).
fof(f11945,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11469,f6630])).
fof(f11469,plain,(
  ~sP5('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f948])).
fof(f11944,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11464,f1196])).
fof(f11464,plain,(
  sP3('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f851])).
fof(f11943,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11463,f1192])).
fof(f11463,plain,(
  sP3('const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0','F_0') | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f798])).
fof(f11942,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11460,f9133])).
fof(f11460,plain,(
  ~sP5('T_0','T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f738])).
fof(f11941,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11459,f6256])).
fof(f11459,plain,(
  ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0') | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f735])).
fof(f11940,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11455,f6478])).
fof(f11455,plain,(
  sP5('F_0','T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f716])).
fof(f11939,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11454,f6290])).
fof(f11454,plain,(
  sP5('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0') | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f714])).
fof(f11938,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11937,f506])).
fof(f11937,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl7)),
  inference(forward_demodulation,[],[f11451,f1169])).
fof(f11451,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f675])).
fof(f11934,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11440,f1169])).
fof(f11440,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f11332,f642])).
fof(f11435,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f11434,f484])).
fof(f11434,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl7)),
  inference(forward_demodulation,[],[f11433,f675])).
fof(f11433,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl7),
  inference(forward_demodulation,[],[f11325,f345])).
fof(f11325,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.bool.COND_3'(s(bool,'F_0'),s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f5955,f355])).
fof(f5772,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK13(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK13(X0,X1,bool,bool))))) ) | $spl74),
  inference(cnf_transformation,[],[f5772_D])).
fof(f5772_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK13(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK13(X0,X1,bool,bool))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f5774,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl77),
  inference(cnf_transformation,[],[f5774_D])).
fof(f5774_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f5762,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK13(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK13(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1872,f592])).
fof(f592,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f274,f274])).
fof(f1872,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK13(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK13(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f385,f274])).
fof(f385,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK13(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK13(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK13(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK13(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.real.REAL_POS_LE_ZERO',aHL_EXT)).
fof(f5761,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK13(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK13(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1872,f592])).
fof(f5682,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl68),
  inference(cnf_transformation,[],[f5682_D])).
fof(f5682_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f5678,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK13(X0,X1,bool,bool))) ) | $spl64),
  inference(cnf_transformation,[],[f5678_D])).
fof(f5678_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK13(X0,X1,bool,bool))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f5685,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl73),
  inference(cnf_transformation,[],[f5685_D])).
fof(f5685_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f5676,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK13(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1871,f274])).
fof(f1871,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK13(X0,X1,X2,bool))) )),
  inference(superposition,[],[f385,f274])).
fof(f5683,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl70),
  inference(cnf_transformation,[],[f5683_D])).
fof(f5683_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f5675,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK13(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1871,f592])).
fof(f5674,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK13(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1871,f592])).
fof(f5677,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl62),
  inference(cnf_transformation,[],[f5677_D])).
fof(f5677_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f5680,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl67),
  inference(cnf_transformation,[],[f5680_D])).
fof(f5680_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f5673,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK13(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1871,f592])).
fof(f5672,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK13(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1871,f592])).
fof(f5642,plain,(
  ( ! [X2,X3,X1] : (~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) ) | $spl60),
  inference(cnf_transformation,[],[f5642_D])).
fof(f5642_D,plain,(
  ( ! [X2,X3,X1] : (~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f2983,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f2983_D])).
fof(f2983_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f5644,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) )),
  inference(condensation,[],[f5643])).
fof(f5643,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f5630,f693])).
fof(f693,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f688,f484])).
fof(f688,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f349,f274])).
fof(f349,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f144])).
fof(f5630,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f838,f592])).
fof(f838,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X59),s(bool,X60)),X61,X62) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X59),s(bool,X60)),X61,X62)) )),
  inference(superposition,[],[f473,f345])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f441])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f249])).
fof(f249,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f5641,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) )),
  inference(condensation,[],[f5640])).
fof(f5640,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f5629,f693])).
fof(f5629,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~sP8('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f838,f592])).
fof(f5584,plain,(
  ( ! [X2,X3,X1] : (~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) ) | $spl58),
  inference(cnf_transformation,[],[f5584_D])).
fof(f5584_D,plain,(
  ( ! [X2,X3,X1] : (~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f5586,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) )),
  inference(condensation,[],[f5585])).
fof(f5585,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f5572,f693])).
fof(f5572,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X30),s(bool,'T_0')),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f837,f592])).
fof(f837,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),X57,X58) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X55),s(bool,X56)),X57,X58)) )),
  inference(superposition,[],[f472,f345])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f458])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f257])).
fof(f257,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f256])).
fof(f256,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f5583,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,'T_0')),X2,X3)) )),
  inference(condensation,[],[f5582])).
fof(f5582,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f5571,f693])).
fof(f5571,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~sP10('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X25),s(bool,'T_0')),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f837,f592])).
fof(f5535,plain,(
  ( ! [X2,X3,X1] : (~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) ) | $spl56),
  inference(cnf_transformation,[],[f5535_D])).
fof(f5535_D,plain,(
  ( ! [X2,X3,X1] : (~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f5537,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) )),
  inference(condensation,[],[f5536])).
fof(f5536,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f5523,f693])).
fof(f5523,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f785,f592])).
fof(f785,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X52)),X53,X54) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X51),s(bool,X52)),X53,X54)) )),
  inference(superposition,[],[f473,f344])).
fof(f5534,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) )),
  inference(condensation,[],[f5533])).
fof(f5533,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f5522,f693])).
fof(f5522,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~sP8('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f785,f592])).
fof(f5490,plain,(
  ( ! [X2,X3,X1] : (~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) ) | $spl54),
  inference(cnf_transformation,[],[f5490_D])).
fof(f5490_D,plain,(
  ( ! [X2,X3,X1] : (~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f5492,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) )),
  inference(condensation,[],[f5491])).
fof(f5491,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f5478,f693])).
fof(f5478,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X30)),X31,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f784,f592])).
fof(f784,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,X47),s(bool,X48)),X49,X50) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X47),s(bool,X48)),X49,X50)) )),
  inference(superposition,[],[f472,f344])).
fof(f5489,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X1)),X2,X3)) )),
  inference(condensation,[],[f5488])).
fof(f5488,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f5477,f693])).
fof(f5477,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~sP10('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X25)),X26,X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f784,f592])).
fof(f5446,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2,'T_0') | p(s(bool,X2))) ) | $spl52),
  inference(cnf_transformation,[],[f5446_D])).
fof(f5446_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2,'T_0') | p(s(bool,X2))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f5448,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2,'T_0')) )),
  inference(condensation,[],[f5447])).
fof(f5447,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,X30)),X31,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f5434,f693])).
fof(f5434,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,X30)),X31,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1133,f592])).
fof(f1133,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,X16)) | ~p(s(bool,X15)) | p(s(bool,X17)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X14),s(bool,X15)),X17,X16)) )),
  inference(superposition,[],[f469,f345])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f150,f167,f166])).
fof(f166,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f150,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X15,X16,X17] : ((p(s(bool,X17)) <=> (p(s(bool,X16)) => p(s(bool,X15)))) <=> ((p(s(bool,X17)) | p(s(bool,X16))) & (p(s(bool,X17)) | ~p(s(bool,X15))) & (~p(s(bool,X16)) | p(s(bool,X15)) | ~p(s(bool,X17)))))),
  file('thm.real.REAL_POS_LE_ZERO','thm.sat.sat_27_prove_ltb')).
fof(f5445,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X1)),X2,'T_0')) )),
  inference(condensation,[],[f5444])).
fof(f5444,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | p(s(bool,X26)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,X25)),X26,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f5433,f693])).
fof(f5433,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,X25)),X26,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1133,f592])).
fof(f5403,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2,'T_0') | p(s(bool,X2))) ) | $spl50),
  inference(cnf_transformation,[],[f5403_D])).
fof(f5403_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2,'T_0') | p(s(bool,X2))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f5405,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2,'T_0')) )),
  inference(condensation,[],[f5404])).
fof(f5404,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X32)),X31,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f5391,f693])).
fof(f5391,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X32)),X31,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1132,f592])).
fof(f1132,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,X12)) | ~p(s(bool,X10)) | p(s(bool,X13)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)),X13,X12)) )),
  inference(superposition,[],[f469,f344])).
fof(f5402,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X3)),X2,'T_0')) )),
  inference(condensation,[],[f5401])).
fof(f5401,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | p(s(bool,X26)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X27)),X26,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f5390,f693])).
fof(f5390,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | ~sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X27)),X26,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1132,f592])).
fof(f5359,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'T_0')),X2,X1)) ) | $spl48),
  inference(cnf_transformation,[],[f5359_D])).
fof(f5359_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'T_0')),X2,X1)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f5361,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'T_0')),X2,X1)) )),
  inference(condensation,[],[f5360])).
fof(f5360,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,'T_0')),X31,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f5347,f693])).
fof(f5347,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | p(s(bool,X31)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X32),s(bool,'T_0')),X31,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1023,f592])).
fof(f1023,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | p(s(bool,X17)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X14),s(bool,X15)),X17,X16)) )),
  inference(superposition,[],[f452,f345])).
fof(f452,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f126,f164,f163])).
fof(f163,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f126,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X15,X16,X17] : ((p(s(bool,X17)) <=> (p(s(bool,X16)) | p(s(bool,X15)))) <=> ((p(s(bool,X17)) | ~p(s(bool,X16))) & (p(s(bool,X17)) | ~p(s(bool,X15))) & (p(s(bool,X16)) | p(s(bool,X15)) | ~p(s(bool,X17)))))),
  file('thm.real.REAL_POS_LE_ZERO','thm.sat.sat_25_prove_ltb')).
fof(f5358,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'T_0')),X2,X1)) )),
  inference(condensation,[],[f5357])).
fof(f5357,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,'T_0')),X26,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f5346,f693])).
fof(f5346,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | p(s(bool,X26)) | ~sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X27),s(bool,'T_0')),X26,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1023,f592])).
fof(f5321,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X3)),X2,X1)) ) | $spl46),
  inference(cnf_transformation,[],[f5321_D])).
fof(f5321_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X3)),X2,X1)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f5323,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X3)),X2,X1)) )),
  inference(condensation,[],[f5322])).
fof(f5322,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X32)),X31,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f5309,f693])).
fof(f5309,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | p(s(bool,X31)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X32)),X31,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1022,f592])).
fof(f1022,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,X10)) | p(s(bool,X12)) | p(s(bool,X13)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X10),s(bool,X11)),X13,X12)) )),
  inference(superposition,[],[f452,f344])).
fof(f5320,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X3)),X2,X1)) )),
  inference(condensation,[],[f5319])).
fof(f5319,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | p(s(bool,X26)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X27)),X26,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f5308,f693])).
fof(f5308,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | p(s(bool,X26)) | ~sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X27)),X26,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1022,f592])).
fof(f4721,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) ) | $spl44),
  inference(cnf_transformation,[],[f4721_D])).
fof(f4721_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4723,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) )),
  inference(condensation,[],[f4722])).
fof(f4722,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4712,f693])).
fof(f4712,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f836,f592])).
fof(f836,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X53)) | ~p(s(bool,X52)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X51),s(bool,X52)),X53,X54)) )),
  inference(superposition,[],[f468,f345])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f4720,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) )),
  inference(condensation,[],[f4719])).
fof(f4719,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f4711,f693])).
fof(f4711,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f836,f592])).
fof(f4654,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X3,X1)) ) | $spl42),
  inference(cnf_transformation,[],[f4654_D])).
fof(f4654_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X3,X1)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4656,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X3,X1)) )),
  inference(condensation,[],[f4655])).
fof(f4655,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,'T_0')),X32,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4645,f693])).
fof(f4645,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,'T_0')),X32,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f835,f592])).
fof(f835,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X47),s(bool,X48)),X50,X49)) )),
  inference(superposition,[],[f467,f345])).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f263])).
fof(f4653,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,'T_0')),X3,X1)) )),
  inference(condensation,[],[f4652])).
fof(f4652,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,'T_0')),X27,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f4644,f693])).
fof(f4644,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,'T_0')),X27,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f835,f592])).
fof(f4619,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) ) | $spl40),
  inference(cnf_transformation,[],[f4619_D])).
fof(f4619_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4621,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) )),
  inference(condensation,[],[f4620])).
fof(f4620,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4610,f693])).
fof(f4610,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f834,f592])).
fof(f834,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | ~p(s(bool,X44)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X43),s(bool,X44)),X45,X46)) )),
  inference(superposition,[],[f451,f345])).
fof(f451,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f255])).
fof(f4618,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),'T_0',X3)) )),
  inference(condensation,[],[f4617])).
fof(f4617,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f4609,f693])).
fof(f4609,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f834,f592])).
fof(f4589,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,'T_0')) ) | $spl38),
  inference(cnf_transformation,[],[f4589_D])).
fof(f4589_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,'T_0')) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4591,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,'T_0')) )),
  inference(condensation,[],[f4590])).
fof(f4590,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4580,f693])).
fof(f4580,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X31),s(bool,X30)),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f833,f592])).
fof(f833,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X41)) | ~p(s(bool,X40)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X39),s(bool,X40)),X42,X41)) )),
  inference(superposition,[],[f450,f345])).
fof(f450,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f255])).
fof(f4588,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1)),X3,'T_0')) )),
  inference(condensation,[],[f4587])).
fof(f4587,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f4579,f693])).
fof(f4579,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'F_0'),s(bool,X26),s(bool,X25)),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f833,f592])).
fof(f4552,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) ) | $spl36),
  inference(cnf_transformation,[],[f4552_D])).
fof(f4552_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4554,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) )),
  inference(condensation,[],[f4553])).
fof(f4553,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4543,f693])).
fof(f4543,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f783,f592])).
fof(f783,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | ~p(s(bool,X43)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X43),s(bool,X44)),X45,X46)) )),
  inference(superposition,[],[f468,f344])).
fof(f4551,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) )),
  inference(condensation,[],[f4550])).
fof(f4550,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f4542,f693])).
fof(f4542,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f783,f592])).
fof(f4518,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X3,X1)) ) | $spl34),
  inference(cnf_transformation,[],[f4518_D])).
fof(f4518_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X3,X1)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4520,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X3,X1)) )),
  inference(condensation,[],[f4519])).
fof(f4519,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X31)),X32,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4509,f693])).
fof(f4509,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X31)),X32,X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f782,f592])).
fof(f782,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | p(s(bool,X41)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40)),X42,X41)) )),
  inference(superposition,[],[f467,f344])).
fof(f4517,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)),X3,X1)) )),
  inference(condensation,[],[f4516])).
fof(f4516,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X26)),X27,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f4508,f693])).
fof(f4508,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP11('const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X26)),X27,X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f782,f592])).
fof(f4490,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) ) | $spl32),
  inference(cnf_transformation,[],[f4490_D])).
fof(f4490_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4492,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) )),
  inference(condensation,[],[f4491])).
fof(f4491,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4481,f693])).
fof(f4481,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),'T_0',X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f781,f592])).
fof(f781,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | ~p(s(bool,X35)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X35),s(bool,X36)),X37,X38)) )),
  inference(superposition,[],[f451,f344])).
fof(f4489,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),'T_0',X3)) )),
  inference(condensation,[],[f4488])).
fof(f4488,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f4480,f693])).
fof(f4480,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),'T_0',X27) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f781,f592])).
fof(f4461,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,'T_0')) ) | $spl30),
  inference(cnf_transformation,[],[f4461_D])).
fof(f4461_D,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,'T_0')) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4463,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,'T_0')) )),
  inference(condensation,[],[f4462])).
fof(f4462,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4452,f693])).
fof(f4452,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X30),s(bool,X31)),X32,'T_0') | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f780,f592])).
fof(f780,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | ~p(s(bool,X31)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X31),s(bool,X32)),X34,X33)) )),
  inference(superposition,[],[f450,f344])).
fof(f4460,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2)),X3,'T_0')) )),
  inference(condensation,[],[f4459])).
fof(f4459,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f4451,f693])).
fof(f4451,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | sP9('const.bool.COND_3'(s(bool,'T_0'),s(bool,X25),s(bool,X26)),X27,'T_0') | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f780,f592])).
fof(f3345,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3281,f3092])).
fof(f3092,plain,(
  ( ! [X0,X1] : (~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X0)))) ) | ($spl7 | $spl10)),
  inference(backward_demodulation,[],[f2983,f1709])).
fof(f1709,plain,(
  ( ! [X0,X1] : (~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) ) | $spl7),
  inference(unit_resulting_resolution,[],[f345,f1649,f420])).
fof(f3281,plain,(
  ( ! [X0,X1] : (sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X0)))) ) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f3088])).
fof(f3088,plain,(
  ( ! [X0,X1] : (sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X0)))) ) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f2983,f1405])).
fof(f1405,plain,(
  ( ! [X0,X1] : (sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),X0,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X1),s(bool,X0)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f345,f1353,f418])).
fof(f3176,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.pos_1'(s('type.realax.real',X0)))) ) | $spl10),
  inference(forward_demodulation,[],[f3154,f344])).
fof(f3154,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.pos_1'(s('type.realax.real',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,'T_0'),s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl10),
  inference(backward_demodulation,[],[f2983,f272])).
fof(f3344,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3280,f3091])).
fof(f3091,plain,(
  ( ! [X0,X1] : (~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X1)) ) | ($spl7 | $spl10)),
  inference(backward_demodulation,[],[f2983,f1705])).
fof(f1705,plain,(
  ( ! [X0,X1] : (~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f345,f1649,f420])).
fof(f3280,plain,(
  ( ! [X0,X1] : (sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X1)) ) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f3087])).
fof(f3087,plain,(
  ( ! [X0,X1] : (sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X1)) ) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f2983,f1401])).
fof(f1401,plain,(
  ( ! [X0,X1] : (sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,X1)),X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f345,f1353,f418])).
fof(f3343,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3276,f3219])).
fof(f3219,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1403])).
fof(f1403,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f675,f1353,f418])).
fof(f3276,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl3 | $spl7 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1707])).
fof(f3342,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3275,f3216])).
fof(f3216,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1399])).
fof(f3275,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | ($spl3 | $spl7 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1703])).
fof(f1703,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | ($spl3 | $spl7)),
  inference(unit_resulting_resolution,[],[f675,f1649,f420])).
fof(f3341,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3274,f1706])).
fof(f1706,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl3 | $spl7)),
  inference(unit_resulting_resolution,[],[f636,f1649,f420])).
fof(f3274,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl3 | $spl7 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1701])).
fof(f1701,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0','const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl3 | $spl7)),
  inference(unit_resulting_resolution,[],[f642,f1649,f417])).
fof(f3340,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3273,f1702])).
fof(f1702,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | ($spl3 | $spl7)),
  inference(unit_resulting_resolution,[],[f636,f1649,f420])).
fof(f3273,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | ($spl3 | $spl7 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1699])).
fof(f1699,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | ($spl3 | $spl7)),
  inference(unit_resulting_resolution,[],[f642,f1649,f417])).
fof(f3339,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f3270,f725])).
fof(f725,plain,(
  sP4('T_0','T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f723,f400])).
fof(f723,plain,(
  sP5('T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f265,f636,f418])).
fof(f3270,plain,(
  ~sP4('T_0','T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1545])).
fof(f3338,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3269,f1816])).
fof(f1816,plain,(
  ( ! [X0] : (sP11('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)))) ) | ($spl3 | $spl7)),
  inference(unit_resulting_resolution,[],[f1649,f671])).
fof(f671,plain,(
  ( ! [X8,X9] : (sP11('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),X9,X8) | p(s(bool,X8))) ) | $spl3),
  inference(subsumption_resolution,[],[f656,f265])).
fof(f656,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,X8)) | sP11('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),X9,X8)) ) | $spl3),
  inference(superposition,[],[f467,f636])).
fof(f3269,plain,(
  ~sP11('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1513])).
fof(f1513,plain,(
  ~sP11('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | $spl3),
  inference(unit_resulting_resolution,[],[f484,f1353,f1134])).
fof(f1134,plain,(
  ( ! [X6,X7] : (~sP11('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),X7,X6) | p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl3),
  inference(subsumption_resolution,[],[f1130,f265])).
fof(f1130,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),X7,X6)) ) | $spl3),
  inference(superposition,[],[f469,f636])).
fof(f3337,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3265,f1731])).
fof(f1731,plain,(
  ( ! [X0] : (~sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0',X0)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f265,f1649,f439])).
fof(f439,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f250])).
fof(f3265,plain,(
  ( ! [X0] : (sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0',X0)) ) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1509])).
fof(f1509,plain,(
  ( ! [X0] : (sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0',X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f501,f1353,f473])).
fof(f501,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f265,f447])).
fof(f447,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f252])).
fof(f252,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f163])).
fof(f3336,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3264,f1730])).
fof(f1730,plain,(
  ( ! [X0] : (~sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,'T_0')) ) | $spl7),
  inference(unit_resulting_resolution,[],[f265,f1649,f438])).
fof(f438,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f250])).
fof(f3264,plain,(
  ( ! [X0] : (sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,'T_0')) ) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1508])).
fof(f1508,plain,(
  ( ! [X0] : (sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),X0,'T_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f500,f1353,f473])).
fof(f500,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f265,f446])).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f253])).
fof(f3335,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3263,f1781])).
fof(f1781,plain,(
  ( ! [X0] : (~sP11('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0',X0)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f265,f1649,f456])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f258])).
fof(f3263,plain,(
  ( ! [X0] : (sP11('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0',X0)) ) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1507])).
fof(f1507,plain,(
  ( ! [X0] : (sP11('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0',X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f504,f1353,f472])).
fof(f504,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f265,f464])).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f260])).
fof(f260,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f166])).
fof(f3334,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3262,f1775])).
fof(f1775,plain,(
  ( ! [X0] : (~sP11('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,'F_0')) ) | $spl7),
  inference(unit_resulting_resolution,[],[f484,f1649,f455])).
fof(f455,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f258])).
fof(f3262,plain,(
  ( ! [X0] : (sP11('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,'F_0')) ) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1506])).
fof(f1506,plain,(
  ( ! [X0] : (sP11('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),X0,'F_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f503,f1353,f472])).
fof(f503,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f484,f463])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f261])).
fof(f3333,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3260,f1796])).
fof(f1796,plain,(
  sP11('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0','T_0') | $spl7),
  inference(unit_resulting_resolution,[],[f484,f265,f1649,f466])).
fof(f466,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f263])).
fof(f3260,plain,(
  ~sP11('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0','T_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1504])).
fof(f1504,plain,(
  ~sP11('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f265,f484,f1353,f469])).
fof(f3332,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3259,f1798])).
fof(f1798,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)))) ) | $spl7),
  inference(unit_resulting_resolution,[],[f265,f1649,f467])).
fof(f3259,plain,(
  ~sP11('T_0','F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1501])).
fof(f1501,plain,(
  ~sP11('T_0','F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | $spl3),
  inference(unit_resulting_resolution,[],[f265,f484,f1353,f469])).
fof(f3331,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3256,f1777])).
fof(f1777,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)))) ) | $spl7),
  inference(unit_resulting_resolution,[],[f484,f1649,f455])).
fof(f3256,plain,(
  sP11('F_0','F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1493])).
fof(f1493,plain,(
  sP11('F_0','F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | $spl3),
  inference(unit_resulting_resolution,[],[f484,f484,f1353,f466])).
fof(f3330,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3254,f1791])).
fof(f1791,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)))) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1649,f463])).
fof(f3254,plain,(
  ~sP10('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1489])).
fof(f1489,plain,(
  ~sP10('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | $spl3),
  inference(unit_resulting_resolution,[],[f1353,f484,f1353,f462])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f261])).
fof(f3329,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3253,f1791])).
fof(f3253,plain,(
  ~sP10('T_0','F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1488])).
fof(f1488,plain,(
  ~sP10('T_0','F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | $spl3),
  inference(unit_resulting_resolution,[],[f265,f484,f1353,f462])).
fof(f3328,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3252,f1792])).
fof(f1792,plain,(
  ( ! [X0,X1] : (sP10('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1649,f465])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f261])).
fof(f3252,plain,(
  ~sP10('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0','T_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1484])).
fof(f1484,plain,(
  ~sP10('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f484,f265,f1353,f462])).
fof(f3327,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3250,f1765])).
fof(f1765,plain,(
  sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0','F_0') | $spl7),
  inference(unit_resulting_resolution,[],[f484,f484,f1649,f449])).
fof(f449,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f255])).
fof(f3250,plain,(
  ~sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0','F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1470])).
fof(f1470,plain,(
  ~sP9('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f484,f484,f1353,f452])).
fof(f3326,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3243,f1754])).
fof(f1754,plain,(
  ( ! [X0,X1] : (sP8('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1649,f448])).
fof(f448,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f253])).
fof(f3243,plain,(
  ~sP8('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0','F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1455])).
fof(f1455,plain,(
  ~sP8('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f484,f484,f1353,f445])).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f253])).
fof(f3325,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3240,f1719])).
fof(f1719,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)))) ) | $spl7),
  inference(unit_resulting_resolution,[],[f484,f1649,f432])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f124,f161,f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f161,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f124,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X15,X16,X17] : ((p(s(bool,X17)) <=> (p(s(bool,X16)) & p(s(bool,X15)))) <=> ((p(s(bool,X17)) | ~p(s(bool,X16)) | ~p(s(bool,X15))) & (p(s(bool,X16)) | ~p(s(bool,X17))) & (p(s(bool,X15)) | ~p(s(bool,X17)))))),
  file('thm.real.REAL_POS_LE_ZERO','thm.sat.sat_23_prove_ltb')).
fof(f3240,plain,(
  ~sP7('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1445])).
fof(f1445,plain,(
  ~sP7('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | $spl3),
  inference(unit_resulting_resolution,[],[f484,f1353,f1353,f437])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f247])).
fof(f3324,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3239,f1719])).
fof(f3239,plain,(
  ~sP7('F_0','T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1444])).
fof(f1444,plain,(
  ~sP7('F_0','T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | $spl3),
  inference(unit_resulting_resolution,[],[f484,f265,f1353,f437])).
fof(f3323,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3238,f1724])).
fof(f1724,plain,(
  ( ! [X0] : (sP7('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f484,f1649,f433])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f247])).
fof(f3238,plain,(
  ~sP7('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1441])).
fof(f1441,plain,(
  ~sP7('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f484,f265,f1353,f437])).
fof(f3322,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3234,f1729])).
fof(f1729,plain,(
  ~sP7('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0','T_0') | $spl7),
  inference(unit_resulting_resolution,[],[f265,f265,f1649,f437])).
fof(f3234,plain,(
  sP7('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0','T_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1437])).
fof(f1437,plain,(
  sP7('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f265,f265,f1353,f434])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f247])).
fof(f3321,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3233,f1710])).
fof(f1710,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)))) ) | $spl7),
  inference(unit_resulting_resolution,[],[f265,f1649,f422])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f242])).
fof(f242,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f161])).
fof(f3233,plain,(
  sP7('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1434])).
fof(f1434,plain,(
  sP7('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | $spl3),
  inference(unit_resulting_resolution,[],[f265,f1353,f1353,f434])).
fof(f3320,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3232,f1710])).
fof(f3232,plain,(
  sP7('T_0','T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1432])).
fof(f1432,plain,(
  sP7('T_0','T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | $spl3),
  inference(unit_resulting_resolution,[],[f265,f265,f1353,f434])).
fof(f3319,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3231,f1712])).
fof(f1712,plain,(
  ( ! [X0] : (~sP7('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f265,f1649,f423])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f243])).
fof(f3231,plain,(
  sP7('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1427])).
fof(f1427,plain,(
  sP7('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f265,f265,f1353,f434])).
fof(f3318,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3229,f1717])).
fof(f1717,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1649,f430])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f160])).
fof(f3229,plain,(
  ~sP6('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1423])).
fof(f1423,plain,(
  ~sP6('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f1353,f484,f1353,f428])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f245])).
fof(f3317,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3228,f1717])).
fof(f3228,plain,(
  ~sP6('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1422])).
fof(f1422,plain,(
  ~sP6('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f265,f484,f1353,f428])).
fof(f3316,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3227,f1718])).
fof(f1718,plain,(
  ( ! [X0,X1] : (sP6('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1649,f431])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f245])).
fof(f3227,plain,(
  ~sP6('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0','F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1418])).
fof(f1418,plain,(
  ~sP6('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f265,f484,f1353,f428])).
fof(f3315,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3226,f1803])).
fof(f1803,plain,(
  ( ! [X0] : (sP7('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0',X0)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f499,f1649,f474])).
fof(f474,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f427])).
fof(f427,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f243])).
fof(f499,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f484,f431])).
fof(f3226,plain,(
  ( ! [X0] : (~sP7('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0',X0)) ) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1411])).
fof(f1411,plain,(
  ( ! [X0] : (~sP7('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0',X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f484,f1353,f423])).
fof(f3314,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3225,f1802])).
fof(f1802,plain,(
  ( ! [X0] : (sP7('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,'F_0')) ) | $spl7),
  inference(unit_resulting_resolution,[],[f498,f1649,f474])).
fof(f498,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f484,f430])).
fof(f3225,plain,(
  ( ! [X0] : (~sP7('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,'F_0')) ) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1410])).
fof(f1410,plain,(
  ( ! [X0] : (~sP7('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),X0,'F_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f484,f1353,f422])).
fof(f3313,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3223,f1698])).
fof(f1698,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0','T_0') | $spl7),
  inference(unit_resulting_resolution,[],[f506,f1649,f417])).
fof(f3223,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0','T_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1408])).
fof(f1408,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f506,f1353,f419])).
fof(f3312,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3221,f1700])).
fof(f1700,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0','F_0') | $spl7),
  inference(unit_resulting_resolution,[],[f506,f1649,f417])).
fof(f3221,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0','F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1406])).
fof(f1406,plain,(
  ~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f506,f1353,f419])).
fof(f3311,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3220,f1708])).
fof(f1708,plain,(
  ( ! [X0,X1] : (~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl7),
  inference(unit_resulting_resolution,[],[f344,f1649,f420])).
fof(f3220,plain,(
  ( ! [X0,X1] : (sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1404])).
fof(f1404,plain,(
  ( ! [X0,X1] : (sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),X0,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f344,f1353,f418])).
fof(f3310,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3218,f1706])).
fof(f3218,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1402])).
fof(f1402,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f636,f1353,f418])).
fof(f3309,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3217,f1704])).
fof(f1704,plain,(
  ( ! [X0,X1] : (~sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f344,f1649,f420])).
fof(f3217,plain,(
  ( ! [X0,X1] : (sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1400])).
fof(f1400,plain,(
  ( ! [X0,X1] : (sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X1)),X0)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f344,f1353,f418])).
fof(f3308,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3215,f1702])).
fof(f3215,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1398])).
fof(f1398,plain,(
  sP5('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f636,f1353,f418])).
fof(f3307,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3210,f1678])).
fof(f1678,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1649,f411])).
fof(f3210,plain,(
  ~sP3('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1391])).
fof(f1391,plain,(
  ~sP3('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f1353,f484,f1353,f409])).
fof(f3306,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3209,f1678])).
fof(f3209,plain,(
  ~sP3('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1390])).
fof(f1390,plain,(
  ~sP3('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f265,f484,f1353,f409])).
fof(f3305,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3208,f1679])).
fof(f1679,plain,(
  ( ! [X0,X1] : (sP3('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X0,X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1649,f412])).
fof(f3208,plain,(
  ~sP3('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0','F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1386])).
fof(f1386,plain,(
  ~sP3('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f265,f484,f1353,f409])).
fof(f3304,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3206,f1673])).
fof(f1673,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),X1)) ) | $spl7),
  inference(unit_resulting_resolution,[],[f1649,f407])).
fof(f3206,plain,(
  ~sP4('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1382])).
fof(f1382,plain,(
  ~sP4('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f1353,f484,f1353,f405])).
fof(f3303,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3205,f1673])).
fof(f3205,plain,(
  ~sP4('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1381])).
fof(f1381,plain,(
  ~sP4('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f265,f484,f1353,f405])).
fof(f3302,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3204,f1674])).
fof(f3204,plain,(
  ~sP4('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0','F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1377])).
fof(f1377,plain,(
  ~sP4('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f265,f484,f1353,f405])).
fof(f3301,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3202,f1667])).
fof(f1667,plain,(
  sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0') | $spl7),
  inference(unit_resulting_resolution,[],[f492,f489,f494,f1649,f403])).
fof(f3202,plain,(
  ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1375])).
fof(f1375,plain,(
  ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f265,f484,f1353,f401])).
fof(f3300,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3201,f1661])).
fof(f1661,plain,(
  ~sP1('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | $spl7),
  inference(unit_resulting_resolution,[],[f484,f1649,f375])).
fof(f375,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f222])).
fof(f3201,plain,(
  sP1('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1372])).
fof(f1372,plain,(
  sP1('F_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | $spl3),
  inference(unit_resulting_resolution,[],[f484,f1353,f380])).
fof(f380,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f222])).
fof(f3299,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3200,f1659])).
fof(f1659,plain,(
  ~sP1('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0') | $spl7),
  inference(unit_resulting_resolution,[],[f484,f1649,f375])).
fof(f3200,plain,(
  sP1('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1371])).
fof(f1371,plain,(
  sP1('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f484,f1353,f377])).
fof(f377,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f222])).
fof(f3298,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3198,f1663])).
fof(f1663,plain,(
  sP1('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | $spl7),
  inference(unit_resulting_resolution,[],[f265,f1649,f377])).
fof(f3198,plain,(
  ~sP1('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1369])).
fof(f1369,plain,(
  ~sP1('T_0','const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))))) | $spl3),
  inference(unit_resulting_resolution,[],[f265,f1353,f376])).
fof(f3297,plain,(
  $false | ($spl3 | $spl7 | $spl10)),
  inference(subsumption_resolution,[],[f3197,f1664])).
fof(f1664,plain,(
  sP1('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0') | $spl7),
  inference(unit_resulting_resolution,[],[f265,f1649,f380])).
fof(f3197,plain,(
  ~sP1('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)),'T_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1367])).
fof(f1367,plain,(
  ~sP1('const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12)))),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f265,f1353,f376])).
fof(f3296,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f3193,f1204])).
fof(f1204,plain,(
  ( ! [X0] : (sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')) ) | $spl3),
  inference(unit_resulting_resolution,[],[f636,f1152,f418])).
fof(f3193,plain,(
  ( ! [X0] : (~sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')) ) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1215])).
fof(f3295,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f3192,f1208])).
fof(f3192,plain,(
  ( ! [X0] : (~sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f1213])).
fof(f3294,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f3183,f1612])).
fof(f1612,plain,(
  ~sP3('T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f747,f488,f495,f484,f404])).
fof(f495,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f265,f415])).
fof(f488,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f265,f406])).
fof(f747,plain,(
  ~sP5('F_0','T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f484,f636,f420])).
fof(f3183,plain,(
  sP3('T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f851])).
fof(f3293,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f3182,f1539])).
fof(f1539,plain,(
  ~sP3('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f745,f490,f496,f265,f402])).
fof(f745,plain,(
  ~sP5('F_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f484,f636,f420])).
fof(f3182,plain,(
  sP3('const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f798])).
fof(f3292,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f3181,f724])).
fof(f724,plain,(
  sP5('T_0','T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f265,f636,f418])).
fof(f3181,plain,(
  ~sP5('T_0','T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f738])).
fof(f3291,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f3180,f723])).
fof(f3180,plain,(
  ~sP5('T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f735])).
fof(f3290,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f3179,f747])).
fof(f3179,plain,(
  sP5('F_0','T_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f716])).
fof(f3289,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f3178,f745])).
fof(f3178,plain,(
  sP5('F_0','const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f714])).
fof(f3288,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f3177,f2983])).
fof(f3177,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl3 | $spl10)),
  inference(backward_demodulation,[],[f3176,f642])).
fof(f3163,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3059,f265])).
fof(f3059,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(backward_demodulation,[],[f2983,f484])).
fof(f3162,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f3161,f2983])).
fof(f3161,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl10),
  inference(subsumption_resolution,[],[f3056,f285])).
fof(f3056,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl10),
  inference(backward_demodulation,[],[f2983,f289])).
fof(f289,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X7] : (s(bool,'F_0') = s(bool,X7) <=> ~p(s(bool,X7)))),
  file('thm.real.REAL_POS_LE_ZERO','thm.bool.EQ_CLAUSES|split|3')).
fof(f3010,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) ) | $spl28),
  inference(cnf_transformation,[],[f3010_D])).
fof(f3010_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f3054,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f3053])).
fof(f3053,plain,(
  ( ! [X94,X92,X93,X91] : (~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f2681,f693])).
fof(f2681,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f473,f592])).
fof(f3007,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl26),
  inference(cnf_transformation,[],[f3007_D])).
fof(f3007_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3052,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f3051])).
fof(f3051,plain,(
  ( ! [X90,X88,X87,X89] : (~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f2680,f693])).
fof(f2680,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f472,f592])).
fof(f3004,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f3004_D])).
fof(f3004_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3050,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f3049])).
fof(f3049,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f2679,f693])).
fof(f2679,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f469,f592])).
fof(f3001,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f3001_D])).
fof(f3001_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3048,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f3047])).
fof(f3047,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f2678,f693])).
fof(f2678,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f468,f592])).
fof(f2998,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl20),
  inference(cnf_transformation,[],[f2998_D])).
fof(f2998_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3046,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f3045])).
fof(f3045,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f2677,f693])).
fof(f2677,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f467,f592])).
fof(f2995,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl18),
  inference(cnf_transformation,[],[f2995_D])).
fof(f2995_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3044,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f3043])).
fof(f3043,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | p(s(bool,X73)) | ~sP9('T_0',X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f2676,f693])).
fof(f2676,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | p(s(bool,X73)) | ~sP9('T_0',X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f452,f592])).
fof(f2992,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl16),
  inference(cnf_transformation,[],[f2992_D])).
fof(f2992_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3042,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f3041])).
fof(f3041,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X68)) | sP9('T_0',X68,X69) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f2675,f693])).
fof(f2675,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | ~p(s(bool,X68)) | sP9('T_0',X68,X69) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f451,f592])).
fof(f2989,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl14),
  inference(cnf_transformation,[],[f2989_D])).
fof(f2989_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f3040,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f3039])).
fof(f3039,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | sP9('T_0',X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f2674,f693])).
fof(f2674,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | sP9('T_0',X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f450,f592])).
fof(f2986,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl12),
  inference(cnf_transformation,[],[f2986_D])).
fof(f2986_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3038,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f3037])).
fof(f3037,plain,(
  ( ! [X61,X62,X60] : (p(s(bool,X61)) | sP1(X61,'T_0') | s(bool,X60) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f2673,f693])).
fof(f2673,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,X60)) | p(s(bool,X61)) | sP1(X61,'T_0') | s(bool,X60) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f380,f592])).
fof(f2982,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl8),
  inference(cnf_transformation,[],[f2982_D])).
fof(f2982_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f3036,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f3035])).
fof(f3035,plain,(
  ( ! [X59,X57,X58] : (p(s(bool,X58)) | sP1('T_0',X58) | s(bool,X57) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f2672,f693])).
fof(f2672,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,X57)) | p(s(bool,X58)) | sP1('T_0',X58) | s(bool,X57) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f377,f592])).
fof(f3009,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f3008])).
fof(f3008,plain,(
  ( ! [X94,X92,X93,X91] : (~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f2543,f693])).
fof(f2543,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~sP8('T_0',X92,X93) | sP9('T_0',X92,X93) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f473,f592])).
fof(f3006,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f3005])).
fof(f3005,plain,(
  ( ! [X90,X88,X87,X89] : (~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f2542,f693])).
fof(f2542,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~sP10('T_0',X88,X89) | sP11('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f472,f592])).
fof(f3003,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f3002])).
fof(f3002,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f2541,f693])).
fof(f2541,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~p(s(bool,X84)) | p(s(bool,X85)) | ~sP11('T_0',X85,X84) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f469,f592])).
fof(f3000,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f2999])).
fof(f2999,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f2540,f693])).
fof(f2540,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | sP11('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f468,f592])).
fof(f2997,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f2996])).
fof(f2996,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f2539,f693])).
fof(f2539,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | sP11('T_0',X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f467,f592])).
fof(f2994,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2993])).
fof(f2993,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | p(s(bool,X73)) | ~sP9('T_0',X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f2538,f693])).
fof(f2538,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | p(s(bool,X73)) | ~sP9('T_0',X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f452,f592])).
fof(f2991,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2990])).
fof(f2990,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X68)) | sP9('T_0',X68,X69) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f2537,f693])).
fof(f2537,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | ~p(s(bool,X68)) | sP9('T_0',X68,X69) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f451,f592])).
fof(f2988,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2987])).
fof(f2987,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | sP9('T_0',X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f2536,f693])).
fof(f2536,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | sP9('T_0',X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f450,f592])).
fof(f2985,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f2984])).
fof(f2984,plain,(
  ( ! [X61,X62,X60] : (p(s(bool,X61)) | sP1(X61,'T_0') | s(bool,X60) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f2535,f693])).
fof(f2535,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,X60)) | p(s(bool,X61)) | sP1(X61,'T_0') | s(bool,X60) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f380,f592])).
fof(f2981,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f2980])).
fof(f2980,plain,(
  ( ! [X59,X57,X58] : (p(s(bool,X58)) | sP1('T_0',X58) | s(bool,X57) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f2534,f693])).
fof(f2534,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,X57)) | p(s(bool,X58)) | sP1('T_0',X58) | s(bool,X57) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f377,f592])).
fof(f1645,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)))) | s('type.realax.real',sK12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(subsumption_resolution,[],[f1634,f265])).
fof(f1634,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK12)))) | s('type.realax.real',sK12) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(superposition,[],[f353,f636])).
fof(f633,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl0),
  inference(cnf_transformation,[],[f633_D])).
fof(f633_D,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f635,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(cnf_transformation,[],[f635_D])).
fof(f635_D,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f626,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK12),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pos_1'(s('type.realax.real',sK12))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f264,f274])).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 641
% Passive clauses: 2646
% Generated clauses: 13529
% Final active clauses: 340
% Final passive clauses: 581
% Input formulas: 54
% Initial clauses: 208
% 
% Duplicate literals: 315
% Trivial inequalities: 30
% Fw subsumption resolutions: 580
% Fw demodulations: 1228
% Bw demodulations: 1219
% Condensations: 576
% Global subsumptions: 4
% 
% Simple tautologies: 101
% Equational tautologies: 67
% Forward subsumptions: 7795
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 2859
% Factoring: 3
% Forward superposition: 3719
% Backward superposition: 2116
% Self superposition: 42
% Equality factoring: 28
% Equality resolution: 3
% 
% Split clauses: 63
% Split components: 131
% SAT solver clauses: 2902
% SAT solver unit clauses: 255
% SAT solver binary clauses: 1182
% 
% Sat splits: 63
% Sat splitting refutations: 124
% 
% Memory used [KB]: 6524
% Time elapsed: 0.435 s
% ------------------------------
% ------------------------------
% Success in time 1.763 s
