Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (85819)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 65
% Passive clauses: 2166
% Generated clauses: 11966
% Final active clauses: 63
% Final passive clauses: 2048
% Input formulas: 20
% Initial clauses: 21
% 
% Duplicate literals: 733
% Trivial inequalities: 107
% Fw subsumption resolutions: 948
% Fw demodulations: 609
% 
% Simple tautologies: 243
% Equational tautologies: 95
% Forward subsumptions: 5684
% 
% Binary resolution: 25
% Forward superposition: 4625
% Backward superposition: 4205
% Self superposition: 354
% Equality factoring: 262
% Equality resolution: 6
% 
% Split clauses: 400
% Split components: 810
% SAT solver clauses: 1263
% SAT solver unit clauses: 428
% SAT solver binary clauses: 401
% 
% Sat splits: 399
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1259
% TWLsolver calls for satisfiability: 44
% 
% Memory used [KB]: 8059
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (85834)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 98
% Passive clauses: 1129
% Generated clauses: 15411
% Final active clauses: 86
% Final passive clauses: 620
% Input formulas: 20
% Initial clauses: 21
% 
% Duplicate literals: 679
% Trivial inequalities: 62
% Fw subsumption resolutions: 700
% Fw demodulations: 1915
% Bw demodulations: 313
% Condensations: 856
% 
% Simple tautologies: 163
% Equational tautologies: 19
% Forward subsumptions: 8141
% Bw demodulations to eq. taut.: 106
% 
% Binary resolution: 17
% Forward superposition: 5300
% Backward superposition: 4192
% Self superposition: 166
% Equality factoring: 47
% Equality resolution: 8
% 
% Split clauses: 119
% Split components: 243
% SAT solver clauses: 1730
% SAT solver unit clauses: 36
% SAT solver binary clauses: 24
% 
% Sat splits: 119
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1730
% TWLsolver calls for satisfiability: 114
% 
% Memory used [KB]: 16630
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% (85853)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 269
% Passive clauses: 7402
% Generated clauses: 158056
% Final active clauses: 69
% Final passive clauses: 546
% Input formulas: 20
% Initial clauses: 21
% Discarded non-redundant clauses: 6637
% 
% Duplicate literals: 6745
% Trivial inequalities: 1093
% Fw subsumption resolutions: 10571
% Fw demodulations: 18685
% Bw demodulations: 4200
% Condensations: 15339
% Global subsumptions: 88
% 
% Simple tautologies: 1030
% Equational tautologies: 570
% Forward subsumptions: 93393
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 540
% 
% Unit resulting resolution: 1264
% Forward superposition: 31610
% Backward superposition: 30189
% Self superposition: 723
% Equality factoring: 323
% Equality resolution: 33
% 
% Split clauses: 260
% Split components: 528
% SAT solver clauses: 8195
% SAT solver unit clauses: 45
% SAT solver binary clauses: 353
% 
% Sat splits: 260
% Sat splitting refutations: 183
% 
% Memory used [KB]: 69082
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 7
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 7
% Input formulas: 20
% Initial clauses: 9
% 
% Selected by SInE selection: 8
% SInE iterations: 1
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 1
% 
% Memory used [KB]: 767
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.ringNorm
% (86132)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 85
% Passive clauses: 1324
% Generated clauses: 13432
% Final active clauses: 63
% Final passive clauses: 759
% Input formulas: 20
% Initial clauses: 21
% 
% Duplicate literals: 241
% Trivial inequalities: 174
% Fw subsumption resolutions: 923
% Fw demodulations: 2197
% Bw demodulations: 200
% Global subsumptions: 3
% 
% Simple tautologies: 247
% Equational tautologies: 65
% Forward subsumptions: 8215
% Bw demodulations to eq. taut.: 23
% 
% Unit resulting resolution: 320
% Forward superposition: 4516
% Backward superposition: 4286
% Self superposition: 152
% Equality factoring: 42
% Equality resolution: 13
% 
% Split clauses: 13
% Split components: 26
% SAT solver clauses: 1385
% SAT solver unit clauses: 37
% SAT solver binary clauses: 168
% 
% Sat splits: 13
% Sat splitting refutations: 43
% 
% Memory used [KB]: 13688
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 176 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 19
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 19
% Input formulas: 20
% Initial clauses: 21
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 3
% SAT solver unit clauses: 1
% 
% Memory used [KB]: 895
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% remaining time: 176 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 19
% Passive clauses: 1
% Generated clauses: 2
% Final active clauses: 19
% Input formulas: 20
% Initial clauses: 21
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 5
% SAT solver unit clauses: 2
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 895
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% remaining time: 175 next slice time: 85
lrs+1_5:4_cond=on:fsr=off:fde=none:gs=on:gsem=on:lwlo=on:nm=64:nwc=1:stl=60:sos=all:spl=off_80 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 19
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 19
% Input formulas: 20
% Initial clauses: 21
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 1
% SAT solver unit clauses: 1
% 
% TWLsolver clauses: 1
% TWLsolver calls for satisfiability: 1
% 
% Memory used [KB]: 895
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% remaining time: 175 next slice time: 115
lrs+11_28_cond=on:gs=on:gsaa=from_current:gsem=on:gsssp=full:lwlo=on:nm=64:nwc=1.7:stl=30:sac=on:sdd=off:sfr=on:ssfp=40000:ssfq=1.1:smm=off:ssnc=none:urr=on:updr=off_109 on thm.ringNorm
% (86159)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 118
% Passive clauses: 817
% Generated clauses: 23943
% Final active clauses: 117
% Final passive clauses: 235
% Input formulas: 20
% Initial clauses: 21
% Discarded non-redundant clauses: 33501
% 
% Duplicate literals: 2909
% Trivial inequalities: 222
% Fw subsumption resolutions: 868
% Fw demodulations: 243
% Bw demodulations: 1
% Condensations: 5623
% 
% Simple tautologies: 392
% Equational tautologies: 315
% Forward subsumptions: 11186
% 
% Binary resolution: 30
% Unit resulting resolution: 209
% Forward superposition: 6835
% Backward superposition: 5832
% Self superposition: 490
% Equality factoring: 834
% Equality resolution: 36
% 
% Split clauses: 14
% Split components: 28
% SAT solver clauses: 857
% SAT solver unit clauses: 33
% SAT solver binary clauses: 51
% SAT solver learnt clauses: 15
% SAT solver learnt literals: 50
% 
% Sat splits: 14
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 834
% TWLsolver calls for satisfiability: 34732
% 
% Memory used [KB]: 12537
% Time elapsed: 11.600 s
% ------------------------------
% ------------------------------
% remaining time: 59 next slice time: 59
ott+11_3:1_br=off:gs=on:gsem=on:nwc=1:sas=minisat:sos=all:spl=off:urr=on:updr=off_135 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 19
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 19
% Input formulas: 20
% Initial clauses: 21
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 2
% SAT solver unit clauses: 1
% 
% Memory used [KB]: 1023
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% remaining time: 59 next slice time: 59
lrs+3_5_bd=preordered:bsr=unit_only:gsp=input_only:gs=on:gsem=on:gsssp=full:lwlo=on:nm=64:nwc=1:sas=minisat:stl=90:spl=off:sp=occurrence:urr=ec_only:updr=off_815 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f15889,plain,(
  $false),
  inference(subsumption_resolution,[],[f62,f15888])).
fof(f15888,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X3),s('type.quote.index',X2),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5))))) = s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X2),s('type.quote.varmap'(X0),X3))),s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X3),s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5)))))) )),
  inference(forward_demodulation,[],[f15887,f15886])).
fof(f15886,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4)))) )),
  inference(forward_demodulation,[],[f15657,f80])).
fof(f80,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X18,X19,X20,X21] : s(X6,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X6),X18),s('type.quote.varmap'(X6),X19),s('type.quote.index',X20),s('type.list.list'('type.quote.index'),X21))) = s(X6,happ(s(fun('type.list.list'('type.quote.index'),X6),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6)),happ(s(fun('type.quote.varmap'(X6),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X6),X18))),s('type.quote.varmap'(X6),X19))),s('type.quote.index',X20))),s('type.list.list'('type.quote.index'),X21)))),
  file('thm.ringNorm.ivl_aux_def|split|1','ARITY_1_4_const.ringNorm.r_ivl_aux')).
fof(f15657,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))) )),
  inference(superposition,[],[f79,f66])).
fof(f66,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : ! [X1] : s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7] : s(fun('type.quote.varmap'(X6),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6))),'const.ringNorm.r_ivl_aux_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.quote.varmap'(X6),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.ivl_aux_def|split|1','thm.ringNorm.r_ivl_aux_def')).
fof(f79,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),X4))) = s(X0,happ(s(fun('type.list.list'('type.quote.index'),X0),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0)),happ(s(fun('type.quote.varmap'(X0),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X0))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.quote.index',X3))),s('type.list.list'('type.quote.index'),X4)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X18,X19,X20,X21] : s(X6,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X6),X18),s('type.quote.varmap'(X6),X19),s('type.quote.index',X20),s('type.list.list'('type.quote.index'),X21))) = s(X6,happ(s(fun('type.list.list'('type.quote.index'),X6),happ(s(fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6)),happ(s(fun('type.quote.varmap'(X6),fun('type.quote.index',fun('type.list.list'('type.quote.index'),X6))),'const.canonical.ivl_aux_1'(s('type.semi_ring.semi_ring'(X6),X18))),s('type.quote.varmap'(X6),X19))),s('type.quote.index',X20))),s('type.list.list'('type.quote.index'),X21)))),
  file('thm.ringNorm.ivl_aux_def|split|1','ARITY_1_4_const.canonical.ivl_aux')).
fof(f15887,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X3),s('type.quote.index',X2),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5))))) = s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X2),s('type.quote.varmap'(X0),X3))),s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X3),s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5)))))) )),
  inference(backward_demodulation,[],[f15886,f12435])).
fof(f12435,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X3),s('type.quote.index',X2),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5))))) = s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X2),s('type.quote.varmap'(X0),X3))),s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X3),s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5)))))) )),
  inference(superposition,[],[f81,f11397])).
fof(f11397,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s(X0,X3)))) )),
  inference(forward_demodulation,[],[f11372,f72])).
fof(f72,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X18,X19,X20] : s(X6,'const.ring.ring_RM_3'(s('type.ring.ring'(X6),X18),s(X6,X19),s(X6,X20))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X18))),s(X6,X19))),s(X6,X20)))),
  file('thm.ringNorm.ivl_aux_def|split|1','ARITY_1_3_const.ring.ring_RM')).
fof(f11372,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(superposition,[],[f71,f10614])).
fof(f10614,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(superposition,[],[f78,f67])).
fof(f67,plain,(
  ( ! [X0,X1] : (s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1] : s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0] : ! [X1] : s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X7] : s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7))) = s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X7))),s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X7))),s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X7))),s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.ivl_aux_def|split|1','thm.ring.semi_ring_of_def')).
fof(f78,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X14,X15,X16,X17] : s(fun(X6,fun(X6,X6)),X17) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X14),s(X6,X15),s(fun(X6,fun(X6,X6)),X16),s(fun(X6,fun(X6,X6)),X17)))))),
  file('thm.ringNorm.ivl_aux_def|split|1','thm.semi_ring.semi_ring_SRM')).
fof(f71,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X18,X19,X20] : s(X6,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X6),X18),s(X6,X19),s(X6,X20))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X18))),s(X6,X19))),s(X6,X20)))),
  file('thm.ringNorm.ivl_aux_def|split|1','ARITY_1_3_const.semi_ring.semi_ring_SRM')).
fof(f81,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5))))) = s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))),s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5)))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5))))) = s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))),s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5)))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5] : s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5))))) = s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))),s(X0,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X13,X8,X9,X10,X11] : s(X6,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X6),X13),s('type.quote.varmap'(X6),X8),s('type.quote.index',X9),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X10),s('type.list.list'('type.quote.index'),X11))))) = s(X6,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X6),X13),s(X6,'const.quote.varmap_find_2'(s('type.quote.index',X9),s('type.quote.varmap'(X6),X8))),s(X6,'const.canonical.ivl_aux_4'(s('type.semi_ring.semi_ring'(X6),X13),s('type.quote.varmap'(X6),X8),s('type.quote.index',X10),s('type.list.list'('type.quote.index'),X11)))))),
  file('thm.ringNorm.ivl_aux_def|split|1','thm.canonical.ivl_aux_def|split|1')).
fof(f62,plain,(
  s(sK0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.quote.index',sK3),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',sK4),s('type.list.list'('type.quote.index'),sK5))))) != s(sK0,'const.ring.ring_RM_3'(s('type.ring.ring'(sK0),sK1),s(sK0,'const.quote.varmap_find_2'(s('type.quote.index',sK3),s('type.quote.varmap'(sK0),sK2))),s(sK0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.quote.index',sK4),s('type.list.list'('type.quote.index'),sK5)))))),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  s(sK0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.quote.index',sK3),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',sK4),s('type.list.list'('type.quote.index'),sK5))))) != s(sK0,'const.ring.ring_RM_3'(s('type.ring.ring'(sK0),sK1),s(sK0,'const.quote.varmap_find_2'(s('type.quote.index',sK3),s('type.quote.varmap'(sK0),sK2))),s(sK0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.quote.index',sK4),s('type.list.list'('type.quote.index'),sK5)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5])],[f57])).
fof(f57,plain,(
  ? [X0,X1,X2,X3,X4,X5] : s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5))))) != s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))),s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5)))))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5))))) = s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))),s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5)))))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ~! [X0] : ! [X1,X2,X3,X4,X5] : s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X3),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5))))) = s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,'const.quote.varmap_find_2'(s('type.quote.index',X3),s('type.quote.varmap'(X0),X2))),s(X0,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.quote.index',X4),s('type.list.list'('type.quote.index'),X5)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10,X11] : s(X6,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.quote.index',X9),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X10),s('type.list.list'('type.quote.index'),X11))))) = s(X6,'const.ring.ring_RM_3'(s('type.ring.ring'(X6),X7),s(X6,'const.quote.varmap_find_2'(s('type.quote.index',X9),s('type.quote.varmap'(X6),X8))),s(X6,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.quote.index',X10),s('type.list.list'('type.quote.index'),X11)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10,X11] : s(X6,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.quote.index',X9),s('type.list.list'('type.quote.index'),'const.list.CONS_2'(s('type.quote.index',X10),s('type.list.list'('type.quote.index'),X11))))) = s(X6,'const.ring.ring_RM_3'(s('type.ring.ring'(X6),X7),s(X6,'const.quote.varmap_find_2'(s('type.quote.index',X9),s('type.quote.varmap'(X6),X8))),s(X6,'const.ringNorm.r_ivl_aux_4'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.quote.index',X10),s('type.list.list'('type.quote.index'),X11)))))),
  file('thm.ringNorm.ivl_aux_def|split|1','thm.ringNorm.ivl_aux_def|split|1')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 68
% Passive clauses: 773
% Generated clauses: 15415
% Final active clauses: 67
% Final passive clauses: 704
% Input formulas: 20
% Initial clauses: 21
% 
% Duplicate literals: 1098
% Trivial inequalities: 175
% Fw subsumption resolutions: 1420
% Bw subsumption resolutions: 1
% Fw demodulations: 377
% Bw demodulations: 1
% 
% Simple tautologies: 298
% Equational tautologies: 125
% Forward subsumptions: 10920
% 
% Binary resolution: 47
% Forward superposition: 5769
% Backward superposition: 5729
% Self superposition: 416
% Equality factoring: 349
% Equality resolution: 12
% 
% SAT solver clauses: 774
% SAT solver unit clauses: 29
% SAT solver binary clauses: 21
% 
% Memory used [KB]: 8699
% Time elapsed: 1.980 s
% ------------------------------
% ------------------------------
% Success in time 25.992 s
