Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.arithmetic
% (64346)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 879
% Passive clauses: 12027
% Generated clauses: 24041
% Final active clauses: 869
% Final passive clauses: 10625
% Input formulas: 48
% Initial clauses: 193
% 
% Duplicate literals: 531
% Trivial inequalities: 21
% Fw subsumption resolutions: 2210
% Fw demodulations: 534
% Bw demodulations: 10
% 
% Simple tautologies: 252
% Equational tautologies: 56
% Forward subsumptions: 8051
% 
% Binary resolution: 14045
% Factoring: 1439
% Forward superposition: 3290
% Backward superposition: 1505
% Self superposition: 26
% Equality factoring: 26
% Equality resolution: 7
% 
% Split clauses: 438
% Split components: 886
% SAT solver clauses: 864
% SAT solver unit clauses: 20
% SAT solver binary clauses: 170
% 
% Sat splits: 438
% 
% TWLsolver clauses: 864
% TWLsolver calls for satisfiability: 579
% 
% Memory used [KB]: 11129
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.arithmetic
% (64378)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 599
% Passive clauses: 3866
% Generated clauses: 17789
% Final active clauses: 364
% Final passive clauses: 1721
% Input formulas: 48
% Initial clauses: 193
% 
% Duplicate literals: 236
% Trivial inequalities: 32
% Fw subsumption resolutions: 1535
% Fw demodulations: 1277
% Bw demodulations: 1594
% Condensations: 188
% Global subsumptions: 17
% 
% Simple tautologies: 274
% Equational tautologies: 17
% Forward subsumptions: 6091
% Bw demodulations to eq. taut.: 142
% 
% Binary resolution: 4774
% Factoring: 192
% Forward superposition: 1829
% Backward superposition: 1321
% Self superposition: 26
% Equality factoring: 26
% Equality resolution: 11
% 
% Split clauses: 49
% Split components: 102
% SAT solver clauses: 3615
% SAT solver unit clauses: 206
% SAT solver binary clauses: 86
% SAT solver learnt clauses: 34
% SAT solver learnt literals: 10
% 
% Sat splits: 49
% Sat splitting refutations: 21
% 
% TWLsolver clauses: 3547
% TWLsolver calls for satisfiability: 1960
% 
% Memory used [KB]: 12792
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.arithmetic
% (64403)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 6488
% Passive clauses: 19019
% Generated clauses: 155786
% Final active clauses: 1785
% Final passive clauses: 3885
% Input formulas: 48
% Initial clauses: 193
% 
% Duplicate literals: 785
% Trivial inequalities: 90
% Fw subsumption resolutions: 3251
% Fw demodulations: 38472
% Bw demodulations: 2149
% Condensations: 416
% Global subsumptions: 526
% 
% Simple tautologies: 120
% Equational tautologies: 1372
% Forward subsumptions: 89303
% Bw demodulations to eq. taut.: 15
% 
% Unit resulting resolution: 87205
% Factoring: 1
% Forward superposition: 11021
% Backward superposition: 4881
% Self superposition: 42
% Equality factoring: 42
% Equality resolution: 12
% 
% Split clauses: 89
% Split components: 187
% SAT solver clauses: 21938
% SAT solver unit clauses: 3898
% SAT solver binary clauses: 8859
% 
% Sat splits: 89
% Sat splitting refutations: 2212
% 
% Memory used [KB]: 84689
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f240147,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f449,f456,f460,f464,f465,f467,f468,f454,f466,f2963,f2965,f3083,f3073,f3140,f3141,f3142,f3840,f4189,f4453,f4463,f4580,f4775,f4776,f5087,f5202,f5403,f5404,f5546,f5547,f5813,f5814,f5885,f5886,f6004,f6005,f17015,f23959,f24753,f25285,f25637,f25664,f25666,f26204,f26373,f26548,f27687,f27688,f27953,f27952,f27951,f27947,f27945,f27943,f27939,f27931,f27930,f28027,f28028,f28029,f28030,f28031,f28032,f28075,f28078,f28232,f28233,f28243,f28244,f28386,f28392,f28664,f28668,f28679,f28669,f28670,f28671,f28673,f28674,f28680,f28681,f28693,f28694,f28692,f28685,f28686,f28687,f28688,f28689,f28695,f28696,f28709,f28710,f28708,f28700,f28701,f28702,f28703,f28704,f28705,f28711,f28712,f28723,f28716,f28717,f28718,f28719,f28720,f28726,f28727,f28738,f28731,f28732,f28733,f28734,f28735,f28741,f28757,f28758,f28745,f28746,f28747,f28749,f28750,f28751,f28752,f28753,f28754,f28763,f28766,f28781,f28768,f28769,f28772,f28773,f28774,f28775,f28776,f29004,f29005,f29024,f29079,f29080,f29090,f29091,f29224,f29230,f29518,f29875,f28762,f29876,f28761,f29968,f29970,f26219,f29971,f29998,f30000,f30103,f34494,f42962,f43087,f43085,f43206,f43207,f43227,f43268,f29502,f29499,f29486,f29483,f29416,f29400,f29523,f29515,f43358,f43359,f43388,f43389,f43520,f43527,f43559,f43563,f43564,f43565,f43566,f43568,f43569,f43570,f43578,f43582,f43583,f43584,f43585,f43586,f43588,f43589,f43590,f43598,f43602,f43603,f43604,f43605,f43606,f43608,f43609,f43610,f43618,f43633,f43635,f43622,f43623,f43624,f43625,f43626,f43627,f43628,f43636,f43651,f43640,f43641,f43642,f43643,f43644,f43645,f43653,f43668,f43657,f43658,f43659,f43660,f43661,f43662,f43670,f43691,f43674,f43675,f43676,f43677,f43678,f43679,f43680,f43681,f43682,f43683,f43693,f43711,f43697,f43698,f43699,f43700,f43701,f43702,f43703,f43704,f43714,f43715,f43718,f43719,f43720,f43721,f43722,f43723,f43728,f43729,f43730,f43731,f28586,f44068,f44071,f44078,f44079,f44080,f44081,f44094,f44095,f44069,f44098,f44195,f4789,f44196,f4788,f44215,f29041,f44253,f43282,f458,f44514,f5417,f44515,f5416,f44535,f44553,f44575,f44588,f44861,f44882,f44868,f26564,f44990,f44991,f45011,f45051,f46314,f46315,f46699,f48179,f48202,f48510,f48511,f49823,f57679,f57707,f58025,f58053,f58409,f58529,f58528,f58760,f58883,f58881,f58901,f58920,f58902,f58919,f64564,f64592,f70673,f70788,f71141,f71142,f71629,f71630,f77363,f77490,f77488,f78323,f78557,f78730,f78740,f78801,f78802,f78815,f78816,f78817,f79093,f79099,f79132,f5204,f79308,f79312,f79313,f79314,f79315,f79317,f79319,f79327,f79331,f79332,f79333,f79334,f79346,f79335,f79337,f79339,f79347,f79351,f79352,f79353,f79354,f79366,f79355,f79357,f79359,f79367,f79382,f79383,f79371,f79373,f79374,f79376,f79377,f79384,f79397,f79388,f79389,f79390,f79391,f79421,f79434,f79425,f79426,f79427,f79428,f79435,f79454,f79439,f79440,f79441,f79442,f79443,f79444,f79445,f79446,f79455,f79471,f79459,f79460,f79461,f79462,f79463,f79464,f79473,f79474,f79477,f79478,f79479,f79480,f79481,f79486,f79487,f79488,f79495,f79496,f79497,f79498,f79712,f79713,f80228,f80416,f80648,f80995,f80996,f81011,f81016,f78940,f78938,f81341,f78601,f81503,f5827,f81504,f5826,f81522,f81538,f81558,f81569,f5201,f81662,f6018,f81663,f6017,f81681,f81697,f81763,f81774,f452,f81810,f5560,f81811,f5559,f82271,f82272,f82290,f82340,f70790,f82457,f82708,f82709,f82726,f83118,f83355,f83356,f85341,f86712,f86713,f88100,f92992,f92998,f96568,f96655,f96670,f96718,f96931,f96949,f96932,f96948,f97264,f97265,f98644,f98666,f103626,f98406,f96835,f96833,f103878,f103880,f104018,f104021,f104039,f104040,f104042,f104046,f104060,f104047,f104061,f104048,f104049,f104051,f104053,f104063,f104067,f104081,f104068,f104069,f104070,f104072,f104074,f104083,f104087,f104101,f104088,f104089,f104090,f104092,f104094,f104103,f104118,f104120,f104107,f104119,f104109,f104110,f104112,f104113,f104122,f104135,f104136,f104126,f104127,f104128,f104129,f104138,f104151,f104152,f104142,f104143,f104144,f104145,f104154,f104173,f104174,f104158,f104159,f104160,f104161,f104162,f104163,f104164,f104165,f104176,f104192,f104193,f104180,f104181,f104182,f104183,f104184,f104185,f104196,f104197,f104214,f104200,f104201,f104202,f104203,f104204,f104207,f104208,f104209,f104220,f104221,f104222,f104223,f104243,f104247,f104248,f104249,f104250,f104252,f104254,f78600,f81340,f104426,f104020,f104427,f104862,f96669,f104898,f104902,f104261,f104903,f104904,f104905,f104907,f104909,f104019,f104966,f104967,f104982,f104983,f105122,f28404,f457,f105128,f26220,f105163,f105167,f43577,f105168,f79326,f105169,f105170,f105172,f105174,f105181,f105185,f43597,f105186,f105187,f105188,f105190,f105192,f105199,f105203,f43617,f105204,f105205,f105206,f105208,f105210,f105217,f105232,f105233,f105221,f43634,f105223,f105224,f105226,f105227,f105252,f105265,f43652,f105256,f105257,f105258,f105259,f105266,f105279,f43669,f105270,f105271,f105272,f105273,f105280,f105299,f43692,f105284,f105285,f105286,f105287,f105288,f105289,f105290,f105291,f105300,f105316,f43712,f105304,f105305,f105306,f105307,f105308,f105309,f105318,f105319,f43736,f4578,f105322,f105323,f105324,f105325,f26389,f105326,f105329,f105330,f105331,f105355,f105356,f105357,f105358,f105680,f5899,f105681,f5898,f105699,f45029,f105735,f45064,f450,f105756,f3963,f3961,f106196,f106480,f106481,f106582,f106629,f106644,f106772,f106846,f106861,f107372,f104218,f108570,f108576,f108614,f108617,f108616,f108867,f109511,f109512,f109509,f109507,f96590,f110171,f97282,f110172,f97281,f110191,f110208,f110263,f110275,f96832,f111509,f111510,f111511,f111884,f120491,f120492,f121316,f121317,f451,f126105,f124744,f124736,f124343,f124335,f124295,f124287,f126169,f124223,f126173,f124103,f126177,f124049,f126181,f123995,f126185,f123940,f126189,f123885,f123155,f122874,f122472,f122400,f113818,f113803,f113792,f113793,f113794,f113723,f113708,f113697,f113698,f113699,f113654,f113632,f113621,f113622,f113623,f126316,f103275,f126320,f103225,f126324,f103175,f126328,f103125,f126332,f103070,f126336,f103014,f88446,f88435,f88426,f88427,f88428,f88352,f88341,f88332,f88333,f88334,f88304,f88288,f88277,f88278,f88279,f126597,f81465,f3960,f126904,f129421,f130284,f79736,f130285,f79735,f130404,f106605,f131002,f106790,f462,f137844,f93016,f455,f137849,f46723,f171877,f23984,f172320,f172354,f172321,f172353,f174964,f25311,f176401,f176403,f176430,f176431,f176402,f176980,f177014,f177340,f177342,f177369,f177370,f177341,f177630,f177664,f177755,f177791,f178366,f178400,f178740,f178776,f179249,f179285,f179627,f179663,f180928,f173788,f173680,f173580,f173505,f180939,f180943,f172365,f172117,f180947,f180951,f171545,f171536,f180955,f170736,f169736,f169675,f169548,f180958,f168733,f168635,f168379,f168315,f180961,f180965,f167378,f167314,f180969,f180973,f167149,f167141,f180977,f166875,f166629,f166576,f165393,f165381,f165319,f165307,f165267,f165258,f164738,f164525,f164138,f164072,f162864,f162852,f162790,f162778,f162738,f162729,f162210,f162001,f161492,f161426,f159925,f159913,f159842,f159830,f159384,f159375,f158416,f158404,f158262,f158133,f158121,f158081,f158072,f157715,f157352,f157151,f156648,f156582,f155417,f155345,f154068,f154056,f153994,f153982,f153942,f153933,f153291,f153099,f152575,f152509,f180983,f149814,f180987,f149698,f180999,f149117,f137041,f137037,f137027,f137020,f137016,f137006,f137001,f135247,f135111,f135001,f134739,f133185,f132944,f181015,f181019,f132410,f132403,f181023,f132192,f132001,f131951,f4577,f182177,f182178,f182185,f182186,f182193,f182194,f182195,f182196,f182206,f28076,f182199,f182200,f182210,f182201,f182203,f182217,f182218,f44070,f182219,f182227,f182228,f182223,f182224,f182232,f182243,f182244,f182236,f28365,f182237,f182238,f182239,f182245,f182249,f182261,f182250,f182252,f182253,f182255,f182256,f182262,f182266,f182282,f182267,f182283,f182268,f182269,f182270,f182271,f182273,f182274,f182276,f182277,f182284,f182288,f182300,f182289,f182291,f182292,f182294,f182295,f182301,f182318,f182319,f182305,f182306,f182307,f182308,f182309,f182310,f182311,f182312,f182320,f182321,f182322,f182323,f182328,f182344,f182345,f182332,f182333,f182334,f182335,f182336,f182337,f182338,f182339,f182347,f182364,f182366,f182351,f182365,f182353,f182354,f182355,f182356,f182358,f182359,f182360,f182361,f28077,f182558,f182559,f182580,f182590,f182635,f182636,f182778,f182782,f182796,f182783,f182785,f182786,f182788,f182789,f182828,f182829,f182830,f182831,f182832,f182838,f182844,f182805,f182804,f182845,f182847,f182849,f182863,f182864,f182873,f182874,f182905,f79249,f183023,f183427,f183439,f183617,f183629,f453,f183621,f126918,f183637,f80245,f183675,f183676,f183691,f183692,f183690,f183681,f183682,f183683,f183684,f183685,f183696,f183697,f183708,f183709,f183707,f183701,f183702,f183710,f183711,f183725,f183715,f183716,f183728,f183717,f183718,f183719,f183720,f183729,f183730,f183744,f183734,f183735,f183747,f183736,f183737,f183738,f183739,f183748,f183766,f183767,f183752,f183753,f183770,f183754,f183756,f183771,f183757,f183758,f183759,f183760,f183761,f183772,f183775,f183792,f183777,f183778,f183781,f183782,f183783,f183784,f183785,f183795,f183809,f183810,f183799,f183800,f183801,f183802,f183803,f183813,f183815,f183827,f43745,f183817,f183818,f183819,f183820,f183821,f184440,f71651,f184441,f71650,f184459,f82306,f184559,f83129,f70787,f185220,f185224,f185238,f185225,f185227,f185228,f185230,f185231,f82456,f185462,f71163,f185463,f71162,f186059,f186077,f461,f186063,f17039,f186273,f186349,f186434,f186438,f186455,f186439,f186441,f186442,f186443,f186445,f186446,f186456,f186457,f186473,f186474,f186472,f186462,f186463,f186464,f186465,f186466,f186467,f186475,f186476,f186487,f186488,f186486,f186480,f186481,f186588,f186589,f186603,f186593,f186594,f186606,f186595,f186596,f186597,f186598,f186607,f186608,f186622,f186612,f186613,f186625,f186614,f186615,f186616,f186617,f186626,f186644,f186645,f186630,f186631,f186648,f186632,f186634,f186649,f186635,f186636,f186637,f186638,f186639,f186650,f186653,f186670,f186655,f186656,f186659,f186660,f186661,f186662,f186663,f186671,f186685,f186686,f186675,f186676,f186677,f186678,f186679,f186758,f186760,f186772,f186451,f186762,f186763,f186764,f186765,f186766,f186774,f186778,f186795,f186779,f186781,f186782,f186783,f186785,f186786,f186796,f186797,f186813,f186814,f186812,f186802,f186803,f186804,f186805,f186806,f186807,f186815,f186816,f186827,f186828,f186826,f186820,f186821,f186829,f186830,f186844,f186834,f186835,f186847,f186836,f186837,f186838,f186839,f186909,f186910,f186924,f186914,f186915,f186927,f186916,f186917,f186918,f186919,f186928,f186946,f186947,f186932,f186933,f186950,f186934,f186936,f186951,f186937,f186938,f186939,f186940,f186941,f186952,f186955,f186972,f186957,f28782,f186958,f186961,f186962,f186963,f186964,f186965,f186973,f186987,f186988,f186977,f186978,f186979,f186980,f186981,f186991,f186993,f187005,f186791,f186995,f186996,f186997,f186998,f186999,f187087,f187090,f187092,f187146,f186792,f187149,f187151,f186452,f187912,f187972,f188044,f188110,f43744,f191425,f191426,f191427,f210807,f212763,f212796,f213393,f213413,f192173,f214456,f210808,f4790,f233198,f233223,f238266,f238391,f238407,f238533,f104262,f239607,f2964,f24781,f240146])).
fof(f240146,plain,(
  $false | ($spl12 | $spl24 | $spl72 | $spl131 | $spl155)),
  inference(subsumption_resolution,[],[f182768,f237424])).
fof(f237424,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl12 | $spl24 | $spl131 | $spl155)),
  inference(subsumption_resolution,[],[f237423,f386])).
fof(f386,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f222])).
fof(f222,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X13)))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.bool.IMP_CLAUSES|split|4')).
fof(f237423,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl12 | $spl24 | $spl131 | $spl155)),
  inference(forward_demodulation,[],[f237422,f188316])).
fof(f188316,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | $spl155),
  inference(resolution,[],[f43744,f416])).
fof(f416,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.bool.EQ_CLAUSES|split|3')).
fof(f237422,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl12 | $spl24 | $spl131)),
  inference(forward_demodulation,[],[f237421,f4580])).
fof(f237421,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | ($spl12 | $spl24 | $spl131)),
  inference(subsumption_resolution,[],[f237420,f181127])).
fof(f181127,plain,(
  s('type.num.num',sK14) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl24 | $spl131)),
  inference(backward_demodulation,[],[f4580,f43086])).
fof(f43086,plain,(
  s('type.num.num',sK15) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl131),
  inference(cnf_transformation,[],[f43086_D])).
fof(f43086_D,plain,(
  s('type.num.num',sK15) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f237420,plain,(
  s('type.num.num',sK14) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | ($spl12 | $spl24)),
  inference(forward_demodulation,[],[f237382,f4580])).
fof(f237382,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.num.num',sK15) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | ($spl12 | $spl24)),
  inference(superposition,[],[f78597,f4580])).
fof(f78597,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X0))))) ) | $spl12),
  inference(resolution,[],[f467,f437])).
fof(f437,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1))) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X9,X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X8)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X8)))) | s('type.num.num',X8) = s('type.num.num',X9)))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.arithmetic.LESS_OR_EQ')).
fof(f182768,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl72),
  inference(forward_demodulation,[],[f182655,f262])).
fof(f262,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.arithmetic.ONE')).
fof(f182655,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl72),
  inference(resolution,[],[f28077,f260])).
fof(f260,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  inference(ennf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X9,X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X9)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.arithmetic.SUB_LESS_OR')).
fof(f24781,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl61),
  inference(cnf_transformation,[],[f24781_D])).
fof(f24781_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f2964,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl14),
  inference(cnf_transformation,[],[f2964_D])).
fof(f2964_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f239607,plain,(
  ( ! [X15] : (p(s(bool,X15)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl8),
  inference(resolution,[],[f14997,f3311])).
fof(f3311,plain,(
  ( ! [X2,X3] : (~sP6(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) )),
  inference(duplicate_literal_removal,[],[f3305])).
fof(f3305,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | ~sP6(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) )),
  inference(resolution,[],[f3298,f337])).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP9(X2,X1,X0)) & (~sP9(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP9(X2,X1,X0))),
  inference(definition_folding,[],[f63,f155,f154,f153,f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f153,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f154,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f155,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f63,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X17,X18,X10] : ((p(s(bool,X10)) <=> s(bool,X17) = s(bool,X18)) <=> ((p(s(bool,X10)) | p(s(bool,X18)) | p(s(bool,X17))) & (p(s(bool,X10)) | ~p(s(bool,X17)) | ~p(s(bool,X18))) & (p(s(bool,X18)) | ~p(s(bool,X17)) | ~p(s(bool,X10))) & (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X10)))))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.sat.sat_2|49|_prove_ltb')).
fof(f3298,plain,(
  ( ! [X6,X7] : (sP9(X6,'F_0',X7) | p(s(bool,X6)) | ~sP6(X7,'F_0',X6)) )),
  inference(subsumption_resolution,[],[f3296,f325])).
fof(f325,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f197])).
fof(f197,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f154])).
fof(f3296,plain,(
  ( ! [X6,X7] : (sP9(X6,'F_0',X7) | p(s(bool,X6)) | ~sP8(X6,X7,'F_0') | ~sP6(X7,'F_0',X6)) )),
  inference(resolution,[],[f3265,f321])).
fof(f321,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP6(X2,X1,X0) & sP7(X1,X2,X0) & sP8(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X2,X1,X0) | ~sP7(X1,X2,X0) | ~sP8(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f194])).
fof(f194,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP9(X2,X1,X0)))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP9(X2,X1,X0)))),
  inference(nnf_transformation,[],[f155])).
fof(f3265,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7)) )),
  inference(subsumption_resolution,[],[f3264,f328])).
fof(f328,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f200])).
fof(f200,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f199])).
fof(f199,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f153])).
fof(f3264,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | sP7('F_0',X6,X7)) )),
  inference(subsumption_resolution,[],[f3259,f327])).
fof(f327,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f201])).
fof(f3259,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f3254,f316])).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (sP7(X1,X2,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f195])).
fof(f3254,plain,(
  ( ! [X4,X5] : (sP9(X5,'F_0',X4) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(resolution,[],[f3249,f325])).
fof(f3249,plain,(
  ( ! [X2,X3] : (~sP8(X3,X2,'F_0') | ~p(s(bool,X2)) | sP9(X3,'F_0',X2)) )),
  inference(subsumption_resolution,[],[f3248,f420])).
fof(f420,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.bool.EQ_CLAUSES|split|0')).
fof(f3248,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) )),
  inference(subsumption_resolution,[],[f3245,f3233])).
fof(f3233,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f3227,f331])).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f202])).
fof(f202,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f152])).
fof(f3227,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) )),
  inference(resolution,[],[f1271,f315])).
fof(f315,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f195])).
fof(f1271,plain,(
  ( ! [X6,X5] : (sP9(X6,'F_0',X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) )),
  inference(superposition,[],[f1264,f334])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f204])).
fof(f1264,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(subsumption_resolution,[],[f1257,f250])).
fof(f250,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) & (! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK14)))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X5))))) | (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) & ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X4))))))) & ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) & ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16])],[f165])).
fof(f165,plain,(
  ? [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) & (! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X1)))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X5))))) | (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) & ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X4))))))) & (? [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X1)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ? [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))))))),
  inference(rectify,[],[f164])).
fof(f164,plain,(
  ? [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) & (! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X1)))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) | (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) & ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3))))))) & (? [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X1)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ? [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3)))))))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ? [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) & ((! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X1)))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) | (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) & ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3))))))) & (? [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X1)))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) | (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ? [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) & ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3)))))))))),
  inference(nnf_transformation,[],[f113])).
fof(f113,plain,(
  ? [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) & (! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X1)))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) <~> (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) & ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3))))))))),
  inference(ennf_transformation,[],[f50])).
fof(f50,plain,(
  ~! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) => (! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X1)))) => p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) <=> (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) & ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) => p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))) => (! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) => p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X8))))) <=> (p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) & ! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) => p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X8))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))) => (! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) => p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X8))))) <=> (p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) & ! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) => p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X8))))))))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.arithmetic.BOUNDED_FORALL_THM')).
fof(f1257,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))),
  inference(superposition,[],[f560,f420])).
fof(f560,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))),
  inference(resolution,[],[f250,f415])).
fof(f415,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f242])).
fof(f3245,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) )),
  inference(resolution,[],[f3234,f319])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f3234,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) )),
  inference(subsumption_resolution,[],[f3228,f327])).
fof(f3228,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) )),
  inference(resolution,[],[f1271,f316])).
fof(f14997,plain,(
  ( ! [X114,X113] : (sP6('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X113,X114)) ) | $spl8),
  inference(resolution,[],[f4820,f333])).
fof(f333,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f203])).
fof(f4820,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl8),
  inference(resolution,[],[f462,f438])).
fof(f438,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f249])).
fof(f104262,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',X4))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl408),
  inference(cnf_transformation,[],[f104262_D])).
fof(f104262_D,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',X4))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f238533,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0'))))) ) | $spl656),
  inference(cnf_transformation,[],[f238533_D])).
fof(f238533_D,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0'))))) ) <=> ~$spl656),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl656])])).
fof(f238407,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl12),
  inference(resolution,[],[f78696,f258])).
fof(f258,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2)))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))))))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2)))))) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))))))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X9,X8,X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X9),s('type.num.num',X8))),s('type.num.num',X10)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X10)))))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10))))))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.arithmetic.SUB_RIGHT_LESS')).
fof(f78696,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))))) ) | $spl12),
  inference(superposition,[],[f467,f437])).
fof(f238391,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6))))) ) | $spl654),
  inference(cnf_transformation,[],[f238391_D])).
fof(f238391_D,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6))))) ) <=> ~$spl654),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl654])])).
fof(f238266,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl12),
  inference(resolution,[],[f78695,f258])).
fof(f78695,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl12),
  inference(superposition,[],[f467,f437])).
fof(f233223,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) | $spl651),
  inference(cnf_transformation,[],[f233223_D])).
fof(f233223_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) <=> ~$spl651),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl651])])).
fof(f233198,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))))) ) | $spl284),
  inference(resolution,[],[f232827,f3311])).
fof(f232827,plain,(
  ( ! [X112,X113] : (sP6(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))),X112,X113)) ) | $spl284),
  inference(resolution,[],[f232617,f333])).
fof(f232617,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))))) | $spl284),
  inference(resolution,[],[f81465,f176142])).
fof(f176142,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X0))))))) )),
  inference(superposition,[],[f175698,f442])).
fof(f442,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X9,X8] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X8)))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.arithmetic.ADD_SYM')).
fof(f175698,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK14))))))) )),
  inference(resolution,[],[f175560,f257])).
fof(f257,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f168])).
fof(f175560,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2))),s('type.num.num',sK14))))) )),
  inference(resolution,[],[f175206,f261])).
fof(f261,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X0))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8,X9] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X9))),s('type.num.num',X8))))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.arithmetic.SUB_LESS_EQ')).
fof(f175206,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK14))))) )),
  inference(subsumption_resolution,[],[f567,f738])).
fof(f738,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK14))))) )),
  inference(resolution,[],[f470,f436])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(ennf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X9,X8,X10] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X8)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X10))))) => p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10)))))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.arithmetic.LESS_LESS_EQ_TRANS')).
fof(f470,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))),
  inference(resolution,[],[f250,f438])).
fof(f567,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(superposition,[],[f250,f437])).
fof(f4790,plain,(
  ( ! [X9] : (~p(s(bool,X9))) ) | $spl30),
  inference(cnf_transformation,[],[f4790_D])).
fof(f4790_D,plain,(
  ( ! [X9] : (~p(s(bool,X9))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f210808,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X5))))) ) | $spl648),
  inference(cnf_transformation,[],[f210808_D])).
fof(f210808_D,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X5))))) ) <=> ~$spl648),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl648])])).
fof(f214456,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X6)))) | ~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f214412])).
fof(f214412,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X6)))) | ~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,X7))) ) | $spl12),
  inference(resolution,[],[f78712,f335])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f204])).
fof(f78712,plain,(
  ( ! [X10,X11,X9] : (~sP9(X11,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X10) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X9)))) | ~p(s(bool,X11)) | ~p(s(bool,X10))) ) | $spl12),
  inference(superposition,[],[f467,f336])).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f204])).
fof(f192173,plain,(
  ( ! [X18] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X18))))) ) | $spl646),
  inference(cnf_transformation,[],[f192173_D])).
fof(f192173_D,plain,(
  ( ! [X18] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X18))))) ) <=> ~$spl646),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl646])])).
fof(f213413,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X63)))) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X63),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X65))))) ) | $spl12),
  inference(subsumption_resolution,[],[f213355,f467])).
fof(f213355,plain,(
  ( ! [X64,X65,X63] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X63)))) | ~p(s(bool,X64)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X65),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | s(bool,X64) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X63),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X65))))) ) | $spl12),
  inference(resolution,[],[f78710,f78661])).
fof(f78661,plain,(
  ( ! [X182,X183,X181] : (~sP9('const.prim_rec.<_2'(s('type.num.num',X181),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X183,X182) | s(bool,X182) != s(bool,X183) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X181))))) ) | $spl12),
  inference(resolution,[],[f467,f337])).
fof(f78710,plain,(
  ( ! [X4,X5,X3] : (sP9(X5,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X4) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X3)))) | ~p(s(bool,X4)) | p(s(bool,X5))) ) | $spl12),
  inference(superposition,[],[f467,f334])).
fof(f213393,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X3)))) | ~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f213328])).
fof(f213328,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X3)))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl12),
  inference(resolution,[],[f78710,f337])).
fof(f212796,plain,(
  ( ! [X70,X72,X71] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X70)))) | ~p(s(bool,X71)) | s(bool,X71) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X70),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X72))))) ) | $spl12),
  inference(subsumption_resolution,[],[f212718,f467])).
fof(f212718,plain,(
  ( ! [X70,X72,X71] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X70)))) | ~p(s(bool,X71)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X72),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | s(bool,X71) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X70),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X72))))) ) | $spl12),
  inference(resolution,[],[f78709,f78661])).
fof(f78709,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X0)))) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl12),
  inference(superposition,[],[f467,f334])).
fof(f212763,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X6)))) | ~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f212687])).
fof(f212687,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X6)))) | ~p(s(bool,X7)) | p(s(bool,X8)) | p(s(bool,X8)) | s(bool,X7) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl12),
  inference(resolution,[],[f78709,f337])).
fof(f210807,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X5)))) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f210752])).
fof(f210752,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X5)))) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,X3))) ) | $spl12),
  inference(resolution,[],[f78648,f335])).
fof(f78648,plain,(
  ( ! [X144,X142,X143] : (~sP9(X143,X144,'const.prim_rec.<_2'(s('type.num.num',X142),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,X143)) | ~p(s(bool,X144)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X142))))) ) | $spl12),
  inference(resolution,[],[f467,f318])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f195])).
fof(f191427,plain,(
  $false | ($spl72 | $spl107)),
  inference(subsumption_resolution,[],[f28761,f182745])).
fof(f182745,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | $spl72),
  inference(resolution,[],[f28077,f420])).
fof(f191426,plain,(
  $false | ($spl32 | $spl45)),
  inference(subsumption_resolution,[],[f5826,f182997])).
fof(f182997,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl32),
  inference(resolution,[],[f5201,f420])).
fof(f191425,plain,(
  $false | ($spl8 | $spl41)),
  inference(subsumption_resolution,[],[f5559,f78531])).
fof(f78531,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl8),
  inference(resolution,[],[f78197,f420])).
fof(f78197,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | $spl8),
  inference(resolution,[],[f4821,f261])).
fof(f4821,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',X0))))) ) | $spl8),
  inference(resolution,[],[f462,f436])).
fof(f43744,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | $spl155),
  inference(cnf_transformation,[],[f43744_D])).
fof(f43744_D,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f188110,plain,(
  ( ! [X44] : (p(s(bool,X44)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl2 | $spl24)),
  inference(resolution,[],[f126402,f181112])).
fof(f181112,plain,(
  ( ! [X107,X106] : (sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X106,X107)) ) | ($spl1 | $spl24)),
  inference(backward_demodulation,[],[f4580,f4394])).
fof(f4394,plain,(
  ( ! [X107,X106] : (sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X106,X107)) ) | $spl1),
  inference(resolution,[],[f451,f329])).
fof(f329,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f201])).
fof(f126402,plain,(
  ( ! [X105,X104] : (~sP7(X104,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X105) | p(s(bool,X105)) | ~p(s(bool,X104))) ) | $spl2),
  inference(resolution,[],[f452,f326])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f201])).
fof(f188044,plain,(
  ( ! [X43] : (p(s(bool,X43)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl2 | $spl24)),
  inference(resolution,[],[f126401,f181111])).
fof(f181111,plain,(
  ( ! [X105,X104] : (sP7(X104,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X105)) ) | ($spl1 | $spl24)),
  inference(backward_demodulation,[],[f4580,f4393])).
fof(f4393,plain,(
  ( ! [X105,X104] : (sP7(X104,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X105)) ) | $spl1),
  inference(resolution,[],[f451,f328])).
fof(f126401,plain,(
  ( ! [X103,X102] : (~sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X102,X103) | p(s(bool,X103)) | ~p(s(bool,X102))) ) | $spl2),
  inference(resolution,[],[f452,f326])).
fof(f187972,plain,(
  ( ! [X42] : (p(s(bool,X42)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl2 | $spl24)),
  inference(resolution,[],[f126399,f181110])).
fof(f181110,plain,(
  ( ! [X101,X100] : (sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X100,X101)) ) | ($spl1 | $spl24)),
  inference(backward_demodulation,[],[f4580,f4391])).
fof(f4391,plain,(
  ( ! [X101,X100] : (sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X100,X101)) ) | $spl1),
  inference(resolution,[],[f451,f325])).
fof(f126399,plain,(
  ( ! [X99,X98] : (~sP8(X98,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X99) | p(s(bool,X99)) | ~p(s(bool,X98))) ) | $spl2),
  inference(resolution,[],[f452,f322])).
fof(f322,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f198])).
fof(f187912,plain,(
  ( ! [X41] : (p(s(bool,X41)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl2 | $spl24)),
  inference(resolution,[],[f126398,f181109])).
fof(f181109,plain,(
  ( ! [X99,X98] : (sP8(X98,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X99)) ) | ($spl1 | $spl24)),
  inference(backward_demodulation,[],[f4580,f4390])).
fof(f4390,plain,(
  ( ! [X99,X98] : (sP8(X98,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X99)) ) | $spl1),
  inference(resolution,[],[f451,f324])).
fof(f324,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f198])).
fof(f126398,plain,(
  ( ! [X97,X96] : (~sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X96,X97) | p(s(bool,X97)) | ~p(s(bool,X96))) ) | $spl2),
  inference(resolution,[],[f452,f322])).
fof(f186452,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | $spl589),
  inference(cnf_transformation,[],[f186452_D])).
fof(f186452_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) <=> ~$spl589),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl589])])).
fof(f187151,plain,(
  ( ! [X43] : (p(s(bool,X43)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl2 | $spl24 | $spl369)),
  inference(resolution,[],[f126386,f181181])).
fof(f181181,plain,(
  ( ! [X74,X75] : (sP4('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X74,X75)) ) | ($spl24 | $spl369)),
  inference(backward_demodulation,[],[f4580,f97042])).
fof(f97042,plain,(
  ( ! [X74,X75] : (sP4('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X74,X75)) ) | $spl369),
  inference(resolution,[],[f96833,f308])).
fof(f308,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f126386,plain,(
  ( ! [X72,X73] : (~sP4(X72,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X73) | p(s(bool,X73)) | ~p(s(bool,X72))) ) | $spl2),
  inference(resolution,[],[f452,f305])).
fof(f305,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f190])).
fof(f187149,plain,(
  ( ! [X41] : (p(s(bool,X41)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl2 | $spl24)),
  inference(resolution,[],[f126386,f181108])).
fof(f181108,plain,(
  ( ! [X74,X75] : (sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X74,X75)) ) | ($spl1 | $spl24)),
  inference(backward_demodulation,[],[f4580,f4378])).
fof(f4378,plain,(
  ( ! [X74,X75] : (sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X74,X75)) ) | $spl1),
  inference(resolution,[],[f451,f308])).
fof(f186792,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | $spl617),
  inference(cnf_transformation,[],[f186792_D])).
fof(f186792_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) <=> ~$spl617),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl617])])).
fof(f187146,plain,(
  ( ! [X35] : (p(s(bool,X35)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl2 | $spl24 | $spl365)),
  inference(resolution,[],[f126386,f181191])).
fof(f181191,plain,(
  ( ! [X76,X75] : (sP4('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X75,X76)) ) | ($spl24 | $spl365)),
  inference(backward_demodulation,[],[f4580,f105994])).
fof(f105994,plain,(
  ( ! [X76,X75] : (sP4('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X75,X76)) ) | $spl365),
  inference(resolution,[],[f96590,f308])).
fof(f187092,plain,(
  ( ! [X43] : (p(s(bool,X43)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl2 | $spl24 | $spl369)),
  inference(resolution,[],[f126385,f181180])).
fof(f181180,plain,(
  ( ! [X72,X73] : (sP4(X72,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X73)) ) | ($spl24 | $spl369)),
  inference(backward_demodulation,[],[f4580,f97041])).
fof(f97041,plain,(
  ( ! [X72,X73] : (sP4(X72,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X73)) ) | $spl369),
  inference(resolution,[],[f96833,f307])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f190])).
fof(f126385,plain,(
  ( ! [X70,X71] : (~sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X70,X71) | p(s(bool,X71)) | ~p(s(bool,X70))) ) | $spl2),
  inference(resolution,[],[f452,f305])).
fof(f187090,plain,(
  ( ! [X41] : (p(s(bool,X41)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl2 | $spl24)),
  inference(resolution,[],[f126385,f181107])).
fof(f181107,plain,(
  ( ! [X72,X73] : (sP4(X72,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X73)) ) | ($spl1 | $spl24)),
  inference(backward_demodulation,[],[f4580,f4377])).
fof(f4377,plain,(
  ( ! [X72,X73] : (sP4(X72,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X73)) ) | $spl1),
  inference(resolution,[],[f451,f307])).
fof(f187087,plain,(
  ( ! [X35] : (p(s(bool,X35)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl2 | $spl24 | $spl365)),
  inference(resolution,[],[f126385,f181190])).
fof(f181190,plain,(
  ( ! [X74,X73] : (sP4(X73,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X74)) ) | ($spl24 | $spl365)),
  inference(backward_demodulation,[],[f4580,f105993])).
fof(f105993,plain,(
  ( ! [X74,X73] : (sP4(X73,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X74)) ) | $spl365),
  inference(resolution,[],[f96590,f307])).
fof(f186999,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365 | $spl367)),
  inference(resolution,[],[f181195,f109078])).
fof(f109078,plain,(
  ( ! [X105,X106] : (~sP7(X105,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X106) | p(s(bool,X106)) | ~p(s(bool,X105))) ) | $spl367),
  inference(resolution,[],[f108612,f326])).
fof(f108612,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl367),
  inference(trivial_inequality_removal,[],[f108606])).
fof(f108606,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl367),
  inference(superposition,[],[f96669,f418])).
fof(f418,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.bool.EQ_CLAUSES|split|2')).
fof(f181195,plain,(
  ( ! [X107,X108] : (sP7('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X107,X108)) ) | ($spl24 | $spl365)),
  inference(backward_demodulation,[],[f4580,f106010])).
fof(f106010,plain,(
  ( ! [X107,X108] : (sP7('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X107,X108)) ) | $spl365),
  inference(resolution,[],[f96590,f329])).
fof(f186998,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181195,f791])).
fof(f791,plain,(
  ( ! [X105,X106] : (~sP7(X105,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X106) | p(s(bool,X106)) | ~p(s(bool,X105))) )),
  inference(resolution,[],[f470,f326])).
fof(f186997,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl24 | $spl365)),
  inference(resolution,[],[f181195,f5278])).
fof(f5278,plain,(
  ( ! [X105,X104] : (~sP7(X104,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X105) | p(s(bool,X105)) | ~p(s(bool,X104))) ) | $spl4),
  inference(resolution,[],[f454,f326])).
fof(f186996,plain,(
  ( ! [X9] : (p(s(bool,X9)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181195,f525])).
fof(f525,plain,(
  ( ! [X107,X108] : (~sP7(X107,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X108) | p(s(bool,X108)) | ~p(s(bool,X107))) )),
  inference(resolution,[],[f250,f326])).
fof(f186995,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl221 | $spl365)),
  inference(resolution,[],[f181195,f58599])).
fof(f58599,plain,(
  ( ! [X105,X106] : (~sP7(X105,X106,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X106)) | ~p(s(bool,X105))) ) | $spl221),
  inference(resolution,[],[f58528,f326])).
fof(f186791,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | $spl616),
  inference(cnf_transformation,[],[f186791_D])).
fof(f186791_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) <=> ~$spl616),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl616])])).
fof(f187005,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X5) | ~sP6(X5,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X4) | ~sP8(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | $spl642),
  inference(cnf_transformation,[],[f187005_D])).
fof(f187005_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X5) | ~sP6(X5,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X4) | ~sP8(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl642),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl642])])).
fof(f186993,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~sP8(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~sP6(X5,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X4)) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181195,f320])).
fof(f320,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f195])).
fof(f186991,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181195,f326])).
fof(f186981,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365 | $spl367)),
  inference(resolution,[],[f181194,f109077])).
fof(f109077,plain,(
  ( ! [X103,X104] : (~sP7('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X103,X104) | p(s(bool,X104)) | ~p(s(bool,X103))) ) | $spl367),
  inference(resolution,[],[f108612,f326])).
fof(f181194,plain,(
  ( ! [X105,X106] : (sP7(X105,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X106)) ) | ($spl24 | $spl365)),
  inference(backward_demodulation,[],[f4580,f106009])).
fof(f106009,plain,(
  ( ! [X105,X106] : (sP7(X105,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X106)) ) | $spl365),
  inference(resolution,[],[f96590,f328])).
fof(f186980,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181194,f790])).
fof(f790,plain,(
  ( ! [X103,X104] : (~sP7('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X103,X104) | p(s(bool,X104)) | ~p(s(bool,X103))) )),
  inference(resolution,[],[f470,f326])).
fof(f186979,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl24 | $spl365)),
  inference(resolution,[],[f181194,f5277])).
fof(f5277,plain,(
  ( ! [X103,X102] : (~sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X102,X103) | p(s(bool,X103)) | ~p(s(bool,X102))) ) | $spl4),
  inference(resolution,[],[f454,f326])).
fof(f186978,plain,(
  ( ! [X9] : (p(s(bool,X9)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181194,f524])).
fof(f524,plain,(
  ( ! [X105,X106] : (~sP7('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X105,X106) | p(s(bool,X106)) | ~p(s(bool,X105))) )),
  inference(resolution,[],[f250,f326])).
fof(f186977,plain,(
  ( ! [X8] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X8))) ) | ($spl24 | $spl221 | $spl365)),
  inference(resolution,[],[f181194,f58599])).
fof(f186988,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | $spl640),
  inference(cnf_transformation,[],[f186988_D])).
fof(f186988_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl640),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl640])])).
fof(f186987,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(subsumption_resolution,[],[f186986,f181192])).
fof(f181192,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X100)) ) | ($spl24 | $spl365)),
  inference(backward_demodulation,[],[f4580,f106006])).
fof(f106006,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X100)) ) | $spl365),
  inference(resolution,[],[f96590,f324])).
fof(f186986,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X3)) ) | ($spl24 | $spl365)),
  inference(subsumption_resolution,[],[f186974,f333])).
fof(f186974,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X3) | ~sP6('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X3,X2)) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181194,f319])).
fof(f186973,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181194,f326])).
fof(f186965,plain,(
  ( ! [X17] : (p(s(bool,X17)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365 | $spl367)),
  inference(resolution,[],[f181193,f109075])).
fof(f109075,plain,(
  ( ! [X99,X100] : (~sP8(X99,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X100) | p(s(bool,X100)) | ~p(s(bool,X99))) ) | $spl367),
  inference(resolution,[],[f108612,f322])).
fof(f181193,plain,(
  ( ! [X101,X102] : (sP8('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X101,X102)) ) | ($spl24 | $spl365)),
  inference(backward_demodulation,[],[f4580,f106007])).
fof(f106007,plain,(
  ( ! [X101,X102] : (sP8('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X101,X102)) ) | $spl365),
  inference(resolution,[],[f96590,f325])).
fof(f186964,plain,(
  ( ! [X16] : (p(s(bool,X16)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181193,f788])).
fof(f788,plain,(
  ( ! [X99,X100] : (~sP8(X99,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X100) | p(s(bool,X100)) | ~p(s(bool,X99))) )),
  inference(resolution,[],[f470,f322])).
fof(f186963,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl24 | $spl365)),
  inference(resolution,[],[f181193,f5275])).
fof(f5275,plain,(
  ( ! [X99,X98] : (~sP8(X98,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X99) | p(s(bool,X99)) | ~p(s(bool,X98))) ) | $spl4),
  inference(resolution,[],[f454,f322])).
fof(f186962,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181193,f522])).
fof(f522,plain,(
  ( ! [X101,X102] : (~sP8(X101,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X102) | p(s(bool,X102)) | ~p(s(bool,X101))) )),
  inference(resolution,[],[f250,f322])).
fof(f186961,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl221 | $spl365)),
  inference(resolution,[],[f181193,f58596])).
fof(f58596,plain,(
  ( ! [X99,X100] : (~sP8(X99,X100,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X100)) | ~p(s(bool,X99))) ) | $spl221),
  inference(resolution,[],[f58528,f322])).
fof(f186958,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181193,f3425])).
fof(f3425,plain,(
  ( ! [X12,X13] : (~sP8(X12,X13,'F_0') | ~p(s(bool,X13)) | ~p(s(bool,X12))) )),
  inference(subsumption_resolution,[],[f3424,f333])).
fof(f3424,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | ~sP8(X12,X13,'F_0') | ~sP6(X13,'F_0',X12)) )),
  inference(subsumption_resolution,[],[f3423,f327])).
fof(f3423,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | ~sP8(X12,X13,'F_0') | ~sP7('F_0',X13,X12) | ~sP6(X13,'F_0',X12)) )),
  inference(subsumption_resolution,[],[f3414,f420])).
fof(f3414,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP8(X12,X13,'F_0') | ~sP7('F_0',X13,X12) | ~sP6(X13,'F_0',X12)) )),
  inference(resolution,[],[f1273,f319])).
fof(f1273,plain,(
  ( ! [X10,X9] : (~sP9(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) )),
  inference(superposition,[],[f1264,f336])).
fof(f28782,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,X9)) ) | $spl112),
  inference(cnf_transformation,[],[f28782_D])).
fof(f28782_D,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,X9)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f186957,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,X9) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181193,f3428])).
fof(f3428,plain,(
  ( ! [X14,X15] : (~sP8(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) )),
  inference(subsumption_resolution,[],[f3427,f331])).
fof(f3427,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f3426,f386])).
fof(f3426,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) )),
  inference(subsumption_resolution,[],[f3415,f329])).
fof(f3415,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP7('F_0',X15,X14) | ~sP6(X15,'F_0',X14)) )),
  inference(resolution,[],[f1273,f320])).
fof(f186972,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X6,X7) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | $spl638),
  inference(cnf_transformation,[],[f186972_D])).
fof(f186972_D,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X6,X7) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl638),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl638])])).
fof(f186955,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181193,f321])).
fof(f186952,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181193,f322])).
fof(f186941,plain,(
  ( ! [X17] : (p(s(bool,X17)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365 | $spl367)),
  inference(resolution,[],[f181192,f109074])).
fof(f109074,plain,(
  ( ! [X97,X98] : (~sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X97,X98) | p(s(bool,X98)) | ~p(s(bool,X97))) ) | $spl367),
  inference(resolution,[],[f108612,f322])).
fof(f186940,plain,(
  ( ! [X16] : (p(s(bool,X16)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181192,f787])).
fof(f787,plain,(
  ( ! [X97,X98] : (~sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X97,X98) | p(s(bool,X98)) | ~p(s(bool,X97))) )),
  inference(resolution,[],[f470,f322])).
fof(f186939,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl24 | $spl365)),
  inference(resolution,[],[f181192,f5274])).
fof(f5274,plain,(
  ( ! [X97,X96] : (~sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X96,X97) | p(s(bool,X97)) | ~p(s(bool,X96))) ) | $spl4),
  inference(resolution,[],[f454,f322])).
fof(f186938,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181192,f521])).
fof(f521,plain,(
  ( ! [X99,X100] : (~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X99,X100) | p(s(bool,X100)) | ~p(s(bool,X99))) )),
  inference(resolution,[],[f250,f322])).
fof(f186937,plain,(
  ( ! [X13] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X13))) ) | ($spl24 | $spl221 | $spl365)),
  inference(resolution,[],[f181192,f58596])).
fof(f186951,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | $spl636),
  inference(cnf_transformation,[],[f186951_D])).
fof(f186951_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl636),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl636])])).
fof(f186936,plain,(
  ( ! [X12] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | sP9(X12,'F_0','const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181192,f3249])).
fof(f186934,plain,(
  ( ! [X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181192,f3425])).
fof(f186950,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | $spl635),
  inference(cnf_transformation,[],[f186950_D])).
fof(f186950_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) <=> ~$spl635),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl635])])).
fof(f186933,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181192,f3428])).
fof(f186932,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181192,f3243])).
fof(f3243,plain,(
  ( ! [X4,X5] : (~sP8(X5,X4,'F_0') | sP9(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) )),
  inference(subsumption_resolution,[],[f3242,f429])).
fof(f429,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.arithmetic.BOUNDED_FORALL_THM',aHL_FALSITY)).
fof(f3242,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0')) )),
  inference(subsumption_resolution,[],[f3239,f329])).
fof(f3239,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP8(X5,X4,'F_0') | ~sP7('F_0',X4,X5)) )),
  inference(resolution,[],[f3233,f320])).
fof(f186947,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2)) ) | $spl632),
  inference(cnf_transformation,[],[f186947_D])).
fof(f186947_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2)) ) <=> ~$spl632),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl632])])).
fof(f186946,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2)) ) | ($spl24 | $spl365)),
  inference(subsumption_resolution,[],[f186929,f333])).
fof(f186929,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2) | ~sP6('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X3,X2)) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181192,f319])).
fof(f186928,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181192,f322])).
fof(f186919,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365 | $spl367)),
  inference(resolution,[],[f181191,f109062])).
fof(f109062,plain,(
  ( ! [X74,X73] : (~sP4(X73,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X74) | p(s(bool,X74)) | ~p(s(bool,X73))) ) | $spl367),
  inference(resolution,[],[f108612,f305])).
fof(f186918,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181191,f775])).
fof(f775,plain,(
  ( ! [X74,X73] : (~sP4(X73,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X74) | p(s(bool,X74)) | ~p(s(bool,X73))) )),
  inference(resolution,[],[f470,f305])).
fof(f186917,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl24 | $spl365)),
  inference(resolution,[],[f181191,f5262])).
fof(f5262,plain,(
  ( ! [X72,X73] : (~sP4(X72,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X73) | p(s(bool,X73)) | ~p(s(bool,X72))) ) | $spl4),
  inference(resolution,[],[f454,f305])).
fof(f186916,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181191,f509])).
fof(f509,plain,(
  ( ! [X76,X75] : (~sP4(X75,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X76) | p(s(bool,X76)) | ~p(s(bool,X75))) )),
  inference(resolution,[],[f250,f305])).
fof(f186927,plain,(
  ( ! [X11] : (sP5(X11,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl630),
  inference(cnf_transformation,[],[f186927_D])).
fof(f186927_D,plain,(
  ( ! [X11] : (sP5(X11,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) <=> ~$spl630),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl630])])).
fof(f186915,plain,(
  ( ! [X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | sP5(X11,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181191,f505])).
fof(f505,plain,(
  ( ! [X68,X67] : (~sP4(X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X67) | ~p(s(bool,X68)) | sP5(X67,X68,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) )),
  inference(resolution,[],[f250,f301])).
fof(f301,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP4(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f187])).
fof(f187,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f186914,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl221 | $spl365)),
  inference(resolution,[],[f181191,f58583])).
fof(f58583,plain,(
  ( ! [X74,X73] : (~sP4(X73,X74,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X74)) | ~p(s(bool,X73))) ) | $spl221),
  inference(resolution,[],[f58528,f305])).
fof(f186924,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X3) | ~p(s(bool,X3))) ) | $spl628),
  inference(cnf_transformation,[],[f186924_D])).
fof(f186924_D,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X3) | ~p(s(bool,X3))) ) <=> ~$spl628),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl628])])).
fof(f186910,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X3) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X3))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181191,f301])).
fof(f186909,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181191,f305])).
fof(f186839,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365 | $spl367)),
  inference(resolution,[],[f181190,f109061])).
fof(f109061,plain,(
  ( ! [X72,X71] : (~sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X71,X72) | p(s(bool,X72)) | ~p(s(bool,X71))) ) | $spl367),
  inference(resolution,[],[f108612,f305])).
fof(f186838,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181190,f774])).
fof(f774,plain,(
  ( ! [X72,X71] : (~sP4('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X71,X72) | p(s(bool,X72)) | ~p(s(bool,X71))) )),
  inference(resolution,[],[f470,f305])).
fof(f186837,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl24 | $spl365)),
  inference(resolution,[],[f181190,f5261])).
fof(f5261,plain,(
  ( ! [X70,X71] : (~sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X70,X71) | p(s(bool,X71)) | ~p(s(bool,X70))) ) | $spl4),
  inference(resolution,[],[f454,f305])).
fof(f186836,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181190,f508])).
fof(f508,plain,(
  ( ! [X74,X73] : (~sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X73,X74) | p(s(bool,X74)) | ~p(s(bool,X73))) )),
  inference(resolution,[],[f250,f305])).
fof(f186847,plain,(
  ( ! [X11] : (sP5(X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | $spl626),
  inference(cnf_transformation,[],[f186847_D])).
fof(f186847_D,plain,(
  ( ! [X11] : (sP5(X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl626),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl626])])).
fof(f186835,plain,(
  ( ! [X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | sP5(X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181190,f504])).
fof(f504,plain,(
  ( ! [X66,X65] : (~sP4('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X66,X65) | ~p(s(bool,X66)) | sP5(X65,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X66)) )),
  inference(resolution,[],[f250,f301])).
fof(f186834,plain,(
  ( ! [X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | ($spl24 | $spl221 | $spl365)),
  inference(resolution,[],[f181190,f58583])).
fof(f186844,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X3))) ) | $spl624),
  inference(cnf_transformation,[],[f186844_D])).
fof(f186844_D,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X3))) ) <=> ~$spl624),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl624])])).
fof(f186830,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X3)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181190,f301])).
fof(f186829,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181190,f305])).
fof(f186821,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl221 | $spl365)),
  inference(resolution,[],[f181189,f58570])).
fof(f58570,plain,(
  ( ! [X47,X48] : (~sP2(X47,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X48) | p(s(bool,X48)) | ~p(s(bool,X47))) ) | $spl221),
  inference(resolution,[],[f58528,f288])).
fof(f288,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f182])).
fof(f182,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f181189,plain,(
  ( ! [X47,X48] : (sP2('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X47,X48)) ) | ($spl24 | $spl365)),
  inference(backward_demodulation,[],[f4580,f105980])).
fof(f105980,plain,(
  ( ! [X47,X48] : (sP2('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X47,X48)) ) | $spl365),
  inference(resolution,[],[f96590,f291])).
fof(f291,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f183])).
fof(f186820,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl221 | $spl365)),
  inference(resolution,[],[f181189,f58569])).
fof(f58569,plain,(
  ( ! [X45,X46] : (~sP2(X45,X46,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X46)) | ~p(s(bool,X45))) ) | $spl221),
  inference(resolution,[],[f58528,f288])).
fof(f186826,plain,(
  ( ! [X2,X3] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2,X3)) ) | $spl622),
  inference(cnf_transformation,[],[f186826_D])).
fof(f186826_D,plain,(
  ( ! [X2,X3] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2,X3)) ) <=> ~$spl622),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl622])])).
fof(f186828,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X6,X7) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(subsumption_resolution,[],[f186818,f294])).
fof(f294,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP3(X2,X0,X1))),
  inference(definition_folding,[],[f59,f147,f146])).
fof(f147,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f59,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X17,X18,X10] : ((p(s(bool,X10)) <=> (p(s(bool,X18)) | p(s(bool,X17)))) <=> ((p(s(bool,X10)) | ~p(s(bool,X18))) & (p(s(bool,X10)) | ~p(s(bool,X17))) & (p(s(bool,X18)) | p(s(bool,X17)) | ~p(s(bool,X10)))))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.sat.sat_25_prove_ltb')).
fof(f186818,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181189,f286])).
fof(f286,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP2(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP2(X0,X1,X2) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f179])).
fof(f179,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f178])).
fof(f178,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1)) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f147])).
fof(f186827,plain,(
  ( ! [X4,X5] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X4,X5) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(subsumption_resolution,[],[f186817,f293])).
fof(f293,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f185])).
fof(f186817,plain,(
  ( ! [X4,X5] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X4,X5) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X5))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181189,f285])).
fof(f285,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f180])).
fof(f186816,plain,(
  ( ! [X2,X3] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2,X3) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181189,f444])).
fof(f444,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f284])).
fof(f284,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f180])).
fof(f186815,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181189,f288])).
fof(f186807,plain,(
  ( ! [X16] : (~p(s(bool,X16)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl221 | $spl365)),
  inference(resolution,[],[f181188,f58553])).
fof(f58553,plain,(
  ( ! [X14,X13] : (~sP0(X13,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X14) | ~p(s(bool,X14)) | ~p(s(bool,X13))) ) | $spl221),
  inference(resolution,[],[f58528,f271])).
fof(f271,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f174])).
fof(f174,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f181188,plain,(
  ( ! [X15,X16] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X15,X16)) ) | ($spl24 | $spl365)),
  inference(backward_demodulation,[],[f4580,f105964])).
fof(f105964,plain,(
  ( ! [X15,X16] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X15,X16)) ) | $spl365),
  inference(resolution,[],[f96590,f274])).
fof(f274,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f175])).
fof(f186806,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365 | $spl367)),
  inference(resolution,[],[f181188,f109032])).
fof(f109032,plain,(
  ( ! [X14,X13] : (~sP0(X13,X14,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X14)) | ~p(s(bool,X13))) ) | $spl367),
  inference(resolution,[],[f108612,f271])).
fof(f186805,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181188,f745])).
fof(f745,plain,(
  ( ! [X14,X13] : (~sP0(X13,X14,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | p(s(bool,X14)) | ~p(s(bool,X13))) )),
  inference(resolution,[],[f470,f271])).
fof(f186804,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl2 | $spl24 | $spl365)),
  inference(resolution,[],[f181188,f126356])).
fof(f126356,plain,(
  ( ! [X12,X13] : (~sP0(X12,X13,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | p(s(bool,X13)) | ~p(s(bool,X12))) ) | $spl2),
  inference(resolution,[],[f452,f271])).
fof(f186803,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl24 | $spl365)),
  inference(resolution,[],[f181188,f5232])).
fof(f5232,plain,(
  ( ! [X12,X13] : (~sP0(X12,X13,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,X13)) | ~p(s(bool,X12))) ) | $spl4),
  inference(resolution,[],[f454,f271])).
fof(f186802,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181188,f479])).
fof(f479,plain,(
  ( ! [X15,X16] : (~sP0(X15,X16,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | p(s(bool,X16)) | ~p(s(bool,X15))) )),
  inference(resolution,[],[f250,f271])).
fof(f186812,plain,(
  ( ! [X2,X3] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2,X3)) ) | $spl620),
  inference(cnf_transformation,[],[f186812_D])).
fof(f186812_D,plain,(
  ( ! [X2,X3] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2,X3)) ) <=> ~$spl620),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl620])])).
fof(f186814,plain,(
  ( ! [X6,X7] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X6,X7) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(subsumption_resolution,[],[f186799,f277])).
fof(f277,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP1(X2,X0,X1))),
  inference(definition_folding,[],[f116,f144,f143])).
fof(f144,plain,(
  ! [X2,X0,X1] : (sP1(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f116,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X17,X18,X10] : ((p(s(bool,X10)) <=> (p(s(bool,X18)) => p(s(bool,X17)))) <=> ((p(s(bool,X10)) | p(s(bool,X18))) & (p(s(bool,X10)) | ~p(s(bool,X17))) & (~p(s(bool,X18)) | p(s(bool,X17)) | ~p(s(bool,X10)))))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.sat.sat_27_prove_ltb')).
fof(f186799,plain,(
  ( ! [X6,X7] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181188,f269])).
fof(f269,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP0(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP0(X0,X1,X2) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f171])).
fof(f171,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1) | sP1(X2,X0,X1)))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1)) | sP1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f144])).
fof(f186813,plain,(
  ( ! [X4,X5] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X4,X5) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(subsumption_resolution,[],[f186798,f276])).
fof(f276,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f177])).
fof(f186798,plain,(
  ( ! [X4,X5] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X4,X5) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X5))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181188,f268])).
fof(f268,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f172])).
fof(f186797,plain,(
  ( ! [X2,X3] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2,X3) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181188,f443])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f267])).
fof(f267,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f172])).
fof(f186796,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181188,f271])).
fof(f186786,plain,(
  ( ! [X17] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X17))) ) | ($spl24 | $spl365 | $spl367)),
  inference(resolution,[],[f181187,f109031])).
fof(f109031,plain,(
  ( ! [X12,X11] : (~sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X11,X12) | ~p(s(bool,X12)) | p(s(bool,X11))) ) | $spl367),
  inference(resolution,[],[f108612,f271])).
fof(f181187,plain,(
  ( ! [X14,X13] : (sP0(X13,X14,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | ($spl24 | $spl365)),
  inference(backward_demodulation,[],[f4580,f105963])).
fof(f105963,plain,(
  ( ! [X14,X13] : (sP0(X13,X14,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) | $spl365),
  inference(resolution,[],[f96590,f272])).
fof(f272,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f175])).
fof(f186785,plain,(
  ( ! [X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X16))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181187,f744])).
fof(f744,plain,(
  ( ! [X12,X11] : (~sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X11,X12) | ~p(s(bool,X12)) | p(s(bool,X11))) )),
  inference(resolution,[],[f470,f271])).
fof(f186783,plain,(
  ( ! [X14] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X14))) ) | ($spl2 | $spl24 | $spl365)),
  inference(resolution,[],[f181187,f126355])).
fof(f126355,plain,(
  ( ! [X10,X11] : (~sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X10,X11) | ~p(s(bool,X11)) | p(s(bool,X10))) ) | $spl2),
  inference(resolution,[],[f452,f271])).
fof(f186782,plain,(
  ( ! [X13] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X13))) ) | ($spl4 | $spl24 | $spl365)),
  inference(resolution,[],[f181187,f5231])).
fof(f5231,plain,(
  ( ! [X10,X11] : (~sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X10,X11) | ~p(s(bool,X11)) | p(s(bool,X10))) ) | $spl4),
  inference(resolution,[],[f454,f271])).
fof(f186781,plain,(
  ( ! [X12] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X12))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181187,f478])).
fof(f478,plain,(
  ( ! [X14,X13] : (~sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X13,X14) | ~p(s(bool,X14)) | p(s(bool,X13))) )),
  inference(resolution,[],[f250,f271])).
fof(f186779,plain,(
  ( ! [X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | ($spl24 | $spl221 | $spl365)),
  inference(resolution,[],[f181187,f58553])).
fof(f186795,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9))) ) | $spl618),
  inference(cnf_transformation,[],[f186795_D])).
fof(f186795_D,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9))) ) <=> ~$spl618),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl618])])).
fof(f186778,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181187,f270])).
fof(f270,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f172])).
fof(f186774,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl365)),
  inference(resolution,[],[f181187,f271])).
fof(f186766,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl367 | $spl369)),
  inference(resolution,[],[f181185,f109078])).
fof(f181185,plain,(
  ( ! [X107,X106] : (sP7('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X106,X107)) ) | ($spl24 | $spl369)),
  inference(backward_demodulation,[],[f4580,f97058])).
fof(f97058,plain,(
  ( ! [X107,X106] : (sP7('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X106,X107)) ) | $spl369),
  inference(resolution,[],[f96833,f329])).
fof(f186765,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181185,f791])).
fof(f186764,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl24 | $spl369)),
  inference(resolution,[],[f181185,f5278])).
fof(f186763,plain,(
  ( ! [X9] : (p(s(bool,X9)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181185,f525])).
fof(f186762,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl221 | $spl369)),
  inference(resolution,[],[f181185,f58599])).
fof(f186451,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | $spl588),
  inference(cnf_transformation,[],[f186451_D])).
fof(f186451_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) <=> ~$spl588),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl588])])).
fof(f186772,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X4) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | $spl614),
  inference(cnf_transformation,[],[f186772_D])).
fof(f186772_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X4) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl614),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl614])])).
fof(f186760,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X4)) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181185,f320])).
fof(f186758,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181185,f326])).
fof(f186679,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl367 | $spl369)),
  inference(resolution,[],[f181184,f109077])).
fof(f181184,plain,(
  ( ! [X105,X104] : (sP7(X104,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X105)) ) | ($spl24 | $spl369)),
  inference(backward_demodulation,[],[f4580,f97057])).
fof(f97057,plain,(
  ( ! [X105,X104] : (sP7(X104,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X105)) ) | $spl369),
  inference(resolution,[],[f96833,f328])).
fof(f186678,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181184,f790])).
fof(f186677,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl24 | $spl369)),
  inference(resolution,[],[f181184,f5277])).
fof(f186676,plain,(
  ( ! [X9] : (p(s(bool,X9)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181184,f524])).
fof(f186675,plain,(
  ( ! [X8] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X8))) ) | ($spl24 | $spl221 | $spl369)),
  inference(resolution,[],[f181184,f58599])).
fof(f186686,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | $spl612),
  inference(cnf_transformation,[],[f186686_D])).
fof(f186686_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl612),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl612])])).
fof(f186685,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(subsumption_resolution,[],[f186684,f181182])).
fof(f181182,plain,(
  ( ! [X99,X98] : (sP8(X98,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X99)) ) | ($spl24 | $spl369)),
  inference(backward_demodulation,[],[f4580,f97054])).
fof(f97054,plain,(
  ( ! [X99,X98] : (sP8(X98,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X99)) ) | $spl369),
  inference(resolution,[],[f96833,f324])).
fof(f186684,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~sP8(X2,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X3)) ) | ($spl24 | $spl369)),
  inference(subsumption_resolution,[],[f186672,f333])).
fof(f186672,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~sP8(X2,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X3) | ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X3,X2)) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181184,f319])).
fof(f186671,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181184,f326])).
fof(f186663,plain,(
  ( ! [X17] : (p(s(bool,X17)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl367 | $spl369)),
  inference(resolution,[],[f181183,f109075])).
fof(f181183,plain,(
  ( ! [X101,X100] : (sP8('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X100,X101)) ) | ($spl24 | $spl369)),
  inference(backward_demodulation,[],[f4580,f97055])).
fof(f97055,plain,(
  ( ! [X101,X100] : (sP8('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X100,X101)) ) | $spl369),
  inference(resolution,[],[f96833,f325])).
fof(f186662,plain,(
  ( ! [X16] : (p(s(bool,X16)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181183,f788])).
fof(f186661,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl24 | $spl369)),
  inference(resolution,[],[f181183,f5275])).
fof(f186660,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181183,f522])).
fof(f186659,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl221 | $spl369)),
  inference(resolution,[],[f181183,f58596])).
fof(f186656,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181183,f3425])).
fof(f186655,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,X9) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181183,f3428])).
fof(f186670,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | $spl610),
  inference(cnf_transformation,[],[f186670_D])).
fof(f186670_D,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl610),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl610])])).
fof(f186653,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181183,f321])).
fof(f186650,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181183,f322])).
fof(f186639,plain,(
  ( ! [X17] : (p(s(bool,X17)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl367 | $spl369)),
  inference(resolution,[],[f181182,f109074])).
fof(f186638,plain,(
  ( ! [X16] : (p(s(bool,X16)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181182,f787])).
fof(f186637,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl24 | $spl369)),
  inference(resolution,[],[f181182,f5274])).
fof(f186636,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181182,f521])).
fof(f186635,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X13))) ) | ($spl24 | $spl221 | $spl369)),
  inference(resolution,[],[f181182,f58596])).
fof(f186649,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | $spl608),
  inference(cnf_transformation,[],[f186649_D])).
fof(f186649_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl608),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl608])])).
fof(f186634,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | sP9(X12,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181182,f3249])).
fof(f186632,plain,(
  ( ! [X10] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181182,f3425])).
fof(f186648,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | $spl607),
  inference(cnf_transformation,[],[f186648_D])).
fof(f186648_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) <=> ~$spl607),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl607])])).
fof(f186631,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181182,f3428])).
fof(f186630,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181182,f3243])).
fof(f186645,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~sP7(X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2)) ) | $spl604),
  inference(cnf_transformation,[],[f186645_D])).
fof(f186645_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~sP7(X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2)) ) <=> ~$spl604),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl604])])).
fof(f186644,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~sP7(X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2)) ) | ($spl24 | $spl369)),
  inference(subsumption_resolution,[],[f186627,f333])).
fof(f186627,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~sP7(X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2) | ~sP6('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X3,X2)) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181182,f319])).
fof(f186626,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181182,f322])).
fof(f186617,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl367 | $spl369)),
  inference(resolution,[],[f181181,f109062])).
fof(f186616,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181181,f775])).
fof(f186615,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl24 | $spl369)),
  inference(resolution,[],[f181181,f5262])).
fof(f186614,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181181,f509])).
fof(f186625,plain,(
  ( ! [X11] : (sP5(X11,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl602),
  inference(cnf_transformation,[],[f186625_D])).
fof(f186625_D,plain,(
  ( ! [X11] : (sP5(X11,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) <=> ~$spl602),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl602])])).
fof(f186613,plain,(
  ( ! [X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | sP5(X11,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181181,f505])).
fof(f186612,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl221 | $spl369)),
  inference(resolution,[],[f181181,f58583])).
fof(f186622,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X3) | ~p(s(bool,X3))) ) | $spl600),
  inference(cnf_transformation,[],[f186622_D])).
fof(f186622_D,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X3) | ~p(s(bool,X3))) ) <=> ~$spl600),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl600])])).
fof(f186608,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X3))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181181,f301])).
fof(f186607,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181181,f305])).
fof(f186598,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl367 | $spl369)),
  inference(resolution,[],[f181180,f109061])).
fof(f186597,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181180,f774])).
fof(f186596,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl24 | $spl369)),
  inference(resolution,[],[f181180,f5261])).
fof(f186595,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181180,f508])).
fof(f186606,plain,(
  ( ! [X11] : (sP5(X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | $spl598),
  inference(cnf_transformation,[],[f186606_D])).
fof(f186606_D,plain,(
  ( ! [X11] : (sP5(X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl598),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl598])])).
fof(f186594,plain,(
  ( ! [X11] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | sP5(X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181180,f504])).
fof(f186593,plain,(
  ( ! [X10] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | ($spl24 | $spl221 | $spl369)),
  inference(resolution,[],[f181180,f58583])).
fof(f186603,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X3))) ) | $spl596),
  inference(cnf_transformation,[],[f186603_D])).
fof(f186603_D,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X3))) ) <=> ~$spl596),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl596])])).
fof(f186589,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X3)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181180,f301])).
fof(f186588,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181180,f305])).
fof(f186481,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl221 | $spl369)),
  inference(resolution,[],[f181179,f58570])).
fof(f181179,plain,(
  ( ! [X47,X46] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X46,X47)) ) | ($spl24 | $spl369)),
  inference(backward_demodulation,[],[f4580,f97028])).
fof(f97028,plain,(
  ( ! [X47,X46] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X46,X47)) ) | $spl369),
  inference(resolution,[],[f96833,f291])).
fof(f186480,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl221 | $spl369)),
  inference(resolution,[],[f181179,f58569])).
fof(f186486,plain,(
  ( ! [X2,X3] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2,X3)) ) | $spl594),
  inference(cnf_transformation,[],[f186486_D])).
fof(f186486_D,plain,(
  ( ! [X2,X3] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2,X3)) ) <=> ~$spl594),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl594])])).
fof(f186488,plain,(
  ( ! [X6,X7] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X6,X7) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(subsumption_resolution,[],[f186478,f294])).
fof(f186478,plain,(
  ( ! [X6,X7] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181179,f286])).
fof(f186487,plain,(
  ( ! [X4,X5] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X4,X5) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(subsumption_resolution,[],[f186477,f293])).
fof(f186477,plain,(
  ( ! [X4,X5] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X4,X5) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X5))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181179,f285])).
fof(f186476,plain,(
  ( ! [X2,X3] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2,X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181179,f444])).
fof(f186475,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181179,f288])).
fof(f186467,plain,(
  ( ! [X16] : (~p(s(bool,X16)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl221 | $spl369)),
  inference(resolution,[],[f181178,f58553])).
fof(f181178,plain,(
  ( ! [X14,X15] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X14,X15)) ) | ($spl24 | $spl369)),
  inference(backward_demodulation,[],[f4580,f97012])).
fof(f97012,plain,(
  ( ! [X14,X15] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X14,X15)) ) | $spl369),
  inference(resolution,[],[f96833,f274])).
fof(f186466,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl367 | $spl369)),
  inference(resolution,[],[f181178,f109032])).
fof(f186465,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181178,f745])).
fof(f186464,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl2 | $spl24 | $spl369)),
  inference(resolution,[],[f181178,f126356])).
fof(f186463,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl24 | $spl369)),
  inference(resolution,[],[f181178,f5232])).
fof(f186462,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181178,f479])).
fof(f186472,plain,(
  ( ! [X2,X3] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2,X3)) ) | $spl592),
  inference(cnf_transformation,[],[f186472_D])).
fof(f186472_D,plain,(
  ( ! [X2,X3] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2,X3)) ) <=> ~$spl592),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl592])])).
fof(f186474,plain,(
  ( ! [X6,X7] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X6,X7) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(subsumption_resolution,[],[f186459,f277])).
fof(f186459,plain,(
  ( ! [X6,X7] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181178,f269])).
fof(f186473,plain,(
  ( ! [X4,X5] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X4,X5) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(subsumption_resolution,[],[f186458,f276])).
fof(f186458,plain,(
  ( ! [X4,X5] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X4,X5) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X5))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181178,f268])).
fof(f186457,plain,(
  ( ! [X2,X3] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')),X2,X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181178,f443])).
fof(f186456,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181178,f271])).
fof(f186446,plain,(
  ( ! [X17] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X17))) ) | ($spl24 | $spl367 | $spl369)),
  inference(resolution,[],[f181177,f109031])).
fof(f181177,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) ) | ($spl24 | $spl369)),
  inference(backward_demodulation,[],[f4580,f97011])).
fof(f97011,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) | $spl369),
  inference(resolution,[],[f96833,f272])).
fof(f186445,plain,(
  ( ! [X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X16))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181177,f744])).
fof(f186443,plain,(
  ( ! [X14] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X14))) ) | ($spl2 | $spl24 | $spl369)),
  inference(resolution,[],[f181177,f126355])).
fof(f186442,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X13))) ) | ($spl4 | $spl24 | $spl369)),
  inference(resolution,[],[f181177,f5231])).
fof(f186441,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X12))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181177,f478])).
fof(f186439,plain,(
  ( ! [X10] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | ($spl24 | $spl221 | $spl369)),
  inference(resolution,[],[f181177,f58553])).
fof(f186455,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9))) ) | $spl590),
  inference(cnf_transformation,[],[f186455_D])).
fof(f186455_D,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9))) ) <=> ~$spl590),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl590])])).
fof(f186438,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181177,f270])).
fof(f186434,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num','const.num.|48|_0'))))) ) | ($spl24 | $spl369)),
  inference(resolution,[],[f181177,f271])).
fof(f186349,plain,(
  ( ! [X40] : (p(s(bool,X40)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl2 | $spl24)),
  inference(resolution,[],[f126356,f181105])).
fof(f181105,plain,(
  ( ! [X14,X15] : (sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X14,X15)) ) | ($spl1 | $spl24)),
  inference(backward_demodulation,[],[f4580,f4348])).
fof(f4348,plain,(
  ( ! [X14,X15] : (sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X14,X15)) ) | $spl1),
  inference(resolution,[],[f451,f274])).
fof(f186273,plain,(
  ( ! [X15] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | ($spl1 | $spl2 | $spl24)),
  inference(resolution,[],[f126355,f181104])).
fof(f181104,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) | ($spl1 | $spl24)),
  inference(backward_demodulation,[],[f4580,f4347])).
fof(f4347,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) | $spl1),
  inference(resolution,[],[f451,f272])).
fof(f17039,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl57),
  inference(cnf_transformation,[],[f17039_D])).
fof(f17039_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f186063,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl8),
  inference(resolution,[],[f4879,f3311])).
fof(f4879,plain,(
  ( ! [X114,X113] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X113,X114)) ) | $spl8),
  inference(resolution,[],[f462,f333])).
fof(f461,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl9),
  inference(cnf_transformation,[],[f461_D])).
fof(f461_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f186077,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X2) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X3)) ) | $spl586),
  inference(cnf_transformation,[],[f186077_D])).
fof(f186077_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X2) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X3)) ) <=> ~$spl586),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl586])])).
fof(f186059,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X3) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X2)) ) | $spl8),
  inference(resolution,[],[f4879,f319])).
fof(f71162,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | $spl239),
  inference(cnf_transformation,[],[f71162_D])).
fof(f71162_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) <=> ~$spl239),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl239])])).
fof(f185463,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | $spl235),
  inference(resolution,[],[f82079,f3428])).
fof(f82079,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X100)) ) | $spl235),
  inference(resolution,[],[f71229,f324])).
fof(f71229,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | $spl235),
  inference(resolution,[],[f70788,f438])).
fof(f71163,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) | $spl240),
  inference(cnf_transformation,[],[f71163_D])).
fof(f71163_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f185462,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) | $spl235),
  inference(resolution,[],[f82079,f3243])).
fof(f82456,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | $spl339),
  inference(cnf_transformation,[],[f82456_D])).
fof(f82456_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) <=> ~$spl339),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl339])])).
fof(f185231,plain,(
  ( ! [X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X16))) ) | ($spl235 | $spl367)),
  inference(resolution,[],[f82036,f109031])).
fof(f82036,plain,(
  ( ! [X14,X13] : (sP0(X13,X14,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) | $spl235),
  inference(resolution,[],[f71229,f272])).
fof(f185230,plain,(
  ( ! [X15] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X15))) ) | $spl235),
  inference(resolution,[],[f82036,f744])).
fof(f185228,plain,(
  ( ! [X13] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X13))) ) | ($spl4 | $spl235)),
  inference(resolution,[],[f82036,f5231])).
fof(f185227,plain,(
  ( ! [X12] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X12))) ) | $spl235),
  inference(resolution,[],[f82036,f478])).
fof(f185225,plain,(
  ( ! [X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | ($spl221 | $spl235)),
  inference(resolution,[],[f82036,f58553])).
fof(f185238,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9))) ) | $spl566),
  inference(cnf_transformation,[],[f185238_D])).
fof(f185238_D,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9))) ) <=> ~$spl566),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl566])])).
fof(f185224,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))))) ) | $spl235),
  inference(resolution,[],[f82036,f270])).
fof(f185220,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))))) ) | $spl235),
  inference(resolution,[],[f82036,f271])).
fof(f70787,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | $spl234),
  inference(cnf_transformation,[],[f70787_D])).
fof(f70787_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f83129,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X4) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) | $spl342),
  inference(cnf_transformation,[],[f83129_D])).
fof(f83129_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X4) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f184559,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X4)) ) | $spl235),
  inference(resolution,[],[f71284,f320])).
fof(f71284,plain,(
  ( ! [X107,X106] : (sP7('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X106,X107)) ) | $spl235),
  inference(resolution,[],[f70788,f329])).
fof(f82306,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) | $spl336),
  inference(cnf_transformation,[],[f82306_D])).
fof(f82306_D,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f184459,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) | $spl235),
  inference(resolution,[],[f71281,f321])).
fof(f71281,plain,(
  ( ! [X101,X100] : (sP8('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X100,X101)) ) | $spl235),
  inference(resolution,[],[f70788,f325])).
fof(f71650,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | $spl243),
  inference(cnf_transformation,[],[f71650_D])).
fof(f71650_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) <=> ~$spl243),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl243])])).
fof(f184441,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | $spl235),
  inference(resolution,[],[f71280,f3428])).
fof(f71280,plain,(
  ( ! [X99,X98] : (sP8(X98,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X99)) ) | $spl235),
  inference(resolution,[],[f70788,f324])).
fof(f71651,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) | $spl244),
  inference(cnf_transformation,[],[f71651_D])).
fof(f71651_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f184440,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) | $spl235),
  inference(resolution,[],[f71280,f3243])).
fof(f183821,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24 | $spl367)),
  inference(resolution,[],[f181112,f109078])).
fof(f183820,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181112,f791])).
fof(f183819,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl4 | $spl24)),
  inference(resolution,[],[f181112,f5278])).
fof(f183818,plain,(
  ( ! [X9] : (p(s(bool,X9)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181112,f525])).
fof(f183817,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24 | $spl221)),
  inference(resolution,[],[f181112,f58599])).
fof(f43745,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | $spl154),
  inference(cnf_transformation,[],[f43745_D])).
fof(f43745_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f183827,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X5) | ~sP6(X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X4) | ~sP8(X4,X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) | $spl564),
  inference(cnf_transformation,[],[f183827_D])).
fof(f183827_D,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X5) | ~sP6(X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X4) | ~sP8(X4,X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) <=> ~$spl564),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl564])])).
fof(f183815,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X5) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~sP8(X4,X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | ~sP6(X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X4)) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181112,f320])).
fof(f183813,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181112,f326])).
fof(f183803,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24 | $spl367)),
  inference(resolution,[],[f181111,f109077])).
fof(f183802,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181111,f790])).
fof(f183801,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl4 | $spl24)),
  inference(resolution,[],[f181111,f5277])).
fof(f183800,plain,(
  ( ! [X9] : (p(s(bool,X9)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181111,f524])).
fof(f183799,plain,(
  ( ! [X8] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~p(s(bool,X8))) ) | ($spl1 | $spl24 | $spl221)),
  inference(resolution,[],[f181111,f58599])).
fof(f183810,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) | $spl562),
  inference(cnf_transformation,[],[f183810_D])).
fof(f183810_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) <=> ~$spl562),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl562])])).
fof(f183809,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(subsumption_resolution,[],[f183808,f181109])).
fof(f183808,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X3)) ) | ($spl1 | $spl24)),
  inference(subsumption_resolution,[],[f183796,f333])).
fof(f183796,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X3) | ~sP6(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X3,X2)) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181111,f319])).
fof(f183795,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181111,f326])).
fof(f183785,plain,(
  ( ! [X17] : (p(s(bool,X17)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24 | $spl367)),
  inference(resolution,[],[f181110,f109075])).
fof(f183784,plain,(
  ( ! [X16] : (p(s(bool,X16)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181110,f788])).
fof(f183783,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl4 | $spl24)),
  inference(resolution,[],[f181110,f5275])).
fof(f183782,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181110,f522])).
fof(f183781,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24 | $spl221)),
  inference(resolution,[],[f181110,f58596])).
fof(f183778,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181110,f3425])).
fof(f183777,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,X9) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181110,f3428])).
fof(f183792,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X6,X7) | ~sP6(X7,X6,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | ~sP7(X6,X7,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) | $spl560),
  inference(cnf_transformation,[],[f183792_D])).
fof(f183792_D,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X6,X7) | ~sP6(X7,X6,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | ~sP7(X6,X7,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) <=> ~$spl560),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl560])])).
fof(f183775,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X6,X7) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~sP7(X6,X7,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | ~sP6(X7,X6,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181110,f321])).
fof(f183772,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181110,f322])).
fof(f183761,plain,(
  ( ! [X17] : (p(s(bool,X17)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24 | $spl367)),
  inference(resolution,[],[f181109,f109074])).
fof(f183760,plain,(
  ( ! [X16] : (p(s(bool,X16)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181109,f787])).
fof(f183759,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl4 | $spl24)),
  inference(resolution,[],[f181109,f5274])).
fof(f183758,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181109,f521])).
fof(f183757,plain,(
  ( ! [X13] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl1 | $spl24 | $spl221)),
  inference(resolution,[],[f181109,f58596])).
fof(f183771,plain,(
  ( ! [X8] : (sP9(X8,'F_0',happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) | $spl558),
  inference(cnf_transformation,[],[f183771_D])).
fof(f183771_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0',happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) <=> ~$spl558),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl558])])).
fof(f183756,plain,(
  ( ! [X12] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | sP9(X12,'F_0',happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181109,f3249])).
fof(f183754,plain,(
  ( ! [X10] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181109,f3425])).
fof(f183770,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | $spl557),
  inference(cnf_transformation,[],[f183770_D])).
fof(f183770_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) <=> ~$spl557),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl557])])).
fof(f183753,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | ~p(s(bool,X9))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181109,f3428])).
fof(f183752,plain,(
  ( ! [X8] : (sP9(X8,'F_0',happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181109,f3243])).
fof(f183767,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X2)) ) | $spl554),
  inference(cnf_transformation,[],[f183767_D])).
fof(f183767_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X2)) ) <=> ~$spl554),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl554])])).
fof(f183766,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X2)) ) | ($spl1 | $spl24)),
  inference(subsumption_resolution,[],[f183749,f333])).
fof(f183749,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X2) | ~sP6(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X3,X2)) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181109,f319])).
fof(f183748,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181109,f322])).
fof(f183739,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24 | $spl367)),
  inference(resolution,[],[f181108,f109062])).
fof(f183738,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181108,f775])).
fof(f183737,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl4 | $spl24)),
  inference(resolution,[],[f181108,f5262])).
fof(f183736,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181108,f509])).
fof(f183747,plain,(
  ( ! [X11] : (sP5(X11,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl552),
  inference(cnf_transformation,[],[f183747_D])).
fof(f183747_D,plain,(
  ( ! [X11] : (sP5(X11,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) <=> ~$spl552),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl552])])).
fof(f183735,plain,(
  ( ! [X11] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | sP5(X11,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181108,f505])).
fof(f183734,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24 | $spl221)),
  inference(resolution,[],[f181108,f58583])).
fof(f183744,plain,(
  ( ! [X2,X3] : (sP5(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X3) | ~p(s(bool,X3))) ) | $spl550),
  inference(cnf_transformation,[],[f183744_D])).
fof(f183744_D,plain,(
  ( ! [X2,X3] : (sP5(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X3) | ~p(s(bool,X3))) ) <=> ~$spl550),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl550])])).
fof(f183730,plain,(
  ( ! [X2,X3] : (sP5(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X3) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~p(s(bool,X3))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181108,f301])).
fof(f183729,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181108,f305])).
fof(f183720,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24 | $spl367)),
  inference(resolution,[],[f181107,f109061])).
fof(f183719,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181107,f774])).
fof(f183718,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl4 | $spl24)),
  inference(resolution,[],[f181107,f5261])).
fof(f183717,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181107,f508])).
fof(f183728,plain,(
  ( ! [X11] : (sP5(X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) | $spl548),
  inference(cnf_transformation,[],[f183728_D])).
fof(f183728_D,plain,(
  ( ! [X11] : (sP5(X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) <=> ~$spl548),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl548])])).
fof(f183716,plain,(
  ( ! [X11] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | sP5(X11,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181107,f504])).
fof(f183715,plain,(
  ( ! [X10] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl1 | $spl24 | $spl221)),
  inference(resolution,[],[f181107,f58583])).
fof(f183725,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | ~p(s(bool,X3))) ) | $spl546),
  inference(cnf_transformation,[],[f183725_D])).
fof(f183725_D,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | ~p(s(bool,X3))) ) <=> ~$spl546),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl546])])).
fof(f183711,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | ~p(s(bool,X3)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181107,f301])).
fof(f183710,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181107,f305])).
fof(f183702,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24 | $spl221)),
  inference(resolution,[],[f181106,f58570])).
fof(f181106,plain,(
  ( ! [X47,X46] : (sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X46,X47)) ) | ($spl1 | $spl24)),
  inference(backward_demodulation,[],[f4580,f4364])).
fof(f4364,plain,(
  ( ! [X47,X46] : (sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X46,X47)) ) | $spl1),
  inference(resolution,[],[f451,f291])).
fof(f183701,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24 | $spl221)),
  inference(resolution,[],[f181106,f58569])).
fof(f183707,plain,(
  ( ! [X2,X3] : (sP3(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X2,X3)) ) | $spl544),
  inference(cnf_transformation,[],[f183707_D])).
fof(f183707_D,plain,(
  ( ! [X2,X3] : (sP3(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X2,X3)) ) <=> ~$spl544),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl544])])).
fof(f183709,plain,(
  ( ! [X6,X7] : (sP3(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X6,X7) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(subsumption_resolution,[],[f183699,f294])).
fof(f183699,plain,(
  ( ! [X6,X7] : (sP3(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181106,f286])).
fof(f183708,plain,(
  ( ! [X4,X5] : (sP3(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X4,X5) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(subsumption_resolution,[],[f183698,f293])).
fof(f183698,plain,(
  ( ! [X4,X5] : (sP3(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X4,X5) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X5))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181106,f285])).
fof(f183697,plain,(
  ( ! [X2,X3] : (sP3(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X2,X3) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181106,f444])).
fof(f183696,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181106,f288])).
fof(f183685,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24 | $spl221)),
  inference(resolution,[],[f181105,f58553])).
fof(f183684,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24 | $spl367)),
  inference(resolution,[],[f181105,f109032])).
fof(f183683,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181105,f745])).
fof(f183682,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl4 | $spl24)),
  inference(resolution,[],[f181105,f5232])).
fof(f183681,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181105,f479])).
fof(f183690,plain,(
  ( ! [X2,X3] : (sP1(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X2,X3)) ) | $spl542),
  inference(cnf_transformation,[],[f183690_D])).
fof(f183690_D,plain,(
  ( ! [X2,X3] : (sP1(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X2,X3)) ) <=> ~$spl542),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl542])])).
fof(f183692,plain,(
  ( ! [X6,X7] : (sP1(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X6,X7) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(subsumption_resolution,[],[f183678,f277])).
fof(f183678,plain,(
  ( ! [X6,X7] : (sP1(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181105,f269])).
fof(f183691,plain,(
  ( ! [X4,X5] : (sP1(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X4,X5) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(subsumption_resolution,[],[f183677,f276])).
fof(f183677,plain,(
  ( ! [X4,X5] : (sP1(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X4,X5) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~p(s(bool,X5))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181105,f268])).
fof(f183676,plain,(
  ( ! [X2,X3] : (sP1(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X2,X3) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181105,f443])).
fof(f183675,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181105,f271])).
fof(f80245,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl273),
  inference(cnf_transformation,[],[f80245_D])).
fof(f80245_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) <=> ~$spl273),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl273])])).
fof(f183637,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl32),
  inference(resolution,[],[f79190,f3311])).
fof(f79190,plain,(
  ( ! [X114,X113] : (sP6('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X113,X114)) ) | $spl32),
  inference(resolution,[],[f5201,f333])).
fof(f126918,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | $spl465),
  inference(cnf_transformation,[],[f126918_D])).
fof(f126918_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) <=> ~$spl465),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl465])])).
fof(f183621,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) | $spl2),
  inference(resolution,[],[f126406,f3311])).
fof(f126406,plain,(
  ( ! [X112,X113] : (sP6(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X112,X113)) ) | $spl2),
  inference(resolution,[],[f452,f333])).
fof(f453,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) | $spl3),
  inference(cnf_transformation,[],[f453_D])).
fof(f453_D,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f183629,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X2) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X3)) ) | $spl540),
  inference(cnf_transformation,[],[f183629_D])).
fof(f183629_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X2) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X3)) ) <=> ~$spl540),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl540])])).
fof(f183617,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X3) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X2)) ) | $spl2),
  inference(resolution,[],[f126406,f319])).
fof(f183439,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl535),
  inference(cnf_transformation,[],[f183439_D])).
fof(f183439_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) <=> ~$spl535),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl535])])).
fof(f183427,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(resolution,[],[f78498,f3311])).
fof(f78498,plain,(
  ( ! [X114,X113] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X113,X114)) ) | $spl8),
  inference(resolution,[],[f78197,f333])).
fof(f183023,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK16)))) | $spl518),
  inference(cnf_transformation,[],[f183023_D])).
fof(f183023_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK16)))) <=> ~$spl518),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl518])])).
fof(f79249,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | $spl260),
  inference(cnf_transformation,[],[f79249_D])).
fof(f79249_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f182905,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK16)))) | $spl32),
  inference(resolution,[],[f5201,f569])).
fof(f569,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK14)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(superposition,[],[f250,f437])).
fof(f182874,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl8),
  inference(subsumption_resolution,[],[f182866,f404])).
fof(f404,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f234])).
fof(f234,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X13] : ((p(s(bool,X13)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.bool.AND_CLAUSES|split|3')).
fof(f182866,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl8),
  inference(resolution,[],[f78519,f364])).
fof(f364,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f216])).
fof(f216,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f78519,plain,(
  ( ! [X146] : (sP12('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X146)) ) | $spl8),
  inference(resolution,[],[f78197,f366])).
fof(f366,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f217])).
fof(f182873,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f182865,f404])).
fof(f182865,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl8),
  inference(resolution,[],[f78519,f368])).
fof(f368,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f218])).
fof(f218,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f120,f161])).
fof(f120,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X15,X14] : ((~(p(s(bool,X14)) | p(s(bool,X15))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X14)) => p(s(bool,'F_0'))) => (~p(s(bool,X15)) => p(s(bool,'F_0')))))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.sat.sat_9_prove_ltb')).
fof(f182864,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(subsumption_resolution,[],[f182856,f404])).
fof(f182856,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(resolution,[],[f78518,f364])).
fof(f78518,plain,(
  ( ! [X145] : (sP12(X145,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl8),
  inference(resolution,[],[f78197,f365])).
fof(f365,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f217])).
fof(f182863,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(subsumption_resolution,[],[f182855,f404])).
fof(f182855,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(resolution,[],[f78518,f368])).
fof(f182849,plain,(
  ( ! [X39] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X39))) ) | ($spl1 | $spl8 | $spl24)),
  inference(forward_demodulation,[],[f74658,f4580])).
fof(f74658,plain,(
  ( ! [X39] : (p(s(bool,X39)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f4875,f4394])).
fof(f4875,plain,(
  ( ! [X105,X106] : (~sP7(X105,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X106) | p(s(bool,X106)) | ~p(s(bool,X105))) ) | $spl8),
  inference(resolution,[],[f462,f326])).
fof(f182847,plain,(
  ( ! [X38] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X38))) ) | ($spl1 | $spl8 | $spl24)),
  inference(forward_demodulation,[],[f74384,f4580])).
fof(f74384,plain,(
  ( ! [X38] : (p(s(bool,X38)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f4874,f4393])).
fof(f4874,plain,(
  ( ! [X103,X104] : (~sP7('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X103,X104) | p(s(bool,X104)) | ~p(s(bool,X103))) ) | $spl8),
  inference(resolution,[],[f462,f326])).
fof(f182845,plain,(
  ( ! [X30] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X30))) ) | ($spl1 | $spl8 | $spl24)),
  inference(forward_demodulation,[],[f73926,f4580])).
fof(f73926,plain,(
  ( ! [X30] : (p(s(bool,X30)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f4872,f4391])).
fof(f4872,plain,(
  ( ! [X99,X100] : (~sP8(X99,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X100) | p(s(bool,X100)) | ~p(s(bool,X99))) ) | $spl8),
  inference(resolution,[],[f462,f322])).
fof(f182804,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f78544])).
fof(f78544,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,X8)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(superposition,[],[f78197,f377])).
fof(f377,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.bool.IMP_ANTISYM_AX')).
fof(f182805,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f78545])).
fof(f78545,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | p(s(bool,X9))) ) | $spl8),
  inference(superposition,[],[f78197,f377])).
fof(f182844,plain,(
  ( ! [X29] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X29))) ) | ($spl1 | $spl8 | $spl24)),
  inference(forward_demodulation,[],[f71739,f4580])).
fof(f71739,plain,(
  ( ! [X29] : (p(s(bool,X29)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f4858,f4377])).
fof(f4858,plain,(
  ( ! [X72,X71] : (~sP4('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X71,X72) | p(s(bool,X72)) | ~p(s(bool,X71))) ) | $spl8),
  inference(resolution,[],[f462,f305])).
fof(f182838,plain,(
  ( ! [X16] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X16))) ) | ($spl1 | $spl8 | $spl24)),
  inference(forward_demodulation,[],[f69778,f4580])).
fof(f69778,plain,(
  ( ! [X16] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | p(s(bool,X16))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f4828,f4347])).
fof(f4828,plain,(
  ( ! [X12,X11] : (~sP0('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X11,X12) | ~p(s(bool,X12)) | p(s(bool,X11))) ) | $spl8),
  inference(resolution,[],[f462,f271])).
fof(f182832,plain,(
  ( ! [X29] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X29))) ) | ($spl1 | $spl8 | $spl24)),
  inference(forward_demodulation,[],[f71964,f4580])).
fof(f71964,plain,(
  ( ! [X29] : (p(s(bool,X29)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f4859,f4378])).
fof(f4859,plain,(
  ( ! [X74,X73] : (~sP4(X73,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X74) | p(s(bool,X74)) | ~p(s(bool,X73))) ) | $spl8),
  inference(resolution,[],[f462,f305])).
fof(f182831,plain,(
  ( ! [X23] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X23))) ) | ($spl1 | $spl8 | $spl24)),
  inference(forward_demodulation,[],[f73551,f4580])).
fof(f73551,plain,(
  ( ! [X23] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | p(s(bool,X23))) ) | ($spl1 | $spl8)),
  inference(duplicate_literal_removal,[],[f73506])).
fof(f73506,plain,(
  ( ! [X23] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | p(s(bool,X23)) | p(s(bool,X23))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f4867,f44972])).
fof(f44972,plain,(
  ( ! [X8,X9] : (sP5(X8,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X9) | p(s(bool,X8))) ) | $spl1),
  inference(resolution,[],[f4378,f445])).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0)) | ~sP4(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f304])).
fof(f304,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f188])).
fof(f4867,plain,(
  ( ! [X90,X89] : (~sP5(X89,X90,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,X90)) | p(s(bool,X89))) ) | $spl8),
  inference(resolution,[],[f462,f314])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP5(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP5(X2,X0,X1))),
  inference(definition_folding,[],[f61,f150,f149])).
fof(f61,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X17,X18,X10] : ((p(s(bool,X10)) <=> (p(s(bool,X18)) & p(s(bool,X17)))) <=> ((p(s(bool,X10)) | ~p(s(bool,X18)) | ~p(s(bool,X17))) & (p(s(bool,X18)) | ~p(s(bool,X10))) & (p(s(bool,X17)) | ~p(s(bool,X10)))))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.sat.sat_23_prove_ltb')).
fof(f182830,plain,(
  ( ! [X29] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X29))) ) | ($spl1 | $spl8 | $spl24)),
  inference(forward_demodulation,[],[f73856,f4580])).
fof(f73856,plain,(
  ( ! [X29] : (p(s(bool,X29)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f4871,f4390])).
fof(f4871,plain,(
  ( ! [X97,X98] : (~sP8('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X97,X98) | p(s(bool,X98)) | ~p(s(bool,X97))) ) | $spl8),
  inference(resolution,[],[f462,f322])).
fof(f182829,plain,(
  ( ! [X23] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X23))) ) | ($spl1 | $spl8 | $spl24)),
  inference(forward_demodulation,[],[f72983,f4580])).
fof(f72983,plain,(
  ( ! [X23] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | p(s(bool,X23))) ) | ($spl1 | $spl8)),
  inference(duplicate_literal_removal,[],[f72938])).
fof(f72938,plain,(
  ( ! [X23] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | p(s(bool,X23)) | p(s(bool,X23))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f4866,f44954])).
fof(f44954,plain,(
  ( ! [X8,X9] : (sP5(X8,X9,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | p(s(bool,X8))) ) | $spl1),
  inference(resolution,[],[f4377,f445])).
fof(f4866,plain,(
  ( ! [X88,X87] : (~sP5(X87,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X88) | ~p(s(bool,X88)) | p(s(bool,X87))) ) | $spl8),
  inference(resolution,[],[f462,f314])).
fof(f182828,plain,(
  ( ! [X28] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X28))) ) | ($spl1 | $spl8 | $spl24)),
  inference(forward_demodulation,[],[f70133,f4580])).
fof(f70133,plain,(
  ( ! [X28] : (p(s(bool,X28)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))) ) | ($spl1 | $spl8)),
  inference(resolution,[],[f4829,f4348])).
fof(f4829,plain,(
  ( ! [X14,X13] : (~sP0(X13,X14,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,X14)) | ~p(s(bool,X13))) ) | $spl8),
  inference(resolution,[],[f462,f271])).
fof(f182789,plain,(
  ( ! [X16] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X16))) ) | ($spl1 | $spl24 | $spl367)),
  inference(resolution,[],[f181104,f109031])).
fof(f182788,plain,(
  ( ! [X15] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181104,f744])).
fof(f182786,plain,(
  ( ! [X13] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X13))) ) | ($spl1 | $spl4 | $spl24)),
  inference(resolution,[],[f181104,f5231])).
fof(f182785,plain,(
  ( ! [X12] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X12))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181104,f478])).
fof(f182783,plain,(
  ( ! [X10] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl1 | $spl24 | $spl221)),
  inference(resolution,[],[f181104,f58553])).
fof(f182796,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | p(s(bool,X9))) ) | $spl516),
  inference(cnf_transformation,[],[f182796_D])).
fof(f182796_D,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | p(s(bool,X9))) ) <=> ~$spl516),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl516])])).
fof(f182782,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))) | p(s(bool,X9)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181104,f270])).
fof(f182778,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl1 | $spl24)),
  inference(resolution,[],[f181104,f271])).
fof(f182636,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl2 | $spl34)),
  inference(subsumption_resolution,[],[f182628,f404])).
fof(f182628,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl2 | $spl34)),
  inference(resolution,[],[f181846,f364])).
fof(f181846,plain,(
  ( ! [X145] : (sP12(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)),X145)) ) | ($spl2 | $spl34)),
  inference(backward_demodulation,[],[f5204,f126427])).
fof(f126427,plain,(
  ( ! [X145] : (sP12(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X145)) ) | $spl2),
  inference(resolution,[],[f452,f366])).
fof(f182635,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | ($spl2 | $spl34)),
  inference(subsumption_resolution,[],[f182627,f404])).
fof(f182627,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | ($spl2 | $spl34)),
  inference(resolution,[],[f181846,f368])).
fof(f182590,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | $spl515),
  inference(cnf_transformation,[],[f182590_D])).
fof(f182590_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14))) <=> ~$spl515),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl515])])).
fof(f182580,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28051,f3311])).
fof(f28051,plain,(
  ( ! [X114,X113] : (sP6('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X113,X114)) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f25991,f4580])).
fof(f25991,plain,(
  ( ! [X114,X113] : (sP6('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X113,X114)) ) | $spl6),
  inference(resolution,[],[f4200,f333])).
fof(f4200,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | $spl6),
  inference(resolution,[],[f458,f438])).
fof(f182559,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl2 | $spl34)),
  inference(subsumption_resolution,[],[f182551,f404])).
fof(f182551,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl2 | $spl34)),
  inference(resolution,[],[f181845,f364])).
fof(f181845,plain,(
  ( ! [X144] : (sP12(X144,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) ) | ($spl2 | $spl34)),
  inference(backward_demodulation,[],[f5204,f126426])).
fof(f126426,plain,(
  ( ! [X144] : (sP12(X144,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) | $spl2),
  inference(resolution,[],[f452,f365])).
fof(f182558,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl2 | $spl34)),
  inference(subsumption_resolution,[],[f182550,f404])).
fof(f182550,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) ) | ($spl2 | $spl34)),
  inference(resolution,[],[f181845,f368])).
fof(f28077,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | $spl72),
  inference(cnf_transformation,[],[f28077_D])).
fof(f28077_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f182361,plain,(
  ( ! [X19] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X19))) ) | ($spl6 | $spl24 | $spl367)),
  inference(resolution,[],[f28067,f109061])).
fof(f28067,plain,(
  ( ! [X76,X75] : (sP4(X75,X76,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f4240,f4580])).
fof(f4240,plain,(
  ( ! [X76,X75] : (sP4(X75,X76,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f458,f306])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f190])).
fof(f182360,plain,(
  ( ! [X18] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X18))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28067,f774])).
fof(f182359,plain,(
  ( ! [X17] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X17))) ) | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f28067,f5261])).
fof(f182358,plain,(
  ( ! [X16] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X16))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28067,f508])).
fof(f182356,plain,(
  ( ! [X14] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl24 | $spl367)),
  inference(resolution,[],[f28067,f109062])).
fof(f182355,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28067,f775])).
fof(f182354,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f28067,f5262])).
fof(f182353,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28067,f509])).
fof(f182365,plain,(
  ( ! [X4,X5] : (sP5('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X4,X5)) ) | $spl512),
  inference(cnf_transformation,[],[f182365_D])).
fof(f182365_D,plain,(
  ( ! [X4,X5] : (sP5('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X4,X5)) ) <=> ~$spl512),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl512])])).
fof(f182351,plain,(
  ( ! [X8,X9] : (sP5('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X8,X9) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28067,f445])).
fof(f182366,plain,(
  ( ! [X6,X7] : (sP5('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182350,f309])).
fof(f309,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f182350,plain,(
  ( ! [X6,X7] : (sP5('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X7))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28067,f303])).
fof(f303,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2)) | ~sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f188])).
fof(f182364,plain,(
  ( ! [X4,X5] : (sP5('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X4,X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182349,f310])).
fof(f310,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f182349,plain,(
  ( ! [X4,X5] : (sP5('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28067,f302])).
fof(f302,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f188])).
fof(f182347,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28067,f305])).
fof(f182339,plain,(
  ( ! [X15] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X15))) ) | ($spl6 | $spl24 | $spl367)),
  inference(resolution,[],[f28064,f109074])).
fof(f28064,plain,(
  ( ! [X101,X102] : (sP8(X101,X102,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f4253,f4580])).
fof(f4253,plain,(
  ( ! [X101,X102] : (sP8(X101,X102,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f458,f323])).
fof(f323,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f198])).
fof(f182338,plain,(
  ( ! [X14] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28064,f787])).
fof(f182337,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f28064,f5274])).
fof(f182336,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28064,f521])).
fof(f182335,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | ($spl6 | $spl24 | $spl367)),
  inference(resolution,[],[f28064,f109075])).
fof(f182334,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28064,f788])).
fof(f182333,plain,(
  ( ! [X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X9))) ) | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f28064,f5275])).
fof(f182332,plain,(
  ( ! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X8))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28064,f522])).
fof(f182345,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X5)) ) | $spl510),
  inference(cnf_transformation,[],[f182345_D])).
fof(f182345_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X5)) ) <=> ~$spl510),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl510])])).
fof(f182344,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182343,f28057])).
fof(f28057,plain,(
  ( ! [X111,X112] : (sP6(X111,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X112)) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f4258,f4580])).
fof(f4258,plain,(
  ( ! [X111,X112] : (sP6(X111,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X112)) ) | $spl6),
  inference(resolution,[],[f458,f332])).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f203])).
fof(f182343,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~sP6(X5,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X4)) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182330,f329])).
fof(f182330,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~sP7('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X5,X4) | ~sP6(X5,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X4)) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28064,f320])).
fof(f182328,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28064,f322])).
fof(f182323,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28063,f339])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f206])).
fof(f206,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f28063,plain,(
  ( ! [X132,X131] : (sP10(X131,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X132)) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f4268,f4580])).
fof(f4268,plain,(
  ( ! [X132,X131] : (sP10(X131,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X132)) ) | $spl6),
  inference(resolution,[],[f458,f344])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & (~p(s(bool,X1)) | ~p(s(bool,X2)))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | (p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> sP10(X2,X0,X1))),
  inference(definition_folding,[],[f64,f157])).
fof(f64,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X0)) | (p(s(bool,X1)) & p(s(bool,X2)))) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X14,X15,X16] : ((p(s(bool,X14)) | (p(s(bool,X15)) & p(s(bool,X16)))) <=> ((p(s(bool,X14)) | p(s(bool,X15))) & (p(s(bool,X14)) | p(s(bool,X16)))))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.bool.LEFT_OR_OVER_AND')).
fof(f182322,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28063,f338])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X1)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f207])).
fof(f182321,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28063,f347])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X0)) | ~sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f209])).
fof(f182320,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28063,f346])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f209])).
fof(f182312,plain,(
  ( ! [X15] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X15))) ) | ($spl6 | $spl24 | $spl367)),
  inference(resolution,[],[f28062,f109077])).
fof(f28062,plain,(
  ( ! [X107,X108] : (sP7(X107,X108,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f4256,f4580])).
fof(f4256,plain,(
  ( ! [X107,X108] : (sP7(X107,X108,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f458,f327])).
fof(f182311,plain,(
  ( ! [X14] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28062,f790])).
fof(f182310,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f28062,f5277])).
fof(f182309,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28062,f524])).
fof(f182308,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | ($spl6 | $spl24 | $spl367)),
  inference(resolution,[],[f28062,f109078])).
fof(f182307,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28062,f791])).
fof(f182306,plain,(
  ( ! [X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X9))) ) | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f28062,f5278])).
fof(f182305,plain,(
  ( ! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X8))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28062,f525])).
fof(f182319,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7)) ) | $spl508),
  inference(cnf_transformation,[],[f182319_D])).
fof(f182319_D,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7)) ) <=> ~$spl508),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl508])])).
fof(f182318,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182317,f28058])).
fof(f28058,plain,(
  ( ! [X109,X110] : (sP6(X109,X110,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f4257,f4580])).
fof(f4257,plain,(
  ( ! [X109,X110] : (sP6(X109,X110,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f458,f331])).
fof(f182317,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182304,f325])).
fof(f182304,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~sP8('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X7,X6) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28062,f321])).
fof(f182301,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28062,f326])).
fof(f182295,plain,(
  ( ! [X16] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X16))) ) | ($spl6 | $spl24 | $spl367)),
  inference(resolution,[],[f28061,f109045])).
fof(f109045,plain,(
  ( ! [X39,X40] : (~sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X39,X40) | p(s(bool,X40)) | p(s(bool,X39))) ) | $spl367),
  inference(resolution,[],[f108612,f288])).
fof(f28061,plain,(
  ( ! [X41,X42] : (sP2(X41,X42,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f4223,f4580])).
fof(f4223,plain,(
  ( ! [X41,X42] : (sP2(X41,X42,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f458,f289])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f183])).
fof(f182294,plain,(
  ( ! [X15] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28061,f758])).
fof(f758,plain,(
  ( ! [X39,X40] : (~sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X39,X40) | p(s(bool,X40)) | p(s(bool,X39))) )),
  inference(resolution,[],[f470,f288])).
fof(f182292,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X13))) ) | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f28061,f5245])).
fof(f5245,plain,(
  ( ! [X39,X38] : (~sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X38,X39) | p(s(bool,X39)) | p(s(bool,X38))) ) | $spl4),
  inference(resolution,[],[f454,f288])).
fof(f182291,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X12))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28061,f492])).
fof(f492,plain,(
  ( ! [X41,X42] : (~sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X41,X42) | p(s(bool,X42)) | p(s(bool,X41))) )),
  inference(resolution,[],[f250,f288])).
fof(f182289,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl6 | $spl24 | $spl221)),
  inference(resolution,[],[f28061,f58570])).
fof(f182300,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | p(s(bool,X9))) ) | $spl506),
  inference(cnf_transformation,[],[f182300_D])).
fof(f182300_D,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | p(s(bool,X9))) ) <=> ~$spl506),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl506])])).
fof(f182288,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | p(s(bool,X9)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28061,f287])).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f180])).
fof(f182284,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28061,f288])).
fof(f182277,plain,(
  ( ! [X20] : (~p(s(bool,X20)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24 | $spl367)),
  inference(resolution,[],[f28060,f109031])).
fof(f28060,plain,(
  ( ! [X15,X16] : (sP0(X15,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X16)) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f4210,f4580])).
fof(f4210,plain,(
  ( ! [X15,X16] : (sP0(X15,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X16)) ) | $spl6),
  inference(resolution,[],[f458,f273])).
fof(f273,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f175])).
fof(f182276,plain,(
  ( ! [X19] : (~p(s(bool,X19)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28060,f744])).
fof(f182274,plain,(
  ( ! [X17] : (~p(s(bool,X17)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f28060,f5231])).
fof(f182273,plain,(
  ( ! [X16] : (~p(s(bool,X16)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28060,f478])).
fof(f182271,plain,(
  ( ! [X14] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl24 | $spl367)),
  inference(resolution,[],[f28060,f109032])).
fof(f182270,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28060,f745])).
fof(f182269,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f28060,f5232])).
fof(f182268,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28060,f479])).
fof(f182283,plain,(
  ( ! [X10] : (sP1(X10,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl504),
  inference(cnf_transformation,[],[f182283_D])).
fof(f182283_D,plain,(
  ( ! [X10] : (sP1(X10,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) <=> ~$spl504),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl504])])).
fof(f182267,plain,(
  ( ! [X10] : (sP1(X10,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28060,f477])).
fof(f477,plain,(
  ( ! [X12,X11] : (~sP0(X11,X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | sP1(X11,X12,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | p(s(bool,X12))) )),
  inference(resolution,[],[f250,f270])).
fof(f182282,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X9) | ~p(s(bool,X9))) ) | $spl502),
  inference(cnf_transformation,[],[f182282_D])).
fof(f182282_D,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X9) | ~p(s(bool,X9))) ) <=> ~$spl502),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl502])])).
fof(f182266,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X9) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X9))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28060,f270])).
fof(f182262,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28060,f271])).
fof(f182256,plain,(
  ( ! [X16] : (p(s(bool,X16)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24 | $spl367)),
  inference(resolution,[],[f28059,f109045])).
fof(f28059,plain,(
  ( ! [X43,X44] : (sP2(X43,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X44)) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f4224,f4580])).
fof(f4224,plain,(
  ( ! [X43,X44] : (sP2(X43,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X44)) ) | $spl6),
  inference(resolution,[],[f458,f290])).
fof(f290,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f183])).
fof(f182255,plain,(
  ( ! [X15] : (p(s(bool,X15)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28059,f758])).
fof(f182253,plain,(
  ( ! [X13] : (p(s(bool,X13)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f28059,f5245])).
fof(f182252,plain,(
  ( ! [X12] : (p(s(bool,X12)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28059,f492])).
fof(f182250,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl6 | $spl24 | $spl221)),
  inference(resolution,[],[f28059,f58569])).
fof(f182261,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X9) | p(s(bool,X9))) ) | $spl500),
  inference(cnf_transformation,[],[f182261_D])).
fof(f182261_D,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X9) | p(s(bool,X9))) ) <=> ~$spl500),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl500])])).
fof(f182249,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X9) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X9))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28059,f287])).
fof(f182245,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28059,f288])).
fof(f182239,plain,(
  ( ! [X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X11))) ) | ($spl6 | $spl24 | $spl221)),
  inference(resolution,[],[f28058,f58604])).
fof(f58604,plain,(
  ( ! [X116,X115] : (~sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X115,X116) | p(s(bool,X116)) | p(s(bool,X115))) ) | $spl221),
  inference(resolution,[],[f58528,f330])).
fof(f330,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f203])).
fof(f182238,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X10))) ) | ($spl6 | $spl24 | $spl221)),
  inference(resolution,[],[f28058,f58603])).
fof(f58603,plain,(
  ( ! [X114,X113] : (~sP6(X113,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X114) | p(s(bool,X114)) | p(s(bool,X113))) ) | $spl221),
  inference(resolution,[],[f58528,f330])).
fof(f182237,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28058,f3322])).
fof(f3322,plain,(
  ( ! [X2,X3] : (~sP6(X2,'F_0',X3) | p(s(bool,X2)) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f3315,f416])).
fof(f3315,plain,(
  ( ! [X2,X3] : (~sP6(X2,'F_0',X3) | p(s(bool,X2)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f3300,f337])).
fof(f3300,plain,(
  ( ! [X2,X3] : (sP9(X2,'F_0',X3) | ~sP6(X3,'F_0',X2) | p(s(bool,X3))) )),
  inference(resolution,[],[f3297,f324])).
fof(f3297,plain,(
  ( ! [X4,X5] : (~sP8(X4,X5,'F_0') | sP9(X4,'F_0',X5) | ~sP6(X5,'F_0',X4)) )),
  inference(subsumption_resolution,[],[f3295,f429])).
fof(f3295,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) )),
  inference(resolution,[],[f3265,f320])).
fof(f28365,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8)) ) | $spl88),
  inference(cnf_transformation,[],[f28365_D])).
fof(f28365_D,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f182236,plain,(
  ( ! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | s(bool,'F_0') != s(bool,X8)) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28058,f3311])).
fof(f182244,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | $spl498),
  inference(cnf_transformation,[],[f182244_D])).
fof(f182244_D,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) <=> ~$spl498),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl498])])).
fof(f182243,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182235,f325])).
fof(f182235,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~sP8('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X7,X6) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28058,f321])).
fof(f182232,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28058,f330])).
fof(f182224,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24 | $spl221)),
  inference(resolution,[],[f28057,f58604])).
fof(f182223,plain,(
  ( ! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X8))) ) | ($spl6 | $spl24 | $spl221)),
  inference(resolution,[],[f28057,f58602])).
fof(f58602,plain,(
  ( ! [X111,X112] : (~sP6(X111,X112,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X112)) | p(s(bool,X111))) ) | $spl221),
  inference(resolution,[],[f58528,f330])).
fof(f182228,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X5) | ~sP8(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | $spl496),
  inference(cnf_transformation,[],[f182228_D])).
fof(f182228_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X5) | ~sP8(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) <=> ~$spl496),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl496])])).
fof(f182227,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~sP8(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182221,f329])).
fof(f182221,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~sP8(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~sP7('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X5,X4)) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28057,f320])).
fof(f182219,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28057,f330])).
fof(f44070,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | $spl156),
  inference(cnf_transformation,[],[f44070_D])).
fof(f44070_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f182218,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182212,f404])).
fof(f182212,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28046,f364])).
fof(f28046,plain,(
  ( ! [X145] : (sP12(X145,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f26011,f4580])).
fof(f26011,plain,(
  ( ! [X145] : (sP12(X145,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f4200,f365])).
fof(f182217,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182211,f404])).
fof(f182211,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28046,f368])).
fof(f182203,plain,(
  ( ! [X11] : (p(s(bool,X11)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24 | $spl221)),
  inference(resolution,[],[f28056,f58603])).
fof(f28056,plain,(
  ( ! [X114,X113] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X113,X114)) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f4259,f4580])).
fof(f4259,plain,(
  ( ! [X114,X113] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X113,X114)) ) | $spl6),
  inference(resolution,[],[f458,f333])).
fof(f182201,plain,(
  ( ! [X10] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X10))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28056,f3322])).
fof(f182210,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | $spl495),
  inference(cnf_transformation,[],[f182210_D])).
fof(f182210_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) <=> ~$spl495),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl495])])).
fof(f182200,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28056,f3311])).
fof(f182199,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24 | $spl221)),
  inference(resolution,[],[f28056,f58602])).
fof(f28076,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | $spl73),
  inference(cnf_transformation,[],[f28076_D])).
fof(f28076_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f182206,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X2) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X3)) ) | $spl492),
  inference(cnf_transformation,[],[f182206_D])).
fof(f182206_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X2) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X3)) ) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f182196,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X3) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X2)) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28056,f319])).
fof(f182195,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28056,f330])).
fof(f182194,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182188,f404])).
fof(f182188,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28045,f364])).
fof(f28045,plain,(
  ( ! [X146] : (sP12('const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X146)) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f26012,f4580])).
fof(f26012,plain,(
  ( ! [X146] : (sP12('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X146)) ) | $spl6),
  inference(resolution,[],[f4200,f366])).
fof(f182193,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182187,f404])).
fof(f182187,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28045,f368])).
fof(f182186,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182180,f404])).
fof(f182180,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28066,f364])).
fof(f28066,plain,(
  ( ! [X145] : (sP12(X145,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f4279,f4580])).
fof(f4279,plain,(
  ( ! [X145] : (sP12(X145,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f458,f365])).
fof(f182185,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182179,f404])).
fof(f182179,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28066,f368])).
fof(f182178,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182172,f404])).
fof(f182172,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28065,f364])).
fof(f28065,plain,(
  ( ! [X146] : (sP12('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X146)) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f4280,f4580])).
fof(f4280,plain,(
  ( ! [X146] : (sP12('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X146)) ) | $spl6),
  inference(resolution,[],[f458,f366])).
fof(f182177,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f182171,f404])).
fof(f182171,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28065,f368])).
fof(f4577,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | $spl22),
  inference(cnf_transformation,[],[f4577_D])).
fof(f4577_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f131951,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126355,f25942])).
fof(f25942,plain,(
  ( ! [X15,X16] : (sP0(X15,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X16)) ) | $spl6),
  inference(resolution,[],[f4200,f273])).
fof(f132001,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126356,f25942])).
fof(f132192,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(duplicate_literal_removal,[],[f132159])).
fof(f132159,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X23))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126362,f79309])).
fof(f79309,plain,(
  ( ! [X2,X3] : (sP1(X2,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X3) | ~p(s(bool,X2))) ) | $spl6),
  inference(resolution,[],[f25942,f443])).
fof(f126362,plain,(
  ( ! [X24,X25] : (~sP1(X25,X24,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | ~p(s(bool,X25)) | p(s(bool,X24))) ) | $spl2),
  inference(resolution,[],[f452,f278])).
fof(f278,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP1(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f177])).
fof(f181023,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f181022,f389])).
fof(f389,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.bool.IMP_CLAUSES|split|3')).
fof(f181022,plain,(
  ( ! [X37] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f132284,f126439])).
fof(f126439,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | $spl2),
  inference(resolution,[],[f452,f420])).
fof(f132284,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126363,f79309])).
fof(f126363,plain,(
  ( ! [X26,X27] : (~sP1(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X26,X27) | ~p(s(bool,X27)) | p(s(bool,X26))) ) | $spl2),
  inference(resolution,[],[f452,f278])).
fof(f132403,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X13))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126369,f25955])).
fof(f25955,plain,(
  ( ! [X41,X42] : (sP2(X41,X42,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f4200,f289])).
fof(f126369,plain,(
  ( ! [X39,X38] : (~sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X38,X39) | p(s(bool,X39)) | p(s(bool,X38))) ) | $spl2),
  inference(resolution,[],[f452,f288])).
fof(f132410,plain,(
  ( ! [X20] : (p(s(bool,X20)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126369,f25956])).
fof(f25956,plain,(
  ( ! [X43,X44] : (sP2(X43,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X44)) ) | $spl6),
  inference(resolution,[],[f4200,f290])).
fof(f181019,plain,(
  ( ! [X25] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X25))) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f181018,f388])).
fof(f388,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f225])).
fof(f181018,plain,(
  ( ! [X25] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X25))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f132639,f126439])).
fof(f132639,plain,(
  ( ! [X25] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X25)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126376,f79328])).
fof(f79328,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,X2))) ) | $spl6),
  inference(resolution,[],[f25955,f444])).
fof(f126376,plain,(
  ( ! [X52,X53] : (~sP3(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X52,X53) | p(s(bool,X53)) | p(s(bool,X52))) ) | $spl2),
  inference(resolution,[],[f452,f295])).
fof(f295,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP3(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f185])).
fof(f181015,plain,(
  ( ! [X34] : (p(s(bool,X34)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f181014,f388])).
fof(f181014,plain,(
  ( ! [X34] : (~p(s(bool,'T_0')) | p(s(bool,X34)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f132648,f126439])).
fof(f132648,plain,(
  ( ! [X34] : (p(s(bool,X34)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126376,f79348])).
fof(f79348,plain,(
  ( ! [X2,X3] : (sP3(X2,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X3) | ~p(s(bool,X2))) ) | $spl6),
  inference(resolution,[],[f25956,f444])).
fof(f132944,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126385,f25972])).
fof(f25972,plain,(
  ( ! [X76,X75] : (sP4(X75,X76,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f4200,f306])).
fof(f133185,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126386,f25972])).
fof(f134739,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126398,f25985])).
fof(f25985,plain,(
  ( ! [X101,X102] : (sP8(X101,X102,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f4200,f323])).
fof(f135001,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126399,f25985])).
fof(f135111,plain,(
  ( ! [X16] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X16))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126401,f25988])).
fof(f25988,plain,(
  ( ! [X107,X108] : (sP7(X107,X108,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f4200,f327])).
fof(f135247,plain,(
  ( ! [X17] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X17))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126402,f25988])).
fof(f137001,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f136963])).
fof(f136963,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X8)) | ~p(s(bool,X9))) ) | $spl6),
  inference(resolution,[],[f79368,f314])).
fof(f79368,plain,(
  ( ! [X2,X3] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X2,X3) | ~p(s(bool,X2)) | ~p(s(bool,X3))) ) | $spl6),
  inference(resolution,[],[f25972,f301])).
fof(f137006,plain,(
  ( ! [X16] : (~p(s(bool,X16)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f137005,f389])).
fof(f137005,plain,(
  ( ! [X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f137000,f562])).
fof(f562,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))),
  inference(resolution,[],[f250,f420])).
fof(f137000,plain,(
  ( ! [X16] : (~p(s(bool,X16)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f136967])).
fof(f136967,plain,(
  ( ! [X16] : (~p(s(bool,X16)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,X16)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79368,f517])).
fof(f517,plain,(
  ( ! [X92,X91] : (~sP5(X91,X92,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | ~p(s(bool,X92)) | p(s(bool,X91))) )),
  inference(resolution,[],[f250,f314])).
fof(f137016,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f137015,f389])).
fof(f137015,plain,(
  ( ! [X23] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f136999,f126439])).
fof(f136999,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(duplicate_literal_removal,[],[f136974])).
fof(f136974,plain,(
  ( ! [X23] : (~p(s(bool,X23)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) | ~p(s(bool,X23)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f79368,f126394])).
fof(f126394,plain,(
  ( ! [X88,X89] : (~sP5(X88,X89,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | ~p(s(bool,X89)) | p(s(bool,X88))) ) | $spl2),
  inference(resolution,[],[f452,f314])).
fof(f137020,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f137019,f389])).
fof(f137019,plain,(
  ( ! [X24] : (~p(s(bool,'T_0')) | ~p(s(bool,X24)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f136998,f828])).
fof(f828,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))),
  inference(resolution,[],[f470,f420])).
fof(f136998,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f136975])).
fof(f136975,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,X24)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79368,f783])).
fof(f783,plain,(
  ( ! [X90,X89] : (~sP5(X89,X90,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | ~p(s(bool,X90)) | p(s(bool,X89))) )),
  inference(resolution,[],[f470,f314])).
fof(f137027,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f137026,f389])).
fof(f137026,plain,(
  ( ! [X28] : (~p(s(bool,'T_0')) | ~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f136997,f562])).
fof(f136997,plain,(
  ( ! [X28] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f136979])).
fof(f136979,plain,(
  ( ! [X28] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,X28)) | ~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79368,f516])).
fof(f516,plain,(
  ( ! [X90,X89] : (~sP5(X89,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X90) | ~p(s(bool,X90)) | p(s(bool,X89))) )),
  inference(resolution,[],[f250,f314])).
fof(f137037,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f137036,f389])).
fof(f137036,plain,(
  ( ! [X35] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f136996,f126439])).
fof(f136996,plain,(
  ( ! [X35] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) | ~p(s(bool,X35)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(duplicate_literal_removal,[],[f136986])).
fof(f136986,plain,(
  ( ! [X35] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) | ~p(s(bool,X35)) | ~p(s(bool,X35)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f79368,f126393])).
fof(f126393,plain,(
  ( ! [X87,X86] : (~sP5(X86,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X87) | ~p(s(bool,X87)) | p(s(bool,X86))) ) | $spl2),
  inference(resolution,[],[f452,f314])).
fof(f137041,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f137040,f389])).
fof(f137040,plain,(
  ( ! [X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f136995,f828])).
fof(f136995,plain,(
  ( ! [X36] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,X36)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f136987])).
fof(f136987,plain,(
  ( ! [X36] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,X36)) | ~p(s(bool,X36)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79368,f782])).
fof(f782,plain,(
  ( ! [X88,X87] : (~sP5(X87,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X88) | ~p(s(bool,X88)) | p(s(bool,X87))) )),
  inference(resolution,[],[f470,f314])).
fof(f149117,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f148903,f126363])).
fof(f148903,plain,(
  ( ! [X27] : (sP1(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X27)) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126441,f25942])).
fof(f126441,plain,(
  ( ! [X151,X150] : (~sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X150,X151) | sP1(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X150,X151)) ) | $spl2),
  inference(resolution,[],[f452,f443])).
fof(f180999,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f180998,f389])).
fof(f180998,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f149119,f126439])).
fof(f149119,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f148903,f278])).
fof(f149698,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f149252,f126376])).
fof(f149252,plain,(
  ( ! [X14] : (sP3(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X14,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126442,f25955])).
fof(f126442,plain,(
  ( ! [X152,X153] : (~sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X152,X153) | sP3(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X152,X153)) ) | $spl2),
  inference(resolution,[],[f452,f444])).
fof(f180987,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f180986,f388])).
fof(f180986,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f149700,f126439])).
fof(f149700,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f149252,f295])).
fof(f149814,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f149260,f126376])).
fof(f149260,plain,(
  ( ! [X22] : (sP3(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X22)) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126442,f25956])).
fof(f180983,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f180982,f388])).
fof(f180982,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f149816,f126439])).
fof(f149816,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f149260,f295])).
fof(f152509,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f82490,f25956])).
fof(f82490,plain,(
  ( ! [X43,X44] : (~sP2(X43,X44,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X44)) | ~p(s(bool,X43))) ) | $spl235),
  inference(resolution,[],[f71229,f288])).
fof(f152575,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f82491,f25955])).
fof(f82491,plain,(
  ( ! [X45,X46] : (~sP2(X45,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X46) | p(s(bool,X46)) | ~p(s(bool,X45))) ) | $spl235),
  inference(resolution,[],[f71229,f288])).
fof(f153099,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl235)),
  inference(duplicate_literal_removal,[],[f153047])).
fof(f153047,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X26))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f82496,f79348])).
fof(f82496,plain,(
  ( ! [X56,X55] : (~sP3(X56,X55,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X56)) | p(s(bool,X55))) ) | $spl235),
  inference(resolution,[],[f71229,f295])).
fof(f153291,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl235)),
  inference(duplicate_literal_removal,[],[f153239])).
fof(f153239,plain,(
  ( ! [X26] : (~p(s(bool,X26)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X26))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f82497,f79328])).
fof(f82497,plain,(
  ( ! [X57,X58] : (~sP3(X58,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X57) | ~p(s(bool,X58)) | p(s(bool,X57))) ) | $spl235),
  inference(resolution,[],[f71229,f295])).
fof(f153933,plain,(
  ( ! [X17] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X17))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f82523,f25990])).
fof(f25990,plain,(
  ( ! [X111,X112] : (sP6(X111,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X112)) ) | $spl6),
  inference(resolution,[],[f4200,f332])).
fof(f82523,plain,(
  ( ! [X109,X110] : (~sP6(X109,X110,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X110)) | p(s(bool,X109))) ) | $spl235),
  inference(resolution,[],[f71229,f330])).
fof(f153942,plain,(
  ( ! [X26] : (p(s(bool,X26)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f82523,f25991])).
fof(f153982,plain,(
  ( ! [X19] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X19))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f82524,f25989])).
fof(f25989,plain,(
  ( ! [X109,X110] : (sP6(X109,X110,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f4200,f331])).
fof(f82524,plain,(
  ( ! [X111,X112] : (~sP6(X111,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X112) | p(s(bool,X112)) | p(s(bool,X111))) ) | $spl235),
  inference(resolution,[],[f71229,f330])).
fof(f153994,plain,(
  ( ! [X31] : (p(s(bool,X31)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f82524,f25991])).
fof(f154056,plain,(
  ( ! [X19] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X19))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f82525,f25989])).
fof(f82525,plain,(
  ( ! [X114,X113] : (~sP6('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X113,X114) | p(s(bool,X114)) | p(s(bool,X113))) ) | $spl235),
  inference(resolution,[],[f71229,f330])).
fof(f154068,plain,(
  ( ! [X31] : (p(s(bool,X31)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f82525,f25990])).
fof(f155345,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78840,f25956])).
fof(f78840,plain,(
  ( ! [X43,X44] : (~sP2(X43,X44,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,X44)) | ~p(s(bool,X43))) ) | $spl9),
  inference(resolution,[],[f461,f288])).
fof(f155417,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78841,f25955])).
fof(f78841,plain,(
  ( ! [X45,X46] : (~sP2(X45,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X46) | p(s(bool,X46)) | ~p(s(bool,X45))) ) | $spl9),
  inference(resolution,[],[f461,f288])).
fof(f156582,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f83443,f25956])).
fof(f83443,plain,(
  ( ! [X43,X42] : (~sP2(X42,X43,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X43)) | ~p(s(bool,X42))) ) | $spl235),
  inference(resolution,[],[f82142,f288])).
fof(f82142,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | $spl235),
  inference(subsumption_resolution,[],[f82028,f71230])).
fof(f71230,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK16) | $spl235),
  inference(resolution,[],[f70788,f439])).
fof(f439,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f82028,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK16) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | $spl235),
  inference(resolution,[],[f71229,f437])).
fof(f156648,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f83444,f25955])).
fof(f83444,plain,(
  ( ! [X45,X44] : (~sP2(X44,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X45) | p(s(bool,X45)) | ~p(s(bool,X44))) ) | $spl235),
  inference(resolution,[],[f82142,f288])).
fof(f157151,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl235)),
  inference(duplicate_literal_removal,[],[f157093])).
fof(f157093,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X27))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f83449,f79348])).
fof(f83449,plain,(
  ( ! [X54,X55] : (~sP3(X55,X54,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X55)) | p(s(bool,X54))) ) | $spl235),
  inference(resolution,[],[f82142,f295])).
fof(f157352,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl235)),
  inference(duplicate_literal_removal,[],[f157294])).
fof(f157294,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X27))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f83450,f79328])).
fof(f83450,plain,(
  ( ! [X57,X56] : (~sP3(X57,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X56) | ~p(s(bool,X57)) | p(s(bool,X56))) ) | $spl235),
  inference(resolution,[],[f82142,f295])).
fof(f157715,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl9)),
  inference(duplicate_literal_removal,[],[f157652])).
fof(f157652,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X27))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78846,f79348])).
fof(f78846,plain,(
  ( ! [X56,X55] : (~sP3(X56,X55,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,X56)) | p(s(bool,X55))) ) | $spl9),
  inference(resolution,[],[f461,f295])).
fof(f158072,plain,(
  ( ! [X17] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X17))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f83476,f25990])).
fof(f83476,plain,(
  ( ! [X109,X108] : (~sP6(X108,X109,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X109)) | p(s(bool,X108))) ) | $spl235),
  inference(resolution,[],[f82142,f330])).
fof(f158081,plain,(
  ( ! [X26] : (p(s(bool,X26)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f83476,f25991])).
fof(f158121,plain,(
  ( ! [X19] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X19))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f83477,f25989])).
fof(f83477,plain,(
  ( ! [X111,X110] : (~sP6(X110,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X111) | p(s(bool,X111)) | p(s(bool,X110))) ) | $spl235),
  inference(resolution,[],[f82142,f330])).
fof(f158133,plain,(
  ( ! [X31] : (p(s(bool,X31)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f83477,f25991])).
fof(f158262,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl9)),
  inference(duplicate_literal_removal,[],[f158199])).
fof(f158199,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X27))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78847,f79328])).
fof(f78847,plain,(
  ( ! [X57,X58] : (~sP3(X58,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X57) | ~p(s(bool,X58)) | p(s(bool,X57))) ) | $spl9),
  inference(resolution,[],[f461,f295])).
fof(f158404,plain,(
  ( ! [X19] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X19))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f83478,f25989])).
fof(f83478,plain,(
  ( ! [X112,X113] : (~sP6('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X112,X113) | p(s(bool,X113)) | p(s(bool,X112))) ) | $spl235),
  inference(resolution,[],[f82142,f330])).
fof(f158416,plain,(
  ( ! [X31] : (p(s(bool,X31)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl235)),
  inference(resolution,[],[f83478,f25990])).
fof(f159375,plain,(
  ( ! [X15] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78873,f25990])).
fof(f78873,plain,(
  ( ! [X109,X110] : (~sP6(X109,X110,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,X110)) | p(s(bool,X109))) ) | $spl9),
  inference(resolution,[],[f461,f330])).
fof(f159384,plain,(
  ( ! [X24] : (p(s(bool,X24)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78873,f25991])).
fof(f159830,plain,(
  ( ! [X19] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X19))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78874,f25989])).
fof(f78874,plain,(
  ( ! [X111,X112] : (~sP6(X111,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X112) | p(s(bool,X112)) | p(s(bool,X111))) ) | $spl9),
  inference(resolution,[],[f461,f330])).
fof(f159842,plain,(
  ( ! [X31] : (p(s(bool,X31)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78874,f25991])).
fof(f159913,plain,(
  ( ! [X19] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X19))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78875,f25989])).
fof(f78875,plain,(
  ( ! [X114,X113] : (~sP6('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X113,X114) | p(s(bool,X114)) | p(s(bool,X113))) ) | $spl9),
  inference(resolution,[],[f461,f330])).
fof(f159925,plain,(
  ( ! [X31] : (p(s(bool,X31)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78875,f25990])).
fof(f161426,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl365)),
  inference(resolution,[],[f109599,f25956])).
fof(f109599,plain,(
  ( ! [X43,X44] : (~sP2(X43,X44,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X44)) | ~p(s(bool,X43))) ) | $spl365),
  inference(resolution,[],[f96590,f288])).
fof(f161492,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl365)),
  inference(resolution,[],[f109600,f25955])).
fof(f109600,plain,(
  ( ! [X45,X46] : (~sP2(X45,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X46) | p(s(bool,X46)) | ~p(s(bool,X45))) ) | $spl365),
  inference(resolution,[],[f96590,f288])).
fof(f162001,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl365)),
  inference(duplicate_literal_removal,[],[f161942])).
fof(f161942,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X28))) ) | ($spl6 | $spl365)),
  inference(resolution,[],[f109605,f79348])).
fof(f109605,plain,(
  ( ! [X56,X55] : (~sP3(X56,X55,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X56)) | p(s(bool,X55))) ) | $spl365),
  inference(resolution,[],[f96590,f295])).
fof(f162210,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl365)),
  inference(duplicate_literal_removal,[],[f162151])).
fof(f162151,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X28))) ) | ($spl6 | $spl365)),
  inference(resolution,[],[f109606,f79328])).
fof(f109606,plain,(
  ( ! [X57,X58] : (~sP3(X58,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X57) | ~p(s(bool,X58)) | p(s(bool,X57))) ) | $spl365),
  inference(resolution,[],[f96590,f295])).
fof(f162729,plain,(
  ( ! [X17] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X17))) ) | ($spl6 | $spl365)),
  inference(resolution,[],[f109632,f25990])).
fof(f109632,plain,(
  ( ! [X109,X110] : (~sP6(X109,X110,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X110)) | p(s(bool,X109))) ) | $spl365),
  inference(resolution,[],[f96590,f330])).
fof(f162738,plain,(
  ( ! [X26] : (p(s(bool,X26)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl365)),
  inference(resolution,[],[f109632,f25991])).
fof(f162778,plain,(
  ( ! [X19] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X19))) ) | ($spl6 | $spl365)),
  inference(resolution,[],[f109633,f25989])).
fof(f109633,plain,(
  ( ! [X111,X112] : (~sP6(X111,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X112) | p(s(bool,X112)) | p(s(bool,X111))) ) | $spl365),
  inference(resolution,[],[f96590,f330])).
fof(f162790,plain,(
  ( ! [X31] : (p(s(bool,X31)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl365)),
  inference(resolution,[],[f109633,f25991])).
fof(f162852,plain,(
  ( ! [X19] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X19))) ) | ($spl6 | $spl365)),
  inference(resolution,[],[f109634,f25989])).
fof(f109634,plain,(
  ( ! [X114,X113] : (~sP6('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X113,X114) | p(s(bool,X114)) | p(s(bool,X113))) ) | $spl365),
  inference(resolution,[],[f96590,f330])).
fof(f162864,plain,(
  ( ! [X31] : (p(s(bool,X31)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl365)),
  inference(resolution,[],[f109634,f25990])).
fof(f164072,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl365 | $spl371)),
  inference(resolution,[],[f110522,f25956])).
fof(f110522,plain,(
  ( ! [X43,X42] : (~sP2(X42,X43,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X43)) | ~p(s(bool,X42))) ) | ($spl365 | $spl371)),
  inference(resolution,[],[f106069,f288])).
fof(f106069,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | ($spl365 | $spl371)),
  inference(subsumption_resolution,[],[f105955,f96834])).
fof(f96834,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | $spl371),
  inference(cnf_transformation,[],[f96834_D])).
fof(f96834_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) <=> ~$spl371),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl371])])).
fof(f105955,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | $spl365),
  inference(resolution,[],[f96590,f437])).
fof(f164138,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl365 | $spl371)),
  inference(resolution,[],[f110523,f25955])).
fof(f110523,plain,(
  ( ! [X45,X44] : (~sP2(X44,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X45) | p(s(bool,X45)) | ~p(s(bool,X44))) ) | ($spl365 | $spl371)),
  inference(resolution,[],[f106069,f288])).
fof(f164525,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl365 | $spl371)),
  inference(duplicate_literal_removal,[],[f164462])).
fof(f164462,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X28))) ) | ($spl6 | $spl365 | $spl371)),
  inference(resolution,[],[f110528,f79348])).
fof(f110528,plain,(
  ( ! [X54,X55] : (~sP3(X55,X54,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X55)) | p(s(bool,X54))) ) | ($spl365 | $spl371)),
  inference(resolution,[],[f106069,f295])).
fof(f164738,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl365 | $spl371)),
  inference(duplicate_literal_removal,[],[f164675])).
fof(f164675,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X28))) ) | ($spl6 | $spl365 | $spl371)),
  inference(resolution,[],[f110529,f79328])).
fof(f110529,plain,(
  ( ! [X57,X56] : (~sP3(X57,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X56) | ~p(s(bool,X57)) | p(s(bool,X56))) ) | ($spl365 | $spl371)),
  inference(resolution,[],[f106069,f295])).
fof(f165258,plain,(
  ( ! [X17] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X17))) ) | ($spl6 | $spl365 | $spl371)),
  inference(resolution,[],[f110555,f25990])).
fof(f110555,plain,(
  ( ! [X109,X108] : (~sP6(X108,X109,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X109)) | p(s(bool,X108))) ) | ($spl365 | $spl371)),
  inference(resolution,[],[f106069,f330])).
fof(f165267,plain,(
  ( ! [X26] : (p(s(bool,X26)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl365 | $spl371)),
  inference(resolution,[],[f110555,f25991])).
fof(f165307,plain,(
  ( ! [X19] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X19))) ) | ($spl6 | $spl365 | $spl371)),
  inference(resolution,[],[f110556,f25989])).
fof(f110556,plain,(
  ( ! [X111,X110] : (~sP6(X110,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X111) | p(s(bool,X111)) | p(s(bool,X110))) ) | ($spl365 | $spl371)),
  inference(resolution,[],[f106069,f330])).
fof(f165319,plain,(
  ( ! [X31] : (p(s(bool,X31)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl365 | $spl371)),
  inference(resolution,[],[f110556,f25991])).
fof(f165381,plain,(
  ( ! [X19] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X19))) ) | ($spl6 | $spl365 | $spl371)),
  inference(resolution,[],[f110557,f25989])).
fof(f110557,plain,(
  ( ! [X112,X113] : (~sP6('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X112,X113) | p(s(bool,X113)) | p(s(bool,X112))) ) | ($spl365 | $spl371)),
  inference(resolution,[],[f106069,f330])).
fof(f165393,plain,(
  ( ! [X31] : (p(s(bool,X31)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl365 | $spl371)),
  inference(resolution,[],[f110557,f25990])).
fof(f166576,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109031,f25942])).
fof(f166629,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109032,f25942])).
fof(f166875,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl367)),
  inference(duplicate_literal_removal,[],[f166815])).
fof(f166815,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X28))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109038,f79309])).
fof(f109038,plain,(
  ( ! [X26,X25] : (~sP1(X26,X25,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X26)) | p(s(bool,X25))) ) | $spl367),
  inference(resolution,[],[f108612,f278])).
fof(f180977,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl367)),
  inference(subsumption_resolution,[],[f180976,f389])).
fof(f180976,plain,(
  ( ! [X46] : (~p(s(bool,'T_0')) | ~p(s(bool,X46)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl367)),
  inference(forward_demodulation,[],[f167021,f108611])).
fof(f108611,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl367),
  inference(trivial_inequality_removal,[],[f108610])).
fof(f108610,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl367),
  inference(superposition,[],[f96669,f441])).
fof(f441,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.arithmetic.BOUNDED_FORALL_THM',aHL_BOOL_CASES)).
fof(f167021,plain,(
  ( ! [X46] : (~p(s(bool,X46)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109039,f79309])).
fof(f109039,plain,(
  ( ! [X28,X27] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X27,X28) | ~p(s(bool,X28)) | p(s(bool,X27))) ) | $spl367),
  inference(resolution,[],[f108612,f278])).
fof(f167141,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X14))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109045,f25955])).
fof(f167149,plain,(
  ( ! [X22] : (p(s(bool,X22)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109045,f25956])).
fof(f180973,plain,(
  ( ! [X30] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X30))) ) | ($spl6 | $spl367)),
  inference(subsumption_resolution,[],[f180972,f388])).
fof(f180972,plain,(
  ( ! [X30] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X30))) ) | ($spl6 | $spl367)),
  inference(forward_demodulation,[],[f167191,f108611])).
fof(f167191,plain,(
  ( ! [X30] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X30)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109052,f79328])).
fof(f109052,plain,(
  ( ! [X54,X53] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X53,X54) | p(s(bool,X54)) | p(s(bool,X53))) ) | $spl367),
  inference(resolution,[],[f108612,f295])).
fof(f180969,plain,(
  ( ! [X47] : (p(s(bool,X47)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl367)),
  inference(subsumption_resolution,[],[f180968,f388])).
fof(f180968,plain,(
  ( ! [X47] : (~p(s(bool,'T_0')) | p(s(bool,X47)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl367)),
  inference(forward_demodulation,[],[f167208,f108611])).
fof(f167208,plain,(
  ( ! [X47] : (p(s(bool,X47)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109052,f79348])).
fof(f167314,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109061,f25972])).
fof(f167378,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109062,f25972])).
fof(f180965,plain,(
  ( ! [X52] : (~p(s(bool,X52)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl367)),
  inference(subsumption_resolution,[],[f180964,f389])).
fof(f180964,plain,(
  ( ! [X52] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl367)),
  inference(forward_demodulation,[],[f167853,f108611])).
fof(f167853,plain,(
  ( ! [X52] : (~p(s(bool,X52)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl6 | $spl367)),
  inference(duplicate_literal_removal,[],[f167844])).
fof(f167844,plain,(
  ( ! [X52] : (~p(s(bool,X52)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X52))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109069,f79368])).
fof(f109069,plain,(
  ( ! [X88,X87] : (~sP5(X87,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X88) | ~p(s(bool,X88)) | p(s(bool,X87))) ) | $spl367),
  inference(resolution,[],[f108612,f314])).
fof(f180961,plain,(
  ( ! [X52] : (~p(s(bool,X52)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl367)),
  inference(subsumption_resolution,[],[f180960,f389])).
fof(f180960,plain,(
  ( ! [X52] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl367)),
  inference(forward_demodulation,[],[f168005,f108611])).
fof(f168005,plain,(
  ( ! [X52] : (~p(s(bool,X52)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl6 | $spl367)),
  inference(duplicate_literal_removal,[],[f167996])).
fof(f167996,plain,(
  ( ! [X52] : (~p(s(bool,X52)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X52)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109070,f79368])).
fof(f109070,plain,(
  ( ! [X90,X89] : (~sP5(X89,X90,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X90)) | p(s(bool,X89))) ) | $spl367),
  inference(resolution,[],[f108612,f314])).
fof(f168315,plain,(
  ( ! [X16] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X16))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109074,f25985])).
fof(f168379,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109075,f25985])).
fof(f168635,plain,(
  ( ! [X19] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X19))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109077,f25988])).
fof(f168733,plain,(
  ( ! [X19] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X19))) ) | ($spl6 | $spl367)),
  inference(resolution,[],[f109078,f25988])).
fof(f180958,plain,(
  ( ! [X14] : (sP1(X14,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) | $spl488),
  inference(cnf_transformation,[],[f180958_D])).
fof(f180958_D,plain,(
  ( ! [X14] : (sP1(X14,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f169548,plain,(
  ( ! [X14] : (sP1(X14,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f126354,f25942])).
fof(f126354,plain,(
  ( ! [X8,X9] : (~sP0(X8,X9,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | sP1(X8,X9,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | p(s(bool,X9))) ) | $spl2),
  inference(resolution,[],[f452,f270])).
fof(f169675,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5231,f25942])).
fof(f169736,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5232,f25942])).
fof(f170736,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl4 | $spl6)),
  inference(duplicate_literal_removal,[],[f170667])).
fof(f170667,plain,(
  ( ! [X28] : (~p(s(bool,X28)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X28))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5238,f79309])).
fof(f5238,plain,(
  ( ! [X24,X25] : (~sP1(X25,X24,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,X25)) | p(s(bool,X24))) ) | $spl4),
  inference(resolution,[],[f454,f278])).
fof(f180955,plain,(
  ( ! [X47] : (~p(s(bool,X47)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f180954,f389])).
fof(f180954,plain,(
  ( ! [X47] : (~p(s(bool,'T_0')) | ~p(s(bool,X47)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f171015,f5315])).
fof(f5315,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl4),
  inference(resolution,[],[f454,f420])).
fof(f171015,plain,(
  ( ! [X47] : (~p(s(bool,X47)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5239,f79309])).
fof(f5239,plain,(
  ( ! [X26,X27] : (~sP1(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X26,X27) | ~p(s(bool,X27)) | p(s(bool,X26))) ) | $spl4),
  inference(resolution,[],[f454,f278])).
fof(f171536,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X14))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5245,f25955])).
fof(f171545,plain,(
  ( ! [X23] : (p(s(bool,X23)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5245,f25956])).
fof(f180951,plain,(
  ( ! [X30] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X30))) ) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f180950,f388])).
fof(f180950,plain,(
  ( ! [X30] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X30))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f171741,f5315])).
fof(f171741,plain,(
  ( ! [X30] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X30)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5252,f79328])).
fof(f5252,plain,(
  ( ! [X52,X53] : (~sP3(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X52,X53) | p(s(bool,X53)) | p(s(bool,X52))) ) | $spl4),
  inference(resolution,[],[f454,f295])).
fof(f180947,plain,(
  ( ! [X47] : (p(s(bool,X47)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f180946,f388])).
fof(f180946,plain,(
  ( ! [X47] : (~p(s(bool,'T_0')) | p(s(bool,X47)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f171758,f5315])).
fof(f171758,plain,(
  ( ! [X47] : (p(s(bool,X47)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5252,f79348])).
fof(f172117,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5261,f25972])).
fof(f172365,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5262,f25972])).
fof(f180943,plain,(
  ( ! [X54] : (~p(s(bool,X54)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f180942,f389])).
fof(f180942,plain,(
  ( ! [X54] : (~p(s(bool,'T_0')) | ~p(s(bool,X54)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f172995,f5315])).
fof(f172995,plain,(
  ( ! [X54] : (~p(s(bool,X54)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl4 | $spl6)),
  inference(duplicate_literal_removal,[],[f172978])).
fof(f172978,plain,(
  ( ! [X54] : (~p(s(bool,X54)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,X54))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5269,f79368])).
fof(f5269,plain,(
  ( ! [X87,X86] : (~sP5(X86,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X87) | ~p(s(bool,X87)) | p(s(bool,X86))) ) | $spl4),
  inference(resolution,[],[f454,f314])).
fof(f180939,plain,(
  ( ! [X54] : (~p(s(bool,X54)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f180938,f389])).
fof(f180938,plain,(
  ( ! [X54] : (~p(s(bool,'T_0')) | ~p(s(bool,X54)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f173162,f5315])).
fof(f173162,plain,(
  ( ! [X54] : (~p(s(bool,X54)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl4 | $spl6)),
  inference(duplicate_literal_removal,[],[f173145])).
fof(f173145,plain,(
  ( ! [X54] : (~p(s(bool,X54)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X54)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5270,f79368])).
fof(f5270,plain,(
  ( ! [X88,X89] : (~sP5(X88,X89,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,X89)) | p(s(bool,X88))) ) | $spl4),
  inference(resolution,[],[f454,f314])).
fof(f173505,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5274,f25985])).
fof(f173580,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5275,f25985])).
fof(f173680,plain,(
  ( ! [X19] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X19))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5277,f25988])).
fof(f173788,plain,(
  ( ! [X19] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X19))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f5278,f25988])).
fof(f180928,plain,(
  $false | ($spl22 | $spl25 | $spl129 | $spl131)),
  inference(subsumption_resolution,[],[f180927,f79701])).
fof(f79701,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl22 | $spl25)),
  inference(forward_demodulation,[],[f79586,f262])).
fof(f79586,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ($spl22 | $spl25)),
  inference(resolution,[],[f29374,f260])).
fof(f29374,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ($spl22 | $spl25)),
  inference(subsumption_resolution,[],[f29258,f4579])).
fof(f4579,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK15) | $spl25),
  inference(cnf_transformation,[],[f4579_D])).
fof(f4579_D,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK15) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f29258,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | $spl22),
  inference(resolution,[],[f4577,f437])).
fof(f180927,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl129 | $spl131)),
  inference(subsumption_resolution,[],[f180807,f43086])).
fof(f180807,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl129),
  inference(resolution,[],[f43084,f437])).
fof(f43084,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl129),
  inference(cnf_transformation,[],[f43084_D])).
fof(f43084_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl129),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl129])])).
fof(f179663,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X31)))))) ) | $spl486),
  inference(cnf_transformation,[],[f179663_D])).
fof(f179663_D,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X31)))))) ) <=> ~$spl486),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl486])])).
fof(f179627,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X31)))))) )),
  inference(resolution,[],[f178769,f3311])).
fof(f178769,plain,(
  ( ! [X2,X0,X1] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X0)))),X1,X2)) )),
  inference(superposition,[],[f177146,f442])).
fof(f177146,plain,(
  ( ! [X173,X174,X172] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X172),s('type.num.num',sK14)))),X173,X174)) )),
  inference(resolution,[],[f176036,f333])).
fof(f176036,plain,(
  ( ! [X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK14))))))) )),
  inference(resolution,[],[f175698,f438])).
fof(f179285,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X31)))))) ) | $spl484),
  inference(cnf_transformation,[],[f179285_D])).
fof(f179285_D,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X31)))))) ) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f179249,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X31)))))) )),
  inference(resolution,[],[f177784,f3311])).
fof(f177784,plain,(
  ( ! [X2,X0,X1] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X0)))),X1,X2)) )),
  inference(superposition,[],[f176095,f442])).
fof(f176095,plain,(
  ( ! [X177,X178,X176] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X176),s('type.num.num',sK14)))),X177,X178)) )),
  inference(resolution,[],[f175698,f333])).
fof(f178776,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',sK14)))))) ) | $spl482),
  inference(cnf_transformation,[],[f178776_D])).
fof(f178776_D,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',sK14)))))) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f178740,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',sK14)))))) )),
  inference(resolution,[],[f177146,f3311])).
fof(f178400,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X31))),s('type.num.num',sK14)))) ) | $spl480),
  inference(cnf_transformation,[],[f178400_D])).
fof(f178400_D,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X31))),s('type.num.num',sK14)))) ) <=> ~$spl480),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl480])])).
fof(f178366,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X31))),s('type.num.num',sK14)))) )),
  inference(resolution,[],[f176491,f3311])).
fof(f176491,plain,(
  ( ! [X173,X174,X172] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X172))),s('type.num.num',sK14)),X173,X174)) )),
  inference(resolution,[],[f175700,f333])).
fof(f175700,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))),s('type.num.num',sK14))))) )),
  inference(resolution,[],[f175560,f438])).
fof(f177791,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',sK14)))))) ) | $spl478),
  inference(cnf_transformation,[],[f177791_D])).
fof(f177791_D,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',sK14)))))) ) <=> ~$spl478),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl478])])).
fof(f177755,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',sK14)))))) )),
  inference(resolution,[],[f176095,f3311])).
fof(f177664,plain,(
  ( ! [X14] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X14)))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl476),
  inference(cnf_transformation,[],[f177664_D])).
fof(f177664_D,plain,(
  ( ! [X14] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X14)))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f177630,plain,(
  ( ! [X14,X15] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X14)))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,X15))) ) | $spl12),
  inference(resolution,[],[f78652,f3428])).
fof(f78652,plain,(
  ( ! [X156,X154,X155] : (sP8(X155,'const.prim_rec.<_2'(s('type.num.num',X154),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X156) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X154))))) ) | $spl12),
  inference(resolution,[],[f467,f324])).
fof(f177341,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X1))))))) ) | $spl474),
  inference(cnf_transformation,[],[f177341_D])).
fof(f177341_D,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X1))))))) ) <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f177370,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X2)))))) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f177344,f404])).
fof(f177344,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X2)))))) | p(s(bool,X3))) )),
  inference(resolution,[],[f177269,f364])).
fof(f177269,plain,(
  ( ! [X0,X1] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X0)))),X1)) )),
  inference(superposition,[],[f177167,f442])).
fof(f177167,plain,(
  ( ! [X227,X226] : (sP12('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X226),s('type.num.num',sK14)))),X227)) )),
  inference(resolution,[],[f176036,f366])).
fof(f177369,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X0)))))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f177343,f404])).
fof(f177343,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X0)))))) | p(s(bool,X1))) )),
  inference(resolution,[],[f177269,f368])).
fof(f177342,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X3))))))) )),
  inference(subsumption_resolution,[],[f177315,f404])).
fof(f177315,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X3))))))) )),
  inference(resolution,[],[f177241,f364])).
fof(f177241,plain,(
  ( ! [X0,X1] : (sP12(X1,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X0)))))) )),
  inference(superposition,[],[f177166,f442])).
fof(f177166,plain,(
  ( ! [X225,X224] : (sP12(X224,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X225),s('type.num.num',sK14)))))) )),
  inference(resolution,[],[f176036,f365])).
fof(f177340,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X1))))))) )),
  inference(subsumption_resolution,[],[f177314,f404])).
fof(f177314,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X1))))))) )),
  inference(resolution,[],[f177241,f368])).
fof(f177014,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X31))),s('type.num.num',sK14)))) ) | $spl472),
  inference(cnf_transformation,[],[f177014_D])).
fof(f177014_D,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X31))),s('type.num.num',sK14)))) ) <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f176980,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X31))),s('type.num.num',sK14)))) )),
  inference(resolution,[],[f175759,f3311])).
fof(f175759,plain,(
  ( ! [X175,X173,X174] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X173))),s('type.num.num',sK14)),X174,X175)) )),
  inference(resolution,[],[f175560,f333])).
fof(f176402,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X1))))))) ) | $spl470),
  inference(cnf_transformation,[],[f176402_D])).
fof(f176402_D,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X1))))))) ) <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f176431,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X2)))))) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f176405,f404])).
fof(f176405,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X2)))))) | p(s(bool,X3))) )),
  inference(resolution,[],[f176219,f364])).
fof(f176219,plain,(
  ( ! [X0,X1] : (sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X0)))),X1)) )),
  inference(superposition,[],[f176116,f442])).
fof(f176116,plain,(
  ( ! [X231,X230] : (sP12('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X230),s('type.num.num',sK14)))),X231)) )),
  inference(resolution,[],[f175698,f366])).
fof(f176430,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X0)))))) | p(s(bool,X1))) )),
  inference(subsumption_resolution,[],[f176404,f404])).
fof(f176404,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X0)))))) | p(s(bool,X1))) )),
  inference(resolution,[],[f176219,f368])).
fof(f176403,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X3))))))) )),
  inference(subsumption_resolution,[],[f176376,f404])).
fof(f176376,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X3))))))) )),
  inference(resolution,[],[f176191,f364])).
fof(f176191,plain,(
  ( ! [X0,X1] : (sP12(X1,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X0)))))) )),
  inference(superposition,[],[f176115,f442])).
fof(f176115,plain,(
  ( ! [X229,X228] : (sP12(X228,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X229),s('type.num.num',sK14)))))) )),
  inference(resolution,[],[f175698,f365])).
fof(f176401,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X1))))))) )),
  inference(subsumption_resolution,[],[f176375,f404])).
fof(f176375,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X1))))))) )),
  inference(resolution,[],[f176191,f368])).
fof(f25311,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl63),
  inference(cnf_transformation,[],[f25311_D])).
fof(f25311_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f174964,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl18),
  inference(resolution,[],[f24105,f3311])).
fof(f24105,plain,(
  ( ! [X116,X115] : (sP6('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X115,X116)) ) | $spl18),
  inference(resolution,[],[f3961,f333])).
fof(f172353,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl467),
  inference(cnf_transformation,[],[f172353_D])).
fof(f172353_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl467),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl467])])).
fof(f172321,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,X9))) ) | $spl251),
  inference(resolution,[],[f79998,f3428])).
fof(f79998,plain,(
  ( ! [X99,X98] : (sP8(X98,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X99)) ) | $spl251),
  inference(resolution,[],[f78938,f324])).
fof(f172354,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl468),
  inference(cnf_transformation,[],[f172354_D])).
fof(f172354_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) <=> ~$spl468),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl468])])).
fof(f172320,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl251),
  inference(resolution,[],[f79998,f3243])).
fof(f23984,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl59),
  inference(cnf_transformation,[],[f23984_D])).
fof(f23984_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f171877,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(resolution,[],[f3898,f3311])).
fof(f3898,plain,(
  ( ! [X114,X113] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X113,X114)) )),
  inference(resolution,[],[f587,f333])).
fof(f587,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(forward_demodulation,[],[f472,f262])).
fof(f472,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(resolution,[],[f250,f260])).
fof(f46723,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl213),
  inference(cnf_transformation,[],[f46723_D])).
fof(f46723_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl213),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl213])])).
fof(f137849,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl4),
  inference(resolution,[],[f5282,f3311])).
fof(f5282,plain,(
  ( ! [X112,X113] : (sP6(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X112,X113)) ) | $spl4),
  inference(resolution,[],[f454,f333])).
fof(f455,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl5),
  inference(cnf_transformation,[],[f455_D])).
fof(f455_D,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f93016,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X2) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X3)) ) | $spl362),
  inference(cnf_transformation,[],[f93016_D])).
fof(f93016_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X2) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X3)) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f137844,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X3) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X2)) ) | $spl4),
  inference(resolution,[],[f5282,f319])).
fof(f462,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl8),
  inference(cnf_transformation,[],[f462_D])).
fof(f462_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f106790,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X5) | ~sP6(X5,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X4) | ~sP8(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl436),
  inference(cnf_transformation,[],[f106790_D])).
fof(f106790_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X5) | ~sP6(X5,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X4) | ~sP8(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f131002,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~sP8(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~sP6(X5,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X4)) ) | $spl9),
  inference(resolution,[],[f78872,f320])).
fof(f78872,plain,(
  ( ! [X107,X108] : (sP7('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X107,X108)) ) | $spl9),
  inference(resolution,[],[f461,f329])).
fof(f106605,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X6,X7) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl430),
  inference(cnf_transformation,[],[f106605_D])).
fof(f106605_D,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X6,X7) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f130404,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl9),
  inference(resolution,[],[f78869,f321])).
fof(f78869,plain,(
  ( ! [X101,X102] : (sP8('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X101,X102)) ) | $spl9),
  inference(resolution,[],[f461,f325])).
fof(f79735,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl269),
  inference(cnf_transformation,[],[f79735_D])).
fof(f79735_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl269),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl269])])).
fof(f130285,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,X9))) ) | $spl9),
  inference(resolution,[],[f78868,f3428])).
fof(f78868,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X100)) ) | $spl9),
  inference(resolution,[],[f461,f324])).
fof(f79736,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl270),
  inference(cnf_transformation,[],[f79736_D])).
fof(f79736_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f130284,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl9),
  inference(resolution,[],[f78868,f3243])).
fof(f129421,plain,(
  $false | ($spl2 | $spl53)),
  inference(subsumption_resolution,[],[f6017,f126439])).
fof(f126904,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) | $spl2),
  inference(resolution,[],[f126406,f3311])).
fof(f3960,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl19),
  inference(cnf_transformation,[],[f3960_D])).
fof(f3960_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f81465,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X5),s('type.num.num',X6))))))) ) | $spl284),
  inference(cnf_transformation,[],[f81465_D])).
fof(f81465_D,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X5),s('type.num.num',X6))))))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f126597,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X5),s('type.num.num',X6)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))) ) | $spl12),
  inference(resolution,[],[f467,f259])).
fof(f259,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))))))) )),
  inference(cnf_transformation,[],[f168])).
fof(f88279,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f88251])).
fof(f88251,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl6),
  inference(resolution,[],[f79309,f278])).
fof(f88278,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f88257])).
fof(f88257,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,X10)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79309,f485])).
fof(f485,plain,(
  ( ! [X28,X27] : (~sP1(X28,X27,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | ~p(s(bool,X28)) | p(s(bool,X27))) )),
  inference(resolution,[],[f250,f278])).
fof(f88277,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f88262])).
fof(f88262,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | ~p(s(bool,X15)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79309,f751])).
fof(f751,plain,(
  ( ! [X26,X25] : (~sP1(X26,X25,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | ~p(s(bool,X26)) | p(s(bool,X25))) )),
  inference(resolution,[],[f470,f278])).
fof(f88288,plain,(
  ( ! [X16] : (~p(s(bool,X16)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f88287,f389])).
fof(f88287,plain,(
  ( ! [X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f88263,f562])).
fof(f88263,plain,(
  ( ! [X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,X16)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79309,f486])).
fof(f486,plain,(
  ( ! [X30,X29] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X29,X30) | ~p(s(bool,X30)) | p(s(bool,X29))) )),
  inference(resolution,[],[f250,f278])).
fof(f88304,plain,(
  ( ! [X21] : (~p(s(bool,X21)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f88303,f389])).
fof(f88303,plain,(
  ( ! [X21] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f88272,f828])).
fof(f88272,plain,(
  ( ! [X21] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,X21)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79309,f752])).
fof(f752,plain,(
  ( ! [X28,X27] : (~sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X27,X28) | ~p(s(bool,X28)) | p(s(bool,X27))) )),
  inference(resolution,[],[f470,f278])).
fof(f88334,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f88307])).
fof(f88307,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X0))) ) | $spl6),
  inference(resolution,[],[f79328,f295])).
fof(f88333,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(duplicate_literal_removal,[],[f88314])).
fof(f88314,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,X11)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f79328,f29588])).
fof(f29588,plain,(
  ( ! [X57,X58] : (~sP3(X58,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X57) | ~p(s(bool,X58)) | p(s(bool,X57))) ) | $spl73),
  inference(resolution,[],[f28076,f295])).
fof(f88332,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl221)),
  inference(duplicate_literal_removal,[],[f88315])).
fof(f88315,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | ~p(s(bool,X12)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f79328,f58576])).
fof(f58576,plain,(
  ( ! [X59,X60] : (~sP3(X60,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X59) | ~p(s(bool,X60)) | p(s(bool,X59))) ) | $spl221),
  inference(resolution,[],[f58528,f295])).
fof(f88341,plain,(
  ( ! [X15] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | $spl6),
  inference(subsumption_resolution,[],[f88340,f388])).
fof(f88340,plain,(
  ( ! [X15] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | $spl6),
  inference(forward_demodulation,[],[f88318,f562])).
fof(f88318,plain,(
  ( ! [X15] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | $spl6),
  inference(resolution,[],[f79328,f499])).
fof(f499,plain,(
  ( ! [X56,X55] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X55,X56) | p(s(bool,X56)) | p(s(bool,X55))) )),
  inference(resolution,[],[f250,f295])).
fof(f88352,plain,(
  ( ! [X20] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X20))) ) | $spl6),
  inference(subsumption_resolution,[],[f88351,f388])).
fof(f88351,plain,(
  ( ! [X20] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X20))) ) | $spl6),
  inference(forward_demodulation,[],[f88327,f828])).
fof(f88327,plain,(
  ( ! [X20] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X20))) ) | $spl6),
  inference(resolution,[],[f79328,f765])).
fof(f765,plain,(
  ( ! [X54,X53] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X53,X54) | p(s(bool,X54)) | p(s(bool,X53))) )),
  inference(resolution,[],[f470,f295])).
fof(f88428,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f88401])).
fof(f88401,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl6),
  inference(resolution,[],[f79348,f295])).
fof(f88427,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(duplicate_literal_removal,[],[f88408])).
fof(f88408,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,X11)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f79348,f29587])).
fof(f29587,plain,(
  ( ! [X56,X55] : (~sP3(X56,X55,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~p(s(bool,X56)) | p(s(bool,X55))) ) | $spl73),
  inference(resolution,[],[f28076,f295])).
fof(f88426,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl221)),
  inference(duplicate_literal_removal,[],[f88409])).
fof(f88409,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | ~p(s(bool,X12)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f79348,f58575])).
fof(f58575,plain,(
  ( ! [X57,X58] : (~sP3(X58,X57,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X58)) | p(s(bool,X57))) ) | $spl221),
  inference(resolution,[],[f58528,f295])).
fof(f88435,plain,(
  ( ! [X15] : (p(s(bool,X15)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f88434,f388])).
fof(f88434,plain,(
  ( ! [X15] : (~p(s(bool,'T_0')) | p(s(bool,X15)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f88412,f562])).
fof(f88412,plain,(
  ( ! [X15] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | p(s(bool,X15)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79348,f499])).
fof(f88446,plain,(
  ( ! [X20] : (p(s(bool,X20)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f88445,f388])).
fof(f88445,plain,(
  ( ! [X20] : (~p(s(bool,'T_0')) | p(s(bool,X20)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f88421,f828])).
fof(f88421,plain,(
  ( ! [X20] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | p(s(bool,X20)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79348,f765])).
fof(f103014,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79316,f486])).
fof(f79316,plain,(
  ( ! [X13] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X13)) ) | $spl6),
  inference(resolution,[],[f25942,f564])).
fof(f564,plain,(
  ( ! [X154,X153] : (~sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X153,X154) | sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X153,X154)) )),
  inference(resolution,[],[f250,f443])).
fof(f126336,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f126335,f389])).
fof(f126335,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | $spl6),
  inference(forward_demodulation,[],[f103016,f562])).
fof(f103016,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79316,f278])).
fof(f103070,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79318,f752])).
fof(f79318,plain,(
  ( ! [X15] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X15)) ) | $spl6),
  inference(resolution,[],[f25942,f830])).
fof(f830,plain,(
  ( ! [X152,X151] : (~sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X151,X152) | sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X151,X152)) )),
  inference(resolution,[],[f470,f443])).
fof(f126332,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f126331,f389])).
fof(f126331,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | $spl6),
  inference(forward_demodulation,[],[f103072,f828])).
fof(f103072,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79318,f278])).
fof(f103125,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl6),
  inference(resolution,[],[f79336,f499])).
fof(f79336,plain,(
  ( ! [X14] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X14,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f25955,f565])).
fof(f565,plain,(
  ( ! [X156,X155] : (~sP2('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X155,X156) | sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X155,X156)) )),
  inference(resolution,[],[f250,f444])).
fof(f126328,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f126327,f388])).
fof(f126327,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f103127,f562])).
fof(f103127,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79336,f295])).
fof(f103175,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl6),
  inference(resolution,[],[f79338,f765])).
fof(f79338,plain,(
  ( ! [X16] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X16,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f25955,f831])).
fof(f831,plain,(
  ( ! [X154,X153] : (~sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X153,X154) | sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X153,X154)) )),
  inference(resolution,[],[f470,f444])).
fof(f126324,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f126323,f388])).
fof(f126323,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f103177,f828])).
fof(f103177,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79338,f295])).
fof(f103225,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79356,f499])).
fof(f79356,plain,(
  ( ! [X14] : (sP3('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X14)) ) | $spl6),
  inference(resolution,[],[f25956,f565])).
fof(f126320,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f126319,f388])).
fof(f126319,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(forward_demodulation,[],[f103227,f562])).
fof(f103227,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79356,f295])).
fof(f103275,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79358,f765])).
fof(f79358,plain,(
  ( ! [X16] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X16)) ) | $spl6),
  inference(resolution,[],[f25956,f831])).
fof(f126316,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f126315,f388])).
fof(f126315,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(forward_demodulation,[],[f103277,f828])).
fof(f103277,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79358,f295])).
fof(f113623,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f113583])).
fof(f113583,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl6),
  inference(resolution,[],[f79309,f278])).
fof(f113622,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f113589])).
fof(f113589,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,X10)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79309,f485])).
fof(f113621,plain,(
  ( ! [X17] : (~p(s(bool,X17)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f113596])).
fof(f113596,plain,(
  ( ! [X17] : (~p(s(bool,X17)) | ~p(s(bool,X17)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79309,f751])).
fof(f113632,plain,(
  ( ! [X20] : (~p(s(bool,X20)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f113631,f389])).
fof(f113631,plain,(
  ( ! [X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f113599,f562])).
fof(f113599,plain,(
  ( ! [X20] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,X20)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79309,f486])).
fof(f113654,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f113653,f389])).
fof(f113653,plain,(
  ( ! [X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X27)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f113612,f828])).
fof(f113612,plain,(
  ( ! [X27] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,X27)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79309,f752])).
fof(f113699,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f113663])).
fof(f113663,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X0))) ) | $spl6),
  inference(resolution,[],[f79328,f295])).
fof(f113698,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(duplicate_literal_removal,[],[f113670])).
fof(f113670,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,X11)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f79328,f29588])).
fof(f113697,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl221)),
  inference(duplicate_literal_removal,[],[f113671])).
fof(f113671,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | ~p(s(bool,X12)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f79328,f58576])).
fof(f113708,plain,(
  ( ! [X15] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | $spl6),
  inference(subsumption_resolution,[],[f113707,f388])).
fof(f113707,plain,(
  ( ! [X15] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | $spl6),
  inference(forward_demodulation,[],[f113674,f562])).
fof(f113674,plain,(
  ( ! [X15] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | $spl6),
  inference(resolution,[],[f79328,f499])).
fof(f113723,plain,(
  ( ! [X22] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X22))) ) | $spl6),
  inference(subsumption_resolution,[],[f113722,f388])).
fof(f113722,plain,(
  ( ! [X22] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X22))) ) | $spl6),
  inference(forward_demodulation,[],[f113687,f828])).
fof(f113687,plain,(
  ( ! [X22] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X22))) ) | $spl6),
  inference(resolution,[],[f79328,f765])).
fof(f113794,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f113758])).
fof(f113758,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl6),
  inference(resolution,[],[f79348,f295])).
fof(f113793,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(duplicate_literal_removal,[],[f113765])).
fof(f113765,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,X11)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f79348,f29587])).
fof(f113792,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl221)),
  inference(duplicate_literal_removal,[],[f113766])).
fof(f113766,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | ~p(s(bool,X12)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f79348,f58575])).
fof(f113803,plain,(
  ( ! [X15] : (p(s(bool,X15)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f113802,f388])).
fof(f113802,plain,(
  ( ! [X15] : (~p(s(bool,'T_0')) | p(s(bool,X15)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f113769,f562])).
fof(f113769,plain,(
  ( ! [X15] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | p(s(bool,X15)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79348,f499])).
fof(f113818,plain,(
  ( ! [X22] : (p(s(bool,X22)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f113817,f388])).
fof(f113817,plain,(
  ( ! [X22] : (~p(s(bool,'T_0')) | p(s(bool,X22)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f113782,f828])).
fof(f113782,plain,(
  ( ! [X22] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | p(s(bool,X22)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79348,f765])).
fof(f122400,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78840,f25956])).
fof(f122472,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78841,f25955])).
fof(f122874,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl9)),
  inference(duplicate_literal_removal,[],[f122837])).
fof(f122837,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X24))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78846,f79348])).
fof(f123155,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl9)),
  inference(duplicate_literal_removal,[],[f123118])).
fof(f123118,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X24))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78847,f79328])).
fof(f123885,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79316,f486])).
fof(f126189,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f126188,f389])).
fof(f126188,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | $spl6),
  inference(forward_demodulation,[],[f123887,f562])).
fof(f123887,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79316,f278])).
fof(f123940,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79318,f752])).
fof(f126185,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f126184,f389])).
fof(f126184,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | $spl6),
  inference(forward_demodulation,[],[f123942,f828])).
fof(f123942,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79318,f278])).
fof(f123995,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl6),
  inference(resolution,[],[f79336,f499])).
fof(f126181,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f126180,f388])).
fof(f126180,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f123997,f562])).
fof(f123997,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79336,f295])).
fof(f124049,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl6),
  inference(resolution,[],[f79338,f765])).
fof(f126177,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f126176,f388])).
fof(f126176,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(forward_demodulation,[],[f124051,f828])).
fof(f124051,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79338,f295])).
fof(f124103,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79356,f499])).
fof(f126173,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f126172,f388])).
fof(f126172,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(forward_demodulation,[],[f124105,f562])).
fof(f124105,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79356,f295])).
fof(f124223,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79358,f765])).
fof(f126169,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f126168,f388])).
fof(f126168,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(forward_demodulation,[],[f124225,f828])).
fof(f124225,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f79358,f295])).
fof(f124287,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X14))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78873,f25990])).
fof(f124295,plain,(
  ( ! [X22] : (p(s(bool,X22)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78873,f25991])).
fof(f124335,plain,(
  ( ! [X15] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78874,f25989])).
fof(f124343,plain,(
  ( ! [X23] : (p(s(bool,X23)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78874,f25991])).
fof(f124736,plain,(
  ( ! [X15] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78875,f25989])).
fof(f124744,plain,(
  ( ! [X23] : (p(s(bool,X23)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl9)),
  inference(resolution,[],[f78875,f25990])).
fof(f126105,plain,(
  $false | ($spl4 | $spl49 | $spl130)),
  inference(subsumption_resolution,[],[f125927,f5898])).
fof(f125927,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | ($spl4 | $spl130)),
  inference(backward_demodulation,[],[f43087,f5315])).
fof(f451,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | $spl1),
  inference(cnf_transformation,[],[f451_D])).
fof(f451_D,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f121317,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))) ) | $spl49),
  inference(subsumption_resolution,[],[f121306,f420])).
fof(f121306,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))) ) | $spl49),
  inference(superposition,[],[f5898,f380])).
fof(f380,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f130])).
fof(f121316,plain,(
  ( ! [X10] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | ~p(s(bool,X10))) ) | $spl49),
  inference(subsumption_resolution,[],[f121305,f420])).
fof(f121305,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | ~p(s(bool,X10))) ) | $spl49),
  inference(superposition,[],[f5898,f380])).
fof(f120492,plain,(
  $false | ($spl227 | $spl382)),
  inference(subsumption_resolution,[],[f120485,f104020])).
fof(f120485,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | $spl227),
  inference(trivial_inequality_removal,[],[f120481])).
fof(f120481,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | $spl227),
  inference(superposition,[],[f58882,f422])).
fof(f422,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.bool.EQ_CLAUSES|split|1')).
fof(f58882,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | $spl227),
  inference(cnf_transformation,[],[f58882_D])).
fof(f58882_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) <=> ~$spl227),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl227])])).
fof(f120491,plain,(
  $false | ($spl227 | $spl382)),
  inference(subsumption_resolution,[],[f120486,f104020])).
fof(f120486,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | $spl227),
  inference(trivial_inequality_removal,[],[f120480])).
fof(f120480,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | $spl227),
  inference(superposition,[],[f58882,f420])).
fof(f111884,plain,(
  $false | ($spl365 | $spl439)),
  inference(subsumption_resolution,[],[f106861,f109665])).
fof(f109665,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | $spl365),
  inference(resolution,[],[f96590,f416])).
fof(f111511,plain,(
  $false | ($spl377 | $spl433)),
  inference(subsumption_resolution,[],[f111502,f97281])).
fof(f111502,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | $spl433),
  inference(trivial_inequality_removal,[],[f111501])).
fof(f111501,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | $spl433),
  inference(superposition,[],[f106644,f441])).
fof(f111510,plain,(
  $false | ($spl365 | $spl371 | $spl433)),
  inference(subsumption_resolution,[],[f111503,f106069])).
fof(f111503,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | $spl433),
  inference(trivial_inequality_removal,[],[f111497])).
fof(f111497,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | $spl433),
  inference(superposition,[],[f106644,f418])).
fof(f111509,plain,(
  $false | ($spl365 | $spl371 | $spl433)),
  inference(subsumption_resolution,[],[f111504,f106069])).
fof(f111504,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | $spl433),
  inference(trivial_inequality_removal,[],[f111496])).
fof(f111496,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | $spl433),
  inference(superposition,[],[f106644,f416])).
fof(f96832,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | $spl368),
  inference(cnf_transformation,[],[f96832_D])).
fof(f96832_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f110275,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X4) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) | $spl462),
  inference(cnf_transformation,[],[f110275_D])).
fof(f110275_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X4) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f110263,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X4)) ) | $spl369),
  inference(resolution,[],[f97058,f320])).
fof(f110208,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) | $spl458),
  inference(cnf_transformation,[],[f110208_D])).
fof(f110208_D,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f110191,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) | $spl369),
  inference(resolution,[],[f97055,f321])).
fof(f97281,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | $spl377),
  inference(cnf_transformation,[],[f97281_D])).
fof(f97281_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) <=> ~$spl377),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl377])])).
fof(f110172,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | $spl369),
  inference(resolution,[],[f97054,f3428])).
fof(f97282,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) | $spl378),
  inference(cnf_transformation,[],[f97282_D])).
fof(f97282_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f110171,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) | $spl369),
  inference(resolution,[],[f97054,f3243])).
fof(f96590,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | $spl365),
  inference(cnf_transformation,[],[f96590_D])).
fof(f96590_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) <=> ~$spl365),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl365])])).
fof(f109507,plain,(
  ( ! [X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | $spl373),
  inference(subsumption_resolution,[],[f109496,f420])).
fof(f109496,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | $spl373),
  inference(superposition,[],[f96948,f380])).
fof(f109509,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))))) ) | $spl373),
  inference(subsumption_resolution,[],[f109497,f420])).
fof(f109497,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))))) ) | $spl373),
  inference(superposition,[],[f96948,f380])).
fof(f109512,plain,(
  $false | ($spl364 | $spl373)),
  inference(subsumption_resolution,[],[f109505,f96589])).
fof(f96589,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | $spl364),
  inference(cnf_transformation,[],[f96589_D])).
fof(f96589_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f109505,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | $spl373),
  inference(trivial_inequality_removal,[],[f109501])).
fof(f109501,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | $spl373),
  inference(superposition,[],[f96948,f422])).
fof(f109511,plain,(
  $false | ($spl364 | $spl373)),
  inference(subsumption_resolution,[],[f109506,f96589])).
fof(f109506,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | $spl373),
  inference(trivial_inequality_removal,[],[f109500])).
fof(f109500,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | $spl373),
  inference(superposition,[],[f96948,f420])).
fof(f108867,plain,(
  $false | $spl366),
  inference(subsumption_resolution,[],[f108856,f429])).
fof(f108856,plain,(
  p(s(bool,'F_0')) | $spl366),
  inference(trivial_inequality_removal,[],[f108633])).
fof(f108633,plain,(
  p(s(bool,'F_0')) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | $spl366),
  inference(superposition,[],[f439,f96668])).
fof(f96668,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl366),
  inference(cnf_transformation,[],[f96668_D])).
fof(f96668_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f108616,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl440),
  inference(cnf_transformation,[],[f108616_D])).
fof(f108616_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f108617,plain,(
  ( ! [X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X9))) ) | $spl367),
  inference(subsumption_resolution,[],[f108602,f416])).
fof(f108602,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X9))) ) | $spl367),
  inference(superposition,[],[f96669,f377])).
fof(f108614,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl367),
  inference(subsumption_resolution,[],[f108601,f416])).
fof(f108601,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))))) ) | $spl367),
  inference(superposition,[],[f96669,f377])).
fof(f108576,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl4),
  inference(resolution,[],[f5282,f3311])).
fof(f108570,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X3) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X2)) ) | $spl4),
  inference(resolution,[],[f5282,f319])).
fof(f104218,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | $spl405),
  inference(cnf_transformation,[],[f104218_D])).
fof(f104218_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) <=> ~$spl405),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl405])])).
fof(f107372,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | $spl382),
  inference(resolution,[],[f107091,f3311])).
fof(f107091,plain,(
  ( ! [X112,X113] : (sP6(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X112,X113)) ) | $spl382),
  inference(resolution,[],[f104020,f333])).
fof(f106861,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | $spl439),
  inference(cnf_transformation,[],[f106861_D])).
fof(f106861_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) <=> ~$spl439),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl439])])).
fof(f106846,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) | $spl364),
  inference(resolution,[],[f106320,f3311])).
fof(f106320,plain,(
  ( ! [X114,X113] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X113,X114)) ) | $spl364),
  inference(resolution,[],[f96589,f333])).
fof(f106772,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~sP8(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~sP6(X5,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X4)) ) | $spl9),
  inference(resolution,[],[f78872,f320])).
fof(f106644,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | $spl433),
  inference(cnf_transformation,[],[f106644_D])).
fof(f106644_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) <=> ~$spl433),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl433])])).
fof(f106629,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) | $spl368),
  inference(resolution,[],[f106137,f3311])).
fof(f106137,plain,(
  ( ! [X114,X113] : (sP6('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X113,X114)) ) | $spl368),
  inference(resolution,[],[f96832,f333])).
fof(f106582,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl9),
  inference(resolution,[],[f78869,f321])).
fof(f106481,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,X9))) ) | $spl9),
  inference(resolution,[],[f78868,f3428])).
fof(f106480,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl9),
  inference(resolution,[],[f78868,f3243])).
fof(f106196,plain,(
  $false | ($spl365 | $spl368 | $spl371)),
  inference(subsumption_resolution,[],[f106195,f96590])).
fof(f106195,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | ($spl368 | $spl371)),
  inference(subsumption_resolution,[],[f106079,f96834])).
fof(f106079,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | $spl368),
  inference(resolution,[],[f96832,f437])).
fof(f3961,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl18),
  inference(cnf_transformation,[],[f3961_D])).
fof(f3961_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3963,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl20),
  inference(cnf_transformation,[],[f3963_D])).
fof(f3963_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f105756,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(resolution,[],[f587,f437])).
fof(f450,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | $spl0),
  inference(cnf_transformation,[],[f450_D])).
fof(f450_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f45064,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X5) | ~sP6(X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X4) | ~sP8(X4,X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) | $spl210),
  inference(cnf_transformation,[],[f45064_D])).
fof(f45064_D,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X5) | ~sP6(X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X4) | ~sP8(X4,X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f105735,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X5) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | ~sP8(X4,X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | ~sP6(X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X4)) ) | $spl1),
  inference(resolution,[],[f4394,f320])).
fof(f45029,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X6,X7) | ~sP6(X7,X6,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | ~sP7(X6,X7,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) | $spl206),
  inference(cnf_transformation,[],[f45029_D])).
fof(f45029_D,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X6,X7) | ~sP6(X7,X6,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | ~sP7(X6,X7,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f105699,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X6,X7) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | ~sP7(X6,X7,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | ~sP6(X7,X6,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) | $spl1),
  inference(resolution,[],[f4391,f321])).
fof(f5898,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | $spl49),
  inference(cnf_transformation,[],[f5898_D])).
fof(f5898_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f105681,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | ~p(s(bool,X9))) ) | $spl1),
  inference(resolution,[],[f4390,f3428])).
fof(f5899,plain,(
  ( ! [X8] : (sP9(X8,'F_0',happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) | $spl50),
  inference(cnf_transformation,[],[f5899_D])).
fof(f5899_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0',happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f105680,plain,(
  ( ! [X8] : (sP9(X8,'F_0',happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) | $spl1),
  inference(resolution,[],[f4390,f3243])).
fof(f105358,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26000,f339])).
fof(f26000,plain,(
  ( ! [X132,X131] : (sP10(X131,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X132)) ) | $spl6),
  inference(resolution,[],[f4200,f344])).
fof(f105357,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26000,f338])).
fof(f105356,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26000,f347])).
fof(f105355,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26000,f346])).
fof(f105331,plain,(
  ( ! [X15] : (p(s(bool,X15)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25991,f4681])).
fof(f4681,plain,(
  ( ! [X111,X110] : (~sP6(X110,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X111) | p(s(bool,X111)) | p(s(bool,X110))) ) | $spl3),
  inference(resolution,[],[f453,f330])).
fof(f105330,plain,(
  ( ! [X14] : (p(s(bool,X14)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25991,f58603])).
fof(f105329,plain,(
  ( ! [X13] : (p(s(bool,X13)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25991,f29615])).
fof(f29615,plain,(
  ( ! [X111,X112] : (~sP6(X111,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X112) | p(s(bool,X112)) | p(s(bool,X111))) ) | $spl73),
  inference(resolution,[],[f28076,f330])).
fof(f105326,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X12))) ) | $spl6),
  inference(resolution,[],[f25991,f3322])).
fof(f26389,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | $spl69),
  inference(cnf_transformation,[],[f26389_D])).
fof(f26389_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f105325,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f25991,f3311])).
fof(f105324,plain,(
  ( ! [X10] : (p(s(bool,X10)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25991,f4680])).
fof(f4680,plain,(
  ( ! [X109,X108] : (~sP6(X108,X109,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | p(s(bool,X109)) | p(s(bool,X108))) ) | $spl3),
  inference(resolution,[],[f453,f330])).
fof(f105323,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25991,f58602])).
fof(f105322,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25991,f29614])).
fof(f29614,plain,(
  ( ! [X109,X110] : (~sP6(X109,X110,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | p(s(bool,X110)) | p(s(bool,X109))) ) | $spl73),
  inference(resolution,[],[f28076,f330])).
fof(f4578,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | $spl23),
  inference(cnf_transformation,[],[f4578_D])).
fof(f4578_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f43736,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP7(X3,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X2) | ~sP8(X2,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X3)) ) | $spl152),
  inference(cnf_transformation,[],[f43736_D])).
fof(f43736_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP7(X3,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X2) | ~sP8(X2,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X3)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f105319,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X2,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X3) | ~sP7(X3,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X2)) ) | $spl6),
  inference(resolution,[],[f25991,f319])).
fof(f105318,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f25991,f330])).
fof(f105309,plain,(
  ( ! [X13] : (p(s(bool,X13)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25990,f4682])).
fof(f4682,plain,(
  ( ! [X112,X113] : (~sP6(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X112,X113) | p(s(bool,X113)) | p(s(bool,X112))) ) | $spl3),
  inference(resolution,[],[f453,f330])).
fof(f105308,plain,(
  ( ! [X12] : (p(s(bool,X12)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25990,f58604])).
fof(f105307,plain,(
  ( ! [X11] : (p(s(bool,X11)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25990,f29616])).
fof(f29616,plain,(
  ( ! [X114,X113] : (~sP6('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X113,X114) | p(s(bool,X114)) | p(s(bool,X113))) ) | $spl73),
  inference(resolution,[],[f28076,f330])).
fof(f105306,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X10))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25990,f4680])).
fof(f105305,plain,(
  ( ! [X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X9))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25990,f58602])).
fof(f105304,plain,(
  ( ! [X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X8))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25990,f29614])).
fof(f43712,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5)) ) | $spl150),
  inference(cnf_transformation,[],[f43712_D])).
fof(f43712_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f105316,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f105315,f25985])).
fof(f105315,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(subsumption_resolution,[],[f105302,f329])).
fof(f105302,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5,X4)) ) | $spl6),
  inference(resolution,[],[f25990,f320])).
fof(f105300,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f25990,f330])).
fof(f105291,plain,(
  ( ! [X15] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25989,f4682])).
fof(f105290,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X14))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25989,f58604])).
fof(f105289,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X13))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25989,f29616])).
fof(f105288,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X12))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25989,f4681])).
fof(f105287,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X11))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25989,f58603])).
fof(f105286,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X10))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25989,f29615])).
fof(f105285,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25989,f3322])).
fof(f105284,plain,(
  ( ! [X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | s(bool,'F_0') != s(bool,X8)) ) | $spl6),
  inference(resolution,[],[f25989,f3311])).
fof(f43692,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7)) ) | $spl148),
  inference(cnf_transformation,[],[f43692_D])).
fof(f43692_D,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f105299,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f105298,f25988])).
fof(f105298,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(subsumption_resolution,[],[f105283,f325])).
fof(f105283,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X7,X6) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f25989,f321])).
fof(f105280,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25989,f330])).
fof(f105273,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | $spl6),
  inference(resolution,[],[f25988,f790])).
fof(f105272,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | $spl6),
  inference(resolution,[],[f25988,f524])).
fof(f105271,plain,(
  ( ! [X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X9))) ) | $spl6),
  inference(resolution,[],[f25988,f791])).
fof(f105270,plain,(
  ( ! [X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X8))) ) | $spl6),
  inference(resolution,[],[f25988,f525])).
fof(f43669,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl146),
  inference(cnf_transformation,[],[f43669_D])).
fof(f43669_D,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f105279,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(subsumption_resolution,[],[f105269,f325])).
fof(f105269,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X7,X6) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f25988,f321])).
fof(f105266,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25988,f326])).
fof(f105259,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | $spl6),
  inference(resolution,[],[f25985,f787])).
fof(f105258,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | $spl6),
  inference(resolution,[],[f25985,f521])).
fof(f105257,plain,(
  ( ! [X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X9))) ) | $spl6),
  inference(resolution,[],[f25985,f788])).
fof(f105256,plain,(
  ( ! [X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X8))) ) | $spl6),
  inference(resolution,[],[f25985,f522])).
fof(f43652,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4)) ) | $spl144),
  inference(cnf_transformation,[],[f43652_D])).
fof(f43652_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f105265,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4)) ) | $spl6),
  inference(subsumption_resolution,[],[f105254,f329])).
fof(f105254,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5,X4) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4)) ) | $spl6),
  inference(resolution,[],[f25985,f320])).
fof(f105252,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25985,f322])).
fof(f105227,plain,(
  ( ! [X15] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X15))) ) | $spl6),
  inference(resolution,[],[f25972,f774])).
fof(f105226,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | $spl6),
  inference(resolution,[],[f25972,f508])).
fof(f105224,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | $spl6),
  inference(resolution,[],[f25972,f775])).
fof(f105223,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | $spl6),
  inference(resolution,[],[f25972,f509])).
fof(f43634,plain,(
  ( ! [X4,X5] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4,X5)) ) | $spl142),
  inference(cnf_transformation,[],[f43634_D])).
fof(f43634_D,plain,(
  ( ! [X4,X5] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4,X5)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f105221,plain,(
  ( ! [X8,X9] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X8,X9) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25972,f445])).
fof(f105233,plain,(
  ( ! [X6,X7] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f105220,f309])).
fof(f105220,plain,(
  ( ! [X6,X7] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X7))) ) | $spl6),
  inference(resolution,[],[f25972,f303])).
fof(f105232,plain,(
  ( ! [X4,X5] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4,X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f105219,f310])).
fof(f105219,plain,(
  ( ! [X4,X5] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25972,f302])).
fof(f105217,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25972,f305])).
fof(f105210,plain,(
  ( ! [X16] : (p(s(bool,X16)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25956,f758])).
fof(f105208,plain,(
  ( ! [X14] : (p(s(bool,X14)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25956,f492])).
fof(f105206,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25956,f4647])).
fof(f4647,plain,(
  ( ! [X43,X42] : (~sP2(X42,X43,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | p(s(bool,X43)) | ~p(s(bool,X42))) ) | $spl3),
  inference(resolution,[],[f453,f288])).
fof(f105205,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25956,f58569])).
fof(f105204,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25956,f29581])).
fof(f29581,plain,(
  ( ! [X43,X44] : (~sP2(X43,X44,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | p(s(bool,X44)) | ~p(s(bool,X43))) ) | $spl73),
  inference(resolution,[],[f28076,f288])).
fof(f43617,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X9) | p(s(bool,X9))) ) | $spl140),
  inference(cnf_transformation,[],[f43617_D])).
fof(f43617_D,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X9) | p(s(bool,X9))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f105203,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X9) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X9))) ) | $spl6),
  inference(resolution,[],[f25956,f287])).
fof(f105199,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f25956,f288])).
fof(f105192,plain,(
  ( ! [X16] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X16))) ) | $spl6),
  inference(resolution,[],[f25955,f758])).
fof(f105190,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X14))) ) | $spl6),
  inference(resolution,[],[f25955,f492])).
fof(f105188,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25955,f4648])).
fof(f4648,plain,(
  ( ! [X45,X44] : (~sP2(X44,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X45) | p(s(bool,X45)) | ~p(s(bool,X44))) ) | $spl3),
  inference(resolution,[],[f453,f288])).
fof(f105187,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25955,f58570])).
fof(f105186,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25955,f29582])).
fof(f29582,plain,(
  ( ! [X45,X46] : (~sP2(X45,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X46) | p(s(bool,X46)) | ~p(s(bool,X45))) ) | $spl73),
  inference(resolution,[],[f28076,f288])).
fof(f43597,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | p(s(bool,X9))) ) | $spl138),
  inference(cnf_transformation,[],[f43597_D])).
fof(f43597_D,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | p(s(bool,X9))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f105185,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25955,f287])).
fof(f105181,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25955,f288])).
fof(f105174,plain,(
  ( ! [X16] : (~p(s(bool,X16)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25942,f744])).
fof(f105172,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25942,f478])).
fof(f105170,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | $spl6),
  inference(resolution,[],[f25942,f745])).
fof(f105169,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | $spl6),
  inference(resolution,[],[f25942,f479])).
fof(f79326,plain,(
  ( ! [X10] : (sP1(X10,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl262),
  inference(cnf_transformation,[],[f79326_D])).
fof(f79326_D,plain,(
  ( ! [X10] : (sP1(X10,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f105168,plain,(
  ( ! [X10] : (sP1(X10,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25942,f477])).
fof(f43577,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X9) | ~p(s(bool,X9))) ) | $spl136),
  inference(cnf_transformation,[],[f43577_D])).
fof(f43577_D,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X9) | ~p(s(bool,X9))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f105167,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X9) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X9))) ) | $spl6),
  inference(resolution,[],[f25942,f270])).
fof(f105163,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f25942,f271])).
fof(f26220,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | $spl67),
  inference(cnf_transformation,[],[f26220_D])).
fof(f26220_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f105128,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f4259,f3311])).
fof(f457,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | $spl7),
  inference(cnf_transformation,[],[f457_D])).
fof(f457_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f28404,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X2) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X3)) ) | $spl92),
  inference(cnf_transformation,[],[f28404_D])).
fof(f28404_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X2) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X3)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f105122,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X3) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X2)) ) | $spl6),
  inference(resolution,[],[f4259,f319])).
fof(f104983,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f104969,f404])).
fof(f104969,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f26012,f364])).
fof(f104982,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f104968,f404])).
fof(f104968,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl6),
  inference(resolution,[],[f26012,f368])).
fof(f104967,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f104953,f404])).
fof(f104953,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26011,f364])).
fof(f104966,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f104952,f404])).
fof(f104952,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26011,f368])).
fof(f104019,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | $spl383),
  inference(cnf_transformation,[],[f104019_D])).
fof(f104019_D,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) <=> ~$spl383),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl383])])).
fof(f104909,plain,(
  ( ! [X16] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X16))) ) | ($spl1 | $spl370)),
  inference(resolution,[],[f104583,f744])).
fof(f104583,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | ($spl1 | $spl370)),
  inference(forward_demodulation,[],[f4347,f96835])).
fof(f104907,plain,(
  ( ! [X14] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X14))) ) | ($spl1 | $spl370)),
  inference(resolution,[],[f104583,f478])).
fof(f104905,plain,(
  ( ! [X12] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X12))) ) | ($spl1 | $spl3 | $spl370)),
  inference(resolution,[],[f104583,f4631])).
fof(f4631,plain,(
  ( ! [X10,X11] : (~sP0(X10,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X11) | ~p(s(bool,X11)) | ~p(s(bool,X10))) ) | $spl3),
  inference(resolution,[],[f453,f271])).
fof(f104904,plain,(
  ( ! [X11] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X11))) ) | ($spl1 | $spl221 | $spl370)),
  inference(resolution,[],[f104583,f58553])).
fof(f104903,plain,(
  ( ! [X10] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | ($spl1 | $spl73 | $spl370)),
  inference(resolution,[],[f104583,f29565])).
fof(f29565,plain,(
  ( ! [X12,X11] : (~sP0(X11,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X12) | ~p(s(bool,X12)) | ~p(s(bool,X11))) ) | $spl73),
  inference(resolution,[],[f28076,f271])).
fof(f104261,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9))) ) | $spl406),
  inference(cnf_transformation,[],[f104261_D])).
fof(f104261_D,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9))) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f104902,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl1 | $spl370)),
  inference(resolution,[],[f104583,f270])).
fof(f104898,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl1 | $spl370)),
  inference(resolution,[],[f104583,f271])).
fof(f96669,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl367),
  inference(cnf_transformation,[],[f96669_D])).
fof(f96669_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) <=> ~$spl367),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl367])])).
fof(f104862,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | ($spl368 | $spl370)),
  inference(resolution,[],[f104652,f3311])).
fof(f104652,plain,(
  ( ! [X114,X113] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X113,X114)) ) | ($spl368 | $spl370)),
  inference(resolution,[],[f104593,f333])).
fof(f104593,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl368 | $spl370)),
  inference(forward_demodulation,[],[f96832,f96835])).
fof(f104427,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X5),s('type.num.num',X6)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))) ) | $spl12),
  inference(resolution,[],[f467,f259])).
fof(f104020,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | $spl382),
  inference(cnf_transformation,[],[f104020_D])).
fof(f104020_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f104426,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',X4))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl12),
  inference(resolution,[],[f467,f258])).
fof(f81340,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',X4))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl12),
  inference(resolution,[],[f467,f258])).
fof(f78600,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',X4))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl12),
  inference(resolution,[],[f467,f258])).
fof(f104254,plain,(
  ( ! [X16] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X16))) ) | ($spl1 | $spl20 | $spl130)),
  inference(resolution,[],[f103639,f744])).
fof(f103639,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | ($spl1 | $spl20 | $spl130)),
  inference(backward_demodulation,[],[f103625,f4347])).
fof(f103625,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | ($spl20 | $spl130)),
  inference(forward_demodulation,[],[f43087,f3963])).
fof(f104252,plain,(
  ( ! [X14] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X14))) ) | ($spl1 | $spl20 | $spl130)),
  inference(resolution,[],[f103639,f478])).
fof(f104250,plain,(
  ( ! [X12] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X12))) ) | ($spl1 | $spl3 | $spl20 | $spl130)),
  inference(resolution,[],[f103639,f4631])).
fof(f104249,plain,(
  ( ! [X11] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X11))) ) | ($spl1 | $spl20 | $spl130 | $spl221)),
  inference(resolution,[],[f103639,f58553])).
fof(f104248,plain,(
  ( ! [X10] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | ($spl1 | $spl20 | $spl73 | $spl130)),
  inference(resolution,[],[f103639,f29565])).
fof(f104247,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl1 | $spl20 | $spl130)),
  inference(resolution,[],[f103639,f270])).
fof(f104243,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl1 | $spl20 | $spl130)),
  inference(resolution,[],[f103639,f271])).
fof(f104223,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96011,f339])).
fof(f96011,plain,(
  ( ! [X130,X131] : (sP10(X130,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X131)) ) | ($spl4 | $spl20)),
  inference(backward_demodulation,[],[f3963,f5291])).
fof(f5291,plain,(
  ( ! [X130,X131] : (sP10(X130,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X131)) ) | $spl4),
  inference(resolution,[],[f454,f344])).
fof(f104222,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96011,f338])).
fof(f104221,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96011,f347])).
fof(f104220,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96011,f346])).
fof(f104209,plain,(
  ( ! [X15] : (p(s(bool,X15)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl3 | $spl4 | $spl20)),
  inference(resolution,[],[f96010,f4681])).
fof(f96010,plain,(
  ( ! [X112,X113] : (sP6(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X112,X113)) ) | ($spl4 | $spl20)),
  inference(backward_demodulation,[],[f3963,f5282])).
fof(f104208,plain,(
  ( ! [X14] : (p(s(bool,X14)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20 | $spl221)),
  inference(resolution,[],[f96010,f58603])).
fof(f104207,plain,(
  ( ! [X13] : (p(s(bool,X13)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20 | $spl73)),
  inference(resolution,[],[f96010,f29615])).
fof(f104204,plain,(
  ( ! [X12] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X12))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96010,f3322])).
fof(f104203,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96010,f3311])).
fof(f104202,plain,(
  ( ! [X10] : (p(s(bool,X10)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl3 | $spl4 | $spl20)),
  inference(resolution,[],[f96010,f4680])).
fof(f104201,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20 | $spl221)),
  inference(resolution,[],[f96010,f58602])).
fof(f104200,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20 | $spl73)),
  inference(resolution,[],[f96010,f29614])).
fof(f104214,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X2) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X3)) ) | $spl402),
  inference(cnf_transformation,[],[f104214_D])).
fof(f104214_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X2) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X3)) ) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f104197,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X3) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X2)) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96010,f319])).
fof(f104196,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96010,f330])).
fof(f104185,plain,(
  ( ! [X13] : (p(s(bool,X13)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl3 | $spl4 | $spl20)),
  inference(resolution,[],[f96009,f4682])).
fof(f96009,plain,(
  ( ! [X111,X110] : (sP6(X110,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X111)) ) | ($spl4 | $spl20)),
  inference(backward_demodulation,[],[f3963,f5281])).
fof(f5281,plain,(
  ( ! [X111,X110] : (sP6(X110,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X111)) ) | $spl4),
  inference(resolution,[],[f454,f332])).
fof(f104184,plain,(
  ( ! [X12] : (p(s(bool,X12)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20 | $spl221)),
  inference(resolution,[],[f96009,f58604])).
fof(f104183,plain,(
  ( ! [X11] : (p(s(bool,X11)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20 | $spl73)),
  inference(resolution,[],[f96009,f29616])).
fof(f104182,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X10))) ) | ($spl3 | $spl4 | $spl20)),
  inference(resolution,[],[f96009,f4680])).
fof(f104181,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X9))) ) | ($spl4 | $spl20 | $spl221)),
  inference(resolution,[],[f96009,f58602])).
fof(f104180,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X8))) ) | ($spl4 | $spl20 | $spl73)),
  inference(resolution,[],[f96009,f29614])).
fof(f104193,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X5)) ) | $spl400),
  inference(cnf_transformation,[],[f104193_D])).
fof(f104193_D,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X5)) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f104192,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(subsumption_resolution,[],[f104191,f96006])).
fof(f96006,plain,(
  ( ! [X101,X100] : (sP8(X100,X101,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | ($spl4 | $spl20)),
  inference(backward_demodulation,[],[f3963,f5276])).
fof(f5276,plain,(
  ( ! [X101,X100] : (sP8(X100,X101,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl4),
  inference(resolution,[],[f454,f323])).
fof(f104191,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~sP8(X4,X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | ($spl4 | $spl20)),
  inference(subsumption_resolution,[],[f104178,f329])).
fof(f104178,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~sP8(X4,X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | ~sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X5,X4)) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96009,f320])).
fof(f104176,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96009,f330])).
fof(f104165,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X15))) ) | ($spl3 | $spl4 | $spl20)),
  inference(resolution,[],[f96008,f4682])).
fof(f96008,plain,(
  ( ! [X109,X108] : (sP6(X108,X109,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | ($spl4 | $spl20)),
  inference(backward_demodulation,[],[f3963,f5280])).
fof(f5280,plain,(
  ( ! [X109,X108] : (sP6(X108,X109,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl4),
  inference(resolution,[],[f454,f331])).
fof(f104164,plain,(
  ( ! [X14] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X14))) ) | ($spl4 | $spl20 | $spl221)),
  inference(resolution,[],[f96008,f58604])).
fof(f104163,plain,(
  ( ! [X13] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X13))) ) | ($spl4 | $spl20 | $spl73)),
  inference(resolution,[],[f96008,f29616])).
fof(f104162,plain,(
  ( ! [X12] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X12))) ) | ($spl3 | $spl4 | $spl20)),
  inference(resolution,[],[f96008,f4681])).
fof(f104161,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X11))) ) | ($spl4 | $spl20 | $spl221)),
  inference(resolution,[],[f96008,f58603])).
fof(f104160,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X10))) ) | ($spl4 | $spl20 | $spl73)),
  inference(resolution,[],[f96008,f29615])).
fof(f104159,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96008,f3322])).
fof(f104158,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | s(bool,'F_0') != s(bool,X8)) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96008,f3311])).
fof(f104174,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X6,X7)) ) | $spl398),
  inference(cnf_transformation,[],[f104174_D])).
fof(f104174_D,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X6,X7)) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f104173,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(subsumption_resolution,[],[f104172,f96007])).
fof(f96007,plain,(
  ( ! [X107,X106] : (sP7(X106,X107,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | ($spl4 | $spl20)),
  inference(backward_demodulation,[],[f3963,f5279])).
fof(f5279,plain,(
  ( ! [X107,X106] : (sP7(X106,X107,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl4),
  inference(resolution,[],[f454,f327])).
fof(f104172,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~sP7(X6,X7,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | ($spl4 | $spl20)),
  inference(subsumption_resolution,[],[f104157,f325])).
fof(f104157,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X7,X6) | ~sP7(X6,X7,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96008,f321])).
fof(f104154,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96008,f330])).
fof(f104145,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X11))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96007,f790])).
fof(f104144,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96007,f524])).
fof(f104143,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X9))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96007,f791])).
fof(f104142,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X8))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96007,f525])).
fof(f104152,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X6,X7) | ~sP6(X7,X6,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | $spl396),
  inference(cnf_transformation,[],[f104152_D])).
fof(f104152_D,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X6,X7) | ~sP6(X7,X6,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f104151,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~sP6(X7,X6,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | ($spl4 | $spl20)),
  inference(subsumption_resolution,[],[f104141,f325])).
fof(f104141,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X7,X6) | ~sP6(X7,X6,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96007,f321])).
fof(f104138,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96007,f326])).
fof(f104129,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X11))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96006,f787])).
fof(f104128,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96006,f521])).
fof(f104127,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X9))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96006,f788])).
fof(f104126,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X8))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96006,f522])).
fof(f104136,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X5) | ~sP6(X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X4)) ) | $spl394),
  inference(cnf_transformation,[],[f104136_D])).
fof(f104136_D,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X5) | ~sP6(X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X4)) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f104135,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~sP6(X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X4)) ) | ($spl4 | $spl20)),
  inference(subsumption_resolution,[],[f104124,f329])).
fof(f104124,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X5,X4) | ~sP6(X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X4)) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96006,f320])).
fof(f104122,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96006,f322])).
fof(f104113,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X15))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96005,f774])).
fof(f96005,plain,(
  ( ! [X74,X75] : (sP4(X74,X75,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | ($spl4 | $spl20)),
  inference(backward_demodulation,[],[f3963,f5263])).
fof(f5263,plain,(
  ( ! [X74,X75] : (sP4(X74,X75,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl4),
  inference(resolution,[],[f454,f306])).
fof(f104112,plain,(
  ( ! [X14] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X14))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96005,f508])).
fof(f104110,plain,(
  ( ! [X12] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X12))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96005,f775])).
fof(f104109,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X11))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96005,f509])).
fof(f104119,plain,(
  ( ! [X4,X5] : (sP5(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X4,X5)) ) | $spl392),
  inference(cnf_transformation,[],[f104119_D])).
fof(f104119_D,plain,(
  ( ! [X4,X5] : (sP5(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X4,X5)) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f104107,plain,(
  ( ! [X8,X9] : (sP5(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X8,X9) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96005,f445])).
fof(f104120,plain,(
  ( ! [X6,X7] : (sP5(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(subsumption_resolution,[],[f104106,f309])).
fof(f104106,plain,(
  ( ! [X6,X7] : (sP5(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X7))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96005,f303])).
fof(f104118,plain,(
  ( ! [X4,X5] : (sP5(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X4,X5) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(subsumption_resolution,[],[f104105,f310])).
fof(f104105,plain,(
  ( ! [X4,X5] : (sP5(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X4,X5) | ~p(s(bool,X4)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96005,f302])).
fof(f104103,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96005,f305])).
fof(f104094,plain,(
  ( ! [X16] : (p(s(bool,X16)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96004,f758])).
fof(f96004,plain,(
  ( ! [X43,X42] : (sP2(X42,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X43)) ) | ($spl4 | $spl20)),
  inference(backward_demodulation,[],[f3963,f5247])).
fof(f5247,plain,(
  ( ! [X43,X42] : (sP2(X42,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X43)) ) | $spl4),
  inference(resolution,[],[f454,f290])).
fof(f104092,plain,(
  ( ! [X14] : (p(s(bool,X14)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96004,f492])).
fof(f104090,plain,(
  ( ! [X12] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X12))) ) | ($spl3 | $spl4 | $spl20)),
  inference(resolution,[],[f96004,f4647])).
fof(f104089,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X11))) ) | ($spl4 | $spl20 | $spl221)),
  inference(resolution,[],[f96004,f58569])).
fof(f104088,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | ($spl4 | $spl20 | $spl73)),
  inference(resolution,[],[f96004,f29581])).
fof(f104101,plain,(
  ( ! [X8,X9] : (sP3(X8,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X9) | p(s(bool,X9))) ) | $spl390),
  inference(cnf_transformation,[],[f104101_D])).
fof(f104101_D,plain,(
  ( ! [X8,X9] : (sP3(X8,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X9) | p(s(bool,X9))) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f104087,plain,(
  ( ! [X8,X9] : (sP3(X8,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X9) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X9))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96004,f287])).
fof(f104083,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96004,f288])).
fof(f104074,plain,(
  ( ! [X16] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X16))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96003,f758])).
fof(f96003,plain,(
  ( ! [X41,X40] : (sP2(X40,X41,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | ($spl4 | $spl20)),
  inference(backward_demodulation,[],[f3963,f5246])).
fof(f5246,plain,(
  ( ! [X41,X40] : (sP2(X40,X41,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl4),
  inference(resolution,[],[f454,f289])).
fof(f104072,plain,(
  ( ! [X14] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X14))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96003,f492])).
fof(f104070,plain,(
  ( ! [X12] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X12))) ) | ($spl3 | $spl4 | $spl20)),
  inference(resolution,[],[f96003,f4648])).
fof(f104069,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X11))) ) | ($spl4 | $spl20 | $spl221)),
  inference(resolution,[],[f96003,f58570])).
fof(f104068,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X10))) ) | ($spl4 | $spl20 | $spl73)),
  inference(resolution,[],[f96003,f29582])).
fof(f104081,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9))) ) | $spl388),
  inference(cnf_transformation,[],[f104081_D])).
fof(f104081_D,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9))) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f104067,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X9)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96003,f287])).
fof(f104063,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96003,f288])).
fof(f104053,plain,(
  ( ! [X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96002,f744])).
fof(f96002,plain,(
  ( ! [X14,X15] : (sP0(X14,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X15)) ) | ($spl4 | $spl20)),
  inference(backward_demodulation,[],[f3963,f5233])).
fof(f5233,plain,(
  ( ! [X14,X15] : (sP0(X14,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X15)) ) | $spl4),
  inference(resolution,[],[f454,f273])).
fof(f104051,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96002,f478])).
fof(f104049,plain,(
  ( ! [X12] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X12))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96002,f745])).
fof(f104048,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X11))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96002,f479])).
fof(f104061,plain,(
  ( ! [X10] : (sP1(X10,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) | $spl386),
  inference(cnf_transformation,[],[f104061_D])).
fof(f104061_D,plain,(
  ( ! [X10] : (sP1(X10,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f104047,plain,(
  ( ! [X10] : (sP1(X10,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96002,f477])).
fof(f104060,plain,(
  ( ! [X8,X9] : (sP1(X8,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X9) | ~p(s(bool,X9))) ) | $spl384),
  inference(cnf_transformation,[],[f104060_D])).
fof(f104060_D,plain,(
  ( ! [X8,X9] : (sP1(X8,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X9) | ~p(s(bool,X9))) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f104046,plain,(
  ( ! [X8,X9] : (sP1(X8,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X9) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X9))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96002,f270])).
fof(f104042,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,X1))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96002,f271])).
fof(f104040,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | ($spl4 | $spl20)),
  inference(subsumption_resolution,[],[f104026,f404])).
fof(f104026,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X1))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96013,f364])).
fof(f96013,plain,(
  ( ! [X145] : (sP12(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')),X145)) ) | ($spl4 | $spl20)),
  inference(backward_demodulation,[],[f3963,f5303])).
fof(f5303,plain,(
  ( ! [X145] : (sP12(happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X145)) ) | $spl4),
  inference(resolution,[],[f454,f366])).
fof(f104039,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X0))) ) | ($spl4 | $spl20)),
  inference(subsumption_resolution,[],[f104025,f404])).
fof(f104025,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X0))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96013,f368])).
fof(f104021,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(subsumption_resolution,[],[f104005,f404])).
fof(f104005,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96012,f364])).
fof(f96012,plain,(
  ( ! [X144] : (sP12(X144,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0')))) ) | ($spl4 | $spl20)),
  inference(backward_demodulation,[],[f3963,f5302])).
fof(f5302,plain,(
  ( ! [X144] : (sP12(X144,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl4),
  inference(resolution,[],[f454,f365])).
fof(f104018,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(subsumption_resolution,[],[f104004,f404])).
fof(f104004,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))))) ) | ($spl4 | $spl20)),
  inference(resolution,[],[f96012,f368])).
fof(f103880,plain,(
  $false | ($spl20 | $spl49 | $spl130 | $spl226)),
  inference(subsumption_resolution,[],[f103879,f1264])).
fof(f103879,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl49 | $spl130 | $spl226)),
  inference(forward_demodulation,[],[f103668,f58883])).
fof(f103668,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | ($spl20 | $spl49 | $spl130)),
  inference(backward_demodulation,[],[f103625,f48507])).
fof(f48507,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | $spl49),
  inference(trivial_inequality_removal,[],[f48505])).
fof(f48505,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | $spl49),
  inference(superposition,[],[f5898,f441])).
fof(f103878,plain,(
  $false | ($spl20 | $spl49 | $spl130 | $spl226)),
  inference(subsumption_resolution,[],[f103649,f58883])).
fof(f103649,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | ($spl20 | $spl49 | $spl130)),
  inference(backward_demodulation,[],[f103625,f5898])).
fof(f96833,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | $spl369),
  inference(cnf_transformation,[],[f96833_D])).
fof(f96833_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) <=> ~$spl369),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl369])])).
fof(f96835,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | $spl370),
  inference(cnf_transformation,[],[f96835_D])).
fof(f96835_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f98406,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | ($spl12 | $spl20 | $spl49)),
  inference(subsumption_resolution,[],[f98171,f386])).
fof(f98171,plain,(
  p(s(bool,'F_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | ($spl12 | $spl20 | $spl49)),
  inference(superposition,[],[f437,f96807])).
fof(f96807,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | ($spl12 | $spl20 | $spl49)),
  inference(resolution,[],[f96499,f416])).
fof(f96499,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | ($spl12 | $spl20 | $spl49)),
  inference(subsumption_resolution,[],[f96472,f386])).
fof(f96472,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | ($spl12 | $spl20 | $spl49)),
  inference(superposition,[],[f96000,f48507])).
fof(f96000,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X4)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0'))))) ) | ($spl12 | $spl20)),
  inference(backward_demodulation,[],[f3963,f467])).
fof(f103626,plain,(
  $false | ($spl20 | $spl130 | $spl369)),
  inference(subsumption_resolution,[],[f103625,f97004])).
fof(f97004,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK15) | $spl369),
  inference(resolution,[],[f96833,f439])).
fof(f98666,plain,(
  ( ! [X20] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num','const.num.|48|_0')))) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X20)))) ) | $spl380),
  inference(cnf_transformation,[],[f98666_D])).
fof(f98666_D,plain,(
  ( ! [X20] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num','const.num.|48|_0')))) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X20)))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f98644,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X21)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X20)))) ) | ($spl12 | $spl20)),
  inference(resolution,[],[f96425,f3311])).
fof(f96425,plain,(
  ( ! [X169,X170,X168] : (sP6(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X168)),X169,X170) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X168),s('type.num.num','const.num.|48|_0'))))) ) | ($spl12 | $spl20)),
  inference(resolution,[],[f96000,f333])).
fof(f97265,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | $spl369),
  inference(resolution,[],[f97054,f3428])).
fof(f97264,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) | $spl369),
  inference(resolution,[],[f97054,f3243])).
fof(f96948,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | $spl373),
  inference(cnf_transformation,[],[f96948_D])).
fof(f96948_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) <=> ~$spl373),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl373])])).
fof(f96932,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | ($spl12 | $spl20 | $spl49)),
  inference(resolution,[],[f96769,f3428])).
fof(f96769,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')),X100)) ) | ($spl12 | $spl20 | $spl49)),
  inference(resolution,[],[f96499,f324])).
fof(f96949,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) | $spl374),
  inference(cnf_transformation,[],[f96949_D])).
fof(f96949_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f96931,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) ) | ($spl12 | $spl20 | $spl49)),
  inference(resolution,[],[f96769,f3243])).
fof(f96718,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK15) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | ($spl12 | $spl20 | $spl49)),
  inference(resolution,[],[f96499,f437])).
fof(f96670,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X0))) ) | ($spl12 | $spl20 | $spl49)),
  inference(resolution,[],[f96446,f48629])).
fof(f48629,plain,(
  ( ! [X143] : (~sP12(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X143) | p(s(bool,X143))) ) | $spl49),
  inference(subsumption_resolution,[],[f48589,f404])).
fof(f48589,plain,(
  ( ! [X143] : (p(s(bool,'F_0')) | p(s(bool,X143)) | ~sP12(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X143)) ) | $spl49),
  inference(resolution,[],[f48508,f364])).
fof(f48508,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | $spl49),
  inference(trivial_inequality_removal,[],[f48504])).
fof(f48504,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | $spl49),
  inference(superposition,[],[f5898,f422])).
fof(f96446,plain,(
  ( ! [X222,X223] : (sP12(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X222)),X223) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X222),s('type.num.num','const.num.|48|_0'))))) ) | ($spl12 | $spl20)),
  inference(resolution,[],[f96000,f366])).
fof(f96655,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl20),
  inference(resolution,[],[f96277,f3311])).
fof(f96277,plain,(
  ( ! [X114,X113] : (sP6('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X113,X114)) ) | $spl20),
  inference(resolution,[],[f96001,f333])).
fof(f96001,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl20),
  inference(backward_demodulation,[],[f3963,f587])).
fof(f96568,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.num.|48|_0')))) | p(s(bool,X0))) ) | ($spl12 | $spl20 | $spl49)),
  inference(resolution,[],[f96445,f48628])).
fof(f48628,plain,(
  ( ! [X142] : (~sP12(X142,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | p(s(bool,X142))) ) | $spl49),
  inference(subsumption_resolution,[],[f48588,f404])).
fof(f48588,plain,(
  ( ! [X142] : (p(s(bool,'F_0')) | p(s(bool,X142)) | ~sP12(X142,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) | $spl49),
  inference(resolution,[],[f48508,f364])).
fof(f96445,plain,(
  ( ! [X220,X221] : (sP12(X221,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X220))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X220),s('type.num.num','const.num.|48|_0'))))) ) | ($spl12 | $spl20)),
  inference(resolution,[],[f96000,f365])).
fof(f92998,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl4),
  inference(resolution,[],[f5282,f3311])).
fof(f92992,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X3) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X2)) ) | $spl4),
  inference(resolution,[],[f5282,f319])).
fof(f88100,plain,(
  $false | ($spl33 | $spl35 | $spl273)),
  inference(subsumption_resolution,[],[f80245,f81992])).
fof(f81992,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ($spl33 | $spl35)),
  inference(resolution,[],[f81259,f416])).
fof(f81259,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ($spl33 | $spl35)),
  inference(subsumption_resolution,[],[f81145,f5203])).
fof(f5203,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK16) | $spl35),
  inference(cnf_transformation,[],[f5203_D])).
fof(f5203_D,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK16) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f81145,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK16) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | $spl33),
  inference(resolution,[],[f5613,f437])).
fof(f5613,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | $spl33),
  inference(resolution,[],[f5202,f438])).
fof(f86713,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,X9))) ) | $spl9),
  inference(resolution,[],[f78868,f3428])).
fof(f86712,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl9),
  inference(resolution,[],[f78868,f3243])).
fof(f85341,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(resolution,[],[f587,f437])).
fof(f83356,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | $spl235),
  inference(resolution,[],[f82079,f3428])).
fof(f83355,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) | $spl235),
  inference(resolution,[],[f82079,f3243])).
fof(f83118,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X4)) ) | $spl235),
  inference(resolution,[],[f71284,f320])).
fof(f82726,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) | $spl235),
  inference(resolution,[],[f71281,f321])).
fof(f82709,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | $spl235),
  inference(resolution,[],[f71280,f3428])).
fof(f82708,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) | $spl235),
  inference(resolution,[],[f71280,f3243])).
fof(f82457,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | $spl338),
  inference(cnf_transformation,[],[f82457_D])).
fof(f82457_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f70790,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK16) | $spl236),
  inference(cnf_transformation,[],[f70790_D])).
fof(f70790_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK16) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f82340,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK16) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | $spl234),
  inference(resolution,[],[f70787,f437])).
fof(f82290,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) | $spl235),
  inference(resolution,[],[f71281,f321])).
fof(f82272,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | $spl235),
  inference(resolution,[],[f71280,f3428])).
fof(f82271,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) | $spl235),
  inference(resolution,[],[f71280,f3243])).
fof(f5559,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl41),
  inference(cnf_transformation,[],[f5559_D])).
fof(f5559_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f81811,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X9))) ) | $spl33),
  inference(resolution,[],[f81196,f3428])).
fof(f81196,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X100)) ) | $spl33),
  inference(resolution,[],[f5613,f324])).
fof(f5560,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl42),
  inference(cnf_transformation,[],[f5560_D])).
fof(f5560_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f81810,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl33),
  inference(resolution,[],[f81196,f3243])).
fof(f452,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) | $spl2),
  inference(cnf_transformation,[],[f452_D])).
fof(f452_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f81774,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X5) | ~sP6(X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X4) | ~sP8(X4,X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) | $spl318),
  inference(cnf_transformation,[],[f81774_D])).
fof(f81774_D,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X5) | ~sP6(X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X4) | ~sP8(X4,X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f81763,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X5) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) | ~sP8(X4,X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | ~sP6(X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X4)) ) | $spl3),
  inference(resolution,[],[f4679,f320])).
fof(f4679,plain,(
  ( ! [X107,X106] : (sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X106,X107)) ) | $spl3),
  inference(resolution,[],[f453,f329])).
fof(f81697,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X6,X7) | ~sP6(X7,X6,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | ~sP7(X6,X7,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) | $spl314),
  inference(cnf_transformation,[],[f81697_D])).
fof(f81697_D,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X6,X7) | ~sP6(X7,X6,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | ~sP7(X6,X7,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f81681,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X6,X7) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) | ~sP7(X6,X7,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | ~sP6(X7,X6,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) | $spl3),
  inference(resolution,[],[f4676,f321])).
fof(f4676,plain,(
  ( ! [X101,X100] : (sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X100,X101)) ) | $spl3),
  inference(resolution,[],[f453,f325])).
fof(f6017,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | $spl53),
  inference(cnf_transformation,[],[f6017_D])).
fof(f6017_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f81663,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | ~p(s(bool,X9))) ) | $spl3),
  inference(resolution,[],[f4675,f3428])).
fof(f4675,plain,(
  ( ! [X99,X98] : (sP8(X98,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X99)) ) | $spl3),
  inference(resolution,[],[f453,f324])).
fof(f6018,plain,(
  ( ! [X8] : (sP9(X8,'F_0',happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) | $spl54),
  inference(cnf_transformation,[],[f6018_D])).
fof(f6018_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0',happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f81662,plain,(
  ( ! [X8] : (sP9(X8,'F_0',happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) | $spl3),
  inference(resolution,[],[f4675,f3243])).
fof(f5201,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | $spl32),
  inference(cnf_transformation,[],[f5201_D])).
fof(f5201_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f81569,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X4) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl296),
  inference(cnf_transformation,[],[f81569_D])).
fof(f81569_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X4) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f81558,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X4)) ) | $spl33),
  inference(resolution,[],[f5668,f320])).
fof(f5668,plain,(
  ( ! [X107,X106] : (sP7('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X106,X107)) ) | $spl33),
  inference(resolution,[],[f5202,f329])).
fof(f81538,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl292),
  inference(cnf_transformation,[],[f81538_D])).
fof(f81538_D,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f81522,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl33),
  inference(resolution,[],[f5665,f321])).
fof(f5665,plain,(
  ( ! [X101,X100] : (sP8('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X100,X101)) ) | $spl33),
  inference(resolution,[],[f5202,f325])).
fof(f5826,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl45),
  inference(cnf_transformation,[],[f5826_D])).
fof(f5826_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f81504,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X9))) ) | $spl33),
  inference(resolution,[],[f5664,f3428])).
fof(f5664,plain,(
  ( ! [X99,X98] : (sP8(X98,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X99)) ) | $spl33),
  inference(resolution,[],[f5202,f324])).
fof(f5827,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl46),
  inference(cnf_transformation,[],[f5827_D])).
fof(f5827_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f81503,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl33),
  inference(resolution,[],[f5664,f3243])).
fof(f78601,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X5),s('type.num.num',X6)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))) ) | $spl12),
  inference(resolution,[],[f467,f259])).
fof(f81341,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X5),s('type.num.num',X6)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))))) ) | $spl12),
  inference(resolution,[],[f467,f259])).
fof(f78938,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl251),
  inference(cnf_transformation,[],[f78938_D])).
fof(f78938_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl251),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl251])])).
fof(f78940,plain,(
  s('type.num.num',sK16) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl252),
  inference(cnf_transformation,[],[f78940_D])).
fof(f78940_D,plain,(
  s('type.num.num',sK16) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f81016,plain,(
  s('type.num.num',sK16) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl9),
  inference(resolution,[],[f461,f437])).
fof(f81011,plain,(
  $false | ($spl33 | $spl34)),
  inference(subsumption_resolution,[],[f5614,f5204])).
fof(f5614,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK16) | $spl33),
  inference(resolution,[],[f5202,f439])).
fof(f80996,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl32 | $spl34)),
  inference(subsumption_resolution,[],[f80988,f404])).
fof(f80988,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl32 | $spl34)),
  inference(resolution,[],[f80826,f364])).
fof(f80826,plain,(
  ( ! [X145] : (sP12(X145,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl32 | $spl34)),
  inference(backward_demodulation,[],[f5204,f79210])).
fof(f79210,plain,(
  ( ! [X145] : (sP12(X145,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl32),
  inference(resolution,[],[f5201,f365])).
fof(f80995,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl32 | $spl34)),
  inference(subsumption_resolution,[],[f80987,f404])).
fof(f80987,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl32 | $spl34)),
  inference(resolution,[],[f80826,f368])).
fof(f80648,plain,(
  $false | ($spl3 | $spl4 | $spl252)),
  inference(subsumption_resolution,[],[f80499,f4715])).
fof(f4715,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | $spl3),
  inference(resolution,[],[f453,f419])).
fof(f419,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f244])).
fof(f80499,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | ($spl4 | $spl252)),
  inference(backward_demodulation,[],[f78940,f5315])).
fof(f80416,plain,(
  $false | ($spl251 | $spl260)),
  inference(subsumption_resolution,[],[f80415,f78938])).
fof(f80415,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl260),
  inference(forward_demodulation,[],[f80300,f262])).
fof(f80300,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl260),
  inference(resolution,[],[f79249,f260])).
fof(f80228,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl32),
  inference(resolution,[],[f79190,f3311])).
fof(f79713,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,X9))) ) | $spl9),
  inference(resolution,[],[f78868,f3428])).
fof(f79712,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl9),
  inference(resolution,[],[f78868,f3243])).
fof(f79498,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26000,f339])).
fof(f79497,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26000,f338])).
fof(f79496,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26000,f347])).
fof(f79495,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26000,f346])).
fof(f79488,plain,(
  ( ! [X15] : (p(s(bool,X15)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25991,f4681])).
fof(f79487,plain,(
  ( ! [X14] : (p(s(bool,X14)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25991,f58603])).
fof(f79486,plain,(
  ( ! [X13] : (p(s(bool,X13)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25991,f29615])).
fof(f79481,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X12))) ) | $spl6),
  inference(resolution,[],[f25991,f3322])).
fof(f79480,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f25991,f3311])).
fof(f79479,plain,(
  ( ! [X10] : (p(s(bool,X10)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25991,f4680])).
fof(f79478,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25991,f58602])).
fof(f79477,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25991,f29614])).
fof(f79474,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X2,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X3) | ~sP7(X3,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X2)) ) | $spl6),
  inference(resolution,[],[f25991,f319])).
fof(f79473,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f25991,f330])).
fof(f79464,plain,(
  ( ! [X13] : (p(s(bool,X13)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25990,f4682])).
fof(f79463,plain,(
  ( ! [X12] : (p(s(bool,X12)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25990,f58604])).
fof(f79462,plain,(
  ( ! [X11] : (p(s(bool,X11)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25990,f29616])).
fof(f79461,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X10))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25990,f4680])).
fof(f79460,plain,(
  ( ! [X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X9))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25990,f58602])).
fof(f79459,plain,(
  ( ! [X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X8))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25990,f29614])).
fof(f79471,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f79470,f25985])).
fof(f79470,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(subsumption_resolution,[],[f79457,f329])).
fof(f79457,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5,X4)) ) | $spl6),
  inference(resolution,[],[f25990,f320])).
fof(f79455,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f25990,f330])).
fof(f79446,plain,(
  ( ! [X15] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25989,f4682])).
fof(f79445,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X14))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25989,f58604])).
fof(f79444,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X13))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25989,f29616])).
fof(f79443,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X12))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25989,f4681])).
fof(f79442,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X11))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25989,f58603])).
fof(f79441,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X10))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25989,f29615])).
fof(f79440,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25989,f3322])).
fof(f79439,plain,(
  ( ! [X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | s(bool,'F_0') != s(bool,X8)) ) | $spl6),
  inference(resolution,[],[f25989,f3311])).
fof(f79454,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f79453,f25988])).
fof(f79453,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(subsumption_resolution,[],[f79438,f325])).
fof(f79438,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X7,X6) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f25989,f321])).
fof(f79435,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25989,f330])).
fof(f79428,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | $spl6),
  inference(resolution,[],[f25988,f790])).
fof(f79427,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | $spl6),
  inference(resolution,[],[f25988,f524])).
fof(f79426,plain,(
  ( ! [X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X9))) ) | $spl6),
  inference(resolution,[],[f25988,f791])).
fof(f79425,plain,(
  ( ! [X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X8))) ) | $spl6),
  inference(resolution,[],[f25988,f525])).
fof(f79434,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(subsumption_resolution,[],[f79424,f325])).
fof(f79424,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X7,X6) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f25988,f321])).
fof(f79421,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25988,f326])).
fof(f79391,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | $spl6),
  inference(resolution,[],[f25985,f787])).
fof(f79390,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | $spl6),
  inference(resolution,[],[f25985,f521])).
fof(f79389,plain,(
  ( ! [X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X9))) ) | $spl6),
  inference(resolution,[],[f25985,f788])).
fof(f79388,plain,(
  ( ! [X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X8))) ) | $spl6),
  inference(resolution,[],[f25985,f522])).
fof(f79397,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4)) ) | $spl6),
  inference(subsumption_resolution,[],[f79386,f329])).
fof(f79386,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5,X4) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4)) ) | $spl6),
  inference(resolution,[],[f25985,f320])).
fof(f79384,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25985,f322])).
fof(f79377,plain,(
  ( ! [X15] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X15))) ) | $spl6),
  inference(resolution,[],[f25972,f774])).
fof(f79376,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | $spl6),
  inference(resolution,[],[f25972,f508])).
fof(f79374,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | $spl6),
  inference(resolution,[],[f25972,f775])).
fof(f79373,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | $spl6),
  inference(resolution,[],[f25972,f509])).
fof(f79371,plain,(
  ( ! [X8,X9] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X8,X9) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25972,f445])).
fof(f79383,plain,(
  ( ! [X6,X7] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f79370,f309])).
fof(f79370,plain,(
  ( ! [X6,X7] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X7))) ) | $spl6),
  inference(resolution,[],[f25972,f303])).
fof(f79382,plain,(
  ( ! [X4,X5] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4,X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f79369,f310])).
fof(f79369,plain,(
  ( ! [X4,X5] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25972,f302])).
fof(f79367,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25972,f305])).
fof(f79359,plain,(
  ( ! [X17] : (p(s(bool,X17)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25956,f758])).
fof(f79357,plain,(
  ( ! [X15] : (p(s(bool,X15)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25956,f492])).
fof(f79355,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25956,f4647])).
fof(f79366,plain,(
  ( ! [X12] : (sP3(X12,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) | $spl266),
  inference(cnf_transformation,[],[f79366_D])).
fof(f79366_D,plain,(
  ( ! [X12] : (sP3(X12,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f79354,plain,(
  ( ! [X12] : (sP3(X12,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25956,f4645])).
fof(f4645,plain,(
  ( ! [X39,X38] : (~sP2(X38,X39,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | sP3(X38,X39,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | p(s(bool,X39))) ) | $spl3),
  inference(resolution,[],[f453,f287])).
fof(f79353,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25956,f58569])).
fof(f79352,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25956,f29581])).
fof(f79351,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X9) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X9))) ) | $spl6),
  inference(resolution,[],[f25956,f287])).
fof(f79347,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f25956,f288])).
fof(f79339,plain,(
  ( ! [X17] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X17))) ) | $spl6),
  inference(resolution,[],[f25955,f758])).
fof(f79337,plain,(
  ( ! [X15] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | $spl6),
  inference(resolution,[],[f25955,f492])).
fof(f79335,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25955,f4648])).
fof(f79346,plain,(
  ( ! [X12] : (sP3(X12,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl264),
  inference(cnf_transformation,[],[f79346_D])).
fof(f79346_D,plain,(
  ( ! [X12] : (sP3(X12,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f79334,plain,(
  ( ! [X12] : (sP3(X12,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25955,f4646])).
fof(f4646,plain,(
  ( ! [X41,X40] : (~sP2(X40,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X41) | sP3(X40,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X41) | p(s(bool,X41))) ) | $spl3),
  inference(resolution,[],[f453,f287])).
fof(f79333,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | ($spl6 | $spl221)),
  inference(resolution,[],[f25955,f58570])).
fof(f79332,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25955,f29582])).
fof(f79331,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25955,f287])).
fof(f79327,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25955,f288])).
fof(f79319,plain,(
  ( ! [X16] : (~p(s(bool,X16)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25942,f744])).
fof(f79317,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25942,f478])).
fof(f79315,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | $spl6),
  inference(resolution,[],[f25942,f745])).
fof(f79314,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | $spl6),
  inference(resolution,[],[f25942,f479])).
fof(f79313,plain,(
  ( ! [X10] : (sP1(X10,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25942,f477])).
fof(f79312,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X9) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X9))) ) | $spl6),
  inference(resolution,[],[f25942,f270])).
fof(f79308,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f25942,f271])).
fof(f5204,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK16) | $spl34),
  inference(cnf_transformation,[],[f5204_D])).
fof(f5204_D,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK16) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f79132,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK16) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | $spl32),
  inference(resolution,[],[f5201,f437])).
fof(f79099,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f4259,f3311])).
fof(f79093,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X3) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X2)) ) | $spl6),
  inference(resolution,[],[f4259,f319])).
fof(f78817,plain,(
  s('type.num.num',sK16) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl9),
  inference(resolution,[],[f461,f437])).
fof(f78816,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f78804,f404])).
fof(f78804,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f26012,f364])).
fof(f78815,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f78803,f404])).
fof(f78803,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl6),
  inference(resolution,[],[f26012,f368])).
fof(f78802,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f78790,f404])).
fof(f78790,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26011,f364])).
fof(f78801,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f78789,f404])).
fof(f78789,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26011,f368])).
fof(f78740,plain,(
  $false | ($spl3 | $spl8 | $spl12)),
  inference(subsumption_resolution,[],[f78739,f429])).
fof(f78739,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl8 | $spl12)),
  inference(forward_demodulation,[],[f78738,f4713])).
fof(f4713,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | $spl3),
  inference(resolution,[],[f453,f416])).
fof(f78738,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f78708,f388])).
fof(f78708,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) | ($spl8 | $spl12)),
  inference(superposition,[],[f467,f4912])).
fof(f4912,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl8),
  inference(resolution,[],[f462,f420])).
fof(f78730,plain,(
  $false | ($spl12 | $spl49 | $spl128)),
  inference(subsumption_resolution,[],[f78729,f429])).
fof(f78729,plain,(
  p(s(bool,'F_0')) | ($spl12 | $spl49 | $spl128)),
  inference(forward_demodulation,[],[f78602,f48507])).
fof(f78602,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | ($spl12 | $spl128)),
  inference(resolution,[],[f467,f43085])).
fof(f78557,plain,(
  $false | ($spl8 | $spl33)),
  inference(subsumption_resolution,[],[f78556,f429])).
fof(f78556,plain,(
  p(s(bool,'F_0')) | ($spl8 | $spl33)),
  inference(forward_demodulation,[],[f78439,f5702])).
fof(f5702,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | $spl33),
  inference(resolution,[],[f5202,f416])).
fof(f78439,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | $spl8),
  inference(resolution,[],[f78197,f438])).
fof(f78323,plain,(
  $false | ($spl3 | $spl8 | $spl10)),
  inference(subsumption_resolution,[],[f78322,f429])).
fof(f78322,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl8 | $spl10)),
  inference(forward_demodulation,[],[f78197,f5176])).
fof(f5176,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ($spl3 | $spl10)),
  inference(resolution,[],[f4625,f416])).
fof(f4625,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ($spl3 | $spl10)),
  inference(resolution,[],[f453,f466])).
fof(f77488,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl246),
  inference(cnf_transformation,[],[f77488_D])).
fof(f77488_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f77490,plain,(
  s('type.num.num',sK16) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl248),
  inference(cnf_transformation,[],[f77490_D])).
fof(f77490_D,plain,(
  s('type.num.num',sK16) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f77363,plain,(
  s('type.num.num',sK16) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl8),
  inference(resolution,[],[f4943,f437])).
fof(f4943,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl8),
  inference(forward_demodulation,[],[f4822,f262])).
fof(f4822,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl8),
  inference(resolution,[],[f462,f260])).
fof(f71630,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | $spl235),
  inference(resolution,[],[f71280,f3428])).
fof(f71629,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) | $spl235),
  inference(resolution,[],[f71280,f3243])).
fof(f71142,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f70724,f3428])).
fof(f70724,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')),X100)) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f70669,f324])).
fof(f70669,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f70642,f392])).
fof(f392,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f226])).
fof(f226,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X13] : ((p(s(bool,'F_0')) => p(s(bool,X13))) <=> p(s(bool,'T_0')))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.bool.IMP_CLAUSES|split|2')).
fof(f70642,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | ($spl3 | $spl10)),
  inference(superposition,[],[f5088,f828])).
fof(f5088,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK14)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',X0))))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f4625,f436])).
fof(f71141,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f70724,f3243])).
fof(f70788,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | $spl235),
  inference(cnf_transformation,[],[f70788_D])).
fof(f70788_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f70673,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK16) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.num.|48|_0')))) | ($spl3 | $spl10)),
  inference(resolution,[],[f70669,f437])).
fof(f64592,plain,(
  ( ! [X19,X20] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X19),s('type.num.num',X20))),s('type.num.num','const.num.|48|_0')))) ) | $spl232),
  inference(cnf_transformation,[],[f64592_D])).
fof(f64592_D,plain,(
  ( ! [X19,X20] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X19),s('type.num.num',X20))),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f64564,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X19),s('type.num.num',X20))),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X21))) ) | $spl221),
  inference(resolution,[],[f64038,f3428])).
fof(f64038,plain,(
  ( ! [X206,X204,X205,X203] : (sP8(X203,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X204),s('type.num.num',X205))),s('type.num.num','const.num.|48|_0')),X206)) ) | $spl221),
  inference(resolution,[],[f58545,f324])).
fof(f58545,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.num.|48|_0'))))) ) | $spl221),
  inference(resolution,[],[f58528,f258])).
fof(f58919,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | $spl229),
  inference(cnf_transformation,[],[f58919_D])).
fof(f58919_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) <=> ~$spl229),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl229])])).
fof(f58902,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,X9))) ) | $spl221),
  inference(resolution,[],[f58597,f3428])).
fof(f58597,plain,(
  ( ! [X101,X102] : (sP8(X101,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X102)) ) | $spl221),
  inference(resolution,[],[f58528,f324])).
fof(f58920,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl230),
  inference(cnf_transformation,[],[f58920_D])).
fof(f58920_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f58901,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl221),
  inference(resolution,[],[f58597,f3243])).
fof(f58881,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',X4))),s('type.num.num',sK14))))) ) | $spl224),
  inference(cnf_transformation,[],[f58881_D])).
fof(f58881_D,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',X4))),s('type.num.num',sK14))))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f58883,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | $spl226),
  inference(cnf_transformation,[],[f58883_D])).
fof(f58883_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f58760,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num',X4))),s('type.num.num',sK14))))) ) | $spl10),
  inference(resolution,[],[f949,f258])).
fof(f949,plain,(
  ( ! [X239] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X239),s('type.num.num',sK14)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X239)))) ) | $spl10),
  inference(resolution,[],[f466,f420])).
fof(f58528,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl221),
  inference(cnf_transformation,[],[f58528_D])).
fof(f58528_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) <=> ~$spl221),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl221])])).
fof(f58529,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X5),s('type.num.num',X6))),s('type.num.num',sK14)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0'))))))) ) | $spl222),
  inference(cnf_transformation,[],[f58529_D])).
fof(f58529_D,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X5),s('type.num.num',X6))),s('type.num.num',sK14)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f58409,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X5),s('type.num.num',X6))),s('type.num.num',sK14)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(resolution,[],[f738,f259])).
fof(f58053,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl219),
  inference(cnf_transformation,[],[f58053_D])).
fof(f58053_D,plain,(
  s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl219),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl219])])).
fof(f58025,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl128),
  inference(resolution,[],[f57224,f3311])).
fof(f57224,plain,(
  ( ! [X114,X113] : (sP6('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X113,X114)) ) | $spl128),
  inference(resolution,[],[f56948,f333])).
fof(f56948,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl128),
  inference(resolution,[],[f43085,f438])).
fof(f57707,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl217),
  inference(cnf_transformation,[],[f57707_D])).
fof(f57707_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl217),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl217])])).
fof(f57679,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl128),
  inference(resolution,[],[f57007,f3311])).
fof(f57007,plain,(
  ( ! [X114,X113] : (sP6('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),X113,X114)) ) | $spl128),
  inference(resolution,[],[f43085,f333])).
fof(f49823,plain,(
  $false | ($spl49 | $spl71)),
  inference(subsumption_resolution,[],[f26564,f48507])).
fof(f48511,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))) ) | $spl49),
  inference(subsumption_resolution,[],[f48500,f420])).
fof(f48500,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))) ) | $spl49),
  inference(superposition,[],[f5898,f380])).
fof(f48510,plain,(
  ( ! [X10] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | ~p(s(bool,X10))) ) | $spl49),
  inference(subsumption_resolution,[],[f48499,f420])).
fof(f48499,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | ~p(s(bool,X10))) ) | $spl49),
  inference(superposition,[],[f5898,f380])).
fof(f48202,plain,(
  ( ! [X22] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num',sK14)))) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X22)))) ) | $spl214),
  inference(cnf_transformation,[],[f48202_D])).
fof(f48202_D,plain,(
  ( ! [X22] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num',sK14)))) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X22)))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f48179,plain,(
  ( ! [X23,X22] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num',sK14)))) | p(s(bool,X23)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X22)))) ) | $spl10),
  inference(resolution,[],[f916,f3311])).
fof(f916,plain,(
  ( ! [X169,X170,X168] : (sP6(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X168)),X169,X170) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X168),s('type.num.num',sK14))))) ) | $spl10),
  inference(resolution,[],[f466,f333])).
fof(f46699,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl4),
  inference(resolution,[],[f5282,f3311])).
fof(f46315,plain,(
  $false | ($spl23 | $spl25 | $spl69)),
  inference(subsumption_resolution,[],[f26389,f44696])).
fof(f44696,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ($spl23 | $spl25)),
  inference(resolution,[],[f44381,f416])).
fof(f44381,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ($spl23 | $spl25)),
  inference(subsumption_resolution,[],[f44267,f4579])).
fof(f44267,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | $spl23),
  inference(resolution,[],[f4962,f437])).
fof(f4962,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | $spl23),
  inference(resolution,[],[f4578,f438])).
fof(f46314,plain,(
  $false | ($spl23 | $spl67)),
  inference(subsumption_resolution,[],[f26220,f44356])).
fof(f44356,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | $spl23),
  inference(resolution,[],[f4962,f416])).
fof(f45051,plain,(
  ( ! [X4,X5] : (sP9(X4,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X5) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | ~sP8(X4,X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | ~sP6(X5,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X4)) ) | $spl1),
  inference(resolution,[],[f4394,f320])).
fof(f45011,plain,(
  ( ! [X6,X7] : (sP9(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X6,X7) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | ~sP7(X6,X7,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | ~sP6(X7,X6,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) | $spl1),
  inference(resolution,[],[f4391,f321])).
fof(f44991,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | ~p(s(bool,X9))) ) | $spl1),
  inference(resolution,[],[f4390,f3428])).
fof(f44990,plain,(
  ( ! [X8] : (sP9(X8,'F_0',happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) | $spl1),
  inference(resolution,[],[f4390,f3243])).
fof(f26564,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | $spl71),
  inference(cnf_transformation,[],[f26564_D])).
fof(f26564_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f44868,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) | $spl0),
  inference(resolution,[],[f25869,f3311])).
fof(f25869,plain,(
  ( ! [X112,X113] : (sP6(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X112,X113)) ) | $spl0),
  inference(resolution,[],[f450,f333])).
fof(f44882,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X2) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X3)) ) | $spl192),
  inference(cnf_transformation,[],[f44882_D])).
fof(f44882_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X2) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X3)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f44861,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | ~sP8(X2,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X3) | ~sP7(X3,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X2)) ) | $spl0),
  inference(resolution,[],[f25869,f319])).
fof(f44588,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl174),
  inference(cnf_transformation,[],[f44588_D])).
fof(f44588_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f44575,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4)) ) | $spl23),
  inference(resolution,[],[f5017,f320])).
fof(f5017,plain,(
  ( ! [X107,X106] : (sP7('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X106,X107)) ) | $spl23),
  inference(resolution,[],[f4578,f329])).
fof(f44553,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl170),
  inference(cnf_transformation,[],[f44553_D])).
fof(f44553_D,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f44535,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl23),
  inference(resolution,[],[f5014,f321])).
fof(f5014,plain,(
  ( ! [X101,X100] : (sP8('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X100,X101)) ) | $spl23),
  inference(resolution,[],[f4578,f325])).
fof(f5416,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | $spl37),
  inference(cnf_transformation,[],[f5416_D])).
fof(f5416_D,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f44515,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,X9))) ) | $spl23),
  inference(resolution,[],[f5013,f3428])).
fof(f5013,plain,(
  ( ! [X99,X98] : (sP8(X98,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X99)) ) | $spl23),
  inference(resolution,[],[f4578,f324])).
fof(f5417,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl38),
  inference(cnf_transformation,[],[f5417_D])).
fof(f5417_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f44514,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl23),
  inference(resolution,[],[f5013,f3243])).
fof(f458,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | $spl6),
  inference(cnf_transformation,[],[f458_D])).
fof(f458_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f43282,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | ~sP6(X5,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4) | ~sP8(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl134),
  inference(cnf_transformation,[],[f43282_D])).
fof(f43282_D,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | ~sP6(X5,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4) | ~sP8(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f44253,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP6(X5,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4)) ) | $spl7),
  inference(resolution,[],[f4518,f320])).
fof(f4518,plain,(
  ( ! [X107,X108] : (sP7('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X107,X108)) ) | $spl7),
  inference(resolution,[],[f457,f329])).
fof(f29041,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl126),
  inference(cnf_transformation,[],[f29041_D])).
fof(f29041_D,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f44215,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f4515,f321])).
fof(f4515,plain,(
  ( ! [X101,X102] : (sP8('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X101,X102)) ) | $spl7),
  inference(resolution,[],[f457,f325])).
fof(f4788,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | $spl27),
  inference(cnf_transformation,[],[f4788_D])).
fof(f4788_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f44196,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,X9))) ) | $spl7),
  inference(resolution,[],[f4514,f3428])).
fof(f4514,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X100)) ) | $spl7),
  inference(resolution,[],[f457,f324])).
fof(f4789,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl28),
  inference(cnf_transformation,[],[f4789_D])).
fof(f4789_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f44195,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f4514,f3243])).
fof(f44098,plain,(
  $false | ($spl23 | $spl24)),
  inference(subsumption_resolution,[],[f4963,f4580])).
fof(f4963,plain,(
  s('type.num.num',sK14) != s('type.num.num',sK15) | $spl23),
  inference(resolution,[],[f4578,f439])).
fof(f44069,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | $spl157),
  inference(cnf_transformation,[],[f44069_D])).
fof(f44069_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f44095,plain,(
  ( ! [X22] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X22))) ) | ($spl3 | $spl10 | $spl23 | $spl24)),
  inference(forward_demodulation,[],[f44072,f4580])).
fof(f44072,plain,(
  ( ! [X22] : (~p(s(bool,X22)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl10 | $spl23)),
  inference(duplicate_literal_removal,[],[f16653])).
fof(f16653,plain,(
  ( ! [X22] : (~p(s(bool,X22)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X22))) ) | ($spl3 | $spl10 | $spl23)),
  inference(resolution,[],[f5100,f4974])).
fof(f4974,plain,(
  ( ! [X21,X20] : (sP1(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,X20))) ) | $spl23),
  inference(resolution,[],[f4578,f276])).
fof(f5100,plain,(
  ( ! [X24,X23] : (~sP1(X24,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X23) | ~p(s(bool,X24)) | ~p(s(bool,X23))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f4625,f278])).
fof(f44094,plain,(
  ( ! [X18] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X18))) ) | ($spl3 | $spl10 | $spl23 | $spl24)),
  inference(forward_demodulation,[],[f44073,f4580])).
fof(f44073,plain,(
  ( ! [X18] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X18))) ) | ($spl3 | $spl10 | $spl23)),
  inference(duplicate_literal_removal,[],[f16581])).
fof(f16581,plain,(
  ( ! [X18] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X18)) | p(s(bool,X18))) ) | ($spl3 | $spl10 | $spl23)),
  inference(resolution,[],[f5098,f4964])).
fof(f4964,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | p(s(bool,X0))) ) | $spl23),
  inference(resolution,[],[f4578,f264])).
fof(f264,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f172])).
fof(f5098,plain,(
  ( ! [X19,X20] : (sP1(X19,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X20) | ~p(s(bool,X20)) | p(s(bool,X19))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f4625,f275])).
fof(f275,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f177])).
fof(f44081,plain,(
  ( ! [X22] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X22))) ) | ($spl3 | $spl10 | $spl23 | $spl24)),
  inference(forward_demodulation,[],[f28783,f4580])).
fof(f28783,plain,(
  ( ! [X22] : (~p(s(bool,X22)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl10 | $spl23)),
  inference(duplicate_literal_removal,[],[f16653])).
fof(f44080,plain,(
  ( ! [X18] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X18))) ) | ($spl3 | $spl10 | $spl23 | $spl24)),
  inference(forward_demodulation,[],[f28784,f4580])).
fof(f28784,plain,(
  ( ! [X18] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X18))) ) | ($spl3 | $spl10 | $spl23)),
  inference(duplicate_literal_removal,[],[f16581])).
fof(f44079,plain,(
  ( ! [X22] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X22))) ) | ($spl3 | $spl10 | $spl23 | $spl24)),
  inference(forward_demodulation,[],[f43099,f4580])).
fof(f43099,plain,(
  ( ! [X22] : (~p(s(bool,X22)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl10 | $spl23)),
  inference(duplicate_literal_removal,[],[f16653])).
fof(f44078,plain,(
  ( ! [X18] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X18))) ) | ($spl3 | $spl10 | $spl23 | $spl24)),
  inference(forward_demodulation,[],[f43100,f4580])).
fof(f43100,plain,(
  ( ! [X18] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X18))) ) | ($spl3 | $spl10 | $spl23)),
  inference(duplicate_literal_removal,[],[f16581])).
fof(f44071,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl22 | $spl24)),
  inference(subsumption_resolution,[],[f44061,f404])).
fof(f44061,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f44057,f364])).
fof(f44057,plain,(
  ( ! [X145] : (sP12(X145,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl22 | $spl24)),
  inference(forward_demodulation,[],[f29336,f4580])).
fof(f29336,plain,(
  ( ! [X145] : (sP12(X145,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl22),
  inference(resolution,[],[f4577,f365])).
fof(f44068,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl22 | $spl24)),
  inference(subsumption_resolution,[],[f44060,f404])).
fof(f44060,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f44057,f368])).
fof(f28586,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ($spl10 | $spl24)),
  inference(superposition,[],[f466,f4580])).
fof(f43731,plain,(
  ( ! [X17] : (p(s(bool,X17)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl33)),
  inference(resolution,[],[f25991,f5670])).
fof(f5670,plain,(
  ( ! [X111,X110] : (~sP6(X110,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X111) | p(s(bool,X111)) | p(s(bool,X110))) ) | $spl33),
  inference(resolution,[],[f5202,f330])).
fof(f43730,plain,(
  ( ! [X16] : (p(s(bool,X16)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25991,f4681])).
fof(f43729,plain,(
  ( ! [X15] : (p(s(bool,X15)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25991,f29615])).
fof(f43728,plain,(
  ( ! [X14] : (p(s(bool,X14)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25991,f5144])).
fof(f5144,plain,(
  ( ! [X111,X112] : (~sP6(X111,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X112) | p(s(bool,X112)) | p(s(bool,X111))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f4625,f330])).
fof(f43723,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X13))) ) | $spl6),
  inference(resolution,[],[f25991,f3322])).
fof(f43722,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f25991,f3311])).
fof(f43721,plain,(
  ( ! [X11] : (p(s(bool,X11)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl33)),
  inference(resolution,[],[f25991,f5669])).
fof(f5669,plain,(
  ( ! [X109,X108] : (~sP6(X108,X109,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X109)) | p(s(bool,X108))) ) | $spl33),
  inference(resolution,[],[f5202,f330])).
fof(f43720,plain,(
  ( ! [X10] : (p(s(bool,X10)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25991,f4680])).
fof(f43719,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25991,f29614])).
fof(f43718,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25991,f5143])).
fof(f5143,plain,(
  ( ! [X109,X110] : (~sP6(X109,X110,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X110)) | p(s(bool,X109))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f4625,f330])).
fof(f43715,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X2,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X3) | ~sP7(X3,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X2)) ) | $spl6),
  inference(resolution,[],[f25991,f319])).
fof(f43714,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f25991,f330])).
fof(f43704,plain,(
  ( ! [X15] : (p(s(bool,X15)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl33)),
  inference(resolution,[],[f25990,f5671])).
fof(f5671,plain,(
  ( ! [X112,X113] : (~sP6('const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X112,X113) | p(s(bool,X113)) | p(s(bool,X112))) ) | $spl33),
  inference(resolution,[],[f5202,f330])).
fof(f43703,plain,(
  ( ! [X14] : (p(s(bool,X14)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25990,f4682])).
fof(f43702,plain,(
  ( ! [X13] : (p(s(bool,X13)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25990,f29616])).
fof(f43701,plain,(
  ( ! [X12] : (p(s(bool,X12)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25990,f5145])).
fof(f5145,plain,(
  ( ! [X114,X113] : (~sP6('const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X113,X114) | p(s(bool,X114)) | p(s(bool,X113))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f4625,f330])).
fof(f43700,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X11))) ) | ($spl6 | $spl33)),
  inference(resolution,[],[f25990,f5669])).
fof(f43699,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X10))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25990,f4680])).
fof(f43698,plain,(
  ( ! [X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X9))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25990,f29614])).
fof(f43697,plain,(
  ( ! [X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X8))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25990,f5143])).
fof(f43711,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f43710,f25985])).
fof(f43710,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(subsumption_resolution,[],[f43695,f329])).
fof(f43695,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X4,X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5,X4)) ) | $spl6),
  inference(resolution,[],[f25990,f320])).
fof(f43693,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f25990,f330])).
fof(f43683,plain,(
  ( ! [X17] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X17))) ) | ($spl6 | $spl33)),
  inference(resolution,[],[f25989,f5671])).
fof(f43682,plain,(
  ( ! [X16] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X16))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25989,f4682])).
fof(f43681,plain,(
  ( ! [X15] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25989,f29616])).
fof(f43680,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X14))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25989,f5145])).
fof(f43679,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X13))) ) | ($spl6 | $spl33)),
  inference(resolution,[],[f25989,f5670])).
fof(f43678,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X12))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25989,f4681])).
fof(f43677,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X11))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25989,f29615])).
fof(f43676,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X10))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25989,f5144])).
fof(f43675,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25989,f3322])).
fof(f43674,plain,(
  ( ! [X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | s(bool,'F_0') != s(bool,X8)) ) | $spl6),
  inference(resolution,[],[f25989,f3311])).
fof(f43691,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f43690,f25988])).
fof(f43690,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(subsumption_resolution,[],[f43673,f325])).
fof(f43673,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X7,X6) | ~sP7(X6,X7,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f25989,f321])).
fof(f43670,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25989,f330])).
fof(f43662,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | $spl6),
  inference(resolution,[],[f25988,f790])).
fof(f43661,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f25988,f25864])).
fof(f25864,plain,(
  ( ! [X103,X102] : (~sP7(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X102,X103) | p(s(bool,X103)) | ~p(s(bool,X102))) ) | $spl0),
  inference(resolution,[],[f450,f326])).
fof(f43660,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | $spl6),
  inference(resolution,[],[f25988,f524])).
fof(f43659,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | $spl6),
  inference(resolution,[],[f25988,f791])).
fof(f43658,plain,(
  ( ! [X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X9))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f25988,f25865])).
fof(f25865,plain,(
  ( ! [X105,X104] : (~sP7(X104,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X105) | p(s(bool,X105)) | ~p(s(bool,X104))) ) | $spl0),
  inference(resolution,[],[f450,f326])).
fof(f43657,plain,(
  ( ! [X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X8))) ) | $spl6),
  inference(resolution,[],[f25988,f525])).
fof(f43668,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(subsumption_resolution,[],[f43656,f325])).
fof(f43656,plain,(
  ( ! [X6,X7] : (sP9('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X7,X6) | ~sP6(X7,X6,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f25988,f321])).
fof(f43653,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25988,f326])).
fof(f43645,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | $spl6),
  inference(resolution,[],[f25985,f787])).
fof(f43644,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f25985,f25861])).
fof(f25861,plain,(
  ( ! [X97,X96] : (~sP8(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X96,X97) | p(s(bool,X97)) | ~p(s(bool,X96))) ) | $spl0),
  inference(resolution,[],[f450,f322])).
fof(f43643,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | $spl6),
  inference(resolution,[],[f25985,f521])).
fof(f43642,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | $spl6),
  inference(resolution,[],[f25985,f788])).
fof(f43641,plain,(
  ( ! [X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X9))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f25985,f25862])).
fof(f25862,plain,(
  ( ! [X99,X98] : (~sP8(X98,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X99) | p(s(bool,X99)) | ~p(s(bool,X98))) ) | $spl0),
  inference(resolution,[],[f450,f322])).
fof(f43640,plain,(
  ( ! [X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X8))) ) | $spl6),
  inference(resolution,[],[f25985,f522])).
fof(f43651,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4)) ) | $spl6),
  inference(subsumption_resolution,[],[f43638,f329])).
fof(f43638,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP7('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5,X4) | ~sP6(X5,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4)) ) | $spl6),
  inference(resolution,[],[f25985,f320])).
fof(f43636,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25985,f322])).
fof(f43628,plain,(
  ( ! [X15] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X15))) ) | $spl6),
  inference(resolution,[],[f25972,f774])).
fof(f43627,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f25972,f25848])).
fof(f25848,plain,(
  ( ! [X70,X71] : (~sP4(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X70,X71) | p(s(bool,X71)) | ~p(s(bool,X70))) ) | $spl0),
  inference(resolution,[],[f450,f305])).
fof(f43626,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | $spl6),
  inference(resolution,[],[f25972,f508])).
fof(f43625,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | $spl6),
  inference(resolution,[],[f25972,f775])).
fof(f43624,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f25972,f25849])).
fof(f25849,plain,(
  ( ! [X72,X73] : (~sP4(X72,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X73) | p(s(bool,X73)) | ~p(s(bool,X72))) ) | $spl0),
  inference(resolution,[],[f450,f305])).
fof(f43623,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | $spl6),
  inference(resolution,[],[f25972,f509])).
fof(f43622,plain,(
  ( ! [X8,X9] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X8,X9) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25972,f445])).
fof(f43635,plain,(
  ( ! [X6,X7] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f43621,f309])).
fof(f43621,plain,(
  ( ! [X6,X7] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X7))) ) | $spl6),
  inference(resolution,[],[f25972,f303])).
fof(f43633,plain,(
  ( ! [X4,X5] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4,X5) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f43620,f310])).
fof(f43620,plain,(
  ( ! [X4,X5] : (sP5('const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4,X5) | ~p(s(bool,X4)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25972,f302])).
fof(f43618,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25972,f305])).
fof(f43610,plain,(
  ( ! [X17] : (p(s(bool,X17)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25956,f758])).
fof(f43609,plain,(
  ( ! [X16] : (p(s(bool,X16)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f25956,f25832])).
fof(f25832,plain,(
  ( ! [X39,X38] : (~sP2(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X38,X39) | p(s(bool,X39)) | p(s(bool,X38))) ) | $spl0),
  inference(resolution,[],[f450,f288])).
fof(f43608,plain,(
  ( ! [X15] : (p(s(bool,X15)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25956,f492])).
fof(f43606,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl6 | $spl33)),
  inference(resolution,[],[f25956,f5636])).
fof(f5636,plain,(
  ( ! [X43,X42] : (~sP2(X42,X43,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X43)) | ~p(s(bool,X42))) ) | $spl33),
  inference(resolution,[],[f5202,f288])).
fof(f43605,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25956,f4647])).
fof(f43604,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25956,f29581])).
fof(f43603,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25956,f5110])).
fof(f5110,plain,(
  ( ! [X43,X44] : (~sP2(X43,X44,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | p(s(bool,X44)) | ~p(s(bool,X43))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f4625,f288])).
fof(f43602,plain,(
  ( ! [X8,X9] : (sP3(X8,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X9) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X9))) ) | $spl6),
  inference(resolution,[],[f25956,f287])).
fof(f43598,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f25956,f288])).
fof(f43590,plain,(
  ( ! [X17] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X17))) ) | $spl6),
  inference(resolution,[],[f25955,f758])).
fof(f43589,plain,(
  ( ! [X16] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X16))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f25955,f25832])).
fof(f43588,plain,(
  ( ! [X15] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | $spl6),
  inference(resolution,[],[f25955,f492])).
fof(f43586,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl6 | $spl33)),
  inference(resolution,[],[f25955,f5637])).
fof(f5637,plain,(
  ( ! [X45,X44] : (~sP2(X44,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X45) | p(s(bool,X45)) | ~p(s(bool,X44))) ) | $spl33),
  inference(resolution,[],[f5202,f288])).
fof(f43585,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | ($spl3 | $spl6)),
  inference(resolution,[],[f25955,f4648])).
fof(f43584,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | ($spl6 | $spl73)),
  inference(resolution,[],[f25955,f29582])).
fof(f43583,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25955,f5111])).
fof(f5111,plain,(
  ( ! [X45,X46] : (~sP2(X45,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X46) | p(s(bool,X46)) | ~p(s(bool,X45))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f4625,f288])).
fof(f43582,plain,(
  ( ! [X8,X9] : (sP3(X8,X9,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25955,f287])).
fof(f43578,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25955,f288])).
fof(f43570,plain,(
  ( ! [X16] : (~p(s(bool,X16)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25942,f744])).
fof(f43569,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f25942,f25818])).
fof(f25818,plain,(
  ( ! [X10,X11] : (~sP0(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)),X10,X11) | ~p(s(bool,X11)) | p(s(bool,X10))) ) | $spl0),
  inference(resolution,[],[f450,f271])).
fof(f43568,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f25942,f478])).
fof(f43566,plain,(
  ( ! [X12] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | $spl6),
  inference(resolution,[],[f25942,f745])).
fof(f43565,plain,(
  ( ! [X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | ($spl0 | $spl6)),
  inference(resolution,[],[f25942,f25819])).
fof(f25819,plain,(
  ( ! [X12,X13] : (~sP0(X12,X13,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | p(s(bool,X13)) | ~p(s(bool,X12))) ) | $spl0),
  inference(resolution,[],[f450,f271])).
fof(f43564,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | $spl6),
  inference(resolution,[],[f25942,f479])).
fof(f43563,plain,(
  ( ! [X8,X9] : (sP1(X8,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X9) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X9))) ) | $spl6),
  inference(resolution,[],[f25942,f270])).
fof(f43559,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f25942,f271])).
fof(f43527,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f4259,f3311])).
fof(f43520,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X3) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X2)) ) | $spl6),
  inference(resolution,[],[f4259,f319])).
fof(f43389,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f43379,f404])).
fof(f43379,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f26012,f364])).
fof(f43388,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f43378,f404])).
fof(f43378,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl6),
  inference(resolution,[],[f26012,f368])).
fof(f43359,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f43349,f404])).
fof(f43349,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26011,f364])).
fof(f43358,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f43348,f404])).
fof(f43348,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26011,f368])).
fof(f29515,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25991,f5143])).
fof(f29523,plain,(
  ( ! [X13] : (p(s(bool,X13)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25991,f5144])).
fof(f29400,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25955,f5111])).
fof(f29416,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25956,f5110])).
fof(f29483,plain,(
  ( ! [X10] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X10))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25989,f5144])).
fof(f29486,plain,(
  ( ! [X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X13))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25989,f5145])).
fof(f29499,plain,(
  ( ! [X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X8))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25990,f5143])).
fof(f29502,plain,(
  ( ! [X11] : (p(s(bool,X11)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | ($spl3 | $spl6 | $spl10)),
  inference(resolution,[],[f25990,f5145])).
fof(f43268,plain,(
  ( ! [X4,X5] : (sP9(X4,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X5) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X4,X5,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP6(X5,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X4)) ) | $spl7),
  inference(resolution,[],[f4518,f320])).
fof(f43227,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f4515,f321])).
fof(f43207,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,X9))) ) | $spl7),
  inference(resolution,[],[f4514,f3428])).
fof(f43206,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f4514,f3243])).
fof(f43085,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl128),
  inference(cnf_transformation,[],[f43085_D])).
fof(f43085_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f43087,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl130),
  inference(cnf_transformation,[],[f43087_D])).
fof(f43087_D,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f42962,plain,(
  s('type.num.num',sK15) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl3 | $spl6)),
  inference(resolution,[],[f29796,f437])).
fof(f29796,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl3 | $spl6)),
  inference(forward_demodulation,[],[f29681,f262])).
fof(f29681,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ($spl3 | $spl6)),
  inference(resolution,[],[f28084,f260])).
fof(f28084,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ($spl3 | $spl6)),
  inference(resolution,[],[f4279,f4739])).
fof(f4739,plain,(
  ( ! [X143] : (~sP12(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X143) | p(s(bool,X143))) ) | $spl3),
  inference(subsumption_resolution,[],[f4702,f404])).
fof(f4702,plain,(
  ( ! [X143] : (p(s(bool,'F_0')) | p(s(bool,X143)) | ~sP12(happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)),X143)) ) | $spl3),
  inference(resolution,[],[f453,f364])).
fof(f34494,plain,(
  $false | ($spl0 | $spl49)),
  inference(subsumption_resolution,[],[f5898,f25902])).
fof(f25902,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | $spl0),
  inference(resolution,[],[f450,f420])).
fof(f30103,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) | $spl0),
  inference(resolution,[],[f25869,f3311])).
fof(f30000,plain,(
  $false | ($spl3 | $spl6 | $spl28)),
  inference(subsumption_resolution,[],[f29999,f28084])).
fof(f29999,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | $spl28),
  inference(subsumption_resolution,[],[f29986,f386])).
fof(f29986,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,'F_0')) | $spl28),
  inference(resolution,[],[f4789,f14361])).
fof(f14361,plain,(
  ( ! [X8,X9] : (~sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X8,X9) | ~p(s(bool,X9)) | p(s(bool,X8))) )),
  inference(resolution,[],[f787,f317])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X2,X1) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f195])).
fof(f29998,plain,(
  $false | ($spl3 | $spl6 | $spl28)),
  inference(subsumption_resolution,[],[f29997,f28084])).
fof(f29997,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | $spl28),
  inference(subsumption_resolution,[],[f29977,f386])).
fof(f29977,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,'F_0')) | $spl28),
  inference(resolution,[],[f4789,f3492])).
fof(f3492,plain,(
  ( ! [X6,X7] : (~sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)),X6,X7) | ~p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(resolution,[],[f521,f317])).
fof(f29971,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl28),
  inference(resolution,[],[f4789,f337])).
fof(f26219,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | $spl66),
  inference(cnf_transformation,[],[f26219_D])).
fof(f26219_D,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f29970,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,X3))) ) | $spl28),
  inference(resolution,[],[f4789,f336])).
fof(f29968,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl28),
  inference(resolution,[],[f4789,f1273])).
fof(f28761,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | $spl107),
  inference(cnf_transformation,[],[f28761_D])).
fof(f28761_D,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f29876,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~p(s(bool,X9))) ) | $spl73),
  inference(resolution,[],[f29609,f3428])).
fof(f29609,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X100)) ) | $spl73),
  inference(resolution,[],[f28076,f324])).
fof(f28762,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | $spl108),
  inference(cnf_transformation,[],[f28762_D])).
fof(f28762_D,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f29875,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | $spl73),
  inference(resolution,[],[f29609,f3243])).
fof(f29518,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f25991,f3311])).
fof(f29230,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f4259,f3311])).
fof(f29224,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X3) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X2)) ) | $spl6),
  inference(resolution,[],[f4259,f319])).
fof(f29091,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f29082,f404])).
fof(f29082,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f26012,f364])).
fof(f29090,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f29081,f404])).
fof(f29081,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl6),
  inference(resolution,[],[f26012,f368])).
fof(f29080,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f29071,f404])).
fof(f29071,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26011,f364])).
fof(f29079,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f29070,f404])).
fof(f29070,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26011,f368])).
fof(f29024,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f4515,f321])).
fof(f29005,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,X9))) ) | $spl7),
  inference(resolution,[],[f4514,f3428])).
fof(f29004,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f4514,f3243])).
fof(f28776,plain,(
  ( ! [X17] : (p(s(bool,X17)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28660,f788])).
fof(f28660,plain,(
  ( ! [X101,X102] : (sP8('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X101,X102)) ) | ($spl7 | $spl24)),
  inference(forward_demodulation,[],[f4515,f4580])).
fof(f28775,plain,(
  ( ! [X16] : (p(s(bool,X16)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28660,f522])).
fof(f28774,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24 | $spl33)),
  inference(resolution,[],[f28660,f5663])).
fof(f5663,plain,(
  ( ! [X97,X96] : (~sP8(X96,X97,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X97)) | ~p(s(bool,X96))) ) | $spl33),
  inference(resolution,[],[f5202,f322])).
fof(f28773,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl3 | $spl7 | $spl24)),
  inference(resolution,[],[f28660,f4674])).
fof(f4674,plain,(
  ( ! [X97,X96] : (~sP8(X96,X97,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | ~p(s(bool,X97)) | ~p(s(bool,X96))) ) | $spl3),
  inference(resolution,[],[f453,f322])).
fof(f28772,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl3 | $spl7 | $spl10 | $spl24)),
  inference(resolution,[],[f28660,f5137])).
fof(f5137,plain,(
  ( ! [X97,X98] : (~sP8(X97,X98,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X98)) | ~p(s(bool,X97))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f4625,f322])).
fof(f28769,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28660,f3425])).
fof(f28768,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,X9) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28660,f3428])).
fof(f28781,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | $spl110),
  inference(cnf_transformation,[],[f28781_D])).
fof(f28781_D,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f28766,plain,(
  ( ! [X6,X7] : (sP9('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~sP7(X6,X7,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~sP6(X7,X6,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28660,f321])).
fof(f28763,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28660,f322])).
fof(f28754,plain,(
  ( ! [X17] : (p(s(bool,X17)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28659,f787])).
fof(f28659,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X100)) ) | ($spl7 | $spl24)),
  inference(forward_demodulation,[],[f4514,f4580])).
fof(f28753,plain,(
  ( ! [X16] : (p(s(bool,X16)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28659,f521])).
fof(f28752,plain,(
  ( ! [X15] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X15))) ) | ($spl7 | $spl24 | $spl33)),
  inference(resolution,[],[f28659,f5663])).
fof(f28751,plain,(
  ( ! [X14] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X14))) ) | ($spl3 | $spl7 | $spl24)),
  inference(resolution,[],[f28659,f4674])).
fof(f28750,plain,(
  ( ! [X13] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X13))) ) | ($spl3 | $spl7 | $spl10 | $spl24)),
  inference(resolution,[],[f28659,f5137])).
fof(f28749,plain,(
  ( ! [X12] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | sP9(X12,'F_0','const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28659,f3249])).
fof(f28747,plain,(
  ( ! [X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28659,f3425])).
fof(f28746,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~p(s(bool,X9))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28659,f3428])).
fof(f28745,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28659,f3243])).
fof(f28758,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X2)) ) | $spl104),
  inference(cnf_transformation,[],[f28758_D])).
fof(f28758_D,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X2)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f28757,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X2)) ) | ($spl7 | $spl24)),
  inference(subsumption_resolution,[],[f28742,f333])).
fof(f28742,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X2) | ~sP6('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X3,X2)) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28659,f319])).
fof(f28741,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28659,f322])).
fof(f28735,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28658,f775])).
fof(f28658,plain,(
  ( ! [X76,X75] : (sP4('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X75,X76)) ) | ($spl7 | $spl24)),
  inference(forward_demodulation,[],[f4502,f4580])).
fof(f4502,plain,(
  ( ! [X76,X75] : (sP4('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X75,X76)) ) | $spl7),
  inference(resolution,[],[f457,f308])).
fof(f28734,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28658,f509])).
fof(f28733,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24 | $spl33)),
  inference(resolution,[],[f28658,f5650])).
fof(f5650,plain,(
  ( ! [X70,X71] : (~sP4(X70,X71,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X71)) | ~p(s(bool,X70))) ) | $spl33),
  inference(resolution,[],[f5202,f305])).
fof(f28732,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl3 | $spl7 | $spl24)),
  inference(resolution,[],[f28658,f4661])).
fof(f4661,plain,(
  ( ! [X70,X71] : (~sP4(X70,X71,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | ~p(s(bool,X71)) | ~p(s(bool,X70))) ) | $spl3),
  inference(resolution,[],[f453,f305])).
fof(f28731,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl3 | $spl7 | $spl10 | $spl24)),
  inference(resolution,[],[f28658,f5124])).
fof(f5124,plain,(
  ( ! [X72,X71] : (~sP4(X71,X72,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X72)) | ~p(s(bool,X71))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f4625,f305])).
fof(f28738,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X3) | ~p(s(bool,X3))) ) | $spl102),
  inference(cnf_transformation,[],[f28738_D])).
fof(f28738_D,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X3) | ~p(s(bool,X3))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f28727,plain,(
  ( ! [X2,X3] : (sP5(X2,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X3) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X3))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28658,f301])).
fof(f28726,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28658,f305])).
fof(f28720,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28657,f774])).
fof(f28657,plain,(
  ( ! [X74,X73] : (sP4(X73,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X74)) ) | ($spl7 | $spl24)),
  inference(forward_demodulation,[],[f4501,f4580])).
fof(f4501,plain,(
  ( ! [X74,X73] : (sP4(X73,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X74)) ) | $spl7),
  inference(resolution,[],[f457,f307])).
fof(f28719,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28657,f508])).
fof(f28718,plain,(
  ( ! [X12] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | ($spl7 | $spl24 | $spl33)),
  inference(resolution,[],[f28657,f5650])).
fof(f28717,plain,(
  ( ! [X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | ($spl3 | $spl7 | $spl24)),
  inference(resolution,[],[f28657,f4661])).
fof(f28716,plain,(
  ( ! [X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl3 | $spl7 | $spl10 | $spl24)),
  inference(resolution,[],[f28657,f5124])).
fof(f28723,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~p(s(bool,X3))) ) | $spl100),
  inference(cnf_transformation,[],[f28723_D])).
fof(f28723_D,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~p(s(bool,X3))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f28712,plain,(
  ( ! [X2,X3] : (sP5(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | ~p(s(bool,X3)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28657,f301])).
fof(f28711,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28657,f305])).
fof(f28705,plain,(
  ( ! [X15] : (p(s(bool,X15)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24 | $spl33)),
  inference(resolution,[],[f28656,f5637])).
fof(f28656,plain,(
  ( ! [X47,X48] : (sP2('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X47,X48)) ) | ($spl7 | $spl24)),
  inference(forward_demodulation,[],[f4488,f4580])).
fof(f4488,plain,(
  ( ! [X47,X48] : (sP2('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X47,X48)) ) | $spl7),
  inference(resolution,[],[f457,f291])).
fof(f28704,plain,(
  ( ! [X14] : (p(s(bool,X14)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl3 | $spl7 | $spl24)),
  inference(resolution,[],[f28656,f4648])).
fof(f28703,plain,(
  ( ! [X13] : (p(s(bool,X13)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl3 | $spl7 | $spl10 | $spl24)),
  inference(resolution,[],[f28656,f5111])).
fof(f28702,plain,(
  ( ! [X12] : (p(s(bool,X12)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24 | $spl33)),
  inference(resolution,[],[f28656,f5636])).
fof(f28701,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl3 | $spl7 | $spl24)),
  inference(resolution,[],[f28656,f4647])).
fof(f28700,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl3 | $spl7 | $spl10 | $spl24)),
  inference(resolution,[],[f28656,f5110])).
fof(f28708,plain,(
  ( ! [X2,X3] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X2,X3)) ) | $spl98),
  inference(cnf_transformation,[],[f28708_D])).
fof(f28708_D,plain,(
  ( ! [X2,X3] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X2,X3)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f28710,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(subsumption_resolution,[],[f28698,f294])).
fof(f28698,plain,(
  ( ! [X6,X7] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28656,f286])).
fof(f28709,plain,(
  ( ! [X4,X5] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X4,X5) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(subsumption_resolution,[],[f28697,f293])).
fof(f28697,plain,(
  ( ! [X4,X5] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X4,X5) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X5))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28656,f285])).
fof(f28696,plain,(
  ( ! [X2,X3] : (sP3('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X2,X3) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28656,f444])).
fof(f28695,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28656,f288])).
fof(f28689,plain,(
  ( ! [X14] : (~p(s(bool,X14)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24 | $spl33)),
  inference(resolution,[],[f28655,f5620])).
fof(f5620,plain,(
  ( ! [X10,X11] : (~sP0(X10,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X11) | ~p(s(bool,X11)) | ~p(s(bool,X10))) ) | $spl33),
  inference(resolution,[],[f5202,f271])).
fof(f28655,plain,(
  ( ! [X15,X16] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X15,X16)) ) | ($spl7 | $spl24)),
  inference(forward_demodulation,[],[f4472,f4580])).
fof(f4472,plain,(
  ( ! [X15,X16] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X15,X16)) ) | $spl7),
  inference(resolution,[],[f457,f274])).
fof(f28688,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl3 | $spl7 | $spl24)),
  inference(resolution,[],[f28655,f4631])).
fof(f28687,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl3 | $spl7 | $spl10 | $spl24)),
  inference(resolution,[],[f28655,f5094])).
fof(f5094,plain,(
  ( ! [X12,X11] : (~sP0(X11,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X12) | ~p(s(bool,X12)) | ~p(s(bool,X11))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f4625,f271])).
fof(f28686,plain,(
  ( ! [X11] : (p(s(bool,X11)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28655,f745])).
fof(f28685,plain,(
  ( ! [X10] : (p(s(bool,X10)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28655,f479])).
fof(f28692,plain,(
  ( ! [X2,X3] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X2,X3)) ) | $spl96),
  inference(cnf_transformation,[],[f28692_D])).
fof(f28692_D,plain,(
  ( ! [X2,X3] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X2,X3)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f28694,plain,(
  ( ! [X6,X7] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(subsumption_resolution,[],[f28683,f277])).
fof(f28683,plain,(
  ( ! [X6,X7] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X6,X7) | p(s(bool,X6)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28655,f269])).
fof(f28693,plain,(
  ( ! [X4,X5] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X4,X5) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(subsumption_resolution,[],[f28682,f276])).
fof(f28682,plain,(
  ( ! [X4,X5] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X4,X5) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X5))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28655,f268])).
fof(f28681,plain,(
  ( ! [X2,X3] : (sP1('const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)),X2,X3) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28655,f443])).
fof(f28680,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28655,f271])).
fof(f28674,plain,(
  ( ! [X15] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X15))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28654,f744])).
fof(f28654,plain,(
  ( ! [X14,X13] : (sP0(X13,X14,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) ) | ($spl7 | $spl24)),
  inference(forward_demodulation,[],[f4471,f4580])).
fof(f4471,plain,(
  ( ! [X14,X13] : (sP0(X13,X14,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f457,f272])).
fof(f28673,plain,(
  ( ! [X14] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X14))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28654,f478])).
fof(f28671,plain,(
  ( ! [X12] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X12))) ) | ($spl7 | $spl24 | $spl33)),
  inference(resolution,[],[f28654,f5620])).
fof(f28670,plain,(
  ( ! [X11] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X11))) ) | ($spl3 | $spl7 | $spl24)),
  inference(resolution,[],[f28654,f4631])).
fof(f28669,plain,(
  ( ! [X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | ($spl3 | $spl7 | $spl10 | $spl24)),
  inference(resolution,[],[f28654,f5094])).
fof(f28679,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | p(s(bool,X9))) ) | $spl94),
  inference(cnf_transformation,[],[f28679_D])).
fof(f28679_D,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | p(s(bool,X9))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f28668,plain,(
  ( ! [X8,X9] : (sP1(X8,X9,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))) | p(s(bool,X9)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28654,f270])).
fof(f28664,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14))))) ) | ($spl7 | $spl24)),
  inference(resolution,[],[f28654,f271])).
fof(f28392,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f4259,f3311])).
fof(f28386,plain,(
  ( ! [X2,X3] : (sP9(X2,X3,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ~sP8(X2,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X3) | ~sP7(X3,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)),X2)) ) | $spl6),
  inference(resolution,[],[f4259,f319])).
fof(f28244,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f28235,f404])).
fof(f28235,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | $spl6),
  inference(resolution,[],[f26012,f364])).
fof(f28243,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f28234,f404])).
fof(f28234,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | $spl6),
  inference(resolution,[],[f26012,f368])).
fof(f28233,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f28224,f404])).
fof(f28224,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26011,f364])).
fof(f28232,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(subsumption_resolution,[],[f28223,f404])).
fof(f28223,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))))) ) | $spl6),
  inference(resolution,[],[f26011,f368])).
fof(f28078,plain,(
  ( ! [X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f28069,f404])).
fof(f28069,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X1))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28065,f364])).
fof(f28075,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f28068,f404])).
fof(f28068,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK14)))) | p(s(bool,X0))) ) | ($spl6 | $spl24)),
  inference(resolution,[],[f28065,f368])).
fof(f28032,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f27929,f4790])).
fof(f27929,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(resolution,[],[f4790,f432])).
fof(f432,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f108])).
fof(f108,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f28])).
fof(f28,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.bool.NOT_CLAUSES|split|1')).
fof(f28031,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f27928,f4790])).
fof(f27928,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(resolution,[],[f4790,f430])).
fof(f430,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f107])).
fof(f107,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f29])).
fof(f29,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.bool.NOT_CLAUSES|split|2')).
fof(f28030,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f27925,f4790])).
fof(f27925,plain,(
  ( ! [X240] : (p(s(bool,X240))) ) | $spl30),
  inference(resolution,[],[f4790,f397])).
fof(f397,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f230])).
fof(f230,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X13] : ((p(s(bool,'T_0')) => p(s(bool,X13))) <=> p(s(bool,X13)))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.bool.IMP_CLAUSES|split|0')).
fof(f28029,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f27924,f4790])).
fof(f27924,plain,(
  ( ! [X239] : (p(s(bool,X239))) ) | $spl30),
  inference(resolution,[],[f4790,f394])).
fof(f394,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X13] : ((p(s(bool,X13)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.bool.IMP_CLAUSES|split|1')).
fof(f28028,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f27922,f4790])).
fof(f27922,plain,(
  p(s(bool,'F_0')) | $spl30),
  inference(resolution,[],[f4790,f391])).
fof(f391,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f227])).
fof(f28027,plain,(
  $false | $spl30),
  inference(subsumption_resolution,[],[f27920,f4790])).
fof(f27920,plain,(
  ( ! [X236] : (p(s(bool,X236))) ) | $spl30),
  inference(resolution,[],[f4790,f388])).
fof(f27930,plain,(
  $false | $spl30),
  inference(resolution,[],[f4790,f434])).
fof(f434,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f15])).
fof(f15,axiom,(
  p(s(bool,'T_0'))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.bool.TRUTH')).
fof(f27931,plain,(
  $false | $spl30),
  inference(resolution,[],[f4790,f435])).
fof(f435,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.arithmetic.BOUNDED_FORALL_THM',aHL_TRUTH)).
fof(f27939,plain,(
  $false | $spl30),
  inference(resolution,[],[f4790,f440])).
fof(f440,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X9,X8] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))))),
  file('thm.arithmetic.BOUNDED_FORALL_THM','thm.arithmetic.LESS_ADD_SUC')).
fof(f27943,plain,(
  $false | ($spl18 | $spl30)),
  inference(resolution,[],[f4790,f3961])).
fof(f27945,plain,(
  $false | ($spl8 | $spl30)),
  inference(resolution,[],[f4790,f462])).
fof(f27947,plain,(
  $false | ($spl4 | $spl30)),
  inference(resolution,[],[f4790,f454])).
fof(f27951,plain,(
  $false | $spl30),
  inference(resolution,[],[f4790,f587])).
fof(f27952,plain,(
  $false | $spl30),
  inference(resolution,[],[f4790,f261])).
fof(f27953,plain,(
  $false | ($spl8 | $spl30)),
  inference(resolution,[],[f4790,f4820])).
fof(f27688,plain,(
  $false | ($spl6 | $spl37)),
  inference(subsumption_resolution,[],[f5416,f26024])).
fof(f26024,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | $spl6),
  inference(resolution,[],[f4200,f420])).
fof(f27687,plain,(
  $false | ($spl6 | $spl27)),
  inference(subsumption_resolution,[],[f4788,f25759])).
fof(f25759,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | $spl6),
  inference(resolution,[],[f458,f420])).
fof(f26548,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) | $spl0),
  inference(resolution,[],[f25869,f3311])).
fof(f26373,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f25991,f3311])).
fof(f26204,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f4259,f3311])).
fof(f25666,plain,(
  $false | ($spl6 | $spl23)),
  inference(subsumption_resolution,[],[f25665,f429])).
fof(f25665,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl23)),
  inference(forward_demodulation,[],[f4200,f5051])).
fof(f5051,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | $spl23),
  inference(resolution,[],[f4578,f416])).
fof(f25664,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))))),s('type.num.num',sK14)))) ) | $spl64),
  inference(cnf_transformation,[],[f25664_D])).
fof(f25664_D,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))))),s('type.num.num',sK14)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f25637,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))))),s('type.num.num',sK14))) | ~p(s(bool,X15))) ) | $spl7),
  inference(resolution,[],[f24970,f3428])).
fof(f24970,plain,(
  ( ! [X151,X149,X150] : (sP8(X149,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num','const.num.SUC_1'(s('type.num.num',X150))))),s('type.num.num',sK14)),X151)) ) | $spl7),
  inference(resolution,[],[f21060,f324])).
fof(f21060,plain,(
  ( ! [X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK15),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))),s('type.num.num',sK14))))) ) | $spl7),
  inference(resolution,[],[f4464,f440])).
fof(f4464,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK14))))) ) | $spl7),
  inference(resolution,[],[f457,f436])).
fof(f25285,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl18),
  inference(resolution,[],[f24105,f3311])).
fof(f24753,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl8),
  inference(resolution,[],[f14997,f3311])).
fof(f23959,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(resolution,[],[f3898,f3311])).
fof(f17015,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl8),
  inference(resolution,[],[f4879,f3311])).
fof(f6005,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | ~p(s(bool,X9))) ) | $spl3),
  inference(resolution,[],[f4675,f3428])).
fof(f6004,plain,(
  ( ! [X8] : (sP9(X8,'F_0',happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16))) | s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) ) | $spl3),
  inference(resolution,[],[f4675,f3243])).
fof(f5886,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | ~p(s(bool,X9))) ) | $spl1),
  inference(resolution,[],[f4390,f3428])).
fof(f5885,plain,(
  ( ! [X8] : (sP9(X8,'F_0',happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) ) | $spl1),
  inference(resolution,[],[f4390,f3243])).
fof(f5814,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X9))) ) | $spl33),
  inference(resolution,[],[f5664,f3428])).
fof(f5813,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | $spl33),
  inference(resolution,[],[f5664,f3243])).
fof(f5547,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | ~p(s(bool,X9))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f5138,f3428])).
fof(f5138,plain,(
  ( ! [X99,X100] : (sP8(X99,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)),X100)) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f4625,f324])).
fof(f5546,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) ) | ($spl3 | $spl10)),
  inference(resolution,[],[f5138,f3243])).
fof(f5404,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,X9))) ) | $spl23),
  inference(resolution,[],[f5013,f3428])).
fof(f5403,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl23),
  inference(resolution,[],[f5013,f3243])).
fof(f5202,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | $spl33),
  inference(cnf_transformation,[],[f5202_D])).
fof(f5202_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f5087,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK16) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK16),s('type.num.num',sK14)))) | ($spl3 | $spl10)),
  inference(resolution,[],[f4625,f437])).
fof(f4776,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | ~p(s(bool,X9))) ) | $spl7),
  inference(resolution,[],[f4514,f3428])).
fof(f4775,plain,(
  ( ! [X8] : (sP9(X8,'F_0','const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) ) | $spl7),
  inference(resolution,[],[f4514,f3243])).
fof(f4580,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) | $spl24),
  inference(cnf_transformation,[],[f4580_D])).
fof(f4580_D,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4463,plain,(
  s('type.num.num',sK14) = s('type.num.num',sK15) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | $spl7),
  inference(resolution,[],[f457,f437])).
fof(f4453,plain,(
  $false | ($spl1 | $spl6 | $spl10)),
  inference(subsumption_resolution,[],[f4340,f458])).
fof(f4340,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14)))) | ($spl1 | $spl10)),
  inference(resolution,[],[f451,f466])).
fof(f4189,plain,(
  $false | ($spl5 | $spl10)),
  inference(subsumption_resolution,[],[f4188,f440])).
fof(f4188,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl5 | $spl10)),
  inference(forward_demodulation,[],[f4187,f442])).
fof(f4187,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK14)))))) | ($spl5 | $spl10)),
  inference(subsumption_resolution,[],[f4186,f392])).
fof(f4186,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK14)))))) | ($spl5 | $spl10)),
  inference(forward_demodulation,[],[f4067,f562])).
fof(f4067,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK14)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK14)))))) | ($spl5 | $spl10)),
  inference(resolution,[],[f859,f259])).
fof(f859,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',sK14)))) | ($spl5 | $spl10)),
  inference(resolution,[],[f466,f455])).
fof(f3840,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(resolution,[],[f587,f437])).
fof(f3142,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3075,f2964])).
fof(f3075,plain,(
  ~p(s(bool,'T_0')) | $spl14),
  inference(resolution,[],[f2964,f433])).
fof(f433,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f247])).
fof(f3141,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3074,f2964])).
fof(f3074,plain,(
  ~p(s(bool,'T_0')) | $spl14),
  inference(resolution,[],[f2964,f431])).
fof(f431,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f246])).
fof(f3140,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f3069,f2964])).
fof(f3069,plain,(
  ( ! [X250] : (~p(s(bool,X250))) ) | $spl14),
  inference(resolution,[],[f2964,f386])).
fof(f3073,plain,(
  $false | $spl14),
  inference(resolution,[],[f2964,f429])).
fof(f3083,plain,(
  $false | ($spl5 | $spl14)),
  inference(resolution,[],[f2964,f455])).
fof(f2965,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl16),
  inference(cnf_transformation,[],[f2965_D])).
fof(f2965_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2963,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f2957])).
fof(f2957,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) )),
  inference(resolution,[],[f1270,f337])).
fof(f1270,plain,(
  ( ! [X4,X3] : (sP9(X4,X3,'F_0') | s(bool,'T_0') != s(bool,X3) | p(s(bool,X4))) )),
  inference(superposition,[],[f1264,f334])).
fof(f466,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X5)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK14))))) ) | $spl10),
  inference(cnf_transformation,[],[f466_D])).
fof(f466_D,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X5)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK14))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f454,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl4),
  inference(cnf_transformation,[],[f454_D])).
fof(f454_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f468,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X5)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK14))))) )),
  inference(forward_demodulation,[],[f251,f262])).
fof(f251,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X5)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK14))))) )),
  inference(cnf_transformation,[],[f166])).
fof(f467,plain,(
  ( ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X4))))) ) | $spl12),
  inference(cnf_transformation,[],[f467_D])).
fof(f467_D,plain,(
  ( ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X4))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f465,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X4)))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X5)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK14))))) )),
  inference(forward_demodulation,[],[f252,f262])).
fof(f252,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X4)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X5)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK14))))) )),
  inference(cnf_transformation,[],[f166])).
fof(f464,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))))),
  inference(forward_demodulation,[],[f463,f262])).
fof(f463,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))))),
  inference(forward_demodulation,[],[f253,f262])).
fof(f253,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))))),
  inference(cnf_transformation,[],[f166])).
fof(f460,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))),
  inference(forward_demodulation,[],[f459,f262])).
fof(f459,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))),
  inference(forward_demodulation,[],[f254,f262])).
fof(f254,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))),
  inference(cnf_transformation,[],[f166])).
fof(f456,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))))),
  inference(forward_demodulation,[],[f255,f262])).
fof(f255,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK15),s('type.num.num',sK14))))),
  inference(cnf_transformation,[],[f166])).
fof(f449,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))),
  inference(forward_demodulation,[],[f256,f262])).
fof(f256,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))))),
  inference(cnf_transformation,[],[f166])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 4144
% Passive clauses: 43665
% Generated clauses: 244204
% Final active clauses: 1486
% Final passive clauses: 16877
% Input formulas: 48
% Initial clauses: 193
% 
% Selected by SInE selection: 48
% SInE iterations: 1
% 
% Duplicate literals: 11932
% Trivial inequalities: 377
% Fw subsumption resolutions: 40457
% Fw demodulations: 53343
% Bw demodulations: 1138
% 
% Simple tautologies: 18072
% Equational tautologies: 6
% Forward subsumptions: 77931
% 
% Binary resolution: 94010
% Factoring: 1
% Forward superposition: 24240
% Backward superposition: 14085
% Equality resolution: 7
% 
% Split clauses: 3449
% Split components: 7227
% SAT solver clauses: 3589
% SAT solver unit clauses: 16
% SAT solver binary clauses: 86
% 
% Sat splits: 3449
% Sat splitting refutations: 60
% 
% Memory used [KB]: 44519
% Time elapsed: 3.067 s
% ------------------------------
% ------------------------------
% Success in time 14.788 s
