Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numeral
% (97257)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 656
% Passive clauses: 9158
% Generated clauses: 19835
% Final active clauses: 603
% Final passive clauses: 7812
% Input formulas: 51
% Initial clauses: 184
% 
% Duplicate literals: 348
% Trivial inequalities: 84
% Fw subsumption resolutions: 1664
% Fw demodulations: 886
% Bw demodulations: 51
% 
% Simple tautologies: 249
% Equational tautologies: 60
% Forward subsumptions: 7261
% 
% Binary resolution: 9631
% Factoring: 996
% Forward superposition: 3428
% Backward superposition: 2330
% Self superposition: 51
% Equality factoring: 26
% Equality resolution: 24
% 
% Split clauses: 360
% Split components: 721
% SAT solver clauses: 742
% SAT solver unit clauses: 20
% SAT solver binary clauses: 183
% 
% Sat splits: 360
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 740
% TWLsolver calls for satisfiability: 436
% 
% Memory used [KB]: 8187
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.numeral
% (97271)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 470
% Passive clauses: 2789
% Generated clauses: 14879
% Final active clauses: 463
% Final passive clauses: 2294
% Input formulas: 51
% Initial clauses: 184
% 
% Duplicate literals: 247
% Trivial inequalities: 19
% Fw subsumption resolutions: 1068
% Fw demodulations: 1302
% Bw demodulations: 30
% Condensations: 359
% Global subsumptions: 18
% 
% Simple tautologies: 257
% Equational tautologies: 26
% Forward subsumptions: 5450
% 
% Binary resolution: 3769
% Factoring: 258
% Forward superposition: 2292
% Backward superposition: 2412
% Self superposition: 59
% Equality factoring: 28
% Equality resolution: 4
% 
% Split clauses: 70
% Split components: 145
% SAT solver clauses: 3515
% SAT solver unit clauses: 200
% SAT solver binary clauses: 72
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 10
% 
% Sat splits: 70
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 3443
% TWLsolver calls for satisfiability: 1690
% 
% Memory used [KB]: 14200
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.numeral
% (97297)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2230
% Passive clauses: 9307
% Generated clauses: 105939
% Final active clauses: 467
% Final passive clauses: 755
% Input formulas: 51
% Initial clauses: 184
% 
% Duplicate literals: 2614
% Trivial inequalities: 160
% Fw subsumption resolutions: 3387
% Fw demodulations: 12177
% Bw demodulations: 4879
% Condensations: 7430
% Global subsumptions: 2994
% 
% Simple tautologies: 161
% Equational tautologies: 626
% Forward subsumptions: 62679
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 582
% 
% Unit resulting resolution: 14504
% Factoring: 1
% Forward superposition: 28554
% Backward superposition: 22544
% Self superposition: 278
% Equality factoring: 78
% Equality resolution: 5
% 
% Split clauses: 250
% Split components: 529
% SAT solver clauses: 16497
% SAT solver unit clauses: 841
% SAT solver binary clauses: 5461
% 
% Sat splits: 250
% Sat splitting refutations: 3693
% 
% Memory used [KB]: 170018
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.numeral
% (97583)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 3400
% Passive clauses: 46398
% Generated clauses: 203604
% Final active clauses: 1946
% Final passive clauses: 23079
% Input formulas: 51
% Initial clauses: 183
% 
% Selected by SInE selection: 50
% SInE iterations: 1
% 
% Duplicate literals: 7271
% Trivial inequalities: 479
% Fw subsumption resolutions: 23088
% Fw demodulations: 39320
% Bw demodulations: 572
% 
% Simple tautologies: 9688
% Equational tautologies: 14
% Forward subsumptions: 78979
% 
% Binary resolution: 72724
% Forward superposition: 37398
% Backward superposition: 19286
% Self superposition: 44
% Equality resolution: 57
% 
% Split clauses: 1931
% Split components: 4064
% SAT solver clauses: 2048
% SAT solver unit clauses: 21
% SAT solver binary clauses: 96
% 
% Sat splits: 1931
% Sat splitting refutations: 112
% 
% Memory used [KB]: 62173
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.numeral
% (97666)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 357
% Passive clauses: 2845
% Generated clauses: 14681
% Final active clauses: 352
% Final passive clauses: 2416
% Input formulas: 51
% Initial clauses: 184
% 
% Duplicate literals: 481
% Trivial inequalities: 205
% Fw subsumption resolutions: 1722
% Fw demodulations: 342
% Bw demodulations: 17
% Global subsumptions: 14
% 
% Simple tautologies: 523
% Equational tautologies: 34
% Forward subsumptions: 8041
% 
% Unit resulting resolution: 1314
% Factoring: 18
% Forward superposition: 5123
% Backward superposition: 5100
% Self superposition: 93
% Equality factoring: 6
% Equality resolution: 22
% 
% Split clauses: 20
% Split components: 40
% SAT solver clauses: 2932
% SAT solver unit clauses: 177
% SAT solver binary clauses: 255
% 
% Sat splits: 20
% Sat splitting refutations: 51
% 
% Memory used [KB]: 7803
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.numeral
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numeral
% SZS output start Proof for thm.numeral
fof(f22705,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f481,f486,f484,f479,f602,f483,f9537,f9627,f11684,f11682,f21390,f21391,f22207,f22224,f22634])).
fof(f22634,plain,(
  $false | ($spl0 | $spl3 | $spl8)),
  inference(subsumption_resolution,[],[f22633,f445])).
fof(f445,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f314])).
fof(f314,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X14] : ! [X15] : (s(X14,X15) = s(X14,X15) <=> p(s(bool,'T_0')))),
  file('thm.numeral.numeral_5|49|_prove_ltb','thm.bool.REFL_CLAUSE')).
fof(f22633,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl3 | $spl8)),
  inference(forward_demodulation,[],[f22148,f21480])).
fof(f21480,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK13))) | ($spl0 | $spl8)),
  inference(backward_demodulation,[],[f11684,f9629])).
fof(f9629,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14))) | $spl0),
  inference(resolution,[],[f484,f273])).
fof(f273,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.numeral.numeral_5|49|_prove_ltb','thm.bool.EQ_CLAUSES|split|0')).
fof(f22148,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK13)))) | ($spl0 | $spl3 | $spl8)),
  inference(backward_demodulation,[],[f11684,f19117])).
fof(f19117,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f19116,f256])).
fof(f256,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.numeral.numeral_5|49|_prove_ltb',aHL_TRUTH)).
fof(f19116,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f19108,f9629])).
fof(f19108,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | ($spl0 | $spl3)),
  inference(trivial_inequality_removal,[],[f19061])).
fof(f19061,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK12) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | ($spl0 | $spl3)),
  inference(superposition,[],[f491,f18637])).
fof(f18637,plain,(
  ( ! [X0,X1] : (s('type.num.num',X1) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X0)))))) ) | $spl0),
  inference(superposition,[],[f18413,f18418])).
fof(f18418,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK13),s('type.num.num',sK14))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X2)))) ) | $spl0),
  inference(superposition,[],[f319,f10414])).
fof(f10414,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK13),s('type.num.num',sK14))),s('type.num.num',X0)))) ) | $spl0),
  inference(forward_demodulation,[],[f10205,f317])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X14] : ! [X11,X12] : s(X14,X11) = s(X14,'const.bool.COND_3'(s(bool,'T_0'),s(X14,X11),s(X14,X12)))),
  file('thm.numeral.numeral_5|49|_prove_ltb','thm.bool.COND_CLAUSES|split|0')).
fof(f10205,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK13),s('type.num.num',sK14))),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK13),s('type.num.num',X0))),s('type.num.num',sK14)))))) ) | $spl0),
  inference(superposition,[],[f358,f9629])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))),s('type.num.num',X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8,X7,X6] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X7))),s('type.num.num',X6))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7))),s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X6))),s('type.num.num',X7)))))),
  file('thm.numeral.numeral_5|49|_prove_ltb','thm.arithmetic.SUB_RIGHT_ADD')).
fof(f319,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X9,X10] : s('type.num.num',X9) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X10))),s('type.num.num',X10)))),
  file('thm.numeral.numeral_5|49|_prove_ltb','thm.arithmetic.ADD_SUB')).
fof(f18413,plain,(
  ( ! [X1] : (s('type.num.num',X1) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK13),s('type.num.num',sK14)))))) ) | $spl0),
  inference(superposition,[],[f10414,f318])).
fof(f318,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8,X7] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X8))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X7)))),
  file('thm.numeral.numeral_5|49|_prove_ltb','thm.arithmetic.ADD_SYM')).
fof(f491,plain,(
  ( ! [X0] : (s('type.num.num',sK12) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',X0))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK13))))) ) | $spl3),
  inference(superposition,[],[f486,f320])).
fof(f320,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',X1))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s('type.num.num',X0) = s('type.num.num',X1))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) => s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X7,X8] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7))))) => s('type.num.num',X7) = s('type.num.num',X8))),
  file('thm.numeral.numeral_5|49|_prove_ltb','thm.arithmetic.LESS_EQUAL_ANTISYM')).
fof(f22224,plain,(
  $false | ($spl0 | $spl5 | $spl8)),
  inference(subsumption_resolution,[],[f22223,f319])).
fof(f22223,plain,(
  s('type.num.num',sK12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK13))),s('type.num.num',sK13))) | ($spl0 | $spl5 | $spl8)),
  inference(forward_demodulation,[],[f21445,f18637])).
fof(f21445,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK13))),s('type.num.num',sK13))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK13),s('type.num.num',sK13))))) | ($spl5 | $spl8)),
  inference(backward_demodulation,[],[f11684,f602])).
fof(f22207,plain,(
  $false | ($spl0 | $spl3 | $spl8)),
  inference(subsumption_resolution,[],[f21418,f18637])).
fof(f21418,plain,(
  s('type.num.num',sK12) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK13),s('type.num.num',sK13))))) | ($spl3 | $spl8)),
  inference(backward_demodulation,[],[f11684,f486])).
fof(f21391,plain,(
  $false | ($spl0 | $spl7)),
  inference(subsumption_resolution,[],[f21285,f21235])).
fof(f21235,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) | ($spl0 | $spl7)),
  inference(factoring,[],[f20557])).
fof(f20557,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK12))))) ) | ($spl0 | $spl7)),
  inference(resolution,[],[f20424,f354])).
fof(f354,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0,X1] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8,X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8)))))),
  file('thm.numeral.numeral_5|49|_prove_ltb','thm.arithmetic.NOT_LESS_EQUAL')).
fof(f20424,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK12)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK12))))) ) | ($spl0 | $spl7)),
  inference(resolution,[],[f19620,f354])).
fof(f19620,plain,(
  ( ! [X10] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',sK12))))) ) | ($spl0 | $spl7)),
  inference(forward_demodulation,[],[f19619,f10414])).
fof(f19619,plain,(
  ( ! [X10] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X10)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK13),s('type.num.num',sK14))),s('type.num.num',X10))),s('type.num.num',sK12))))) ) | ($spl0 | $spl7)),
  inference(forward_demodulation,[],[f19557,f18638])).
fof(f18638,plain,(
  ( ! [X2,X3] : (s('type.num.num',X3) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X2))),s('type.num.num',X3)))) ) | $spl0),
  inference(superposition,[],[f10414,f18418])).
fof(f19557,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X9),s('type.num.num',X9))),s('type.num.num',X10)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK13),s('type.num.num',sK14))),s('type.num.num',X10))),s('type.num.num',sK12))))) ) | ($spl0 | $spl7)),
  inference(superposition,[],[f12525,f18418])).
fof(f12525,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',sK12))))) ) | ($spl0 | $spl7)),
  inference(subsumption_resolution,[],[f10977,f12524])).
fof(f12524,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X0))),s('type.num.num',sK13))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))),s('type.num.num',X0)))) ) | $spl7),
  inference(forward_demodulation,[],[f12312,f316])).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X14] : ! [X11,X12] : s(X14,X12) = s(X14,'const.bool.COND_3'(s(bool,'F_0'),s(X14,X11),s(X14,X12)))),
  file('thm.numeral.numeral_5|49|_prove_ltb','thm.bool.COND_CLAUSES|split|1')).
fof(f12312,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X0))),s('type.num.num',sK13)))))) ) | $spl7),
  inference(superposition,[],[f358,f11689])).
fof(f11689,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13))) | $spl7),
  inference(resolution,[],[f11682,f277])).
fof(f277,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.numeral.numeral_5|49|_prove_ltb','thm.bool.EQ_CLAUSES|split|2')).
fof(f10977,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))))),s('type.num.num',sK13))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',sK12))))) ) | $spl0),
  inference(forward_demodulation,[],[f10976,f318])).
fof(f10976,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num',sK14))),s('type.num.num',sK13))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',sK12))))) ) | $spl0),
  inference(forward_demodulation,[],[f10959,f318])).
fof(f10959,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num',sK14))),s('type.num.num',sK13))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',sK12))))) ) | $spl0),
  inference(superposition,[],[f9744,f318])).
fof(f9744,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK14))),s('type.num.num',sK13))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK12))))) ) | $spl0),
  inference(forward_demodulation,[],[f9737,f317])).
fof(f9737,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK14))),s('type.num.num',sK13))),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK12))))) ) | $spl0),
  inference(backward_demodulation,[],[f9629,f473])).
fof(f473,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK14))),s('type.num.num',sK13))),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',sK12))))) )),
  inference(superposition,[],[f450,f320])).
fof(f450,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK14))),s('type.num.num',sK13))),s('type.num.num',sK12)))),
  inference(forward_demodulation,[],[f449,f318])).
fof(f449,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))),s('type.num.num',sK12))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK14))),s('type.num.num',sK13))),s('type.num.num',sK12)))),
  inference(forward_demodulation,[],[f255,f318])).
fof(f255,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))),s('type.num.num',sK12))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',sK12))),s('type.num.num',sK13))),s('type.num.num',sK12)))),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))),s('type.num.num',sK12))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',sK12))),s('type.num.num',sK13))),s('type.num.num',sK12)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f125])).
fof(f125,plain,(
  ? [X0,X1,X2] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num',X1))),s('type.num.num',X0)))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ~! [X0,X1,X2] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num',X1))),s('type.num.num',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X7))),s('type.num.num',X6))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X6))),s('type.num.num',X7))),s('type.num.num',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X7))),s('type.num.num',X6))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X6))),s('type.num.num',X7))),s('type.num.num',X6)))),
  file('thm.numeral.numeral_5|49|_prove_ltb','thm.numeral.numeral_5|49|_prove_ltb')).
fof(f21285,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) | ($spl0 | $spl7)),
  inference(resolution,[],[f21235,f356])).
fof(f356,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(ennf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ~(p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X7] : ~(p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8)))))),
  file('thm.numeral.numeral_5|49|_prove_ltb','thm.arithmetic.LESS_ANTISYM')).
fof(f21390,plain,(
  $false | ($spl0 | $spl7)),
  inference(subsumption_resolution,[],[f21284,f21235])).
fof(f21284,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK12),s('type.num.num',sK12)))) | ($spl0 | $spl7)),
  inference(resolution,[],[f21235,f356])).
fof(f11682,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl7),
  inference(cnf_transformation,[],[f11682_D])).
fof(f11682_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f11684,plain,(
  s('type.num.num',sK13) = s('type.num.num',sK14) | $spl8),
  inference(cnf_transformation,[],[f11684_D])).
fof(f11684_D,plain,(
  s('type.num.num',sK13) = s('type.num.num',sK14) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f9627,plain,(
  s('type.num.num',sK13) = s('type.num.num',sK14) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl0),
  inference(resolution,[],[f484,f320])).
fof(f9537,plain,(
  $false | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f486,f9494])).
fof(f9494,plain,(
  ( ! [X1] : (s('type.num.num',X1) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13)))))) ) | $spl1),
  inference(superposition,[],[f3079,f318])).
fof(f3079,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))),s('type.num.num',X0)))) ) | $spl1),
  inference(forward_demodulation,[],[f2873,f317])).
fof(f2873,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK14),s('type.num.num',X0))),s('type.num.num',sK13)))))) ) | $spl1),
  inference(superposition,[],[f358,f1342])).
fof(f1342,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13))) | $spl1),
  inference(resolution,[],[f1046,f273])).
fof(f1046,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl1),
  inference(resolution,[],[f734,f354])).
fof(f734,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) | $spl1),
  inference(resolution,[],[f604,f356])).
fof(f604,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl1),
  inference(resolution,[],[f483,f354])).
fof(f483,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) | $spl1),
  inference(cnf_transformation,[],[f483_D])).
fof(f483_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f602,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK14))),s('type.num.num',sK13))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))))) | $spl5),
  inference(cnf_transformation,[],[f602_D])).
fof(f602_D,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK14))),s('type.num.num',sK13))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f479,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK14))),s('type.num.num',sK13))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14))))),
  inference(forward_demodulation,[],[f453,f317])).
fof(f453,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK14))),s('type.num.num',sK13))),s('type.num.num',sK12))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14))))),
  inference(superposition,[],[f450,f273])).
fof(f484,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) | $spl0),
  inference(cnf_transformation,[],[f484_D])).
fof(f484_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f486,plain,(
  s('type.num.num',sK12) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))))) | $spl3),
  inference(cnf_transformation,[],[f486_D])).
fof(f486_D,plain,(
  s('type.num.num',sK12) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f481,plain,(
  s('type.num.num',sK12) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14))))),
  inference(forward_demodulation,[],[f455,f316])).
fof(f455,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK14),s('type.num.num',sK13))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK12),s('type.num.num',sK14))),s('type.num.num',sK13))),s('type.num.num',sK12))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK13),s('type.num.num',sK14))))),
  inference(superposition,[],[f450,f277])).
% SZS output end Proof for thm.numeral
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 770
% Passive clauses: 3486
% Generated clauses: 22356
% Final active clauses: 159
% Final passive clauses: 511
% Input formulas: 51
% Initial clauses: 184
% 
% Duplicate literals: 1402
% Trivial inequalities: 77
% Fw subsumption resolutions: 1258
% Bw subsumption resolutions: 24
% Fw demodulations: 3993
% Bw demodulations: 836
% Condensations: 1338
% 
% Simple tautologies: 2065
% Equational tautologies: 1
% Forward subsumptions: 8851
% Backward subsumptions: 201
% 
% Binary resolution: 7191
% Factoring: 1
% Forward superposition: 4075
% Backward superposition: 2067
% Self superposition: 22
% Equality resolution: 12
% 
% Split clauses: 3
% Split components: 6
% SAT solver clauses: 3497
% SAT solver unit clauses: 2
% SAT solver binary clauses: 161
% 
% Sat splits: 3
% Sat splitting refutations: 6
% 
% Memory used [KB]: 10234
% Time elapsed: 0.762 s
% ------------------------------
% ------------------------------
% Success in time 17.184 s
