Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.extreal
% (54034)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 749
% Passive clauses: 4778
% Generated clauses: 19354
% Final active clauses: 690
% Final passive clauses: 3503
% Input formulas: 36
% Initial clauses: 163
% 
% Duplicate literals: 870
% Trivial inequalities: 122
% Fw subsumption resolutions: 2774
% Fw demodulations: 401
% Bw demodulations: 39
% 
% Simple tautologies: 249
% Equational tautologies: 80
% Forward subsumptions: 8157
% 
% Binary resolution: 5412
% Factoring: 1286
% Forward superposition: 2926
% Backward superposition: 4606
% Self superposition: 120
% Equality factoring: 108
% Equality resolution: 26
% 
% Split clauses: 1818
% Split components: 3834
% SAT solver clauses: 2315
% SAT solver unit clauses: 35
% SAT solver binary clauses: 228
% 
% Sat splits: 1818
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2305
% TWLsolver calls for satisfiability: 501
% 
% Memory used [KB]: 9338
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (54083)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 386
% Passive clauses: 2068
% Generated clauses: 8683
% Final active clauses: 380
% Final passive clauses: 1682
% Input formulas: 36
% Initial clauses: 163
% 
% Duplicate literals: 182
% Trivial inequalities: 34
% Fw subsumption resolutions: 797
% Fw demodulations: 135
% Bw demodulations: 3
% Condensations: 196
% Global subsumptions: 13
% 
% Simple tautologies: 223
% Equational tautologies: 23
% Forward subsumptions: 2570
% 
% Binary resolution: 3422
% Factoring: 307
% Forward superposition: 881
% Backward superposition: 509
% Self superposition: 28
% Equality factoring: 26
% Equality resolution: 5
% 
% Split clauses: 83
% Split components: 171
% SAT solver clauses: 2593
% SAT solver unit clauses: 176
% SAT solver binary clauses: 80
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 83
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2541
% TWLsolver calls for satisfiability: 1747
% 
% Memory used [KB]: 11129
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 287 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f6783,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f209,f208,f381,f207,f379,f206,f385,f383,f1700,f1704,f1707,f1710,f1714,f1717,f1720,f1723,f1726,f1729,f1732,f1735,f1738,f1741,f1744,f1747,f1750,f1753,f1756,f1759,f1762,f1765,f1768,f1771,f1804,f1701,f1806,f1705,f1808,f1708,f1810,f1711,f1813,f1715,f1815,f1718,f1817,f1721,f1819,f1724,f1821,f1727,f1823,f1730,f1825,f1733,f1827,f1736,f1829,f1739,f1831,f1742,f1833,f1745,f1835,f1748,f1837,f1751,f1839,f1754,f1841,f1757,f1843,f1760,f1845,f1763,f1847,f1766,f1849,f1769,f1851,f1702,f1772,f1871,f1872,f1876,f1880,f1882,f1884,f1888,f1892,f2310,f2311,f2318,f2315,f2312,f2313,f2321,f2314,f2323,f2316,f2320,f2328,f2329,f2341,f2339,f3592,f3595,f3593,f3597,f3598,f3599,f3601,f3604,f3606,f3616,f3618,f3628,f3630,f3634,f3638,f3642,f3644,f3649,f3654,f3658,f3660,f3664,f3666,f3670,f3672,f3692,f3724,f3743,f3775,f3867,f3875,f3879,f3883,f3899,f3903,f3909,f3913,f3922,f3930,f3939,f3947,f3949,f3957,f3964,f3972,f3979,f3981,f3983,f3985,f3987,f4912,f4913,f4015,f4017,f4019,f4021,f4023,f4025,f4057,f4059,f4066,f4073,f4083,f4091,f4093,f4106,f4116,f4118,f4129,f4136,f4144,f4151,f4153,f4164,f4166,f4173,f4175,f4188,f4190,f4197,f4199,f4212,f4214,f4221,f4223,f4231,f4238,f4246,f4253,f4276,f4299,f4311,f4329,f4333,f4343,f4345,f4347,f4354,f4356,f4363,f4365,f4372,f4374,f4381,f4383,f4407,f4426,f4436,f4446,f4448,f4476,f4478,f4480,f4482,f4484,f4499,f4501,f4516,f4518,f4531,f4533,f4546,f4548,f4550,f4560,f4569,f4578,f4580,f4589,f4591,f4604,f4623,f4643,f4662,f4671,f4673,f4675,f4677,f4679,f4686,f4688,f4690,f4692,f4694,f4696,f4698,f4700,f4702,f4704,f4706,f4708,f4710,f4712,f4714,f4716,f4719,f4726,f4736,f4742,f4748,f4754,f4760,f4766,f4777,f4780,f4786,f4792,f4795,f4801,f4806,f4819,f4826,f4832,f4838,f4842,f4844,f4848,f4850,f4856,f4858,f4860,f4866,f4868,f4874,f4877,f4884,f4887,f4892,f4895,f4914,f4915,f4916,f4917,f4918,f4919,f4920,f4921,f4922,f4923,f4924,f4925,f4926,f4927,f4928,f4929,f4930,f4931,f4932,f4933,f4934,f4935,f4936,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4950,f4951,f4952,f4953,f4954,f4955,f4956,f4957,f4958,f4959,f4960,f4961,f4962,f4963,f4964,f4965,f4966,f4967,f4968,f4969,f4970,f4971,f4972,f4973,f4974,f4975,f4976,f4977,f4978,f4979,f4980,f4981,f4982,f4983,f4984,f4985,f4986,f4987,f4988,f4989,f4990,f4991,f4992,f4993,f4994,f4995,f4996,f4997,f4998,f4999,f5000,f5001,f5002,f5003,f5004,f5005,f5006,f5007,f5008,f5009,f5010,f5011,f5012,f5013,f5014,f5017,f5019,f5021,f5023,f5025,f5027,f5028,f5029,f5036,f5037,f5038,f5039,f5040,f5041,f5048,f5049,f5050,f5051,f5052,f5053,f5054,f5055,f5056,f5057,f5058,f5059,f5060,f5061,f5062,f5063,f5064,f5065,f5066,f5067,f5068,f5069,f5070,f5071,f5072,f5073,f5074,f5075,f5076,f5077,f5078,f5079,f5080,f5081,f5082,f5083,f5084,f5085,f5086,f5087,f5088,f5089,f5090,f5091,f5092,f5093,f5094,f5095,f5096,f5097,f5098,f5099,f5100,f5101,f5102,f5103,f5104,f5105,f5106,f5107,f5108,f5109,f5110,f5111,f5112,f5113,f5114,f5115,f5126,f5127,f5837,f5839,f5838,f5841,f3594,f5842,f6244,f6247,f6245,f6449,f6778])).
fof(f6778,plain,(
  $false | ($spl0 | $spl4)),
  inference(subsumption_resolution,[],[f6454,f383])).
fof(f6454,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f389,f379,f283])).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X2)))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X2))))) & (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(ennf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : (((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X2))))) | (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) => s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6,X7,X8] : (((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X7)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X8))))) | (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) => s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X8))))))),
  file('thm.extreal.add_rdistrib','thm.extreal.add_ldistrib')).
fof(f389,plain,(
  s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK15)))))),
  inference(forward_demodulation,[],[f388,f252])).
fof(f252,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X7] : s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))) = s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6)))),
  file('thm.extreal.add_rdistrib','thm.extreal.mul_comm')).
fof(f388,plain,(
  s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))),s('type.extreal.extreal',sK13))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK15)))))),
  inference(forward_demodulation,[],[f387,f252])).
fof(f387,plain,(
  s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))),s('type.extreal.extreal',sK13))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK13))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK15)))))),
  inference(backward_demodulation,[],[f252,f210])).
fof(f210,plain,(
  s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))),s('type.extreal.extreal',sK13))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK13))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal',sK13)))))),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK14)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK15))))) | (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK15))),s('type.extreal.extreal',sK13))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',sK13))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal',sK13)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15])],[f90])).
fof(f90,plain,(
  ? [X0,X1,X2] : (((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X2))))) | (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))),s('type.extreal.extreal',X0))) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ~! [X0,X1,X2] : (((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X2))))) | (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) => s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))),s('type.extreal.extreal',X0))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X7)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X8))))) | (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) => s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8))),s('type.extreal.extreal',X6))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X6))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X7)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X8))))) | (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) => s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X8))),s('type.extreal.extreal',X6))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X8),s('type.extreal.extreal',X6))))))),
  file('thm.extreal.add_rdistrib','thm.extreal.add_rdistrib')).
fof(f6449,plain,(
  $false | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f6448,f211])).
fof(f211,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.add_rdistrib','thm.bool.TRUTH')).
fof(f6448,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f6438,f439])).
fof(f439,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(unit_resulting_resolution,[],[f381,f226])).
fof(f226,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.extreal.add_rdistrib','thm.bool.EQ_CLAUSES|split|1')).
fof(f6438,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl6),
  inference(unit_resulting_resolution,[],[f389,f1293])).
fof(f1293,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',sK14))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X2)))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1290,f211])).
fof(f1290,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',sK14))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X2)))))) ) | $spl6),
  inference(superposition,[],[f284,f440])).
fof(f440,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl6),
  inference(unit_resulting_resolution,[],[f385,f226])).
fof(f284,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f6245,plain,(
  ( ! [X158,X159] : (s(fun(bool,bool),X158) = s(fun(bool,bool),X159)) ) | $spl84),
  inference(cnf_transformation,[],[f6245_D])).
fof(f6245_D,plain,(
  ( ! [X158,X159] : (s(fun(bool,bool),X158) = s(fun(bool,bool),X159)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f6247,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | $spl87),
  inference(cnf_transformation,[],[f6247_D])).
fof(f6247_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f6244,plain,(
  ( ! [X158,X159] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X158) = s(fun(bool,bool),X159)) ) | ($spl60 | $spl73)),
  inference(forward_demodulation,[],[f6204,f2399])).
fof(f2399,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl73),
  inference(unit_resulting_resolution,[],[f436,f2341,f680])).
fof(f680,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f218,f218])).
fof(f218,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.extreal.add_rdistrib',aHL_BOOL_CASES)).
fof(f436,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f386,f225])).
fof(f225,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f139])).
fof(f386,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f214,f211])).
fof(f214,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f15])).
fof(f15,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.extreal.add_rdistrib','thm.bool.NOT_CLAUSES|split|2')).
fof(f6204,plain,(
  ( ! [X158,X159] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X158) = s(fun(bool,bool),X159)) ) | $spl60),
  inference(duplicate_literal_removal,[],[f6203])).
fof(f6203,plain,(
  ( ! [X158,X159] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X158) = s(fun(bool,bool),X159) | s(fun(bool,bool),X158) = s(fun(bool,bool),X159)) ) | $spl60),
  inference(superposition,[],[f282,f2316])).
fof(f282,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK16(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK16(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK16(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK16(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.extreal.add_rdistrib',aHL_EXT)).
fof(f5842,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl82),
  inference(cnf_transformation,[],[f5842_D])).
fof(f5842_D,plain,(
  ( ! [X28,X29] : (sP1('F_0',X28) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f3594,plain,(
  ( ! [X9] : (p(s(bool,X9))) ) | $spl76),
  inference(cnf_transformation,[],[f3594_D])).
fof(f3594_D,plain,(
  ( ! [X9] : (p(s(bool,X9))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f5841,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f5833,f225])).
fof(f5833,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1('F_0',X28) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f1777,f218])).
fof(f1777,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X20)) | sP1(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1477,f211])).
fof(f1477,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | p(s(bool,X20)) | sP1(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f277,f680])).
fof(f277,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f162])).
fof(f162,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f5838,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl80),
  inference(cnf_transformation,[],[f5838_D])).
fof(f5838_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X8,X10) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f5839,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f5829,f225])).
fof(f5829,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X12,X14) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1777,f680])).
fof(f5837,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f5828,f225])).
fof(f5828,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X8,X10) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1777,f680])).
fof(f5127,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f5118,f211])).
fof(f5118,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(backward_demodulation,[],[f3600,f386])).
fof(f3600,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f3594,f226])).
fof(f5126,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f5125,f3594])).
fof(f5125,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(subsumption_resolution,[],[f5116,f3600])).
fof(f5116,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl76),
  inference(backward_demodulation,[],[f3600,f229])).
fof(f229,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.extreal.add_rdistrib','thm.bool.EQ_CLAUSES|split|2')).
fof(f5115,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3608,f211])).
fof(f3608,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f482,f3594,f273])).
fof(f273,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f163])).
fof(f482,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f386,f211,f277])).
fof(f5114,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3609,f3594])).
fof(f3609,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f528,f3594,f273])).
fof(f528,plain,(
  sP1('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(unit_resulting_resolution,[],[f386,f485])).
fof(f485,plain,(
  ( ! [X0] : (sP1(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f483,f211])).
fof(f483,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl2),
  inference(superposition,[],[f277,f439])).
fof(f5113,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3610,f3594])).
fof(f3610,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f535,f3594,f273])).
fof(f535,plain,(
  sP1('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl6),
  inference(unit_resulting_resolution,[],[f386,f486])).
fof(f486,plain,(
  ( ! [X1] : (sP1(X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f484,f211])).
fof(f484,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP1(X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl6),
  inference(superposition,[],[f277,f440])).
fof(f5112,plain,(
  $false | ($spl73 | $spl76)),
  inference(subsumption_resolution,[],[f3611,f3594])).
fof(f3611,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f3140,f3594,f273])).
fof(f3140,plain,(
  sP1('F_0',f(s(bool,'T_0'))) | $spl73),
  inference(unit_resulting_resolution,[],[f386,f2760])).
fof(f2760,plain,(
  ( ! [X3] : (sP1(X3,f(s(bool,'T_0'))) | p(s(bool,X3))) ) | $spl73),
  inference(subsumption_resolution,[],[f2711,f211])).
fof(f2711,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | sP1(X3,f(s(bool,'T_0')))) ) | $spl73),
  inference(superposition,[],[f277,f2399])).
fof(f5111,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3619,f211])).
fof(f3619,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f475,f3594,f273])).
fof(f475,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f386,f211,f274])).
fof(f274,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f163])).
fof(f5110,plain,(
  $false | ($spl73 | $spl76)),
  inference(subsumption_resolution,[],[f3624,f3594])).
fof(f3624,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f3112,f3594,f273])).
fof(f3112,plain,(
  sP1(f(s(bool,'T_0')),'F_0') | $spl73),
  inference(unit_resulting_resolution,[],[f386,f2759])).
fof(f2759,plain,(
  ( ! [X2] : (sP1(f(s(bool,'T_0')),X2) | p(s(bool,X2))) ) | $spl73),
  inference(subsumption_resolution,[],[f2710,f211])).
fof(f2710,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP1(f(s(bool,'T_0')),X2)) ) | $spl73),
  inference(superposition,[],[f274,f2399])).
fof(f5109,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3625,f3594])).
fof(f3625,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f504,f3594,f273])).
fof(f504,plain,(
  sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f386,f478])).
fof(f478,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f476,f211])).
fof(f476,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0)) ) | $spl2),
  inference(superposition,[],[f274,f439])).
fof(f5108,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3626,f3594])).
fof(f3626,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f521,f3594,f273])).
fof(f521,plain,(
  sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f386,f479])).
fof(f479,plain,(
  ( ! [X1] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1) | p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f477,f211])).
fof(f477,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl6),
  inference(superposition,[],[f274,f440])).
fof(f5107,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3640,f3594])).
fof(f3640,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f517,f3594,f288])).
fof(f288,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f167])).
fof(f167,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f517,plain,(
  ( ! [X0] : (~sP2('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f386,f386,f286])).
fof(f286,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f168])).
fof(f5106,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3655,f3594])).
fof(f3655,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f508,f3594,f289])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f168])).
fof(f508,plain,(
  ( ! [X0] : (~sP2(X0,'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f386,f386,f285])).
fof(f285,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f168])).
fof(f5105,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3662,f3594])).
fof(f3662,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f517,f3594,f290])).
fof(f290,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f168])).
fof(f5104,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3667,f3594])).
fof(f3667,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f508,f3594,f290])).
fof(f5103,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3673,f211])).
fof(f3673,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f735,f386,f3594,f298])).
fof(f298,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (sP3(X2,X1,X0) & sP4(X1,X2,X0) & sP5(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP3(X2,X1,X0) | ~sP4(X1,X2,X0) | ~sP5(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f172])).
fof(f172,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP6(X2,X1,X0)))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP6(X2,X1,X0)))),
  inference(nnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X2,X1,X0] : (sP6(X2,X1,X0) <=> (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f735,plain,(
  sP6('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f386,f436,f314])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP6(X2,X1,X0)) & (~sP6(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP6(X2,X1,X0))),
  inference(definition_folding,[],[f83,f125,f124,f123,f122])).
fof(f122,plain,(
  ! [X0,X1,X2] : (sP3(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f123,plain,(
  ! [X1,X0,X2] : (sP4(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f124,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f83,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X13,X14,X15] : ((p(s(bool,X15)) <=> s(bool,X13) = s(bool,X14)) <=> ((p(s(bool,X15)) | p(s(bool,X14)) | p(s(bool,X13))) & (p(s(bool,X15)) | ~p(s(bool,X13)) | ~p(s(bool,X14))) & (p(s(bool,X14)) | ~p(s(bool,X13)) | ~p(s(bool,X15))) & (p(s(bool,X13)) | ~p(s(bool,X14)) | ~p(s(bool,X15)))))),
  file('thm.extreal.add_rdistrib','thm.sat.sat_2|49|_prove_ltb')).
fof(f5102,plain,(
  $false | ($spl73 | $spl76)),
  inference(subsumption_resolution,[],[f3674,f3594])).
fof(f3674,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f2398,f386,f3594,f298])).
fof(f2398,plain,(
  sP6('F_0',f(s(bool,'T_0')),'F_0') | $spl73),
  inference(unit_resulting_resolution,[],[f386,f2341,f314])).
fof(f5101,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3675,f3594])).
fof(f3675,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f1245,f386,f3594,f298])).
fof(f1245,plain,(
  sP6('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f411,f405,f400,f386,f301])).
fof(f301,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X2,X1) | p(s(bool,X0)) | sP6(X0,X1,X2) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f400,plain,(
  ( ! [X0,X1] : (sP5(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f386,f304])).
fof(f304,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f175])).
fof(f175,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f124])).
fof(f405,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f386,f308])).
fof(f308,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f178])).
fof(f178,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(nnf_transformation,[],[f123])).
fof(f411,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f381,f312])).
fof(f312,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(nnf_transformation,[],[f122])).
fof(f5100,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3676,f3594])).
fof(f3676,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f1246,f386,f3594,f298])).
fof(f1246,plain,(
  sP6('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f412,f405,f400,f386,f301])).
fof(f412,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f385,f312])).
fof(f5099,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3729,f211])).
fof(f3729,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f1193,f386,f3594,f298])).
fof(f1193,plain,(
  sP6('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f407,f402,f400,f386,f300])).
fof(f300,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X2,X1) | p(s(bool,X1)) | sP6(X0,X1,X2) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f402,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f211,f307])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f179])).
fof(f407,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f211,f311])).
fof(f311,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f181])).
fof(f5098,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3730,f3594])).
fof(f3730,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f1194,f386,f3594,f298])).
fof(f1194,plain,(
  sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f408,f403,f400,f386,f300])).
fof(f403,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f381,f307])).
fof(f408,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f381,f311])).
fof(f5097,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3731,f3594])).
fof(f3731,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f1195,f386,f3594,f298])).
fof(f1195,plain,(
  sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f409,f404,f400,f386,f300])).
fof(f404,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f385,f307])).
fof(f409,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f385,f311])).
fof(f5096,plain,(
  $false | ($spl16 | $spl76)),
  inference(subsumption_resolution,[],[f3791,f4818])).
fof(f4818,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','F_0')) ) | ($spl16 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f3594,f1711])).
fof(f3791,plain,(
  ( ! [X0] : (sP6(X0,'T_0','F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f410,f405,f397,f3594,f299])).
fof(f299,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP6(X0,X1,X2) | ~sP5(X0,X2,X1) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f397,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f211,f303])).
fof(f303,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f176])).
fof(f410,plain,(
  ( ! [X0,X1] : (sP3(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f211,f312])).
fof(f5095,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3810,f4717])).
fof(f4717,plain,(
  ( ! [X0] : (~sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')) ) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f3594,f796])).
fof(f796,plain,(
  ( ! [X0,X1] : (~sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f790,f211])).
fof(f790,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl2),
  inference(superposition,[],[f298,f439])).
fof(f3810,plain,(
  ( ! [X0] : (sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')) ) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f411,f405,f398,f3594,f299])).
fof(f398,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f381,f303])).
fof(f5094,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3829,f4724])).
fof(f4724,plain,(
  ( ! [X0] : (~sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')) ) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f3594,f797])).
fof(f797,plain,(
  ( ! [X2,X3] : (~sP6(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3) | p(s(bool,X3)) | ~p(s(bool,X2))) ) | $spl6),
  inference(subsumption_resolution,[],[f791,f211])).
fof(f791,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP6(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3)) ) | $spl6),
  inference(superposition,[],[f298,f440])).
fof(f3829,plain,(
  ( ! [X0] : (sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')) ) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f412,f405,f399,f3594,f299])).
fof(f399,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f385,f303])).
fof(f5093,plain,(
  $false | ($spl16 | $spl76)),
  inference(subsumption_resolution,[],[f3837,f4818])).
fof(f3837,plain,(
  ( ! [X0] : (sP6(X0,'T_0','F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f410,f405,f400,f3594,f299])).
fof(f5092,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3838,f4717])).
fof(f3838,plain,(
  ( ! [X0] : (sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')) ) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f411,f405,f400,f3594,f299])).
fof(f5091,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3839,f4724])).
fof(f3839,plain,(
  ( ! [X0] : (sP6(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')) ) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f412,f405,f400,f3594,f299])).
fof(f5090,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3841,f757])).
fof(f757,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f211,f436,f316])).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | s(bool,X0) = s(bool,X1) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f182])).
fof(f3841,plain,(
  sP6('T_0','T_0','F_0') | $spl76),
  inference(unit_resulting_resolution,[],[f410,f402,f400,f3594,f299])).
fof(f5089,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3842,f1037])).
fof(f1037,plain,(
  ~sP6('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f211,f386,f796])).
fof(f3842,plain,(
  sP6('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f411,f402,f400,f3594,f299])).
fof(f5088,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3843,f1075])).
fof(f1075,plain,(
  ~sP6('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f211,f386,f797])).
fof(f3843,plain,(
  sP6('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f412,f402,f400,f3594,f299])).
fof(f5087,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3845,f1014])).
fof(f1014,plain,(
  ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f436,f762])).
fof(f762,plain,(
  ( ! [X0,X1] : (~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,X0) | s(bool,X0) = s(bool,X1)) ) | $spl2),
  inference(subsumption_resolution,[],[f758,f211])).
fof(f758,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,X1) | ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,X0)) ) | $spl2),
  inference(superposition,[],[f316,f439])).
fof(f3845,plain,(
  sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0') | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f410,f403,f400,f3594,f299])).
fof(f5086,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3846,f4717])).
fof(f3846,plain,(
  sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f411,f403,f400,f3594,f299])).
fof(f5085,plain,(
  $false | ($spl2 | $spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3847,f4724])).
fof(f3847,plain,(
  sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | ($spl2 | $spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f412,f403,f400,f3594,f299])).
fof(f5084,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3849,f1024])).
fof(f1024,plain,(
  ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f436,f763])).
fof(f763,plain,(
  ( ! [X2,X3] : (~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3,X2) | s(bool,X2) = s(bool,X3)) ) | $spl6),
  inference(subsumption_resolution,[],[f759,f211])).
fof(f759,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s(bool,X2) = s(bool,X3) | ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3,X2)) ) | $spl6),
  inference(superposition,[],[f316,f440])).
fof(f3849,plain,(
  sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','F_0') | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f410,f404,f400,f3594,f299])).
fof(f5083,plain,(
  $false | ($spl2 | $spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3850,f4717])).
fof(f3850,plain,(
  sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | ($spl2 | $spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f411,f404,f400,f3594,f299])).
fof(f5082,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3851,f4724])).
fof(f3851,plain,(
  sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f412,f404,f400,f3594,f299])).
fof(f5081,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3864,f3948])).
fof(f3948,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f3594,f303])).
fof(f3864,plain,(
  ~sP5('T_0','F_0','T_0') | $spl76),
  inference(unit_resulting_resolution,[],[f410,f405,f757,f3594,f299])).
fof(f5080,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3865,f3980])).
fof(f3980,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f3594,f307])).
fof(f3865,plain,(
  ~sP4('T_0','F_0','T_0') | $spl76),
  inference(unit_resulting_resolution,[],[f410,f400,f757,f3594,f299])).
fof(f5079,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3866,f3986])).
fof(f3986,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f3594,f313])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f3866,plain,(
  ~sP3('F_0','T_0','T_0') | $spl76),
  inference(unit_resulting_resolution,[],[f405,f400,f757,f3594,f299])).
fof(f5078,plain,(
  $false | ($spl73 | $spl76)),
  inference(subsumption_resolution,[],[f3872,f3948])).
fof(f3872,plain,(
  ~sP5('T_0','F_0',f(s(bool,'T_0'))) | ($spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f407,f405,f2397,f3594,f299])).
fof(f2397,plain,(
  ~sP6('T_0',f(s(bool,'T_0')),'F_0') | $spl73),
  inference(unit_resulting_resolution,[],[f211,f2341,f316])).
fof(f5077,plain,(
  $false | ($spl73 | $spl76)),
  inference(subsumption_resolution,[],[f3873,f3980])).
fof(f3873,plain,(
  ~sP4(f(s(bool,'T_0')),'F_0','T_0') | ($spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f407,f400,f2397,f3594,f299])).
fof(f5076,plain,(
  $false | ($spl73 | $spl76)),
  inference(subsumption_resolution,[],[f3874,f3986])).
fof(f3874,plain,(
  ~sP3('F_0',f(s(bool,'T_0')),'T_0') | ($spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f405,f400,f2397,f3594,f299])).
fof(f5075,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3876,f3948])).
fof(f3876,plain,(
  ~sP5('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f411,f405,f1037,f3594,f299])).
fof(f5074,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3877,f3980])).
fof(f3877,plain,(
  ~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0') | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f411,f400,f1037,f3594,f299])).
fof(f5073,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3878,f3986])).
fof(f3878,plain,(
  ~sP3('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f405,f400,f1037,f3594,f299])).
fof(f5072,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3880,f3948])).
fof(f3880,plain,(
  ~sP5('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f412,f405,f1075,f3594,f299])).
fof(f5071,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3881,f3980])).
fof(f3881,plain,(
  ~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0') | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f412,f400,f1075,f3594,f299])).
fof(f5070,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3882,f3986])).
fof(f3882,plain,(
  ~sP3('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f405,f400,f1075,f3594,f299])).
fof(f5069,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3896,f3948])).
fof(f3896,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0') | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f410,f405,f1014,f3594,f299])).
fof(f5068,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3897,f3980])).
fof(f3897,plain,(
  ~sP4('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f410,f400,f1014,f3594,f299])).
fof(f5067,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3898,f3986])).
fof(f3898,plain,(
  ~sP3('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f405,f400,f1014,f3594,f299])).
fof(f5066,plain,(
  $false | ($spl2 | $spl73 | $spl76)),
  inference(subsumption_resolution,[],[f3900,f3948])).
fof(f3900,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0',f(s(bool,'T_0'))) | ($spl2 | $spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f408,f405,f2396,f3594,f299])).
fof(f2396,plain,(
  ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),f(s(bool,'T_0')),'F_0') | ($spl2 | $spl73)),
  inference(unit_resulting_resolution,[],[f2341,f762])).
fof(f5065,plain,(
  $false | ($spl2 | $spl73 | $spl76)),
  inference(subsumption_resolution,[],[f3901,f3980])).
fof(f3901,plain,(
  ~sP4(f(s(bool,'T_0')),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f408,f400,f2396,f3594,f299])).
fof(f5064,plain,(
  $false | ($spl2 | $spl73 | $spl76)),
  inference(subsumption_resolution,[],[f3902,f3986])).
fof(f3902,plain,(
  ~sP3('F_0',f(s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f405,f400,f2396,f3594,f299])).
fof(f5063,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3906,f3948])).
fof(f3906,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0') | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f410,f405,f1024,f3594,f299])).
fof(f5062,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3907,f3980])).
fof(f3907,plain,(
  ~sP4('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f410,f400,f1024,f3594,f299])).
fof(f5061,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3908,f3986])).
fof(f3908,plain,(
  ~sP3('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f405,f400,f1024,f3594,f299])).
fof(f5060,plain,(
  $false | ($spl6 | $spl73 | $spl76)),
  inference(subsumption_resolution,[],[f3910,f3948])).
fof(f3910,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0',f(s(bool,'T_0'))) | ($spl6 | $spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f409,f405,f2395,f3594,f299])).
fof(f2395,plain,(
  ~sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),f(s(bool,'T_0')),'F_0') | ($spl6 | $spl73)),
  inference(unit_resulting_resolution,[],[f2341,f763])).
fof(f5059,plain,(
  $false | ($spl6 | $spl73 | $spl76)),
  inference(subsumption_resolution,[],[f3911,f3980])).
fof(f3911,plain,(
  ~sP4(f(s(bool,'T_0')),'F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl6 | $spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f409,f400,f2395,f3594,f299])).
fof(f5058,plain,(
  $false | ($spl6 | $spl73 | $spl76)),
  inference(subsumption_resolution,[],[f3912,f3986])).
fof(f3912,plain,(
  ~sP3('F_0',f(s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl6 | $spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f405,f400,f2395,f3594,f299])).
fof(f5057,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3915,f3594])).
fof(f3915,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f401,f386,f3594,f302])).
fof(f302,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f176])).
fof(f401,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f386,f305])).
fof(f305,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f176])).
fof(f5056,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3929,f3948])).
fof(f3929,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f386,f3594,f3594,f302])).
fof(f5055,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3931,f3594])).
fof(f3931,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f400,f386,f3594,f302])).
fof(f5054,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3938,f3948])).
fof(f3938,plain,(
  ( ! [X0] : (~sP5('T_0',X0,'F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f386,f211,f3594,f302])).
fof(f5053,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3946,f3948])).
fof(f3946,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f386,f3594,f3594,f302])).
fof(f5052,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3951,f3594])).
fof(f3951,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f406,f386,f3594,f306])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f179])).
fof(f406,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f386,f309])).
fof(f309,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f179])).
fof(f5051,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3963,f3980])).
fof(f3963,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f386,f3594,f3594,f306])).
fof(f5050,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3965,f3594])).
fof(f3965,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f405,f386,f3594,f306])).
fof(f5049,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3971,f3980])).
fof(f3971,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f386,f211,f3594,f306])).
fof(f5048,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f3978,f3980])).
fof(f3978,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f386,f3594,f3594,f306])).
fof(f5041,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f3996,f439])).
fof(f3996,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f774,f3594,f315])).
fof(f315,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f774,plain,(
  ~sP6('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(unit_resulting_resolution,[],[f386,f439,f317])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f5040,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f3997,f440])).
fof(f3997,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f775,f3594,f315])).
fof(f775,plain,(
  ~sP6('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl6),
  inference(unit_resulting_resolution,[],[f386,f440,f317])).
fof(f5039,plain,(
  $false | ($spl73 | $spl76)),
  inference(subsumption_resolution,[],[f3999,f2399])).
fof(f3999,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f2692,f3594,f315])).
fof(f2692,plain,(
  ~sP6('F_0','T_0',f(s(bool,'T_0'))) | $spl73),
  inference(unit_resulting_resolution,[],[f386,f2399,f317])).
fof(f5038,plain,(
  $false | ($spl73 | $spl76)),
  inference(subsumption_resolution,[],[f4001,f2399])).
fof(f4001,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f2690,f3594,f315])).
fof(f2690,plain,(
  ~sP6('F_0',f(s(bool,'T_0')),'T_0') | $spl73),
  inference(unit_resulting_resolution,[],[f386,f2399,f317])).
fof(f5037,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f4002,f439])).
fof(f4002,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f770,f3594,f315])).
fof(f770,plain,(
  ~sP6('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f386,f439,f317])).
fof(f5036,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f4003,f440])).
fof(f4003,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f772,f3594,f315])).
fof(f772,plain,(
  ~sP6('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl6),
  inference(unit_resulting_resolution,[],[f386,f440,f317])).
fof(f5029,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4026,f436])).
fof(f4026,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl76),
  inference(unit_resulting_resolution,[],[f735,f3594,f316])).
fof(f5028,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4027,f436])).
fof(f4027,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl76),
  inference(unit_resulting_resolution,[],[f734,f3594,f316])).
fof(f734,plain,(
  sP6('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f386,f436,f314])).
fof(f5027,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f5026,f436])).
fof(f5026,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl76)),
  inference(forward_demodulation,[],[f4028,f439])).
fof(f4028,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f1198,f3594,f316])).
fof(f1198,plain,(
  sP6('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(unit_resulting_resolution,[],[f414,f406,f401,f386,f300])).
fof(f414,plain,(
  ( ! [X0,X1] : (sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f381,f313])).
fof(f5025,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f5024,f436])).
fof(f5024,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl76)),
  inference(forward_demodulation,[],[f4029,f440])).
fof(f4029,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f1199,f3594,f316])).
fof(f1199,plain,(
  sP6('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl6),
  inference(unit_resulting_resolution,[],[f415,f406,f401,f386,f300])).
fof(f415,plain,(
  ( ! [X0,X1] : (sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f385,f313])).
fof(f5023,plain,(
  $false | ($spl73 | $spl76)),
  inference(subsumption_resolution,[],[f5022,f436])).
fof(f5022,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl73 | $spl76)),
  inference(forward_demodulation,[],[f4030,f2399])).
fof(f4030,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f2403,f3594,f316])).
fof(f2403,plain,(
  sP6('F_0','F_0',f(s(bool,'T_0'))) | $spl73),
  inference(unit_resulting_resolution,[],[f386,f2341,f314])).
fof(f5021,plain,(
  $false | ($spl73 | $spl76)),
  inference(subsumption_resolution,[],[f5020,f436])).
fof(f5020,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl73 | $spl76)),
  inference(forward_demodulation,[],[f4031,f2399])).
fof(f4031,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | ($spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f2398,f3594,f316])).
fof(f5019,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f5018,f436])).
fof(f5018,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl76)),
  inference(forward_demodulation,[],[f4032,f439])).
fof(f4032,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f1245,f3594,f316])).
fof(f5017,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f5016,f436])).
fof(f5016,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl6 | $spl76)),
  inference(forward_demodulation,[],[f4033,f440])).
fof(f4033,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f1246,f3594,f316])).
fof(f5014,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4065,f4222])).
fof(f4222,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f3594,f3594,f3594,f331])).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | sP8(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP8(X2,X0,X1))),
  inference(definition_folding,[],[f85,f128,f127])).
fof(f127,plain,(
  ! [X0,X1,X2] : (sP7(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f128,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f85,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X13,X14,X15] : ((p(s(bool,X15)) <=> (p(s(bool,X14)) & p(s(bool,X13)))) <=> ((p(s(bool,X15)) | ~p(s(bool,X14)) | ~p(s(bool,X13))) & (p(s(bool,X14)) | ~p(s(bool,X15))) & (p(s(bool,X13)) | ~p(s(bool,X15)))))),
  file('thm.extreal.add_rdistrib','thm.sat.sat_23_prove_ltb')).
fof(f4065,plain,(
  ( ! [X0,X1] : (~sP8(X0,X1,'F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f386,f3594,f319])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (sP7(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP7(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f184])).
fof(f184,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f128])).
fof(f5013,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4072,f4222])).
fof(f4072,plain,(
  ( ! [X0,X1] : (~sP8(X0,'F_0',X1)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f386,f3594,f320])).
fof(f320,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f185])).
fof(f5012,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4075,f3594])).
fof(f4075,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f420,f550,f3594,f321])).
fof(f321,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f185])).
fof(f550,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f211,f386,f320])).
fof(f420,plain,(
  ( ! [X0,X1] : (sP7('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f386,f328])).
fof(f328,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(nnf_transformation,[],[f127])).
fof(f5011,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4076,f3594])).
fof(f4076,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f416,f550,f3594,f321])).
fof(f416,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f211,f326])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f187])).
fof(f5010,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4082,f4152])).
fof(f4152,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f3594,f326])).
fof(f4082,plain,(
  ~sP7('F_0','T_0','T_0') | $spl76),
  inference(unit_resulting_resolution,[],[f550,f211,f3594,f321])).
fof(f5009,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4090,f4152])).
fof(f4090,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f547,f3594,f3594,f321])).
fof(f547,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f211,f386,f319])).
fof(f5008,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4092,f4152])).
fof(f4092,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f550,f3594,f3594,f321])).
fof(f5007,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4097,f3594])).
fof(f4097,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f419,f547,f3594,f321])).
fof(f419,plain,(
  ( ! [X0,X1] : (sP7(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f386,f327])).
fof(f327,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f187])).
fof(f5006,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4098,f3594])).
fof(f4098,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f416,f547,f3594,f321])).
fof(f5005,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4105,f4152])).
fof(f4105,plain,(
  ~sP7('T_0','F_0','T_0') | $spl76),
  inference(unit_resulting_resolution,[],[f547,f211,f3594,f321])).
fof(f5004,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4115,f4152])).
fof(f4115,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f547,f3594,f3594,f321])).
fof(f5003,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4117,f4152])).
fof(f4117,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f550,f3594,f3594,f321])).
fof(f5002,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4123,f3594])).
fof(f4123,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f420,f386,f3594,f325])).
fof(f325,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f187])).
fof(f5001,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4135,f4152])).
fof(f4135,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f386,f3594,f3594,f325])).
fof(f5000,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4137,f3594])).
fof(f4137,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f419,f386,f3594,f325])).
fof(f4999,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4143,f4152])).
fof(f4143,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f386,f211,f3594,f325])).
fof(f4998,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4150,f4152])).
fof(f4150,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f386,f3594,f3594,f325])).
fof(f4997,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4154,f211])).
fof(f4154,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f550,f211,f3594,f331])).
fof(f4996,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4156,f211])).
fof(f4156,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f547,f3594,f3594,f331])).
fof(f4995,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4157,f211])).
fof(f4157,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f550,f3594,f3594,f331])).
fof(f4994,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4162,f3594])).
fof(f4162,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f550,f211,f3594,f331])).
fof(f4993,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4168,f3594])).
fof(f4168,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f550,f3594,f3594,f331])).
fof(f4992,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4169,f211])).
fof(f4169,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f889,f3594,f3594,f331])).
fof(f889,plain,(
  ~sP8('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f877,f318])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f185])).
fof(f877,plain,(
  ~sP7('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f386,f211,f211,f325])).
fof(f4991,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f4170,f3594])).
fof(f4170,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f2501,f3594,f3594,f331])).
fof(f2501,plain,(
  ~sP8('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(unit_resulting_resolution,[],[f2491,f318])).
fof(f2491,plain,(
  ~sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f211,f386,f884])).
fof(f884,plain,(
  ( ! [X0,X1] : (~sP7(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f878,f211])).
fof(f878,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP7(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl2),
  inference(superposition,[],[f325,f439])).
fof(f4990,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f4171,f3594])).
fof(f4171,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f2517,f3594,f3594,f331])).
fof(f2517,plain,(
  ~sP8('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl6),
  inference(unit_resulting_resolution,[],[f2507,f318])).
fof(f2507,plain,(
  ~sP7('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f211,f386,f885])).
fof(f885,plain,(
  ( ! [X2,X3] : (~sP7(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3) | p(s(bool,X3)) | ~p(s(bool,X2))) ) | $spl6),
  inference(subsumption_resolution,[],[f879,f211])).
fof(f879,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP7(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3)) ) | $spl6),
  inference(superposition,[],[f325,f440])).
fof(f4989,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4176,f3594])).
fof(f4176,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f547,f211,f3594,f331])).
fof(f4988,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4178,f3594])).
fof(f4178,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f547,f3594,f3594,f331])).
fof(f4987,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4180,f211])).
fof(f4180,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f889,f3594,f3594,f331])).
fof(f4986,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f4181,f211])).
fof(f4181,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f2501,f3594,f3594,f331])).
fof(f4985,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f4182,f211])).
fof(f4182,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f2517,f3594,f3594,f331])).
fof(f4984,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4183,f211])).
fof(f4183,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f547,f211,f3594,f331])).
fof(f4983,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4191,f211])).
fof(f4191,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f547,f3594,f3594,f331])).
fof(f4982,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4192,f211])).
fof(f4192,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f550,f3594,f3594,f331])).
fof(f4981,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4201,f211])).
fof(f4201,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f889,f211,f3594,f331])).
fof(f4980,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4202,f3594])).
fof(f4202,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f547,f3594,f3594,f331])).
fof(f4979,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4204,f211])).
fof(f4204,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f889,f3594,f3594,f331])).
fof(f4978,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f4205,f211])).
fof(f4205,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f2501,f3594,f3594,f331])).
fof(f4977,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f4206,f211])).
fof(f4206,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f2517,f3594,f3594,f331])).
fof(f4976,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4208,f211])).
fof(f4208,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f889,f211,f3594,f331])).
fof(f4975,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f4209,f3594])).
fof(f4209,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f2501,f211,f3594,f331])).
fof(f4974,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f4210,f3594])).
fof(f4210,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f2517,f211,f3594,f331])).
fof(f4973,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4216,f3594])).
fof(f4216,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f550,f3594,f3594,f331])).
fof(f4972,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4217,f211])).
fof(f4217,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f889,f3594,f3594,f331])).
fof(f4971,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f4218,f3594])).
fof(f4218,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f2501,f3594,f3594,f331])).
fof(f4970,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f4219,f3594])).
fof(f4219,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f2517,f3594,f3594,f331])).
fof(f4969,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4225,f3594])).
fof(f4225,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f564,f386,f3594,f334])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP8(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f564,plain,(
  ( ! [X0] : (sP8('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f386,f386,f330])).
fof(f330,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f189])).
fof(f4968,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4239,f3594])).
fof(f4239,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f556,f386,f3594,f334])).
fof(f556,plain,(
  ( ! [X0] : (sP8('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f386,f386,f329])).
fof(f329,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f189])).
fof(f4967,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4275,f4697])).
fof(f4697,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f427,f3594,f370])).
fof(f370,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f338])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP9(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f191])).
fof(f191,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1)) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f427,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f386,f345])).
fof(f345,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f194])).
fof(f194,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f4275,plain,(
  ( ! [X0,X1] : (~sP10('F_0',X0,X1)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f386,f3594,f335])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f4966,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4298,f4697])).
fof(f4298,plain,(
  ( ! [X0,X1] : (~sP10('F_0',X0,X1)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f386,f3594,f336])).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f192])).
fof(f4965,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4310,f4346])).
fof(f4346,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f3594,f344])).
fof(f344,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f195])).
fof(f4310,plain,(
  ~sP9('F_0','T_0','F_0') | $spl76),
  inference(unit_resulting_resolution,[],[f572,f386,f3594,f339])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f572,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f211,f386,f336])).
fof(f4964,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4324,f4346])).
fof(f4324,plain,(
  ~sP9('F_0','F_0','T_0') | $spl76),
  inference(unit_resulting_resolution,[],[f571,f386,f3594,f340])).
fof(f340,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f192])).
fof(f571,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f211,f386,f335])).
fof(f4963,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4342,f4346])).
fof(f4342,plain,(
  ( ! [X0] : (~sP9(X0,'F_0','F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f386,f386,f3594,f342])).
fof(f342,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f195])).
fof(f4962,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4350,f211])).
fof(f4350,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f978,f3594,f347])).
fof(f347,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP10(X2,X0,X1))),
  inference(definition_folding,[],[f87,f131,f130])).
fof(f87,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X13,X14,X15] : ((p(s(bool,X15)) <=> (p(s(bool,X14)) | p(s(bool,X13)))) <=> ((p(s(bool,X15)) | ~p(s(bool,X14))) & (p(s(bool,X15)) | ~p(s(bool,X13))) & (p(s(bool,X14)) | p(s(bool,X13)) | ~p(s(bool,X15)))))),
  file('thm.extreal.add_rdistrib','thm.sat.sat_25_prove_ltb')).
fof(f978,plain,(
  ~sP10('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f964,f337])).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f192])).
fof(f964,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f386,f386,f211,f342])).
fof(f4961,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f4351,f3594])).
fof(f4351,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f2637,f3594,f347])).
fof(f2637,plain,(
  ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f2623,f337])).
fof(f2623,plain,(
  ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f386,f386,f973])).
fof(f973,plain,(
  ( ! [X0,X1] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f965,f211])).
fof(f965,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,X1)) ) | $spl2),
  inference(superposition,[],[f342,f439])).
fof(f4960,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f4352,f3594])).
fof(f4352,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f2658,f3594,f347])).
fof(f2658,plain,(
  ~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f2644,f337])).
fof(f2644,plain,(
  ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f386,f386,f974])).
fof(f974,plain,(
  ( ! [X2,X3] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X2,X3) | p(s(bool,X3)) | p(s(bool,X2))) ) | $spl6),
  inference(subsumption_resolution,[],[f966,f211])).
fof(f966,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | p(s(bool,X3)) | ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X2,X3)) ) | $spl6),
  inference(superposition,[],[f342,f440])).
fof(f4959,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4357,f211])).
fof(f4357,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f571,f3594,f347])).
fof(f4958,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4358,f3594])).
fof(f4358,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f572,f3594,f347])).
fof(f4957,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4368,f211])).
fof(f4368,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f978,f3594,f348])).
fof(f348,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f4956,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f4369,f3594])).
fof(f4369,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f2637,f3594,f348])).
fof(f4955,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f4370,f3594])).
fof(f4370,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f2658,f3594,f348])).
fof(f4954,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4375,f3594])).
fof(f4375,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f571,f3594,f348])).
fof(f4953,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4376,f211])).
fof(f4376,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f572,f3594,f348])).
fof(f4952,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4425,f4678])).
fof(f4678,plain,(
  ( ! [X0,X1] : (sP12('F_0',X0,X1)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f432,f3594,f369])).
fof(f369,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f355])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP11(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f199])).
fof(f199,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1) | sP12(X2,X0,X1)))),
  inference(flattening,[],[f198])).
fof(f198,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1)) | sP12(X2,X0,X1)))),
  inference(nnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X2,X0,X1] : (sP12(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f432,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f386,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f202])).
fof(f202,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f201])).
fof(f201,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f4425,plain,(
  ( ! [X0,X1] : (~sP12('F_0',X0,X1)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f386,f3594,f353])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f200])).
fof(f4951,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4428,f3594])).
fof(f4428,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f432,f600,f3594,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f200])).
fof(f600,plain,(
  ( ! [X0] : (~sP12('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f211,f386,f353])).
fof(f4950,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4429,f3594])).
fof(f4429,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f429,f600,f3594,f356])).
fof(f429,plain,(
  ( ! [X0,X1] : (sP11(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f211,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f203])).
fof(f4949,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4435,f4549])).
fof(f4549,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f3594,f361])).
fof(f4435,plain,(
  ~sP11('F_0','T_0','T_0') | $spl76),
  inference(unit_resulting_resolution,[],[f600,f211,f3594,f356])).
fof(f4948,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4445,f4549])).
fof(f4445,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f594,f3594,f3594,f356])).
fof(f594,plain,(
  ( ! [X0] : (~sP12('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f386,f386,f352])).
fof(f352,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f200])).
fof(f4947,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4447,f4549])).
fof(f4447,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f600,f3594,f3594,f356])).
fof(f4946,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4468,f4549])).
fof(f4468,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f594,f3594,f3594,f356])).
fof(f4945,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4469,f4549])).
fof(f4469,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f600,f3594,f3594,f356])).
fof(f4944,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4492,f4549])).
fof(f4492,plain,(
  ~sP11('F_0','F_0','F_0') | $spl76),
  inference(unit_resulting_resolution,[],[f594,f386,f3594,f357])).
fof(f357,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f200])).
fof(f4943,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4509,f4549])).
fof(f4509,plain,(
  ~sP11('F_0','F_0','F_0') | $spl76),
  inference(unit_resulting_resolution,[],[f594,f386,f3594,f358])).
fof(f358,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f200])).
fof(f4942,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4529,f3594])).
fof(f4529,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f432,f386,f3594,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f203])).
fof(f4941,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4532,f4549])).
fof(f4532,plain,(
  ( ! [X0,X1] : (~sP11(X0,'F_0',X1)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f3594,f386,f3594,f359])).
fof(f4940,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4543,f3594])).
fof(f4543,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f428,f386,f3594,f359])).
fof(f428,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f386,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f203])).
fof(f4939,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4545,f4549])).
fof(f4545,plain,(
  ( ! [X0] : (~sP11('T_0','F_0',X0)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f211,f386,f3594,f359])).
fof(f4938,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4547,f4549])).
fof(f4547,plain,(
  ( ! [X0,X1] : (~sP11(X0,'F_0',X1)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f3594,f386,f3594,f359])).
fof(f4937,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4572,f211])).
fof(f4572,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f1056,f3594,f365])).
fof(f365,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP12(X2,X0,X1))),
  inference(definition_folding,[],[f115,f134,f133])).
fof(f115,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X13,X14,X15] : ((p(s(bool,X15)) <=> (p(s(bool,X14)) => p(s(bool,X13)))) <=> ((p(s(bool,X15)) | p(s(bool,X14))) & (p(s(bool,X15)) | ~p(s(bool,X13))) & (~p(s(bool,X14)) | p(s(bool,X13)) | ~p(s(bool,X15)))))),
  file('thm.extreal.add_rdistrib','thm.sat.sat_27_prove_ltb')).
fof(f1056,plain,(
  ~sP12('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1044,f354])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f200])).
fof(f1044,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f211,f386,f211,f359])).
fof(f4936,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f4573,f211])).
fof(f4573,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f2821,f3594,f365])).
fof(f2821,plain,(
  ~sP12('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(unit_resulting_resolution,[],[f2784,f354])).
fof(f2784,plain,(
  ~sP11('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(unit_resulting_resolution,[],[f386,f211,f1051])).
fof(f1051,plain,(
  ( ! [X0,X1] : (~sP11(X1,X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X1)) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f1045,f211])).
fof(f1045,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP11(X1,X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl2),
  inference(superposition,[],[f359,f439])).
fof(f4935,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f4574,f211])).
fof(f4574,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f2826,f3594,f365])).
fof(f2826,plain,(
  ~sP12('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl6),
  inference(unit_resulting_resolution,[],[f2811,f354])).
fof(f2811,plain,(
  ~sP11('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl6),
  inference(unit_resulting_resolution,[],[f386,f211,f1052])).
fof(f1052,plain,(
  ( ! [X2,X3] : (~sP11(X3,X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X3)) | p(s(bool,X2))) ) | $spl6),
  inference(subsumption_resolution,[],[f1046,f211])).
fof(f1046,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | ~p(s(bool,X3)) | ~sP11(X3,X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl6),
  inference(superposition,[],[f359,f440])).
fof(f4934,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f4575,f3594])).
fof(f4575,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f2860,f3594,f365])).
fof(f2860,plain,(
  ~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f211,f386,f1089])).
fof(f1089,plain,(
  ( ! [X0,X1] : (~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,X0) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f1083,f211])).
fof(f1083,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,X0)) ) | $spl2),
  inference(superposition,[],[f366,f439])).
fof(f366,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f205])).
fof(f4933,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f4576,f3594])).
fof(f4576,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f2873,f3594,f365])).
fof(f2873,plain,(
  ~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0') | $spl6),
  inference(unit_resulting_resolution,[],[f211,f386,f1090])).
fof(f1090,plain,(
  ( ! [X2,X3] : (~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3,X2) | p(s(bool,X3)) | ~p(s(bool,X2))) ) | $spl6),
  inference(subsumption_resolution,[],[f1084,f211])).
fof(f1084,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3,X2)) ) | $spl6),
  inference(superposition,[],[f366,f440])).
fof(f4932,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4581,f3594])).
fof(f4581,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl76),
  inference(unit_resulting_resolution,[],[f594,f3594,f365])).
fof(f4931,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4582,f211])).
fof(f4582,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f600,f3594,f365])).
fof(f4930,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4592,f211])).
fof(f4592,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f604,f386,f3594,f366])).
fof(f604,plain,(
  ( ! [X0] : (sP12('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f386,f211,f364])).
fof(f364,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f205])).
fof(f4929,plain,(
  $false | ($spl73 | $spl76)),
  inference(subsumption_resolution,[],[f4596,f3594])).
fof(f4596,plain,(
  ~p(s(bool,f(s(bool,'T_0')))) | ($spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f3278,f386,f3594,f366])).
fof(f3278,plain,(
  ( ! [X0] : (sP12(f(s(bool,'T_0')),X0,'F_0')) ) | $spl73),
  inference(unit_resulting_resolution,[],[f386,f2777])).
fof(f2777,plain,(
  ( ! [X39,X38] : (sP12(f(s(bool,'T_0')),X39,X38) | p(s(bool,X38))) ) | $spl73),
  inference(subsumption_resolution,[],[f2729,f211])).
fof(f2729,plain,(
  ( ! [X39,X38] : (~p(s(bool,'T_0')) | p(s(bool,X38)) | sP12(f(s(bool,'T_0')),X39,X38)) ) | $spl73),
  inference(superposition,[],[f364,f2399])).
fof(f4928,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f4597,f3594])).
fof(f4597,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f925,f386,f3594,f366])).
fof(f925,plain,(
  ( ! [X0] : (sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f386,f607])).
fof(f607,plain,(
  ( ! [X0,X1] : (sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,X0) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f605,f211])).
fof(f605,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,X0)) ) | $spl2),
  inference(superposition,[],[f364,f439])).
fof(f4927,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f4598,f3594])).
fof(f4598,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f935,f386,f3594,f366])).
fof(f935,plain,(
  ( ! [X0] : (sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,'F_0')) ) | $spl6),
  inference(unit_resulting_resolution,[],[f386,f608])).
fof(f608,plain,(
  ( ! [X2,X3] : (sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3,X2) | p(s(bool,X2))) ) | $spl6),
  inference(subsumption_resolution,[],[f606,f211])).
fof(f606,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3,X2)) ) | $spl6),
  inference(superposition,[],[f364,f440])).
fof(f4926,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4625,f211])).
fof(f4625,plain,(
  ~p(s(bool,'T_0')) | $spl76),
  inference(unit_resulting_resolution,[],[f1033,f386,f3594,f366])).
fof(f1033,plain,(
  sP12('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f432,f211,f386,f358])).
fof(f4925,plain,(
  $false | ($spl2 | $spl76)),
  inference(subsumption_resolution,[],[f4626,f3594])).
fof(f4626,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f2828,f386,f3594,f366])).
fof(f2828,plain,(
  sP12('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(unit_resulting_resolution,[],[f386,f386,f1071])).
fof(f1071,plain,(
  ( ! [X0,X1] : (sP12(X1,X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl2),
  inference(subsumption_resolution,[],[f1067,f211])).
fof(f1067,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP12(X1,X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X1))) ) | $spl2),
  inference(superposition,[],[f363,f439])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP12(X2,X0,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f205])).
fof(f4924,plain,(
  $false | ($spl6 | $spl76)),
  inference(subsumption_resolution,[],[f4627,f3594])).
fof(f4627,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f2844,f386,f3594,f366])).
fof(f2844,plain,(
  sP12('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl6),
  inference(unit_resulting_resolution,[],[f386,f386,f1072])).
fof(f1072,plain,(
  ( ! [X2,X3] : (sP12(X3,X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X2)) | p(s(bool,X3))) ) | $spl6),
  inference(subsumption_resolution,[],[f1068,f211])).
fof(f1068,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP12(X3,X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f363,f440])).
fof(f4923,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4663,f428])).
fof(f4663,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f594,f3594,f369])).
fof(f4922,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4664,f429])).
fof(f4664,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f600,f3594,f369])).
fof(f4921,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4680,f421])).
fof(f421,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f211,f343])).
fof(f343,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f195])).
fof(f4680,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) ) | $spl76),
  inference(unit_resulting_resolution,[],[f571,f3594,f370])).
fof(f4920,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4681,f424])).
fof(f424,plain,(
  ( ! [X0,X1] : (sP9(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f211,f344])).
fof(f4681,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) ) | $spl76),
  inference(unit_resulting_resolution,[],[f572,f3594,f370])).
fof(f4919,plain,(
  $false | ($spl28 | $spl76)),
  inference(subsumption_resolution,[],[f4840,f3594])).
fof(f4840,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl28 | $spl76)),
  inference(unit_resulting_resolution,[],[f550,f3594,f1730])).
fof(f4918,plain,(
  $false | ($spl28 | $spl76)),
  inference(subsumption_resolution,[],[f4845,f3594])).
fof(f4845,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl28 | $spl76)),
  inference(unit_resulting_resolution,[],[f547,f3594,f1730])).
fof(f4917,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4883,f3594])).
fof(f4883,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK14)))) | $spl76),
  inference(unit_resulting_resolution,[],[f389,f3594,f283])).
fof(f4916,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4886,f3594])).
fof(f4886,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK15)))) | $spl76),
  inference(unit_resulting_resolution,[],[f389,f3594,f283])).
fof(f4915,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4891,f3594])).
fof(f4891,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl76),
  inference(unit_resulting_resolution,[],[f389,f3594,f284])).
fof(f4914,plain,(
  $false | $spl76),
  inference(subsumption_resolution,[],[f4894,f3594])).
fof(f4894,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl76),
  inference(unit_resulting_resolution,[],[f389,f3594,f284])).
fof(f4895,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f3594,f389,f3594,f284])).
fof(f4892,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f3594,f389,f3594,f284])).
fof(f4887,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f3594,f389,f3594,f283])).
fof(f4884,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f3594,f389,f3594,f283])).
fof(f4877,plain,(
  $false | ($spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f517,f3594,f2762])).
fof(f2762,plain,(
  ( ! [X8,X9] : (sP2(X8,X9,f(s(bool,'T_0'))) | ~p(s(bool,X8))) ) | $spl73),
  inference(subsumption_resolution,[],[f2714,f211])).
fof(f2714,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | sP2(X8,X9,f(s(bool,'T_0')))) ) | $spl73),
  inference(superposition,[],[f290,f2399])).
fof(f4874,plain,(
  $false | ($spl48 | $spl76)),
  inference(unit_resulting_resolution,[],[f604,f386,f3594,f1760])).
fof(f4868,plain,(
  $false | ($spl46 | $spl76)),
  inference(unit_resulting_resolution,[],[f1056,f3594,f1757])).
fof(f4866,plain,(
  $false | ($spl40 | $spl76)),
  inference(unit_resulting_resolution,[],[f432,f386,f3594,f1748])).
fof(f4860,plain,(
  $false | ($spl36 | $spl76)),
  inference(unit_resulting_resolution,[],[f978,f3594,f1742])).
fof(f4858,plain,(
  $false | ($spl34 | $spl76)),
  inference(unit_resulting_resolution,[],[f978,f3594,f1739])).
fof(f4856,plain,(
  $false | ($spl30 | $spl76)),
  inference(unit_resulting_resolution,[],[f564,f386,f3594,f1733])).
fof(f4850,plain,(
  $false | ($spl28 | $spl76)),
  inference(unit_resulting_resolution,[],[f547,f3594,f3594,f1730])).
fof(f4848,plain,(
  $false | ($spl28 | $spl76)),
  inference(unit_resulting_resolution,[],[f547,f211,f3594,f1730])).
fof(f4844,plain,(
  $false | ($spl28 | $spl76)),
  inference(unit_resulting_resolution,[],[f547,f3594,f3594,f1730])).
fof(f4842,plain,(
  $false | ($spl28 | $spl76)),
  inference(unit_resulting_resolution,[],[f550,f211,f3594,f1730])).
fof(f4838,plain,(
  $false | ($spl26 | $spl76)),
  inference(unit_resulting_resolution,[],[f420,f386,f3594,f1727])).
fof(f4832,plain,(
  $false | ($spl22 | $spl76)),
  inference(unit_resulting_resolution,[],[f406,f386,f3594,f1721])).
fof(f4826,plain,(
  $false | ($spl20 | $spl76)),
  inference(unit_resulting_resolution,[],[f401,f386,f3594,f1718])).
fof(f4819,plain,(
  $false | ($spl16 | $spl76)),
  inference(unit_resulting_resolution,[],[f735,f386,f3594,f1711])).
fof(f4806,plain,(
  $false | ($spl14 | $spl76)),
  inference(unit_resulting_resolution,[],[f517,f3594,f1708])).
fof(f4801,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f935,f3594,f1090])).
fof(f4795,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f925,f3594,f1089])).
fof(f4792,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f432,f3594,f1052])).
fof(f4786,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f432,f3594,f1051])).
fof(f4780,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f564,f3594,f919])).
fof(f919,plain,(
  ( ! [X2,X3] : (~sP8(X3,X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X3)) | ~p(s(bool,X2))) ) | $spl6),
  inference(subsumption_resolution,[],[f913,f211])).
fof(f913,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP8(X3,X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl6),
  inference(superposition,[],[f334,f440])).
fof(f4777,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f564,f3594,f918])).
fof(f918,plain,(
  ( ! [X0,X1] : (~sP8(X1,X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f912,f211])).
fof(f912,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP8(X1,X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl2),
  inference(superposition,[],[f334,f439])).
fof(f4766,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f420,f3594,f885])).
fof(f4760,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f420,f3594,f884])).
fof(f4754,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f406,f3594,f835])).
fof(f835,plain,(
  ( ! [X2,X3] : (~sP4(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3) | p(s(bool,X3)) | ~p(s(bool,X2))) ) | $spl6),
  inference(subsumption_resolution,[],[f829,f211])).
fof(f829,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP4(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3)) ) | $spl6),
  inference(superposition,[],[f306,f440])).
fof(f4748,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f406,f3594,f834])).
fof(f834,plain,(
  ( ! [X0,X1] : (~sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f828,f211])).
fof(f828,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl2),
  inference(superposition,[],[f306,f439])).
fof(f4742,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f401,f3594,f815])).
fof(f815,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | ~sP5(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3)) ) | $spl6),
  inference(subsumption_resolution,[],[f809,f211])).
fof(f809,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | p(s(bool,X3)) | ~sP5(X2,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3)) ) | $spl6),
  inference(superposition,[],[f302,f440])).
fof(f4736,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f401,f3594,f814])).
fof(f814,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl2),
  inference(subsumption_resolution,[],[f808,f211])).
fof(f808,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl2),
  inference(superposition,[],[f302,f439])).
fof(f4726,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f1246,f3594,f797])).
fof(f4719,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f386,f1245,f3594,f796])).
fof(f4716,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f2873,f3594,f623])).
fof(f623,plain,(
  ( ! [X2,X3] : (sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X2,X3) | ~p(s(bool,X2))) ) | $spl6),
  inference(subsumption_resolution,[],[f621,f211])).
fof(f621,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X2,X3)) ) | $spl6),
  inference(superposition,[],[f365,f440])).
fof(f4714,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f2860,f3594,f622])).
fof(f622,plain,(
  ( ! [X0,X1] : (sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,X1) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f620,f211])).
fof(f620,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,X1)) ) | $spl2),
  inference(superposition,[],[f365,f439])).
fof(f4712,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f2658,f3594,f589])).
fof(f589,plain,(
  ( ! [X2,X3] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X2,X3) | ~p(s(bool,X2))) ) | $spl6),
  inference(subsumption_resolution,[],[f587,f211])).
fof(f587,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X2,X3)) ) | $spl6),
  inference(superposition,[],[f348,f440])).
fof(f4710,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f2637,f3594,f588])).
fof(f588,plain,(
  ( ! [X0,X1] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,X1) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f586,f211])).
fof(f586,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,X1)) ) | $spl2),
  inference(superposition,[],[f348,f439])).
fof(f4708,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f2658,f3594,f578])).
fof(f578,plain,(
  ( ! [X2,X3] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3,X2) | ~p(s(bool,X2))) ) | $spl6),
  inference(subsumption_resolution,[],[f576,f211])).
fof(f576,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X3,X2)) ) | $spl6),
  inference(superposition,[],[f347,f440])).
fof(f4706,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f2637,f3594,f577])).
fof(f577,plain,(
  ( ! [X0,X1] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,X0) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f575,f211])).
fof(f575,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,X0)) ) | $spl2),
  inference(superposition,[],[f347,f439])).
fof(f4704,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f517,f3594,f542])).
fof(f542,plain,(
  ( ! [X2,X3] : (sP2(X2,X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X2))) ) | $spl6),
  inference(subsumption_resolution,[],[f540,f211])).
fof(f540,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | sP2(X2,X3,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl6),
  inference(superposition,[],[f290,f440])).
fof(f4702,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f517,f3594,f541])).
fof(f541,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f539,f211])).
fof(f539,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP2(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl2),
  inference(superposition,[],[f290,f439])).
fof(f4700,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f508,f3594,f372])).
fof(f372,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(duplicate_literal_removal,[],[f287])).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f168])).
fof(f4698,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f571,f427,f3594,f370])).
fof(f4696,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f571,f426,f3594,f370])).
fof(f426,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f385,f344])).
fof(f4694,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f571,f425,f3594,f370])).
fof(f425,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f381,f344])).
fof(f4692,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f572,f424,f3594,f370])).
fof(f4690,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f572,f423,f3594,f370])).
fof(f423,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f385,f343])).
fof(f4688,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f572,f422,f3594,f370])).
fof(f422,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f381,f343])).
fof(f4686,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f571,f421,f3594,f370])).
fof(f4679,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f594,f432,f3594,f369])).
fof(f4677,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f594,f431,f3594,f369])).
fof(f431,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f385,f361])).
fof(f4675,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f594,f430,f3594,f369])).
fof(f430,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f381,f361])).
fof(f4673,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f600,f429,f3594,f369])).
fof(f4671,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f594,f428,f3594,f369])).
fof(f4662,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f604,f386,f3594,f3594,f366])).
fof(f4643,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f1033,f386,f211,f3594,f366])).
fof(f4623,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f604,f386,f3594,f3594,f366])).
fof(f4604,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f604,f386,f211,f3594,f366])).
fof(f4591,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f594,f3594,f3594,f365])).
fof(f4589,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f600,f211,f3594,f365])).
fof(f4580,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f594,f3594,f3594,f365])).
fof(f4578,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f1056,f211,f3594,f365])).
fof(f4569,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f594,f386,f3594,f364])).
fof(f4560,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f594,f386,f386,f3594,f363])).
fof(f4550,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f1044,f3594,f361])).
fof(f4548,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f428,f3594,f386,f3594,f359])).
fof(f4546,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f428,f211,f386,f3594,f359])).
fof(f4533,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f428,f3594,f386,f3594,f359])).
fof(f4531,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f432,f211,f386,f3594,f359])).
fof(f4518,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f594,f432,f386,f3594,f358])).
fof(f4516,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f594,f428,f386,f3594,f358])).
fof(f4501,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f594,f432,f386,f3594,f357])).
fof(f4499,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f594,f428,f386,f3594,f357])).
fof(f4484,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f594,f432,f3594,f3594,f356])).
fof(f4482,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f594,f431,f3594,f3594,f356])).
fof(f4480,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f594,f430,f3594,f3594,f356])).
fof(f4478,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f600,f429,f3594,f3594,f356])).
fof(f4476,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f594,f428,f3594,f3594,f356])).
fof(f4448,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f432,f600,f3594,f3594,f356])).
fof(f4446,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f432,f594,f3594,f3594,f356])).
fof(f4436,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f432,f600,f211,f3594,f356])).
fof(f4426,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f1033,f386,f3594,f353])).
fof(f4407,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f948,f386,f386,f3594,f349])).
fof(f349,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f948,plain,(
  sP10('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f427,f386,f386,f341])).
fof(f341,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f4383,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f571,f3594,f3594,f348])).
fof(f4381,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f572,f211,f3594,f348])).
fof(f4374,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f571,f3594,f3594,f348])).
fof(f4372,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f978,f211,f3594,f348])).
fof(f4365,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f571,f3594,f3594,f347])).
fof(f4363,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f571,f211,f3594,f347])).
fof(f4356,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f571,f3594,f3594,f347])).
fof(f4354,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f978,f211,f3594,f347])).
fof(f4347,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f964,f3594,f344])).
fof(f4345,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f964,f3594,f343])).
fof(f4343,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f427,f386,f386,f3594,f342])).
fof(f4333,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f571,f427,f386,f3594,f340])).
fof(f4329,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f571,f421,f386,f3594,f340])).
fof(f4311,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f427,f572,f386,f3594,f339])).
fof(f4299,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f948,f386,f3594,f336])).
fof(f4276,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f948,f386,f3594,f335])).
fof(f4253,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f556,f386,f3594,f3594,f334])).
fof(f4246,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f556,f386,f211,f3594,f334])).
fof(f4238,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f556,f386,f3594,f3594,f334])).
fof(f4231,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f564,f386,f211,f3594,f334])).
fof(f4223,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f547,f3594,f3594,f3594,f331])).
fof(f4221,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f550,f211,f3594,f3594,f331])).
fof(f4214,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f547,f3594,f211,f3594,f331])).
fof(f4212,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f889,f211,f211,f3594,f331])).
fof(f4199,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f547,f3594,f3594,f3594,f331])).
fof(f4197,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f547,f211,f3594,f3594,f331])).
fof(f4190,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f547,f3594,f211,f3594,f331])).
fof(f4188,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f547,f211,f211,f3594,f331])).
fof(f4175,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f547,f3594,f3594,f3594,f331])).
fof(f4173,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f550,f211,f3594,f3594,f331])).
fof(f4166,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f547,f3594,f211,f3594,f331])).
fof(f4164,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f550,f211,f211,f3594,f331])).
fof(f4153,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f877,f3594,f326])).
fof(f4151,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f419,f386,f3594,f3594,f325])).
fof(f4144,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f419,f386,f211,f3594,f325])).
fof(f4136,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f419,f386,f3594,f3594,f325])).
fof(f4129,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f420,f386,f211,f3594,f325])).
fof(f4118,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f420,f550,f3594,f3594,f321])).
fof(f4116,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f419,f547,f3594,f3594,f321])).
fof(f4106,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f419,f547,f211,f3594,f321])).
fof(f4093,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f420,f550,f3594,f3594,f321])).
fof(f4091,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f419,f547,f3594,f3594,f321])).
fof(f4083,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f420,f550,f211,f3594,f321])).
fof(f4073,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f564,f386,f3594,f320])).
fof(f4066,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f556,f386,f3594,f319])).
fof(f4059,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f735,f436,f3594,f316])).
fof(f4057,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f734,f436,f3594,f316])).
fof(f4025,plain,(
  $false | ($spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f2692,f2399,f3594,f315])).
fof(f4023,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f775,f440,f3594,f315])).
fof(f4021,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f774,f439,f3594,f315])).
fof(f4019,plain,(
  $false | ($spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f2690,f2399,f3594,f315])).
fof(f4017,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f772,f440,f3594,f315])).
fof(f4015,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f770,f439,f3594,f315])).
fof(f4913,plain,(
  $false | $spl76),
  inference(trivial_inequality_removal,[],[f4000])).
fof(f4000,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl76),
  inference(unit_resulting_resolution,[],[f860,f3594,f315])).
fof(f860,plain,(
  ~sP6('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f847,f295])).
fof(f295,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f847,plain,(
  ~sP3('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f386,f386,f386,f310])).
fof(f310,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f4912,plain,(
  $false | $spl76),
  inference(trivial_inequality_removal,[],[f3998])).
fof(f3998,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl76),
  inference(unit_resulting_resolution,[],[f839,f3594,f315])).
fof(f839,plain,(
  ~sP6('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f827,f296])).
fof(f296,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f827,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f386,f211,f211,f306])).
fof(f3987,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f847,f3594,f313])).
fof(f3985,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f847,f3594,f312])).
fof(f3983,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f847,f3594,f311])).
fof(f3981,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f827,f3594,f307])).
fof(f3979,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f405,f386,f3594,f3594,f306])).
fof(f3972,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f405,f386,f211,f3594,f306])).
fof(f3964,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f405,f386,f3594,f3594,f306])).
fof(f3957,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f406,f386,f211,f3594,f306])).
fof(f3949,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f807,f3594,f303])).
fof(f807,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f386,f211,f211,f302])).
fof(f3947,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f400,f386,f3594,f3594,f302])).
fof(f3939,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f400,f386,f211,f3594,f302])).
fof(f3930,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f400,f386,f3594,f3594,f302])).
fof(f3922,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f401,f386,f211,f3594,f302])).
fof(f3913,plain,(
  $false | ($spl6 | $spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f409,f405,f400,f2395,f3594,f299])).
fof(f3909,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f410,f405,f400,f1024,f3594,f299])).
fof(f3903,plain,(
  $false | ($spl2 | $spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f408,f405,f400,f2396,f3594,f299])).
fof(f3899,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f410,f405,f400,f1014,f3594,f299])).
fof(f3883,plain,(
  $false | ($spl6 | $spl76)),
  inference(unit_resulting_resolution,[],[f412,f405,f400,f1075,f3594,f299])).
fof(f3879,plain,(
  $false | ($spl2 | $spl76)),
  inference(unit_resulting_resolution,[],[f411,f405,f400,f1037,f3594,f299])).
fof(f3875,plain,(
  $false | ($spl73 | $spl76)),
  inference(unit_resulting_resolution,[],[f407,f405,f400,f2397,f3594,f299])).
fof(f3867,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f410,f405,f400,f757,f3594,f299])).
fof(f3775,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f735,f386,f3594,f3594,f298])).
fof(f3743,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f1193,f386,f211,f3594,f298])).
fof(f3724,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f735,f386,f3594,f3594,f298])).
fof(f3692,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f735,f386,f211,f3594,f298])).
fof(f3672,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f508,f3594,f3594,f290])).
fof(f3670,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f508,f211,f3594,f290])).
fof(f3666,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f508,f3594,f3594,f290])).
fof(f3664,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f517,f211,f3594,f290])).
fof(f3660,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f508,f3594,f3594,f289])).
fof(f3658,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f508,f211,f3594,f289])).
fof(f3654,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f508,f3594,f3594,f289])).
fof(f3649,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f508,f3594,f3594,f288])).
fof(f3644,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f508,f3594,f3594,f288])).
fof(f3642,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f517,f211,f3594,f288])).
fof(f3638,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f471,f386,f3594,f277])).
fof(f471,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f386,f386,f272])).
fof(f272,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f163])).
fof(f3634,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f471,f386,f3594,f274])).
fof(f3630,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f475,f3594,f3594,f273])).
fof(f3628,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f475,f211,f3594,f273])).
fof(f3618,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f475,f3594,f3594,f273])).
fof(f3616,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f482,f211,f3594,f273])).
fof(f3606,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f497,f3594,f259])).
fof(f259,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f497,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f386,f386,f390])).
fof(f390,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f257,f386])).
fof(f257,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f156])).
fof(f3604,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f497,f3594,f258])).
fof(f258,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f156])).
fof(f3601,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f436,f3594,f226])).
fof(f3599,plain,(
  $false | $spl76),
  inference(unit_resulting_resolution,[],[f386,f3594])).
fof(f3598,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl78),
  inference(cnf_transformation,[],[f3598_D])).
fof(f3598_D,plain,(
  ( ! [X28,X29] : (sP1(X28,'F_0') | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f3597,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f3588,f225])).
fof(f3588,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | sP1(X28,'F_0') | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f1776,f218])).
fof(f1776,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X17)) | sP1(X16,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1476,f211])).
fof(f1476,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | p(s(bool,X17)) | sP1(X16,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f274,f680])).
fof(f3593,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) | $spl74),
  inference(cnf_transformation,[],[f3593_D])).
fof(f3593_D,plain,(
  ( ! [X10,X8,X11] : (sP1(X10,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X11)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3595,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f3584,f225])).
fof(f3584,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X13)) | sP1(X14,X12) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1776,f680])).
fof(f3592,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f3583,f225])).
fof(f3583,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X9)) | sP1(X10,X8) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f1776,f680])).
fof(f2339,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool))))) ) | $spl70),
  inference(cnf_transformation,[],[f2339_D])).
fof(f2339_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f2341,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl73),
  inference(cnf_transformation,[],[f2341_D])).
fof(f2341_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f2329,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK16(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1261,f680])).
fof(f1261,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK16(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK16(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f282,f218])).
fof(f2328,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1261,f680])).
fof(f2320,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl64),
  inference(cnf_transformation,[],[f2320_D])).
fof(f2320_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2316,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl60),
  inference(cnf_transformation,[],[f2316_D])).
fof(f2316_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f2323,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl69),
  inference(cnf_transformation,[],[f2323_D])).
fof(f2323_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f2314,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK16(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1260,f218])).
fof(f1260,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,X2,bool))) )),
  inference(superposition,[],[f282,f218])).
fof(f2321,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl66),
  inference(cnf_transformation,[],[f2321_D])).
fof(f2321_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2313,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK16(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1260,f680])).
fof(f2312,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK16(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1260,f680])).
fof(f2315,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl58),
  inference(cnf_transformation,[],[f2315_D])).
fof(f2315_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2318,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl63),
  inference(cnf_transformation,[],[f2318_D])).
fof(f2318_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f2311,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK16(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1260,f680])).
fof(f2310,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1260,f680])).
fof(f1892,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f210,f209,f208,f207,f206,f211,f386,f218,f226,f225,f230,f229,f252,f389,f259,f258,f390,f277,f274,f273,f272,f282,f284,f283,f290,f289,f288,f372,f286,f285,f301,f300,f299,f298,f297,f296,f295,f305,f304,f303,f302,f309,f308,f307,f306,f313,f312,f311,f310,f317,f316,f315,f314,f371,f321,f320,f319,f318,f328,f327,f326,f325,f334,f331,f330,f329,f341,f340,f339,f370,f337,f336,f335,f345,f344,f343,f342,f349,f348,f347,f346,f358,f357,f356,f369,f354,f353,f352,f362,f361,f360,f359,f366,f365,f364,f363,f391,f394,f397,f400,f401,f402,f405,f406,f407,f410,f413,f416,f419,f420,f421,f424,f427,f428,f429,f432,f433,f436,f471,f472,f475,f482,f497,f508,f517,f538,f547,f550,f556,f564,f571,f572,f574,f585,f594,f600,f604,f619,f680,f735,f734,f757,f755,f807,f827,f839,f847,f860,f864,f877,f889,f948,f964,f978,f1033,f1044,f1056,f1140,f1193,f1260,f1261,f1272,f1287,f1327,f1389,f1390,f1435,f1436,f1700,f1704,f1707,f1710,f1714,f1717,f1720,f1723,f1726,f1729,f1732,f1735,f1738,f1741,f1744,f1747,f1750,f1753,f1756,f1759,f1762,f1765,f1773,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1798,f1801,f1802,f1509,f1510,f1804,f1806,f1808,f1810,f1813,f1815,f1817,f1819,f1821,f1823,f1825,f1827,f1829,f1831,f1833,f1835,f1837,f1839,f1841,f1843,f1845,f1847,f1702,f1870,f1875,f1879,f1887,f1891])).
fof(f1891,plain,(
  ( ! [X2,X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f1890,f211])).
fof(f1890,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(forward_demodulation,[],[f1889,f1702])).
fof(f1889,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X2)))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f1866,f211])).
fof(f1866,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X2)))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(backward_demodulation,[],[f1702,f283])).
fof(f1887,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f1886,f211])).
fof(f1886,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(forward_demodulation,[],[f1885,f1702])).
fof(f1885,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f1865,f211])).
fof(f1865,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(backward_demodulation,[],[f1702,f1272])).
fof(f1879,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f1878,f211])).
fof(f1878,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(forward_demodulation,[],[f1877,f1702])).
fof(f1877,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f1860,f211])).
fof(f1860,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(backward_demodulation,[],[f1702,f1287])).
fof(f1875,plain,(
  ( ! [X2,X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f1874,f211])).
fof(f1874,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(forward_demodulation,[],[f1873,f1702])).
fof(f1873,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f1859,f211])).
fof(f1859,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X2)))))) ) | $spl10),
  inference(backward_demodulation,[],[f1702,f284])).
fof(f1870,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl10),
  inference(subsumption_resolution,[],[f1852,f1702])).
fof(f1852,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl10),
  inference(backward_demodulation,[],[f1702,f229])).
fof(f1510,plain,(
  ( ! [X140,X138,X141,X139] : (s(bool,'T_0') != s(bool,f(s(X138,sK16(X139,X140,bool,X138)))) | s(fun(X138,bool),X139) = s(fun(X138,bool),X140) | s(bool,X141) = s(bool,g(s(X138,sK16(X139,X140,bool,X138)))) | s(bool,'T_0') = s(bool,X141)) )),
  inference(superposition,[],[f282,f680])).
fof(f1509,plain,(
  ( ! [X134,X136,X135,X137] : (s(X136,f(s(bool,'T_0'))) != s(X136,g(s(bool,'T_0'))) | s(fun(bool,X136),X134) = s(fun(bool,X136),X135) | s(bool,X137) = s(bool,sK16(X134,X135,X136,bool)) | s(bool,'T_0') = s(bool,X137)) )),
  inference(superposition,[],[f282,f680])).
fof(f1802,plain,(
  ( ! [X127,X130,X128,X129] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X128)))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X129),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X127),s('type.extreal.extreal',X128))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X129),s('type.extreal.extreal',X127))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X129),s('type.extreal.extreal',X128))))) | s(bool,X130) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X127))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(subsumption_resolution,[],[f1508,f211])).
fof(f1508,plain,(
  ( ! [X127,X130,X128,X129] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X128)))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X129),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X127),s('type.extreal.extreal',X128))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X129),s('type.extreal.extreal',X127))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X129),s('type.extreal.extreal',X128))))) | s(bool,X130) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X127))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f283,f680])).
fof(f1801,plain,(
  ( ! [X125,X123,X126,X124] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X124),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X125),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X123),s('type.extreal.extreal',X124))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X125),s('type.extreal.extreal',X123))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X125),s('type.extreal.extreal',X124))))) | s(bool,X126) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X123),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(subsumption_resolution,[],[f1507,f211])).
fof(f1507,plain,(
  ( ! [X125,X123,X126,X124] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X124),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X125),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X123),s('type.extreal.extreal',X124))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X125),s('type.extreal.extreal',X123))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X125),s('type.extreal.extreal',X124))))) | s(bool,X126) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X123),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f284,f680])).
fof(f1798,plain,(
  ( ! [X109,X107,X108,X106] : (~sP9(X106,X107,X108) | sP10(X106,X107,X108) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f1499,f211])).
fof(f1499,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,'T_0')) | ~sP9(X106,X107,X108) | sP10(X106,X107,X108) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f370,f680])).
fof(f1797,plain,(
  ( ! [X103,X105,X102,X104] : (~sP11(X102,X103,X104) | sP12(X102,X103,X104) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f1498,f211])).
fof(f1498,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,'T_0')) | ~sP11(X102,X103,X104) | sP12(X102,X103,X104) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f369,f680])).
fof(f1796,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,X99)) | p(s(bool,X100)) | ~sP12(X98,X100,X99) | s(bool,X98) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(subsumption_resolution,[],[f1497,f211])).
fof(f1497,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,'T_0')) | ~p(s(bool,X99)) | p(s(bool,X100)) | ~sP12(X98,X100,X99) | s(bool,X98) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f366,f680])).
fof(f1795,plain,(
  ( ! [X94,X97,X95,X96] : (sP12(X94,X95,X96) | ~p(s(bool,X95)) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f1496,f211])).
fof(f1496,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X95)) | sP12(X94,X95,X96) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f365,f680])).
fof(f1794,plain,(
  ( ! [X92,X90,X93,X91] : (sP12(X90,X92,X91) | p(s(bool,X91)) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f1495,f211])).
fof(f1495,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,'T_0')) | p(s(bool,X91)) | sP12(X90,X92,X91) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f364,f680])).
fof(f1793,plain,(
  ( ! [X88,X87,X89,X86] : (p(s(bool,X87)) | sP12(X88,X87,X86) | p(s(bool,X88)) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f1494,f211])).
fof(f1494,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,'T_0')) | p(s(bool,X87)) | sP12(X88,X87,X86) | p(s(bool,X88)) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f363,f680])).
fof(f1792,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,X83)) | ~p(s(bool,X84)) | ~sP11(X84,X83,X82) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f1493,f211])).
fof(f1493,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,'T_0')) | p(s(bool,X83)) | ~p(s(bool,X84)) | ~sP11(X84,X83,X82) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f359,f680])).
fof(f1791,plain,(
  ( ! [X80,X78,X81,X79] : (p(s(bool,X79)) | p(s(bool,X80)) | ~sP10(X78,X80,X79) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f1492,f211])).
fof(f1492,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,'T_0')) | p(s(bool,X79)) | p(s(bool,X80)) | ~sP10(X78,X80,X79) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f349,f680])).
fof(f1790,plain,(
  ( ! [X76,X74,X77,X75] : (sP10(X74,X75,X76) | ~p(s(bool,X75)) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f1491,f211])).
fof(f1491,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X75)) | sP10(X74,X75,X76) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f348,f680])).
fof(f1789,plain,(
  ( ! [X70,X72,X71,X73] : (sP10(X70,X72,X71) | ~p(s(bool,X71)) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f1490,f211])).
fof(f1490,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X71)) | sP10(X70,X72,X71) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f347,f680])).
fof(f1788,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,X67)) | p(s(bool,X68)) | ~sP9(X66,X67,X68) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f1489,f211])).
fof(f1489,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X67)) | p(s(bool,X68)) | ~sP9(X66,X67,X68) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f342,f680])).
fof(f1787,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X63)) | p(s(bool,X64)) | ~sP8(X64,X63,X62) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f1488,f211])).
fof(f1488,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X63)) | p(s(bool,X64)) | ~sP8(X64,X63,X62) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f334,f680])).
fof(f1786,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | sP8(X58,X59,X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f1487,f211])).
fof(f1487,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,X59)) | ~p(s(bool,X60)) | sP8(X58,X59,X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f331,f680])).
fof(f1785,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | ~sP7(X55,X54,X56) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f1486,f211])).
fof(f1486,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X55)) | p(s(bool,X56)) | ~sP7(X55,X54,X56) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f325,f680])).
fof(f1784,plain,(
  ( ! [X52,X50,X53,X51] : (s(bool,X51) = s(bool,X52) | ~sP6(X50,X52,X51) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f1485,f211])).
fof(f1485,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | s(bool,X51) = s(bool,X52) | ~sP6(X50,X52,X51) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f316,f680])).
fof(f1783,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | p(s(bool,X48)) | ~sP4(X47,X46,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f1484,f211])).
fof(f1484,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP4(X47,X46,X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f306,f680])).
fof(f1782,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X43)) | p(s(bool,X44)) | ~sP5(X43,X42,X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f1483,f211])).
fof(f1483,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP5(X43,X42,X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f302,f680])).
fof(f1781,plain,(
  ( ! [X39,X41,X38,X40] : (sP6(X39,X40,X38) | ~sP5(X39,X38,X40) | ~sP4(X40,X38,X39) | ~sP3(X38,X40,X39) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f1482,f211])).
fof(f1482,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | sP6(X39,X40,X38) | ~sP5(X39,X38,X40) | ~sP4(X40,X38,X39) | ~sP3(X38,X40,X39) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f299,f680])).
fof(f1780,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X35)) | p(s(bool,X36)) | ~sP6(X35,X34,X36) | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f1481,f211])).
fof(f1481,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | p(s(bool,X36)) | ~sP6(X35,X34,X36) | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f298,f680])).
fof(f1779,plain,(
  ( ! [X30,X33,X31,X32] : (sP2(X31,X32,X30) | ~p(s(bool,X31)) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f1480,f211])).
fof(f1480,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | sP2(X31,X32,X30) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f290,f680])).
fof(f1778,plain,(
  ( ! [X28,X26,X29,X27] : (sP2(X27,X26,X28) | ~p(s(bool,X27)) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f1479,f211])).
fof(f1479,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X27)) | sP2(X27,X26,X28) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f289,f680])).
fof(f1775,plain,(
  ( ! [X14,X15,X13] : (sP0(X13,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f1475,f211])).
fof(f1475,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'T_0')) | sP0(X13,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f259,f680])).
fof(f1774,plain,(
  ( ! [X12,X10,X11] : (sP0(X11,X10) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f1474,f211])).
fof(f1474,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,'T_0')) | sP0(X11,X10) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f258,f680])).
fof(f1773,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1472,f211])).
fof(f1472,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f229,f680])).
fof(f1436,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f211,f680])).
fof(f1435,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f225,f680])).
fof(f1390,plain,(
  ( ! [X144,X142,X145,X143] : (s(bool,X145) != s(bool,f(s(X142,sK16(X143,X144,bool,X142)))) | s(fun(X142,bool),X143) = s(fun(X142,bool),X144) | s(bool,'T_0') = s(bool,X145) | s(bool,'T_0') = s(bool,g(s(X142,sK16(X143,X144,bool,X142))))) )),
  inference(superposition,[],[f282,f680])).
fof(f1389,plain,(
  ( ! [X140,X138,X141,X139] : (s(X140,f(s(bool,X141))) != s(X140,g(s(bool,X141))) | s(fun(bool,X140),X138) = s(fun(bool,X140),X139) | s(bool,'T_0') = s(bool,X141) | s(bool,'T_0') = s(bool,sK16(X138,X139,X140,bool))) )),
  inference(superposition,[],[f282,f680])).
fof(f1327,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f680,f680])).
fof(f1287,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X2)))))) )),
  inference(superposition,[],[f284,f252])).
fof(f1272,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X2)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X2)))))) )),
  inference(superposition,[],[f283,f252])).
fof(f1140,plain,(
  sP6('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f413,f402,f397,f211,f299])).
fof(f864,plain,(
  sP8('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f416,f211,f211,f321])).
fof(f755,plain,(
  ~sP6('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f211,f436,f316])).
fof(f619,plain,(
  ( ! [X0] : (sP12('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f211,f211,f365])).
fof(f585,plain,(
  ( ! [X0] : (sP10('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f211,f211,f348])).
fof(f574,plain,(
  ( ! [X0] : (sP10('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f211,f211,f347])).
fof(f538,plain,(
  ( ! [X0] : (sP2('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f211,f211,f290])).
fof(f472,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f211,f211,f273])).
fof(f433,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f211,f372])).
fof(f413,plain,(
  ( ! [X0,X1] : (sP3('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f211,f313])).
fof(f394,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f211,f259])).
fof(f391,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f211,f258])).
fof(f346,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f197])).
fof(f371,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f324])).
fof(f324,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f185])).
fof(f297,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f230,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f141])).
fof(f1888,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f210,f209,f208,f207,f206,f211,f386,f218,f226,f225,f230,f229,f252,f389,f259,f258,f390,f277,f274,f273,f272,f282,f284,f283,f290,f289,f288,f372,f286,f285,f301,f300,f299,f298,f297,f296,f295,f305,f304,f303,f302,f309,f308,f307,f306,f313,f312,f311,f310,f317,f316,f315,f314,f371,f321,f320,f319,f318,f328,f327,f326,f325,f334,f331,f330,f329,f341,f340,f339,f370,f337,f336,f335,f345,f344,f343,f342,f349,f348,f347,f346,f358,f357,f356,f369,f354,f353,f352,f362,f361,f360,f359,f366,f365,f364,f363,f391,f394,f397,f400,f401,f402,f405,f406,f407,f410,f413,f416,f419,f420,f421,f424,f427,f428,f429,f432,f433,f436,f471,f472,f475,f482,f497,f508,f517,f538,f547,f550,f556,f564,f571,f572,f574,f585,f594,f600,f604,f619,f680,f735,f734,f757,f755,f807,f827,f839,f847,f860,f864,f877,f889,f948,f964,f978,f1033,f1044,f1056,f1140,f1193,f1260,f1261,f1272,f1287,f1327,f1389,f1390,f1435,f1436,f1700,f1704,f1707,f1710,f1714,f1717,f1720,f1723,f1726,f1729,f1732,f1735,f1738,f1741,f1744,f1747,f1750,f1753,f1756,f1759,f1762,f1765,f1773,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1798,f1801,f1802,f1509,f1510,f1804,f1806,f1808,f1810,f1813,f1815,f1817,f1819,f1821,f1823,f1825,f1827,f1829,f1831,f1833,f1835,f1837,f1839,f1841,f1843,f1845,f1847,f1702,f1870,f1875,f1879,f1887])).
fof(f1884,plain,(
  $false | ($spl6 | $spl10)),
  inference(global_subsumption,[],[f1883,f210,f209,f208,f207,f206,f211,f386,f218,f226,f225,f230,f229,f252,f389,f259,f258,f390,f277,f274,f273,f272,f282,f284,f283,f290,f289,f288,f372,f286,f285,f301,f300,f299,f298,f297,f296,f295,f305,f304,f303,f302,f309,f308,f307,f306,f313,f312,f311,f310,f317,f316,f315,f314,f371,f321,f320,f319,f318,f328,f327,f326,f325,f334,f331,f330,f329,f341,f340,f339,f370,f337,f336,f335,f345,f344,f343,f342,f349,f348,f347,f346,f358,f357,f356,f369,f354,f353,f352,f362,f361,f360,f359,f366,f365,f364,f363,f391,f394,f397,f400,f401,f402,f405,f406,f407,f410,f413,f416,f419,f420,f421,f424,f427,f428,f429,f432,f433,f436,f471,f472,f475,f482,f497,f508,f517,f538,f547,f550,f556,f564,f571,f572,f574,f585,f594,f600,f604,f619,f680,f735,f734,f757,f755,f807,f827,f839,f847,f860,f864,f877,f889,f948,f964,f978,f1033,f1044,f1056,f1140,f1193,f1260,f1261,f1272,f1287,f1327,f1389,f1390,f1435,f1436,f1700,f1704,f1707,f1710,f1714,f1717,f1720,f1723,f1726,f1729,f1732,f1735,f1738,f1741,f1744,f1747,f1750,f1753,f1756,f1759,f1762,f1765,f1773,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1798,f1801,f1802,f1509,f1510,f1804,f1806,f1808,f1810,f1813,f1815,f1817,f1819,f1821,f1823,f1825,f1827,f1829,f1831,f1833,f1835,f1837,f1839,f1841,f1843,f1845,f1847,f1702,f1870,f1875,f1879])).
fof(f1883,plain,(
  ( ! [X2,X3] : (s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',sK14))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X2)))))) ) | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f1862,f211])).
fof(f1862,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal',X2))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',sK14))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X3),s('type.extreal.extreal',X2)))))) ) | ($spl6 | $spl10)),
  inference(backward_demodulation,[],[f1702,f1293])).
fof(f1882,plain,(
  $false | ($spl2 | $spl10)),
  inference(global_subsumption,[],[f1881,f210,f209,f208,f207,f206,f211,f386,f218,f226,f225,f230,f229,f252,f389,f259,f258,f390,f277,f274,f273,f272,f282,f284,f283,f290,f289,f288,f372,f286,f285,f301,f300,f299,f298,f297,f296,f295,f305,f304,f303,f302,f309,f308,f307,f306,f313,f312,f311,f310,f317,f316,f315,f314,f371,f321,f320,f319,f318,f328,f327,f326,f325,f334,f331,f330,f329,f341,f340,f339,f370,f337,f336,f335,f345,f344,f343,f342,f349,f348,f347,f346,f358,f357,f356,f369,f354,f353,f352,f362,f361,f360,f359,f366,f365,f364,f363,f391,f394,f397,f400,f401,f402,f405,f406,f407,f410,f413,f416,f419,f420,f421,f424,f427,f428,f429,f432,f433,f436,f471,f472,f475,f482,f497,f508,f517,f538,f547,f550,f556,f564,f571,f572,f574,f585,f594,f600,f604,f619,f680,f735,f734,f757,f755,f807,f827,f839,f847,f860,f864,f877,f889,f948,f964,f978,f1033,f1044,f1056,f1140,f1193,f1260,f1261,f1272,f1287,f1327,f1389,f1390,f1435,f1436,f1700,f1704,f1707,f1710,f1714,f1717,f1720,f1723,f1726,f1729,f1732,f1735,f1738,f1741,f1744,f1747,f1750,f1753,f1756,f1759,f1762,f1765,f1773,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1798,f1801,f1802,f1509,f1510,f1804,f1806,f1808,f1810,f1813,f1815,f1817,f1819,f1821,f1823,f1825,f1827,f1829,f1831,f1833,f1835,f1837,f1839,f1841,f1843,f1845,f1847,f1702,f1870,f1875,f1879])).
fof(f1881,plain,(
  ( ! [X0,X1] : (s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal',X0))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',sK15))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))) ) | ($spl2 | $spl10)),
  inference(subsumption_resolution,[],[f1861,f211])).
fof(f1861,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal',X0))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',sK15))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))) ) | ($spl2 | $spl10)),
  inference(backward_demodulation,[],[f1702,f1292])).
fof(f1292,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal',X0))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',sK15))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))) ) | $spl2),
  inference(subsumption_resolution,[],[f1289,f211])).
fof(f1289,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal',X0))))) = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',sK15))),s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))) ) | $spl2),
  inference(superposition,[],[f284,f439])).
fof(f1880,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f210,f209,f208,f207,f206,f211,f386,f218,f226,f225,f230,f229,f252,f389,f259,f258,f390,f277,f274,f273,f272,f282,f284,f283,f290,f289,f288,f372,f286,f285,f301,f300,f299,f298,f297,f296,f295,f305,f304,f303,f302,f309,f308,f307,f306,f313,f312,f311,f310,f317,f316,f315,f314,f371,f321,f320,f319,f318,f328,f327,f326,f325,f334,f331,f330,f329,f341,f340,f339,f370,f337,f336,f335,f345,f344,f343,f342,f349,f348,f347,f346,f358,f357,f356,f369,f354,f353,f352,f362,f361,f360,f359,f366,f365,f364,f363,f391,f394,f397,f400,f401,f402,f405,f406,f407,f410,f413,f416,f419,f420,f421,f424,f427,f428,f429,f432,f433,f436,f471,f472,f475,f482,f497,f508,f517,f538,f547,f550,f556,f564,f571,f572,f574,f585,f594,f600,f604,f619,f680,f735,f734,f757,f755,f807,f827,f839,f847,f860,f864,f877,f889,f948,f964,f978,f1033,f1044,f1056,f1140,f1193,f1260,f1261,f1272,f1287,f1327,f1389,f1390,f1435,f1436,f1700,f1704,f1707,f1710,f1714,f1717,f1720,f1723,f1726,f1729,f1732,f1735,f1738,f1741,f1744,f1747,f1750,f1753,f1756,f1759,f1762,f1765,f1773,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1798,f1801,f1802,f1509,f1510,f1804,f1806,f1808,f1810,f1813,f1815,f1817,f1819,f1821,f1823,f1825,f1827,f1829,f1831,f1833,f1835,f1837,f1839,f1841,f1843,f1845,f1847,f1702,f1870,f1875,f1879])).
fof(f1876,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f210,f209,f208,f207,f206,f211,f386,f218,f226,f225,f230,f229,f252,f389,f259,f258,f390,f277,f274,f273,f272,f282,f284,f283,f290,f289,f288,f372,f286,f285,f301,f300,f299,f298,f297,f296,f295,f305,f304,f303,f302,f309,f308,f307,f306,f313,f312,f311,f310,f317,f316,f315,f314,f371,f321,f320,f319,f318,f328,f327,f326,f325,f334,f331,f330,f329,f341,f340,f339,f370,f337,f336,f335,f345,f344,f343,f342,f349,f348,f347,f346,f358,f357,f356,f369,f354,f353,f352,f362,f361,f360,f359,f366,f365,f364,f363,f391,f394,f397,f400,f401,f402,f405,f406,f407,f410,f413,f416,f419,f420,f421,f424,f427,f428,f429,f432,f433,f436,f471,f472,f475,f482,f497,f508,f517,f538,f547,f550,f556,f564,f571,f572,f574,f585,f594,f600,f604,f619,f680,f735,f734,f757,f755,f807,f827,f839,f847,f860,f864,f877,f889,f948,f964,f978,f1033,f1044,f1056,f1140,f1193,f1260,f1261,f1272,f1287,f1327,f1389,f1390,f1435,f1436,f1700,f1704,f1707,f1710,f1714,f1717,f1720,f1723,f1726,f1729,f1732,f1735,f1738,f1741,f1744,f1747,f1750,f1753,f1756,f1759,f1762,f1765,f1773,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1798,f1801,f1802,f1509,f1510,f1804,f1806,f1808,f1810,f1813,f1815,f1817,f1819,f1821,f1823,f1825,f1827,f1829,f1831,f1833,f1835,f1837,f1839,f1841,f1843,f1845,f1847,f1702,f1870,f1875])).
fof(f1872,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1854,f211])).
fof(f1854,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(backward_demodulation,[],[f1702,f386])).
fof(f1871,plain,(
  $false | $spl10),
  inference(global_subsumption,[],[f210,f209,f208,f207,f206,f211,f386,f218,f226,f225,f230,f229,f252,f389,f259,f258,f390,f277,f274,f273,f272,f282,f284,f283,f290,f289,f288,f372,f286,f285,f301,f300,f299,f298,f297,f296,f295,f305,f304,f303,f302,f309,f308,f307,f306,f313,f312,f311,f310,f317,f316,f315,f314,f371,f321,f320,f319,f318,f328,f327,f326,f325,f334,f331,f330,f329,f341,f340,f339,f370,f337,f336,f335,f345,f344,f343,f342,f349,f348,f347,f346,f358,f357,f356,f369,f354,f353,f352,f362,f361,f360,f359,f366,f365,f364,f363,f391,f394,f397,f400,f401,f402,f405,f406,f407,f410,f413,f416,f419,f420,f421,f424,f427,f428,f429,f432,f433,f436,f471,f472,f475,f482,f497,f508,f517,f538,f547,f550,f556,f564,f571,f572,f574,f585,f594,f600,f604,f619,f680,f735,f734,f757,f755,f807,f827,f839,f847,f860,f864,f877,f889,f948,f964,f978,f1033,f1044,f1056,f1140,f1193,f1260,f1261,f1272,f1287,f1327,f1389,f1390,f1435,f1436,f1700,f1704,f1707,f1710,f1714,f1717,f1720,f1723,f1726,f1729,f1732,f1735,f1738,f1741,f1744,f1747,f1750,f1753,f1756,f1759,f1762,f1765,f1773,f1774,f1775,f1776,f1777,f1778,f1779,f1780,f1781,f1782,f1783,f1784,f1785,f1786,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1798,f1801,f1802,f1509,f1510,f1804,f1806,f1808,f1810,f1813,f1815,f1817,f1819,f1821,f1823,f1825,f1827,f1829,f1831,f1833,f1835,f1837,f1839,f1841,f1843,f1845,f1847,f1702,f1870])).
fof(f1772,plain,(
  ( ! [X1] : (~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1) | p(s(bool,X1))) ) | $spl56),
  inference(cnf_transformation,[],[f1772_D])).
fof(f1772_D,plain,(
  ( ! [X1] : (~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1) | p(s(bool,X1))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1702,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f1702_D])).
fof(f1702_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1851,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl6),
  inference(condensation,[],[f1850])).
fof(f1850,plain,(
  ( ! [X118,X119,X117] : (p(s(bool,X118)) | ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X118) | s(bool,X117) = s(bool,X119) | s(bool,'T_0') = s(bool,X119)) ) | $spl6),
  inference(subsumption_resolution,[],[f1544,f1436])).
fof(f1544,plain,(
  ( ! [X118,X119,X117] : (~p(s(bool,X117)) | p(s(bool,X118)) | ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X118) | s(bool,X117) = s(bool,X119) | s(bool,'T_0') = s(bool,X119)) ) | $spl6),
  inference(superposition,[],[f815,f680])).
fof(f1769,plain,(
  ( ! [X1] : (~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1) | p(s(bool,X1))) ) | $spl54),
  inference(cnf_transformation,[],[f1769_D])).
fof(f1769_D,plain,(
  ( ! [X1] : (~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1) | p(s(bool,X1))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1849,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl2),
  inference(condensation,[],[f1848])).
fof(f1848,plain,(
  ( ! [X116,X114,X115] : (p(s(bool,X115)) | ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X115) | s(bool,X114) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) ) | $spl2),
  inference(subsumption_resolution,[],[f1543,f1436])).
fof(f1543,plain,(
  ( ! [X116,X114,X115] : (~p(s(bool,X114)) | p(s(bool,X115)) | ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X115) | s(bool,X114) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) ) | $spl2),
  inference(superposition,[],[f814,f680])).
fof(f1766,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~sP9('T_0',X1,X2)) ) | $spl52),
  inference(cnf_transformation,[],[f1766_D])).
fof(f1766_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~sP9('T_0',X1,X2)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1847,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f1846])).
fof(f1846,plain,(
  ( ! [X111,X112,X110,X113] : (~sP9('T_0',X111,X112) | sP10('T_0',X111,X112) | s(bool,X110) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f1542,f1436])).
fof(f1542,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,X110)) | ~sP9('T_0',X111,X112) | sP10('T_0',X111,X112) | s(bool,X110) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f370,f680])).
fof(f1763,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~sP11('T_0',X1,X2)) ) | $spl50),
  inference(cnf_transformation,[],[f1763_D])).
fof(f1763_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~sP11('T_0',X1,X2)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1845,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f1844])).
fof(f1844,plain,(
  ( ! [X109,X107,X108,X106] : (~sP11('T_0',X107,X108) | sP12('T_0',X107,X108) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f1541,f1436])).
fof(f1541,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,X106)) | ~sP11('T_0',X107,X108) | sP12('T_0',X107,X108) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f369,f680])).
fof(f1760,plain,(
  ( ! [X2,X1] : (~sP12('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl48),
  inference(cnf_transformation,[],[f1760_D])).
fof(f1760_D,plain,(
  ( ! [X2,X1] : (~sP12('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1843,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f1842])).
fof(f1842,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,X103)) | p(s(bool,X104)) | ~sP12('T_0',X104,X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f1540,f1436])).
fof(f1540,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,X102)) | ~p(s(bool,X103)) | p(s(bool,X104)) | ~sP12('T_0',X104,X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f366,f680])).
fof(f1757,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl46),
  inference(cnf_transformation,[],[f1757_D])).
fof(f1757_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1841,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f1840])).
fof(f1840,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,X99)) | sP12('T_0',X99,X100) | s(bool,X98) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(subsumption_resolution,[],[f1539,f1436])).
fof(f1539,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,X98)) | ~p(s(bool,X99)) | sP12('T_0',X99,X100) | s(bool,X98) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f365,f680])).
fof(f1754,plain,(
  ( ! [X2,X1] : (sP12('T_0',X2,X1) | p(s(bool,X1))) ) | $spl44),
  inference(cnf_transformation,[],[f1754_D])).
fof(f1754_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1839,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f1838])).
fof(f1838,plain,(
  ( ! [X94,X97,X95,X96] : (p(s(bool,X95)) | sP12('T_0',X96,X95) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f1538,f1436])).
fof(f1538,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,X94)) | p(s(bool,X95)) | sP12('T_0',X96,X95) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f364,f680])).
fof(f1751,plain,(
  ( ! [X2,X1] : (sP12(X2,X1,'T_0') | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl42),
  inference(cnf_transformation,[],[f1751_D])).
fof(f1751_D,plain,(
  ( ! [X2,X1] : (sP12(X2,X1,'T_0') | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1837,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12(X2,X1,'T_0') | p(s(bool,X2))) )),
  inference(condensation,[],[f1836])).
fof(f1836,plain,(
  ( ! [X92,X90,X93,X91] : (p(s(bool,X91)) | sP12(X92,X91,'T_0') | p(s(bool,X92)) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f1537,f1436])).
fof(f1537,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,X90)) | p(s(bool,X91)) | sP12(X92,X91,'T_0') | p(s(bool,X92)) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f363,f680])).
fof(f1748,plain,(
  ( ! [X2,X1] : (~sP11(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl40),
  inference(cnf_transformation,[],[f1748_D])).
fof(f1748_D,plain,(
  ( ! [X2,X1] : (~sP11(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1835,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) )),
  inference(condensation,[],[f1834])).
fof(f1834,plain,(
  ( ! [X88,X87,X89,X86] : (p(s(bool,X87)) | ~p(s(bool,X88)) | ~sP11(X88,X87,'T_0') | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f1536,f1436])).
fof(f1536,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,X86)) | p(s(bool,X87)) | ~p(s(bool,X88)) | ~sP11(X88,X87,'T_0') | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f359,f680])).
fof(f1745,plain,(
  ( ! [X2,X1] : (~sP10('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl38),
  inference(cnf_transformation,[],[f1745_D])).
fof(f1745_D,plain,(
  ( ! [X2,X1] : (~sP10('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1833,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f1832])).
fof(f1832,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,X83)) | p(s(bool,X84)) | ~sP10('T_0',X84,X83) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f1535,f1436])).
fof(f1535,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,X82)) | p(s(bool,X83)) | p(s(bool,X84)) | ~sP10('T_0',X84,X83) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f349,f680])).
fof(f1742,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl36),
  inference(cnf_transformation,[],[f1742_D])).
fof(f1742_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1831,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f1830])).
fof(f1830,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X79)) | sP10('T_0',X79,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f1534,f1436])).
fof(f1534,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X78)) | ~p(s(bool,X79)) | sP10('T_0',X79,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f348,f680])).
fof(f1739,plain,(
  ( ! [X2,X1] : (sP10('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f1739_D])).
fof(f1739_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1829,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f1828])).
fof(f1828,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X75)) | sP10('T_0',X76,X75) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f1533,f1436])).
fof(f1533,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X74)) | ~p(s(bool,X75)) | sP10('T_0',X76,X75) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f347,f680])).
fof(f1736,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl32),
  inference(cnf_transformation,[],[f1736_D])).
fof(f1736_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1827,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1826])).
fof(f1826,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,X71)) | p(s(bool,X72)) | ~sP9('T_0',X71,X72) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f1532,f1436])).
fof(f1532,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X70)) | p(s(bool,X71)) | p(s(bool,X72)) | ~sP9('T_0',X71,X72) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f342,f680])).
fof(f1733,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f1733_D])).
fof(f1733_D,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1825,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f1824])).
fof(f1824,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | ~sP8(X68,X67,'T_0') | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f1531,f1436])).
fof(f1531,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X66)) | ~p(s(bool,X67)) | p(s(bool,X68)) | ~sP8(X68,X67,'T_0') | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f334,f680])).
fof(f1730,plain,(
  ( ! [X2,X1] : (sP8('T_0',X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2))) ) | $spl28),
  inference(cnf_transformation,[],[f1730_D])).
fof(f1730_D,plain,(
  ( ! [X2,X1] : (sP8('T_0',X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1823,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f1822])).
fof(f1822,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | sP8('T_0',X63,X64) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f1530,f1436])).
fof(f1530,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X62)) | ~p(s(bool,X63)) | ~p(s(bool,X64)) | sP8('T_0',X63,X64) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f331,f680])).
fof(f1727,plain,(
  ( ! [X2,X1] : (~sP7(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl26),
  inference(cnf_transformation,[],[f1727_D])).
fof(f1727_D,plain,(
  ( ! [X2,X1] : (~sP7(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1821,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X1,'T_0',X2)) )),
  inference(condensation,[],[f1820])).
fof(f1820,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X59)) | p(s(bool,X60)) | ~sP7(X59,'T_0',X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f1529,f1436])).
fof(f1529,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X58)) | ~p(s(bool,X59)) | p(s(bool,X60)) | ~sP7(X59,'T_0',X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f325,f680])).
fof(f1724,plain,(
  ( ! [X2,X1] : (~sP6('T_0',X2,X1) | s(bool,X1) = s(bool,X2)) ) | $spl24),
  inference(cnf_transformation,[],[f1724_D])).
fof(f1724_D,plain,(
  ( ! [X2,X1] : (~sP6('T_0',X2,X1) | s(bool,X1) = s(bool,X2)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1819,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | ~sP6('T_0',X2,X1)) )),
  inference(condensation,[],[f1818])).
fof(f1818,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X55) = s(bool,X56) | ~sP6('T_0',X56,X55) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f1528,f1436])).
fof(f1528,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X54)) | s(bool,X55) = s(bool,X56) | ~sP6('T_0',X56,X55) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f316,f680])).
fof(f1721,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl22),
  inference(cnf_transformation,[],[f1721_D])).
fof(f1721_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1817,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1816])).
fof(f1816,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP4(X51,'T_0',X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f1527,f1436])).
fof(f1527,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP4(X51,'T_0',X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f306,f680])).
fof(f1718,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f1718_D])).
fof(f1718_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1815,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f1814])).
fof(f1814,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | p(s(bool,X48)) | ~sP5(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f1526,f1436])).
fof(f1526,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP5(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f302,f680])).
fof(f1715,plain,(
  ( ! [X2,X1] : (~sP4(X2,'T_0',X1) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2)) ) | $spl18),
  inference(cnf_transformation,[],[f1715_D])).
fof(f1715_D,plain,(
  ( ! [X2,X1] : (~sP4(X2,'T_0',X1) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1813,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f1812,f413])).
fof(f1812,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1) | ~sP3('T_0',X2,X1)) )),
  inference(condensation,[],[f1811])).
fof(f1811,plain,(
  ( ! [X45,X43,X44,X42] : (sP6(X43,X44,'T_0') | ~sP5(X43,'T_0',X44) | ~sP4(X44,'T_0',X43) | ~sP3('T_0',X44,X43) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f1525,f1436])).
fof(f1525,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X42)) | sP6(X43,X44,'T_0') | ~sP5(X43,'T_0',X44) | ~sP4(X44,'T_0',X43) | ~sP3('T_0',X44,X43) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f299,f680])).
fof(f1711,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f1711_D])).
fof(f1711_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1810,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1809])).
fof(f1809,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X39)) | p(s(bool,X40)) | ~sP6(X39,'T_0',X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f1524,f1436])).
fof(f1524,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X38)) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~sP6(X39,'T_0',X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f298,f680])).
fof(f1708,plain,(
  ( ! [X2,X1] : (sP2(X1,X2,'T_0') | ~p(s(bool,X1))) ) | $spl14),
  inference(cnf_transformation,[],[f1708_D])).
fof(f1708_D,plain,(
  ( ! [X2,X1] : (sP2(X1,X2,'T_0') | ~p(s(bool,X1))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1808,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2(X1,X2,'T_0')) )),
  inference(condensation,[],[f1807])).
fof(f1807,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X35)) | sP2(X35,X36,'T_0') | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f1523,f1436])).
fof(f1523,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | sP2(X35,X36,'T_0') | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f290,f680])).
fof(f1705,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl12),
  inference(cnf_transformation,[],[f1705_D])).
fof(f1705_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1806,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1805])).
fof(f1805,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X24)) | sP1(X24,'T_0') | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f1520,f1436])).
fof(f1520,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | p(s(bool,X24)) | sP1(X24,'T_0') | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f277,f680])).
fof(f1701,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl8),
  inference(cnf_transformation,[],[f1701_D])).
fof(f1701_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1804,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1803])).
fof(f1803,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | sP1('T_0',X21) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f1519,f1436])).
fof(f1519,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | sP1('T_0',X21) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f274,f680])).
fof(f1771,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl6),
  inference(condensation,[],[f1770])).
fof(f1770,plain,(
  ( ! [X118,X119,X117] : (p(s(bool,X118)) | ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X118) | s(bool,X117) = s(bool,X119) | s(bool,'T_0') = s(bool,X119)) ) | $spl6),
  inference(subsumption_resolution,[],[f1468,f1436])).
fof(f1468,plain,(
  ( ! [X118,X119,X117] : (~p(s(bool,X117)) | p(s(bool,X118)) | ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X118) | s(bool,X117) = s(bool,X119) | s(bool,'T_0') = s(bool,X119)) ) | $spl6),
  inference(superposition,[],[f815,f680])).
fof(f1768,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1)) ) | $spl2),
  inference(condensation,[],[f1767])).
fof(f1767,plain,(
  ( ! [X116,X114,X115] : (p(s(bool,X115)) | ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X115) | s(bool,X114) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) ) | $spl2),
  inference(subsumption_resolution,[],[f1467,f1436])).
fof(f1467,plain,(
  ( ! [X116,X114,X115] : (~p(s(bool,X114)) | p(s(bool,X115)) | ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),X115) | s(bool,X114) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) ) | $spl2),
  inference(superposition,[],[f814,f680])).
fof(f1765,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f1764])).
fof(f1764,plain,(
  ( ! [X111,X112,X110,X113] : (~sP9('T_0',X111,X112) | sP10('T_0',X111,X112) | s(bool,X110) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(subsumption_resolution,[],[f1466,f1436])).
fof(f1466,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,X110)) | ~sP9('T_0',X111,X112) | sP10('T_0',X111,X112) | s(bool,X110) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f370,f680])).
fof(f1762,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f1761])).
fof(f1761,plain,(
  ( ! [X109,X107,X108,X106] : (~sP11('T_0',X107,X108) | sP12('T_0',X107,X108) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f1465,f1436])).
fof(f1465,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,X106)) | ~sP11('T_0',X107,X108) | sP12('T_0',X107,X108) | s(bool,X106) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f369,f680])).
fof(f1759,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f1758])).
fof(f1758,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,X103)) | p(s(bool,X104)) | ~sP12('T_0',X104,X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f1464,f1436])).
fof(f1464,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,X102)) | ~p(s(bool,X103)) | p(s(bool,X104)) | ~sP12('T_0',X104,X103) | s(bool,X102) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f366,f680])).
fof(f1756,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f1755])).
fof(f1755,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,X99)) | sP12('T_0',X99,X100) | s(bool,X98) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(subsumption_resolution,[],[f1463,f1436])).
fof(f1463,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,X98)) | ~p(s(bool,X99)) | sP12('T_0',X99,X100) | s(bool,X98) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f365,f680])).
fof(f1753,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f1752])).
fof(f1752,plain,(
  ( ! [X94,X97,X95,X96] : (p(s(bool,X95)) | sP12('T_0',X96,X95) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(subsumption_resolution,[],[f1462,f1436])).
fof(f1462,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,X94)) | p(s(bool,X95)) | sP12('T_0',X96,X95) | s(bool,X94) = s(bool,X97) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f364,f680])).
fof(f1750,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12(X2,X1,'T_0') | p(s(bool,X2))) )),
  inference(condensation,[],[f1749])).
fof(f1749,plain,(
  ( ! [X92,X90,X93,X91] : (p(s(bool,X91)) | sP12(X92,X91,'T_0') | p(s(bool,X92)) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f1461,f1436])).
fof(f1461,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,X90)) | p(s(bool,X91)) | sP12(X92,X91,'T_0') | p(s(bool,X92)) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f363,f680])).
fof(f1747,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) )),
  inference(condensation,[],[f1746])).
fof(f1746,plain,(
  ( ! [X88,X87,X89,X86] : (p(s(bool,X87)) | ~p(s(bool,X88)) | ~sP11(X88,X87,'T_0') | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f1460,f1436])).
fof(f1460,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,X86)) | p(s(bool,X87)) | ~p(s(bool,X88)) | ~sP11(X88,X87,'T_0') | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f359,f680])).
fof(f1744,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f1743])).
fof(f1743,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,X83)) | p(s(bool,X84)) | ~sP10('T_0',X84,X83) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f1459,f1436])).
fof(f1459,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,X82)) | p(s(bool,X83)) | p(s(bool,X84)) | ~sP10('T_0',X84,X83) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f349,f680])).
fof(f1741,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f1740])).
fof(f1740,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X79)) | sP10('T_0',X79,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f1458,f1436])).
fof(f1458,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X78)) | ~p(s(bool,X79)) | sP10('T_0',X79,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f348,f680])).
fof(f1738,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f1737])).
fof(f1737,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X75)) | sP10('T_0',X76,X75) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f1457,f1436])).
fof(f1457,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X74)) | ~p(s(bool,X75)) | sP10('T_0',X76,X75) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f347,f680])).
fof(f1735,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1734])).
fof(f1734,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,X71)) | p(s(bool,X72)) | ~sP9('T_0',X71,X72) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f1456,f1436])).
fof(f1456,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X70)) | p(s(bool,X71)) | p(s(bool,X72)) | ~sP9('T_0',X71,X72) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f342,f680])).
fof(f1732,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f1731])).
fof(f1731,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | ~sP8(X68,X67,'T_0') | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f1455,f1436])).
fof(f1455,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X66)) | ~p(s(bool,X67)) | p(s(bool,X68)) | ~sP8(X68,X67,'T_0') | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f334,f680])).
fof(f1729,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f1728])).
fof(f1728,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X63)) | ~p(s(bool,X64)) | sP8('T_0',X63,X64) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f1454,f1436])).
fof(f1454,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X62)) | ~p(s(bool,X63)) | ~p(s(bool,X64)) | sP8('T_0',X63,X64) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f331,f680])).
fof(f1726,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X1,'T_0',X2)) )),
  inference(condensation,[],[f1725])).
fof(f1725,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X59)) | p(s(bool,X60)) | ~sP7(X59,'T_0',X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f1453,f1436])).
fof(f1453,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X58)) | ~p(s(bool,X59)) | p(s(bool,X60)) | ~sP7(X59,'T_0',X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f325,f680])).
fof(f1723,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | ~sP6('T_0',X2,X1)) )),
  inference(condensation,[],[f1722])).
fof(f1722,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X55) = s(bool,X56) | ~sP6('T_0',X56,X55) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f1452,f1436])).
fof(f1452,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X54)) | s(bool,X55) = s(bool,X56) | ~sP6('T_0',X56,X55) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f316,f680])).
fof(f1720,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1719])).
fof(f1719,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP4(X51,'T_0',X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f1451,f1436])).
fof(f1451,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP4(X51,'T_0',X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f306,f680])).
fof(f1717,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f1716])).
fof(f1716,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | p(s(bool,X48)) | ~sP5(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f1450,f1436])).
fof(f1450,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP5(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f302,f680])).
fof(f1714,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f1713,f413])).
fof(f1713,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1) | ~sP3('T_0',X2,X1)) )),
  inference(condensation,[],[f1712])).
fof(f1712,plain,(
  ( ! [X45,X43,X44,X42] : (sP6(X43,X44,'T_0') | ~sP5(X43,'T_0',X44) | ~sP4(X44,'T_0',X43) | ~sP3('T_0',X44,X43) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f1449,f1436])).
fof(f1449,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X42)) | sP6(X43,X44,'T_0') | ~sP5(X43,'T_0',X44) | ~sP4(X44,'T_0',X43) | ~sP3('T_0',X44,X43) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f299,f680])).
fof(f1710,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1709])).
fof(f1709,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X39)) | p(s(bool,X40)) | ~sP6(X39,'T_0',X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f1448,f1436])).
fof(f1448,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X38)) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~sP6(X39,'T_0',X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f298,f680])).
fof(f1707,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP2(X1,X2,'T_0')) )),
  inference(condensation,[],[f1706])).
fof(f1706,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X35)) | sP2(X35,X36,'T_0') | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f1447,f1436])).
fof(f1447,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | sP2(X35,X36,'T_0') | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f290,f680])).
fof(f1704,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1703])).
fof(f1703,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X24)) | sP1(X24,'T_0') | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f1444,f1436])).
fof(f1444,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | p(s(bool,X24)) | sP1(X24,'T_0') | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f277,f680])).
fof(f1700,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1699])).
fof(f1699,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | sP1('T_0',X21) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f1443,f1436])).
fof(f1443,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | sP1('T_0',X21) | s(bool,X20) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f274,f680])).
fof(f383,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK14)))) | $spl4),
  inference(cnf_transformation,[],[f383_D])).
fof(f383_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK14)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f385,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl6),
  inference(cnf_transformation,[],[f385_D])).
fof(f385_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f206,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK14))))),
  inference(cnf_transformation,[],[f136])).
fof(f379,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK15)))) | $spl0),
  inference(cnf_transformation,[],[f379_D])).
fof(f379_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK15)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f207,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK15))))),
  inference(cnf_transformation,[],[f136])).
fof(f381,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl2),
  inference(cnf_transformation,[],[f381_D])).
fof(f381_D,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f208,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK14))))),
  inference(cnf_transformation,[],[f136])).
fof(f209,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK15),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK15))))),
  inference(cnf_transformation,[],[f136])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 465
% Passive clauses: 850
% Generated clauses: 6457
% Final active clauses: 336
% Final passive clauses: 244
% Input formulas: 36
% Initial clauses: 163
% 
% Duplicate literals: 212
% Trivial inequalities: 24
% Fw subsumption resolutions: 438
% Fw demodulations: 496
% Bw demodulations: 29
% Condensations: 149
% Global subsumptions: 9
% 
% Simple tautologies: 53
% Equational tautologies: 103
% Forward subsumptions: 3678
% Bw demodulations to eq. taut.: 15
% 
% Unit resulting resolution: 3754
% Forward superposition: 684
% Backward superposition: 367
% Self superposition: 26
% Equality factoring: 26
% Equality resolution: 2
% 
% Split clauses: 66
% Split components: 137
% SAT solver clauses: 1372
% SAT solver unit clauses: 302
% SAT solver binary clauses: 486
% 
% Sat splits: 66
% Sat splitting refutations: 379
% 
% Memory used [KB]: 3454
% Time elapsed: 0.182 s
% ------------------------------
% ------------------------------
% Success in time 1.48 s
