Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.transc
% (41701)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 435
% Passive clauses: 4615
% Generated clauses: 35205
% Final active clauses: 289
% Final passive clauses: 3420
% Input formulas: 261
% Initial clauses: 391
% 
% Duplicate literals: 172
% Trivial inequalities: 170
% Fw subsumption resolutions: 1411
% Fw demodulations: 5820
% Bw demodulations: 127
% 
% Simple tautologies: 124
% Equational tautologies: 35
% Forward subsumptions: 23008
% Fw demodulations to eq. taut.: 14
% 
% Binary resolution: 613
% Factoring: 1
% Forward superposition: 14395
% Backward superposition: 11790
% Self superposition: 186
% Equality factoring: 30
% Equality resolution: 31
% 
% Split clauses: 83
% Split components: 167
% SAT solver clauses: 196
% SAT solver unit clauses: 6
% SAT solver binary clauses: 63
% 
% Sat splits: 83
% Sat splitting refutations: 77
% 
% TWLsolver clauses: 196
% TWLsolver calls for satisfiability: 31
% 
% Memory used [KB]: 20980
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.transc
% (41735)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 305
% Passive clauses: 2533
% Generated clauses: 29740
% Final active clauses: 116
% Final passive clauses: 684
% Input formulas: 261
% Initial clauses: 391
% 
% Duplicate literals: 91
% Trivial inequalities: 725
% Fw subsumption resolutions: 1542
% Fw demodulations: 5899
% Bw demodulations: 1627
% Condensations: 208
% Global subsumptions: 29
% 
% Simple tautologies: 131
% Equational tautologies: 22
% Forward subsumptions: 15635
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 22
% 
% Binary resolution: 262
% Factoring: 1
% Forward superposition: 7330
% Backward superposition: 7981
% Self superposition: 187
% Equality factoring: 16
% Equality resolution: 19
% 
% Split clauses: 97
% Split components: 195
% SAT solver clauses: 3023
% SAT solver unit clauses: 582
% SAT solver binary clauses: 156
% SAT solver learnt clauses: 58
% 
% Sat splits: 97
% Sat splitting refutations: 87
% 
% TWLsolver clauses: 2907
% TWLsolver calls for satisfiability: 2302
% 
% Memory used [KB]: 19445
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.transc
% (41754)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1231
% Passive clauses: 20134
% Generated clauses: 229842
% Final active clauses: 621
% Final passive clauses: 2726
% Input formulas: 261
% Initial clauses: 391
% Discarded non-redundant clauses: 32051
% 
% Duplicate literals: 48
% Trivial inequalities: 518
% Fw subsumption resolutions: 3290
% Fw demodulations: 67815
% Bw demodulations: 5054
% Condensations: 172
% Global subsumptions: 5
% 
% Simple tautologies: 159
% Equational tautologies: 481
% Forward subsumptions: 127357
% Fw demodulations to eq. taut.: 753
% Bw demodulations to eq. taut.: 19
% 
% Unit resulting resolution: 31707
% Factoring: 2
% Forward superposition: 65222
% Backward superposition: 49062
% Self superposition: 303
% Equality factoring: 8
% Equality resolution: 34
% 
% Split clauses: 49
% Split components: 98
% SAT solver clauses: 20250
% SAT solver unit clauses: 2140
% SAT solver binary clauses: 5054
% 
% Sat splits: 49
% Sat splitting refutations: 11
% 
% Memory used [KB]: 165924
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.transc
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.transc
% SZS output start Proof for thm.transc
fof(f31697,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f632,f643,f686,f645,f646,f647,f694,f662,f663,f702,f700,f734,f743,f744,f796,f797,f814,f854,f837,f838,f861,f876,f971,f973,f879,f977,f975,f905,f988,f913,f916,f917,f991,f921,f996,f951,f959,f1005,f962,f1008,f963,f1009,f1091,f1134,f1150,f1191,f688,f1190,f1194,f1422,f1424,f1196,f1426,f1197,f1198,f1136,f1428,f2175,f2176,f2398,f2190,f2191,f2403,f2192,f2193,f2409,f2668,f2607,f2606,f2672,f2670,f1061,f1060,f1189,f925,f924,f2674,f2605,f2604,f2678,f923,f995,f2754,f2756,f2767,f2768,f2988,f2782,f2783,f2991,f2784,f2785,f2993,f2407,f3194,f3195,f692,f3234,f852,f3463,f3464,f3615,f853,f3655,f2396,f3656,f3685,f783,f3741,f3784,f3782,f4011,f4014,f4029,f4030,f4110,f4113,f4116,f4119,f3671,f2405,f3670,f2401,f3637,f3636,f858,f2683,f811,f4363,f5080,f5456,f5538,f5681,f5682,f2676,f6964,f7002,f7082,f7120,f7446,f7566,f7777,f7567,f7779,f6493,f7961,f6489,f6488,f7963,f5086,f1421,f8187,f9136,f9123,f9141,f9139,f9317,f9508,f9542,f10247,f10985,f994,f11435,f11439,f11482,f11484,f11875,f5882,f12265,f12373,f12617,f12619,f12544,f12623,f12621,f12554,f12625,f12562,f989,f12565,f992,f12566,f12627,f993,f12596,f1004,f12604,f12607,f12608,f1007,f12758,f12762,f13341,f13445,f13355,f13373,f13460,f13376,f13461,f13377,f13453,f13462,f14244,f14247,f14487,f14589,f14587,f14648,f14766,f14791,f15036,f14793,f15037,f15040,f14796,f15021,f987,f15120,f15124,f15121,f15149,f15152,f15154,f15157,f14929,f15172,f14934,f15177,f15182,f15184,f15183,f14945,f15123,f15189,f15224,f15228,f15155,f15230,f15232,f15150,f15236,f15300,f15298,f15276,f12643,f15284,f12644,f15287,f12645,f15288,f15320,f12646,f15322,f15368,f15383,f15370,f15226,f15483,f17030,f17123,f17153,f17709,f17711,f17710,f17974,f18021,f18046,f7004,f18102,f18245,f18301,f18299,f18363,f7122,f18419,f18449,f18505,f18503,f18626,f20187,f20279,f2680,f20285,f20286,f13452,f22452,f23134,f24538,f24539,f24540,f14688,f24543,f24599,f24770,f24805,f17151,f24997,f24998,f25002,f25160,f25216,f25161,f14690,f24596,f26002,f26058,f26794,f26940,f22494,f22492,f28338,f28395,f28393,f28421,f28477,f28582,f23176,f23174,f28600,f28657,f28655,f29250,f14799,f29329,f14795,f684,f29364,f680,f29376,f29377,f685,f29409,f29464,f29462,f29421,f683,f29422,f15038,f690,f29445,f15318,f1003,f29471,f29472,f29493,f778,f776,f29543,f30242,f30282,f30285,f30292,f30332,f30330,f30476,f30478,f30480,f30482,f30483,f30486,f30505,f5577,f5575,f30634,f15039,f678,f30880,f30921,f30919,f31073,f7499,f7497,f31685,f31689])).
fof(f31689,plain,(
  $false | $spl33),
  inference(subsumption_resolution,[],[f31688,f636])).
fof(f636,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num',sK4))))),
  inference(resolution,[],[f473,f573])).
fof(f573,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f446])).
fof(f446,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f330])).
fof(f330,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f329])).
fof(f329,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X19,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num',X9)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X19)))))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.NOT_LESS')).
fof(f473,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))))))),
  inference(cnf_transformation,[],[f414])).
fof(f414,plain,(
  p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK2),s('type.realax.real',sK3))),s(fun('type.num.num','type.realax.real'),sK1)))) & (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4])],[f368])).
fof(f368,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) & ? [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) => ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s(fun('type.num.num','type.realax.real'),X6)))) => ! [X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X6)))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s(fun('type.num.num','type.realax.real'),X6)))) => ! [X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X6)))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))))))))),
  file('thm.transc.DIVISION_LT','thm.transc.DIVISION_LT')).
fof(f31688,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num',sK4)))) | $spl33),
  inference(forward_demodulation,[],[f31635,f534])).
fof(f534,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f173])).
fof(f173,axiom,(
  ! [X9,X19] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X19))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.LESS_EQ_MONO')).
fof(f31635,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))) | $spl33),
  inference(resolution,[],[f973,f522])).
fof(f522,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f433])).
fof(f433,plain,(
  ! [X0,X1] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f288])).
fof(f288,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f143])).
fof(f143,axiom,(
  ! [X9,X19] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num',X19)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X9)))))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.NOT_SUC_LESS_EQ')).
fof(f31685,plain,(
  $false | $spl33),
  inference(subsumption_resolution,[],[f31629,f636])).
fof(f31629,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num',sK4)))) | $spl33),
  inference(resolution,[],[f973,f522])).
fof(f7497,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl131),
  inference(cnf_transformation,[],[f7497_D])).
fof(f7497_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f7499,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl132),
  inference(cnf_transformation,[],[f7499_D])).
fof(f7499_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f31073,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl21),
  inference(resolution,[],[f776,f524])).
fof(f524,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f435])).
fof(f435,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) & ((s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))),
  inference(flattening,[],[f434])).
fof(f434,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))) & ((s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) & ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) <=> (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(rectify,[],[f105])).
fof(f105,axiom,(
  ! [X19,X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))))) <=> (s('type.num.num',X19) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X9))))))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.arithmetic_2|49|7_prove_ltb')).
fof(f30919,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20(sK20(sK4))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl289),
  inference(cnf_transformation,[],[f30919_D])).
fof(f30919_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20(sK20(sK4))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl289),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl289])])).
fof(f30921,plain,(
  s('type.num.num',sK20(sK20(sK4))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl290),
  inference(cnf_transformation,[],[f30921_D])).
fof(f30921_D,plain,(
  s('type.num.num',sK20(sK20(sK4))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f30880,plain,(
  s('type.num.num',sK20(sK20(sK4))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20(sK20(sK4))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl137),
  inference(resolution,[],[f7779,f524])).
fof(f678,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl0),
  inference(cnf_transformation,[],[f678_D])).
fof(f678_D,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.num.num','type.realax.real'),sK1))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f15039,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl190),
  inference(cnf_transformation,[],[f15039_D])).
fof(f15039_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f30634,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl6),
  inference(forward_demodulation,[],[f30589,f704])).
fof(f704,plain,(
  s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')))),
  inference(resolution,[],[f472,f477])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X0) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2))))) )),
  inference(cnf_transformation,[],[f418])).
fof(f418,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) | (s('type.realax.real',X0) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.|48|_0'))) & (! [X8] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',sK5(X2,X1,X0))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X8))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))))) & ! [X7] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X7),s('type.num.num',sK5(X2,X1,X0))))) | s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X7))))))) & (s('type.realax.real',X0) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.|48|_0'))) | ! [X3] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6(X3,X2,X1,X0)),s('type.num.num',X3)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',sK6(X3,X2,X1,X0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK6(X3,X2,X1,X0)))))))))) | (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK7(X3,X2,X1,X0)),s('type.num.num',X3)))) & s('type.realax.real',X1) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',sK7(X3,X2,X1,X0)))))) | p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5,sK6,sK7])],[f417])).
fof(f417,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) | (s('type.realax.real',X0) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.|48|_0'))) & ? [X6] : (! [X8] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X6)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X8))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))))) & ! [X7] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X7),s('type.num.num',X6)))) | s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X7))))))) & (s('type.realax.real',X0) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.|48|_0'))) | ! [X3] : (? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X3)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))))))) | ? [X4] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',X3)))) & s('type.realax.real',X1) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4))))) | p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2))))))),
  inference(rectify,[],[f416])).
fof(f416,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) | (s('type.realax.real',X0) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.|48|_0'))) & ? [X3] : (! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X3)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))))))) & ! [X4] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',X3)))) | s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4))))))) & (s('type.realax.real',X0) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.|48|_0'))) | ! [X3] : (? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X3)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))))))) | ? [X4] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',X3)))) & s('type.realax.real',X1) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4))))) | p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2))))))),
  inference(flattening,[],[f415])).
fof(f415,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) | (s('type.realax.real',X0) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.|48|_0'))) & ? [X3] : (! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X3)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))))))) & ! [X4] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',X3)))) | s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4))))))) & ((s('type.realax.real',X0) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.|48|_0'))) | ! [X3] : (? [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X3)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))))))) | ? [X4] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',X3)))) & s('type.realax.real',X1) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4)))))) | p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2))))))),
  inference(nnf_transformation,[],[f373])).
fof(f373,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) <=> (s('type.realax.real',X0) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.|48|_0'))) & ? [X3] : (! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X3)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))))))) & ! [X4] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',X3)))) | s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4)))))))),
  inference(ennf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2)))) <=> (s('type.realax.real',X0) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.|48|_0'))) & ? [X3] : (! [X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X3)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))))))) & ! [X4] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',X3)))) => s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4)))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X7,X8,X6] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s(fun('type.num.num','type.realax.real'),X6)))) <=> (s('type.realax.real',X7) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num','const.num.|48|_0'))) & ? [X18] : (! [X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X18)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X9))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))))))) & ! [X9] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X9),s('type.num.num',X18)))) => s('type.realax.real',X8) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X9)))))))),
  file('thm.transc.DIVISION_LT','thm.transc.division')).
fof(f472,plain,(
  p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK2),s('type.realax.real',sK3))),s(fun('type.num.num','type.realax.real'),sK1))))),
  inference(cnf_transformation,[],[f414])).
fof(f30589,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl6),
  inference(resolution,[],[f685,f485])).
fof(f485,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X6))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(cnf_transformation,[],[f422])).
fof(f422,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) | (s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) & ! [X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X6))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))))) & ! [X5] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X5),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X5)))))) & (s('type.realax.real',X1) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK8(X2,X1,X0)),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK8(X2,X1,X0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK8(X2,X1,X0)))))))))) | (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK9(X2,X1,X0)),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & s('type.realax.real',X2) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK9(X2,X1,X0))))) | p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9])],[f421])).
fof(f421,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) | (s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) & ! [X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X6))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))))) & ! [X5] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X5),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X5)))))) & (s('type.realax.real',X1) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) | ? [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) | ? [X3] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & s('type.realax.real',X2) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3)))) | p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0))))))),
  inference(rectify,[],[f420])).
fof(f420,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) | (s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) & ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) & ! [X3] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3)))))) & (s('type.realax.real',X1) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) | ? [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) | ? [X3] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & s('type.realax.real',X2) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3)))) | p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0))))))),
  inference(flattening,[],[f419])).
fof(f419,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) | (s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) & ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) & ! [X3] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3)))))) & ((s('type.realax.real',X1) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) | ? [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) | ? [X3] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) & s('type.realax.real',X2) != s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))))) | p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0))))))),
  inference(nnf_transformation,[],[f374])).
fof(f374,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) <=> (s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) & ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) & ! [X3] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))))))),
  inference(ennf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) <=> (s('type.realax.real',X1) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.|48|_0'))) & ! [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))))))) & ! [X3] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) => s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7,X8] : (p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s(fun('type.num.num','type.realax.real'),X6)))) <=> (s('type.realax.real',X7) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num','const.num.|48|_0'))) & ! [X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X6)))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X9))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))))))) & ! [X9] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X9),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X6)))))) => s('type.realax.real',X8) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X9))))))),
  file('thm.transc.DIVISION_LT','thm.transc.DIVISION_THM')).
fof(f5575,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20(sK4)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl117),
  inference(cnf_transformation,[],[f5575_D])).
fof(f5575_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20(sK4)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f5577,plain,(
  s('type.num.num',sK20(sK4)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl118),
  inference(cnf_transformation,[],[f5577_D])).
fof(f5577_D,plain,(
  s('type.num.num',sK20(sK4)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f30505,plain,(
  s('type.num.num',sK20(sK4)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20(sK4)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl75),
  inference(resolution,[],[f1426,f524])).
fof(f30486,plain,(
  $false | ($spl7 | $spl26)),
  inference(subsumption_resolution,[],[f30485,f13531])).
fof(f13531,plain,(
  p(s(bool,'T_0')) | $spl26),
  inference(superposition,[],[f13258,f568])).
fof(f568,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f251])).
fof(f251,axiom,(
  ! [X9] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X9)))),
  file('thm.transc.DIVISION_LT','thm.numeral.numeral_distrib|split|26')).
fof(f13258,plain,(
  ( ! [X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X7))))) ) | $spl26),
  inference(subsumption_resolution,[],[f13227,f12884])).
fof(f12884,plain,(
  ~p(s(bool,'F_0')) | $spl26),
  inference(subsumption_resolution,[],[f12847,f535])).
fof(f535,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X0] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f140])).
fof(f140,axiom,(
  ! [X9] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.SUC_NOT')).
fof(f12847,plain,(
  ( ! [X2] : (~p(s(bool,'F_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK21(X2,'const.num.|48|_0')),s('type.num.num','const.num.|48|_0')))))))) ) | $spl26),
  inference(superposition,[],[f12754,f569])).
fof(f569,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f237])).
fof(f237,axiom,(
  ! [X9] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0')))),
  file('thm.transc.DIVISION_LT','thm.numeral.numeral_distrib|split|20')).
fof(f12754,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK21(X1,X0)),s('type.num.num','const.num.|48|_0')))))))) ) | $spl26),
  inference(forward_demodulation,[],[f12753,f532])).
fof(f532,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0,X1] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X9,X19] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X9))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.ADD_CLAUSES|split|3')).
fof(f12753,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK21(X1,X0)),s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl26),
  inference(forward_demodulation,[],[f12752,f532])).
fof(f12752,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK21(X1,X0)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl26),
  inference(forward_demodulation,[],[f12751,f4194])).
fof(f4194,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | $spl26),
  inference(backward_demodulation,[],[f810,f546])).
fof(f546,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f104])).
fof(f104,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.ONE')).
fof(f810,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') | $spl26),
  inference(cnf_transformation,[],[f810_D])).
fof(f810_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f12751,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK21(X1,X0)),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl26),
  inference(forward_demodulation,[],[f577,f810])).
fof(f577,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK21(X1,X0)),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f449])).
fof(f449,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK21(X1,X0)),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21])],[f398])).
fof(f398,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ? [X2] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(ennf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) => ? [X2] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X19,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X19)))) => ? [X21] : s('type.num.num',X19) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.LESS_ADD_|49|')).
fof(f13227,plain,(
  ( ! [X7] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X7))))) ) | $spl26),
  inference(superposition,[],[f572,f13057])).
fof(f13057,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X2)))) ) | $spl26),
  inference(superposition,[],[f12933,f509])).
fof(f509,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f196])).
fof(f196,axiom,(
  ! [X19,X9] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num',X9))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X19))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X9)))))),
  file('thm.transc.DIVISION_LT','thm.real.REAL_LT')).
fof(f12933,plain,(
  ( ! [X112] : (s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X112),s('type.realax.real',X112)))) ) | $spl26),
  inference(resolution,[],[f12885,f510])).
fof(f510,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0] : ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0] : ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0))))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  ! [X14] : ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',X14))))),
  file('thm.transc.DIVISION_LT','thm.realax.REAL_LT_REFL')).
fof(f12885,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl26),
  inference(resolution,[],[f12884,f593])).
fof(f593,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f406])).
fof(f406,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f405])).
fof(f405,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X35,X36] : ((p(s(bool,X35)) => p(s(bool,X36))) => ((p(s(bool,X36)) => p(s(bool,X35))) => s(bool,X35) = s(bool,X36)))),
  file('thm.transc.DIVISION_LT','thm.bool.IMP_ANTISYM_AX')).
fof(f572,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f446])).
fof(f30485,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl26)),
  inference(forward_demodulation,[],[f30484,f13619])).
fof(f13619,plain,(
  ( ! [X56] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X56)))))) ) | $spl26),
  inference(resolution,[],[f13539,f543])).
fof(f543,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f261])).
fof(f261,axiom,(
  ! [X9] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))))),
  file('thm.transc.DIVISION_LT','thm.prim_rec.LESS_|48|')).
fof(f13539,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(resolution,[],[f13531,f596])).
fof(f596,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f406])).
fof(f30484,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))))) | $spl7),
  inference(subsumption_resolution,[],[f30438,f14643])).
fof(f14643,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))),
  inference(resolution,[],[f698,f557])).
fof(f557,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f443])).
fof(f443,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))))),
  inference(nnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f103])).
fof(f103,axiom,(
  ! [X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X9))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.LESS_EQ_|48|')).
fof(f698,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.|48|_0'))))),
  inference(subsumption_resolution,[],[f661,f553])).
fof(f553,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f315])).
fof(f315,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f314])).
fof(f314,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X9] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0'))))),
  file('thm.transc.DIVISION_LT','thm.prim_rec.NOT_LESS_|48|')).
fof(f661,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f473,f556])).
fof(f556,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f443])).
fof(f30438,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))))) | $spl7),
  inference(resolution,[],[f684,f527])).
fof(f527,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) )),
  inference(cnf_transformation,[],[f437])).
fof(f437,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) & ((s('type.num.num',X0) != s('type.num.num',X1) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))),
  inference(flattening,[],[f436])).
fof(f436,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | (s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))))) & ((s('type.num.num',X0) != s('type.num.num',X1) & ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) <=> (s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(rectify,[],[f141])).
fof(f141,axiom,(
  ! [X19,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))))) <=> (s('type.num.num',X9) = s('type.num.num',X19) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num',X9))))))),
  file('thm.transc.DIVISION_LT','thm.prim_rec.LESS_THM')).
fof(f30483,plain,(
  $false | $spl7),
  inference(subsumption_resolution,[],[f30436,f698])).
fof(f30436,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.|48|_0')))) | $spl7),
  inference(resolution,[],[f684,f572])).
fof(f30482,plain,(
  $false | $spl7),
  inference(subsumption_resolution,[],[f30481,f570])).
fof(f570,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f328])).
fof(f328,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))),
  inference(rectify,[],[f131])).
fof(f131,axiom,(
  ! [X9] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X9))))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.ZERO_LESS_EQ')).
fof(f30481,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl7),
  inference(subsumption_resolution,[],[f30435,f14643])).
fof(f30435,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl7),
  inference(resolution,[],[f684,f574])).
fof(f574,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f448])).
fof(f448,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(flattening,[],[f447])).
fof(f447,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1))) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(rectify,[],[f132])).
fof(f132,axiom,(
  ! [X19,X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X9)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num',X9)))) | s('type.num.num',X9) = s('type.num.num',X19)))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.LESS_OR_EQ')).
fof(f30480,plain,(
  $false | $spl7),
  inference(subsumption_resolution,[],[f30479,f14643])).
fof(f30479,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) | $spl7),
  inference(subsumption_resolution,[],[f30434,f553])).
fof(f30434,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) | $spl7),
  inference(resolution,[],[f684,f579])).
fof(f579,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f334])).
fof(f334,plain,(
  ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X19,X9] : (s('type.num.num',X9) = s('type.num.num',X19) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num',X9)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X19)))))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.LESS_LESS_CASES')).
fof(f30478,plain,(
  $false | $spl7),
  inference(subsumption_resolution,[],[f30477,f14643])).
fof(f30477,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) | $spl7),
  inference(subsumption_resolution,[],[f30433,f553])).
fof(f30433,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) | $spl7),
  inference(resolution,[],[f684,f579])).
fof(f30476,plain,(
  $false | ($spl7 | $spl28)),
  inference(subsumption_resolution,[],[f30432,f14643])).
fof(f30432,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) | ($spl7 | $spl28)),
  inference(resolution,[],[f684,f4022])).
fof(f4022,plain,(
  ( ! [X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)) ) | $spl28),
  inference(forward_demodulation,[],[f4021,f852])).
fof(f4021,plain,(
  ( ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X7))))) ) | $spl28),
  inference(subsumption_resolution,[],[f4020,f553])).
fof(f4020,plain,(
  ( ! [X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X7))))) ) | $spl28),
  inference(forward_demodulation,[],[f3815,f852])).
fof(f3815,plain,(
  ( ! [X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X7))))) ) | $spl28),
  inference(superposition,[],[f852,f579])).
fof(f30330,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))) | $spl285),
  inference(cnf_transformation,[],[f30330_D])).
fof(f30330_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))) <=> ~$spl285),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl285])])).
fof(f30332,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))) = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) | $spl286),
  inference(cnf_transformation,[],[f30332_D])).
fof(f30332_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))) = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f30292,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))) = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))))),
  inference(resolution,[],[f636,f524])).
fof(f30285,plain,(
  $false | ($spl23 | $spl26 | $spl66)),
  inference(subsumption_resolution,[],[f9507,f777])).
fof(f777,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl23),
  inference(cnf_transformation,[],[f777_D])).
fof(f777_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f9507,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl26 | $spl66)),
  inference(forward_demodulation,[],[f1136,f810])).
fof(f30282,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',sK4)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK4))))) ) | $spl282),
  inference(cnf_transformation,[],[f30282_D])).
fof(f30282_D,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',sK4)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK4))))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f30242,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',sK4)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK4)))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),sK1))))) )),
  inference(resolution,[],[f633,f475])).
fof(f475,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X4)))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(cnf_transformation,[],[f370])).
fof(f370,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4)))))))),
  inference(flattening,[],[f369])).
fof(f369,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4)))))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X4)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0))))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4)))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7,X8,X19,X9] : ((p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))),s(fun('type.num.num','type.realax.real'),X6)))) & p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num',X9)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X6))))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X19))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X9)))))))),
  file('thm.transc.DIVISION_LT','thm.transc.DIVISION_LT_GEN')).
fof(f633,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))))))),
  inference(resolution,[],[f473,f578])).
fof(f578,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f399])).
fof(f399,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X19,X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num',X9)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X9)))))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.LESS_IMP_LESS_OR_EQ')).
fof(f29543,plain,(
  $false | ($spl75 | $spl144)),
  inference(subsumption_resolution,[],[f9140,f5533])).
fof(f5533,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK20(sK4)) | $spl75),
  inference(resolution,[],[f1426,f557])).
fof(f9140,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK4)) | $spl144),
  inference(cnf_transformation,[],[f9140_D])).
fof(f9140_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK4)) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f776,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl21),
  inference(cnf_transformation,[],[f776_D])).
fof(f776_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f778,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl22),
  inference(cnf_transformation,[],[f778_D])).
fof(f778_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f29493,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl11),
  inference(resolution,[],[f690,f524])).
fof(f29472,plain,(
  $false | ($spl11 | $spl26 | $spl94)),
  inference(subsumption_resolution,[],[f10252,f729])).
fof(f729,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) | $spl11),
  inference(resolution,[],[f690,f557])).
fof(f10252,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK4) | ($spl26 | $spl94)),
  inference(forward_demodulation,[],[f2671,f810])).
fof(f2671,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',sK4) | $spl94),
  inference(cnf_transformation,[],[f2671_D])).
fof(f2671_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',sK4) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f29471,plain,(
  $false | ($spl11 | $spl100)),
  inference(subsumption_resolution,[],[f2677,f729])).
fof(f2677,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK4) | $spl100),
  inference(cnf_transformation,[],[f2677_D])).
fof(f2677_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK4) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f1003,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl53),
  inference(cnf_transformation,[],[f1003_D])).
fof(f1003_D,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f15318,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',X0)))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl214),
  inference(cnf_transformation,[],[f15318_D])).
fof(f15318_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',X0)))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s(fun('type.num.num','type.realax.real'),sK1))))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f29445,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',X0)))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s(fun('type.num.num','type.realax.real'),sK1))))) )),
  inference(superposition,[],[f29330,f486])).
fof(f486,plain,(
  ( ! [X2,X0,X5,X1] : (s('type.realax.real',X2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X5))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X5),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X0)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(cnf_transformation,[],[f422])).
fof(f29330,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))))))),
  inference(forward_demodulation,[],[f474,f704])).
fof(f474,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))))))),
  inference(cnf_transformation,[],[f414])).
fof(f690,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl11),
  inference(cnf_transformation,[],[f690_D])).
fof(f690_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f15038,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl191),
  inference(cnf_transformation,[],[f15038_D])).
fof(f15038_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl191),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl191])])).
fof(f29422,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f29330,f556])).
fof(f683,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK4)))) | $spl4),
  inference(cnf_transformation,[],[f683_D])).
fof(f683_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK4)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f29421,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK4))))),
  inference(superposition,[],[f29330,f539])).
fof(f539,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f440])).
fof(f440,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20])],[f302])).
fof(f302,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f172])).
fof(f172,axiom,(
  ! [X19] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | ? [X9] : s('type.num.num',X19) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.num_CASES')).
fof(f29462,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) | $spl279),
  inference(cnf_transformation,[],[f29462_D])).
fof(f29462_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) <=> ~$spl279),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl279])])).
fof(f29464,plain,(
  s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))) | $spl280),
  inference(cnf_transformation,[],[f29464_D])).
fof(f29464_D,plain,(
  s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f29409,plain,(
  s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))))))),
  inference(resolution,[],[f29330,f501])).
fof(f501,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) = s('type.realax.real',X1) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f427])).
fof(f427,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & s('type.realax.real',X0) != s('type.realax.real',X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))),
  inference(flattening,[],[f426])).
fof(f426,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) = s('type.realax.real',X1))) & ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & s('type.realax.real',X0) != s('type.realax.real',X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))),
  inference(nnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) <=> (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) = s('type.realax.real',X1)))),
  inference(rectify,[],[f107])).
fof(f107,axiom,(
  ! [X14,X15] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X15)))) <=> (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real',X15)))) | s('type.realax.real',X14) = s('type.realax.real',X15)))),
  file('thm.transc.DIVISION_LT','thm.real.REAL_LE_LT')).
fof(f685,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl6),
  inference(cnf_transformation,[],[f685_D])).
fof(f685_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f29377,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f473,f556])).
fof(f29376,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK4))))),
  inference(superposition,[],[f473,f539])).
fof(f680,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',sK4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) | $spl2),
  inference(cnf_transformation,[],[f680_D])).
fof(f680_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',sK4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f29364,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',sK4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.num.num','type.realax.real'),sK1))))) )),
  inference(resolution,[],[f473,f485])).
fof(f684,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl7),
  inference(cnf_transformation,[],[f684_D])).
fof(f684_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f14795,plain,(
  ( ! [X37,X36] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X36),s('type.realax.real',X37))),s(fun('type.num.num','type.realax.real'),sK1))))) )),
  inference(superposition,[],[f485,f704])).
fof(f29329,plain,(
  ( ! [X45,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',sK2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X44),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X45),s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.transc.fine_2'(s(fun('type.realax.real','type.realax.real'),X44),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),sK1),s(fun('type.num.num','type.realax.real'),X45))))))) ) | $spl276),
  inference(cnf_transformation,[],[f29329_D])).
fof(f29329_D,plain,(
  ( ! [X45,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',sK2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X44),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X45),s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.transc.fine_2'(s(fun('type.realax.real','type.realax.real'),X44),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),sK1),s(fun('type.num.num','type.realax.real'),X45))))))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f14799,plain,(
  ( ! [X45,X44] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',sK2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X44),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X45),s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.transc.fine_2'(s(fun('type.realax.real','type.realax.real'),X44),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),sK1),s(fun('type.num.num','type.realax.real'),X45))))))) )),
  inference(superposition,[],[f581,f704])).
fof(f581,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X4))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X1)))))) | ~p(s(bool,'const.transc.fine_2'(s(fun('type.realax.real','type.realax.real'),X0),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X1),s(fun('type.num.num','type.realax.real'),X2))))))) )),
  inference(cnf_transformation,[],[f452])).
fof(f452,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.transc.fine_2'(s(fun('type.realax.real','type.realax.real'),X0),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X1),s(fun('type.num.num','type.realax.real'),X2)))))) | ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X4))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4)))))))))) & ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK22(X2,X1,X0)),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X1)))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK22(X2,X1,X0)))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',sK22(X2,X1,X0)))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',sK22(X2,X1,X0)))))))))) | p(s(bool,'const.transc.fine_2'(s(fun('type.realax.real','type.realax.real'),X0),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X1),s(fun('type.num.num','type.realax.real'),X2))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22])],[f451])).
fof(f451,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.transc.fine_2'(s(fun('type.realax.real','type.realax.real'),X0),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X1),s(fun('type.num.num','type.realax.real'),X2)))))) | ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X4))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X4)))))))))) & (? [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X1)))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X3))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X3))))))))) | p(s(bool,'const.transc.fine_2'(s(fun('type.realax.real','type.realax.real'),X0),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X1),s(fun('type.num.num','type.realax.real'),X2))))))))),
  inference(rectify,[],[f450])).
fof(f450,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.transc.fine_2'(s(fun('type.realax.real','type.realax.real'),X0),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X1),s(fun('type.num.num','type.realax.real'),X2)))))) | ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X3))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X3)))))))))) & (? [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X1)))))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X3))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X3))))))))) | p(s(bool,'const.transc.fine_2'(s(fun('type.realax.real','type.realax.real'),X0),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X1),s(fun('type.num.num','type.realax.real'),X2))))))))),
  inference(nnf_transformation,[],[f400])).
fof(f400,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.transc.fine_2'(s(fun('type.realax.real','type.realax.real'),X0),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X1),s(fun('type.num.num','type.realax.real'),X2)))))) <=> ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X3))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X3))))))))))),
  inference(ennf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.transc.fine_2'(s(fun('type.realax.real','type.realax.real'),X0),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X1),s(fun('type.num.num','type.realax.real'),X2)))))) <=> ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X1)))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X1),s('type.num.num',X3))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X3))))))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X20,X6,X21] : (p(s(bool,'const.transc.fine_2'(s(fun('type.realax.real','type.realax.real'),X20),s('type.pair.prod'(fun('type.num.num','type.realax.real'),fun('type.num.num','type.realax.real')),'const.pair.,_2'(s(fun('type.num.num','type.realax.real'),X6),s(fun('type.num.num','type.realax.real'),X21)))))) <=> ! [X9] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),X6)))))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X9))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X20),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X21),s('type.num.num',X9))))))))))),
  file('thm.transc.DIVISION_LT','thm.transc.fine')).
fof(f29250,plain,(
  $false | $spl0),
  inference(resolution,[],[f678,f472])).
fof(f28655,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl273),
  inference(cnf_transformation,[],[f28655_D])).
fof(f28655_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl273),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl273])])).
fof(f28657,plain,(
  s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl274),
  inference(cnf_transformation,[],[f28657_D])).
fof(f28657_D,plain,(
  s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f28600,plain,(
  s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl253),
  inference(resolution,[],[f23174,f524])).
fof(f23174,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl253),
  inference(cnf_transformation,[],[f23174_D])).
fof(f23174_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl253),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl253])])).
fof(f23176,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))) | $spl254),
  inference(cnf_transformation,[],[f23176_D])).
fof(f23176_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f28582,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl28 | $spl127)),
  inference(subsumption_resolution,[],[f28538,f12884])).
fof(f28538,plain,(
  p(s(bool,'F_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl28 | $spl127)),
  inference(superposition,[],[f524,f23136])).
fof(f23136,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.|48|_0'))) | ($spl26 | $spl28 | $spl127)),
  inference(resolution,[],[f22962,f12885])).
fof(f22962,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.|48|_0')))) | ($spl26 | $spl28 | $spl127)),
  inference(subsumption_resolution,[],[f22723,f18443])).
fof(f18443,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | $spl127),
  inference(resolution,[],[f7120,f525])).
fof(f525,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f435])).
fof(f22723,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.|48|_0')))) | ($spl26 | $spl28)),
  inference(superposition,[],[f6780,f556])).
fof(f6780,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6779,f535])).
fof(f6779,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl28)),
  inference(forward_demodulation,[],[f6579,f852])).
fof(f6579,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) | $spl26),
  inference(superposition,[],[f6485,f539])).
fof(f6485,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) | $spl26),
  inference(forward_demodulation,[],[f6484,f4194])).
fof(f6484,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl26),
  inference(forward_demodulation,[],[f545,f810])).
fof(f545,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(cnf_transformation,[],[f130])).
fof(f130,axiom,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.TWO')).
fof(f28477,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl270),
  inference(cnf_transformation,[],[f28477_D])).
fof(f28477_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f28421,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl26 | $spl28 | $spl127)),
  inference(resolution,[],[f23135,f574])).
fof(f23135,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl26 | $spl28 | $spl127)),
  inference(resolution,[],[f22962,f522])).
fof(f28393,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl267),
  inference(cnf_transformation,[],[f28393_D])).
fof(f28393_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl267),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl267])])).
fof(f28395,plain,(
  s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl268),
  inference(cnf_transformation,[],[f28395_D])).
fof(f28395_D,plain,(
  s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f28338,plain,(
  s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl249),
  inference(resolution,[],[f22492,f524])).
fof(f22492,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl249),
  inference(cnf_transformation,[],[f22492_D])).
fof(f22492_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl249),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl249])])).
fof(f22494,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))) | $spl250),
  inference(cnf_transformation,[],[f22494_D])).
fof(f22494_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f26940,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl28 | $spl123)),
  inference(resolution,[],[f22280,f524])).
fof(f22280,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.|48|_0')))) | ($spl26 | $spl28 | $spl123)),
  inference(subsumption_resolution,[],[f22041,f18239])).
fof(f18239,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | $spl123),
  inference(resolution,[],[f7002,f525])).
fof(f22041,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.|48|_0')))) | ($spl26 | $spl28)),
  inference(superposition,[],[f4740,f556])).
fof(f4740,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl28)),
  inference(subsumption_resolution,[],[f4739,f535])).
fof(f4739,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl28)),
  inference(forward_demodulation,[],[f4547,f852])).
fof(f4547,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | $spl26),
  inference(superposition,[],[f4194,f539])).
fof(f26794,plain,(
  $false | ($spl26 | $spl28 | $spl250)),
  inference(subsumption_resolution,[],[f26793,f13259])).
fof(f13259,plain,(
  ( ! [X22] : (s('type.num.num',X22) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))) ) | $spl26),
  inference(subsumption_resolution,[],[f13239,f12884])).
fof(f13239,plain,(
  ( ! [X22] : (p(s(bool,'F_0')) | s('type.num.num',X22) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))) ) | $spl26),
  inference(superposition,[],[f528,f13057])).
fof(f528,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f437])).
fof(f26793,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl26 | $spl28 | $spl250)),
  inference(forward_demodulation,[],[f26792,f4194])).
fof(f26792,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl26 | $spl28 | $spl250)),
  inference(forward_demodulation,[],[f26791,f852])).
fof(f26791,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl26 | $spl28 | $spl250)),
  inference(forward_demodulation,[],[f26552,f541])).
fof(f541,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f239])).
fof(f239,axiom,(
  ! [X9] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))))),
  file('thm.transc.DIVISION_LT','thm.numeral.numeral_distrib|split|14')).
fof(f26552,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl26 | $spl28 | $spl250)),
  inference(superposition,[],[f541,f26087])).
fof(f26087,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl26 | $spl28 | $spl250)),
  inference(backward_demodulation,[],[f22494,f4740])).
fof(f26058,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl264),
  inference(cnf_transformation,[],[f26058_D])).
fof(f26058_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f26002,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl26 | $spl28 | $spl123)),
  inference(resolution,[],[f22453,f574])).
fof(f22453,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))))) | ($spl26 | $spl28 | $spl123)),
  inference(resolution,[],[f22280,f522])).
fof(f24596,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl256),
  inference(cnf_transformation,[],[f24596_D])).
fof(f24596_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f14690,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) | $spl182),
  inference(cnf_transformation,[],[f14690_D])).
fof(f14690_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f25161,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ($spl10 | $spl32)),
  inference(resolution,[],[f25025,f574])).
fof(f25025,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ($spl10 | $spl32)),
  inference(forward_demodulation,[],[f972,f2713])).
fof(f2713,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK4) | $spl10),
  inference(resolution,[],[f689,f556])).
fof(f689,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(cnf_transformation,[],[f689_D])).
fof(f689_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f972,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl32),
  inference(cnf_transformation,[],[f972_D])).
fof(f972_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f25216,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl262),
  inference(cnf_transformation,[],[f25216_D])).
fof(f25216_D,plain,(
  ( ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f25160,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | ($spl10 | $spl32)),
  inference(resolution,[],[f25025,f475])).
fof(f25002,plain,(
  $false | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f25001,f698])).
fof(f25001,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl33)),
  inference(forward_demodulation,[],[f24945,f534])).
fof(f24945,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl10 | $spl33)),
  inference(resolution,[],[f24619,f522])).
fof(f24619,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ($spl10 | $spl33)),
  inference(forward_demodulation,[],[f973,f2713])).
fof(f24998,plain,(
  $false | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f24943,f697])).
fof(f697,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(forward_demodulation,[],[f696,f546])).
fof(f696,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(subsumption_resolution,[],[f660,f553])).
fof(f660,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(superposition,[],[f473,f554])).
fof(f554,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) )),
  inference(cnf_transformation,[],[f442])).
fof(f442,plain,(
  ! [X0] : ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(nnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f108])).
fof(f108,axiom,(
  ! [X14] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X14))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.arithmetic_|49|95_prove_ltb')).
fof(f24943,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl10 | $spl33)),
  inference(resolution,[],[f24619,f572])).
fof(f24997,plain,(
  $false | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f24939,f698])).
fof(f24939,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl33)),
  inference(resolution,[],[f24619,f522])).
fof(f17151,plain,(
  s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl222),
  inference(cnf_transformation,[],[f17151_D])).
fof(f17151_D,plain,(
  s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f24805,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | $spl261),
  inference(cnf_transformation,[],[f24805_D])).
fof(f24805_D,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl261),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl261])])).
fof(f24770,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl157),
  inference(superposition,[],[f12643,f539])).
fof(f24599,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl259),
  inference(cnf_transformation,[],[f24599_D])).
fof(f24599_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl259),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl259])])).
fof(f24543,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(resolution,[],[f697,f527])).
fof(f14688,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl181),
  inference(cnf_transformation,[],[f14688_D])).
fof(f14688_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl181),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl181])])).
fof(f24540,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f697,f574])).
fof(f24539,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))),
  inference(resolution,[],[f697,f579])).
fof(f24538,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))),
  inference(resolution,[],[f697,f579])).
fof(f23134,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl28 | $spl127)),
  inference(resolution,[],[f22962,f524])).
fof(f22452,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl28 | $spl123)),
  inference(resolution,[],[f22280,f524])).
fof(f13452,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl166),
  inference(cnf_transformation,[],[f13452_D])).
fof(f13452_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f20286,plain,(
  ( ! [X8] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',X8))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',sK2)))))) ) | $spl246),
  inference(cnf_transformation,[],[f20286_D])).
fof(f20286_D,plain,(
  ( ! [X8] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',X8))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',sK2)))))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f20285,plain,(
  ( ! [X8] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',X8))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real',sK2))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl26 | $spl220)),
  inference(forward_demodulation,[],[f20201,f15370])).
fof(f20201,plain,(
  ( ! [X8] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',X8))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl26 | $spl220)),
  inference(superposition,[],[f15518,f539])).
fof(f15518,plain,(
  ( ! [X0] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',X0))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))))) ) | ($spl26 | $spl220)),
  inference(backward_demodulation,[],[f15370,f13012])).
fof(f13012,plain,(
  ( ! [X0] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))))) ) | $spl26),
  inference(forward_demodulation,[],[f4193,f6485])).
fof(f4193,plain,(
  ( ! [X0] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))))))) ) | $spl26),
  inference(backward_demodulation,[],[f810,f520])).
fof(f520,plain,(
  ( ! [X0] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0] : s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  inference(rectify,[],[f197])).
fof(f197,axiom,(
  ! [X31] : s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X31))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  file('thm.transc.DIVISION_LT','thm.real.REAL_LT_HALF|49|')).
fof(f2680,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl102),
  inference(cnf_transformation,[],[f2680_D])).
fof(f2680_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f20279,plain,(
  ( ! [X8] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',X8))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) ) | $spl244),
  inference(cnf_transformation,[],[f20279_D])).
fof(f20279_D,plain,(
  ( ! [X8] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',X8))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f20187,plain,(
  ( ! [X8] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',X8))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl26 | $spl220)),
  inference(superposition,[],[f15518,f539])).
fof(f18626,plain,(
  $false | ($spl10 | $spl131)),
  inference(subsumption_resolution,[],[f18625,f570])).
fof(f18625,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl10 | $spl131)),
  inference(forward_demodulation,[],[f7497,f2713])).
fof(f18503,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl241),
  inference(cnf_transformation,[],[f18503_D])).
fof(f18503_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl241),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl241])])).
fof(f18505,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl242),
  inference(cnf_transformation,[],[f18505_D])).
fof(f18505_D,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f18449,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl127),
  inference(resolution,[],[f7120,f524])).
fof(f18419,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) | $spl238),
  inference(cnf_transformation,[],[f18419_D])).
fof(f18419_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f7122,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | $spl128),
  inference(cnf_transformation,[],[f7122_D])).
fof(f7122_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f18363,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl28)),
  inference(resolution,[],[f7083,f574])).
fof(f7083,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl28)),
  inference(resolution,[],[f6782,f522])).
fof(f6782,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | ($spl26 | $spl28)),
  inference(subsumption_resolution,[],[f6781,f535])).
fof(f6781,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | ($spl26 | $spl28)),
  inference(forward_demodulation,[],[f6580,f852])).
fof(f6580,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | $spl26),
  inference(superposition,[],[f6485,f556])).
fof(f18299,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl235),
  inference(cnf_transformation,[],[f18299_D])).
fof(f18299_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f18301,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl236),
  inference(cnf_transformation,[],[f18301_D])).
fof(f18301_D,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f18245,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl123),
  inference(resolution,[],[f7002,f524])).
fof(f18102,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | $spl232),
  inference(cnf_transformation,[],[f18102_D])).
fof(f18102_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f7004,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | $spl124),
  inference(cnf_transformation,[],[f7004_D])).
fof(f7004_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f18046,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl28)),
  inference(resolution,[],[f6965,f574])).
fof(f6965,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl28)),
  inference(resolution,[],[f4742,f522])).
fof(f4742,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | ($spl26 | $spl28)),
  inference(subsumption_resolution,[],[f4741,f535])).
fof(f4741,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | ($spl26 | $spl28)),
  inference(forward_demodulation,[],[f4548,f852])).
fof(f4548,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | $spl26),
  inference(superposition,[],[f4194,f556])).
fof(f18021,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))) ) | $spl230),
  inference(cnf_transformation,[],[f18021_D])).
fof(f18021_D,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f17974,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f12197,f539])).
fof(f12197,plain,(
  ( ! [X206] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X206),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f11986,f6000])).
fof(f6000,plain,(
  ( ! [X204] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X204))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(subsumption_resolution,[],[f5999,f535])).
fof(f5999,plain,(
  ( ! [X204] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X204),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X204))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(forward_demodulation,[],[f5790,f532])).
fof(f5790,plain,(
  ( ! [X204] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X204))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X204),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f556,f4192])).
fof(f4192,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f540,f546])).
fof(f540,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(cnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X19] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.ADD|49|')).
fof(f11986,plain,(
  ( ! [X206] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X206))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X206),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f557,f4192])).
fof(f17710,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl226),
  inference(cnf_transformation,[],[f17710_D])).
fof(f17710_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f17711,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | ~p(s(bool,X1))) ) | $spl228),
  inference(cnf_transformation,[],[f17711_D])).
fof(f17711_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | ~p(s(bool,X1))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f17709,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl26 | $spl179)),
  inference(subsumption_resolution,[],[f17675,f13539])).
fof(f17675,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl26 | $spl179)),
  inference(superposition,[],[f14630,f13594])).
fof(f13594,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl26),
  inference(resolution,[],[f13539,f12885])).
fof(f14630,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X5)) | ~p(s(bool,X4))) ) | ($spl26 | $spl179)),
  inference(subsumption_resolution,[],[f14601,f13531])).
fof(f14601,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | ~p(s(bool,'T_0'))) ) | ($spl26 | $spl179)),
  inference(superposition,[],[f14597,f596])).
fof(f14597,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X1))) ) | ($spl26 | $spl179)),
  inference(subsumption_resolution,[],[f14593,f12884])).
fof(f14593,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,'F_0')) | p(s(bool,X1))) ) | $spl179),
  inference(superposition,[],[f14589,f593])).
fof(f17153,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | $spl225),
  inference(cnf_transformation,[],[f17153_D])).
fof(f17153_D,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) <=> ~$spl225),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl225])])).
fof(f17123,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl39),
  inference(superposition,[],[f987,f539])).
fof(f17030,plain,(
  $false | ($spl26 | $spl219 | $spl220)),
  inference(subsumption_resolution,[],[f17029,f13531])).
fof(f17029,plain,(
  ~p(s(bool,'T_0')) | ($spl26 | $spl219 | $spl220)),
  inference(forward_demodulation,[],[f17028,f13676])).
fof(f13676,plain,(
  ( ! [X206] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X206),s('type.realax.real',X206)))) ) | $spl26),
  inference(resolution,[],[f13539,f13111])).
fof(f13111,plain,(
  ( ! [X28] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X28),s('type.realax.real',X28))))) ) | $spl26),
  inference(subsumption_resolution,[],[f13077,f12884])).
fof(f13077,plain,(
  ( ! [X28] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X28),s('type.realax.real',X28))))) ) | $spl26),
  inference(superposition,[],[f508,f12933])).
fof(f508,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(rectify,[],[f168])).
fof(f168,axiom,(
  ! [X14,X15] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X15)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X15),s('type.realax.real',X14)))))),
  file('thm.transc.DIVISION_LT','thm.real.REAL_LET_TOTAL')).
fof(f17028,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK2),s('type.realax.real',sK2)))) | ($spl219 | $spl220)),
  inference(forward_demodulation,[],[f15368,f15370])).
fof(f15483,plain,(
  $false | ($spl209 | $spl217)),
  inference(subsumption_resolution,[],[f15439,f15320])).
fof(f15439,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl209),
  inference(resolution,[],[f15226,f508])).
fof(f15226,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | $spl209),
  inference(cnf_transformation,[],[f15226_D])).
fof(f15226_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) <=> ~$spl209),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl209])])).
fof(f15370,plain,(
  s('type.realax.real',sK2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl220),
  inference(cnf_transformation,[],[f15370_D])).
fof(f15370_D,plain,(
  s('type.realax.real',sK2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f15383,plain,(
  s('type.realax.real',sK2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | $spl195),
  inference(resolution,[],[f15123,f501])).
fof(f15368,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl219),
  inference(cnf_transformation,[],[f15368_D])).
fof(f15368_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl219),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl219])])).
fof(f15322,plain,(
  s('type.realax.real',sK2) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl217),
  inference(resolution,[],[f15320,f501])).
fof(f12646,plain,(
  ( ! [X24,X23] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X23),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X24)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X23),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl162),
  inference(cnf_transformation,[],[f12646_D])).
fof(f12646_D,plain,(
  ( ! [X24,X23] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X23),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X24)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X23),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f15320,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl217),
  inference(cnf_transformation,[],[f15320_D])).
fof(f15320_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl217),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl217])])).
fof(f15288,plain,(
  ( ! [X24,X23] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X23),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X24)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X23),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl10),
  inference(superposition,[],[f14748,f562])).
fof(f562,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X2)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f397])).
fof(f397,plain,(
  ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ! [X1,X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X2)))))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(ennf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) => ! [X1,X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X2)))))))) => s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f174])).
fof(f174,axiom,(
  ! [X43] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X43),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) => ! [X14,X15] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X14))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X15)))))))) => s('type.realax.real',X14) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.transc.DIVISION_LT','thm.seq.ABS_NEG_LEMMA')).
fof(f14748,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl10),
  inference(backward_demodulation,[],[f704,f10653])).
fof(f10653,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl10),
  inference(backward_demodulation,[],[f2713,f474])).
fof(f12645,plain,(
  ( ! [X22] : (s('type.realax.real',X22) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X22)))) ) | $spl160),
  inference(cnf_transformation,[],[f12645_D])).
fof(f12645_D,plain,(
  ( ! [X22] : (s('type.realax.real',X22) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X22)))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f15287,plain,(
  ( ! [X22] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X22) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X22)))) ) | $spl10),
  inference(superposition,[],[f14748,f550])).
fof(f550,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(cnf_transformation,[],[f441])).
fof(f441,plain,(
  ! [X0,X1] : ((s('type.realax.real',X1) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(nnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0,X1] : (s('type.realax.real',X1) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f152])).
fof(f152,axiom,(
  ! [X14,X15] : (s('type.realax.real',X15) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X14),s('type.realax.real',X15))) <=> s('type.realax.real',X14) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.transc.DIVISION_LT','thm.real.REAL_ADD_LID_UNIQ')).
fof(f12644,plain,(
  ( ! [X17,X16] : (s('type.realax.real',X16) = s('type.realax.real',X17) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X16))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X17)))) ) | $spl158),
  inference(cnf_transformation,[],[f12644_D])).
fof(f12644_D,plain,(
  ( ! [X17,X16] : (s('type.realax.real',X16) = s('type.realax.real',X17) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X16))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X17)))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f15284,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X16) = s('type.realax.real',X17) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X16))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X17)))) ) | $spl10),
  inference(superposition,[],[f14748,f549])).
fof(f549,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X1) = s('type.realax.real',X2) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f396])).
fof(f396,plain,(
  ! [X0,X1,X2] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))) | s('type.realax.real',X1) = s('type.realax.real',X2))),
  inference(flattening,[],[f395])).
fof(f395,plain,(
  ! [X0,X1,X2] : ((s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) | s('type.realax.real',X1) = s('type.realax.real',X2))),
  inference(ennf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2] : ((s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) => s('type.realax.real',X1) = s('type.realax.real',X2))),
  inference(flattening,[],[f310])).
fof(f310,plain,(
  ! [X0,X1,X2] : ((~s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2)))) => s('type.realax.real',X1) = s('type.realax.real',X2))),
  inference(rectify,[],[f191])).
fof(f191,axiom,(
  ! [X14,X15,X40] : ((~s('type.realax.real',X14) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X15))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X40)))) => s('type.realax.real',X15) = s('type.realax.real',X40))),
  file('thm.transc.DIVISION_LT','thm.real.REAL_EQ_LMUL_IMP')).
fof(f12643,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl157),
  inference(cnf_transformation,[],[f12643_D])).
fof(f12643_D,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f15276,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',X0)))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl10),
  inference(superposition,[],[f14748,f486])).
fof(f15298,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl211),
  inference(cnf_transformation,[],[f15298_D])).
fof(f15298_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl211),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl211])])).
fof(f15300,plain,(
  s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl212),
  inference(cnf_transformation,[],[f15300_D])).
fof(f15300_D,plain,(
  s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f15236,plain,(
  s('type.realax.real',sK2) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl10),
  inference(resolution,[],[f14748,f501])).
fof(f15150,plain,(
  ( ! [X300] : (s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X300),s('type.realax.real',sK2))),s('type.realax.real',X300)))))) ) | $spl196),
  inference(cnf_transformation,[],[f15150_D])).
fof(f15150_D,plain,(
  ( ! [X300] : (s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X300),s('type.realax.real',sK2))),s('type.realax.real',X300)))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f15232,plain,(
  ( ! [X476] : (s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X476),s('type.realax.real',sK2))),s('type.realax.real',X476))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2))))) )),
  inference(forward_demodulation,[],[f15231,f704])).
fof(f15231,plain,(
  ( ! [X476] : (s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X476),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X476))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2))))) )),
  inference(forward_demodulation,[],[f15001,f622])).
fof(f622,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f360])).
fof(f360,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X14,X15] : s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X15),s('type.realax.real',X14))))) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X14),s('type.realax.real',X15)))))),
  file('thm.transc.DIVISION_LT','thm.topology.MR|49|_DEF')).
fof(f15001,plain,(
  ( ! [X476] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X476),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X476),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f621,f704])).
fof(f621,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f411])).
fof(f411,plain,(
  ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real',X1) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))))),
  inference(ennf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) => s('type.realax.real',X1) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X14,X31] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X31)))) => s('type.realax.real',X31) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X14),s('type.realax.real',X31))))))))),
  file('thm.transc.DIVISION_LT','thm.topology.MR|49|_SUB_LE')).
fof(f15230,plain,(
  ( ! [X474] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X474),s('type.realax.real',sK2))),s('type.realax.real',X474)))))) )),
  inference(forward_demodulation,[],[f15229,f704])).
fof(f15229,plain,(
  ( ! [X474] : (s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X474),s('type.realax.real',sK2))),s('type.realax.real',X474))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f14999,f622])).
fof(f14999,plain,(
  ( ! [X474] : (s('type.realax.real',sK2) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X474),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X474),s('type.realax.real',sK2))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f621,f704])).
fof(f15155,plain,(
  ( ! [X303] : (s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X303),s('type.realax.real',sK2))),s('type.realax.real',X303)))))) ) | $spl198),
  inference(cnf_transformation,[],[f15155_D])).
fof(f15155_D,plain,(
  ( ! [X303] : (s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X303),s('type.realax.real',sK2))),s('type.realax.real',X303)))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f15228,plain,(
  ( ! [X473] : (s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X473),s('type.realax.real',sK2))),s('type.realax.real',X473))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2))))) )),
  inference(forward_demodulation,[],[f15227,f704])).
fof(f15227,plain,(
  ( ! [X473] : (s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X473),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X473))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2))))) )),
  inference(forward_demodulation,[],[f14998,f622])).
fof(f14998,plain,(
  ( ! [X473] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X473),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X473),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f620,f704])).
fof(f620,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f410])).
fof(f410,plain,(
  ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real',X1) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))))),
  inference(ennf_transformation,[],[f358])).
fof(f358,plain,(
  ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) => s('type.realax.real',X1) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X14,X31] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X31)))) => s('type.realax.real',X31) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X14),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X14),s('type.realax.real',X31))))))))),
  file('thm.transc.DIVISION_LT','thm.topology.MR|49|_ADD_POS')).
fof(f15224,plain,(
  ( ! [X471] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X471),s('type.realax.real',sK2))),s('type.realax.real',X471)))))) )),
  inference(forward_demodulation,[],[f15223,f704])).
fof(f15223,plain,(
  ( ! [X471] : (s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X471),s('type.realax.real',sK2))),s('type.realax.real',X471))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f14996,f622])).
fof(f14996,plain,(
  ( ! [X471] : (s('type.realax.real',sK2) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X471),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X471),s('type.realax.real',sK2))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f620,f704])).
fof(f15189,plain,(
  ( ! [X372,X374,X373] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X372),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X373),s('type.realax.real',X374)))))) | sP0(X372) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X372),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK17(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X372)),s('type.realax.real',sK18(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X372))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X373),s('type.realax.real',X374))))) ) | $spl206),
  inference(cnf_transformation,[],[f15189_D])).
fof(f15189_D,plain,(
  ( ! [X372,X374,X373] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X372),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X373),s('type.realax.real',X374)))))) | sP0(X372) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X372),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK17(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X372)),s('type.realax.real',sK18(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X372))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X373),s('type.realax.real',X374))))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f15123,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | $spl195),
  inference(cnf_transformation,[],[f15123_D])).
fof(f15123_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) <=> ~$spl195),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl195])])).
fof(f14945,plain,(
  ( ! [X372,X374,X373] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X372),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X373),s('type.realax.real',X374)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X373),s('type.realax.real',X374)))) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X372),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK17(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X372)),s('type.realax.real',sK18(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X372))))))) | sP0(X372)) )),
  inference(superposition,[],[f519,f704])).
fof(f519,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK17(X4,X0)),s('type.realax.real',sK18(X4,X0))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | sP0(X0)) )),
  inference(cnf_transformation,[],[f432])).
fof(f432,plain,(
  ! [X0] : (sP0(X0) | ! [X4] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X4,X0)),s('type.realax.real',sK16(X0))))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X0)),s('type.realax.real',sK18(X4,X0))))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK18(X4,X0)),s('type.realax.real',sK17(X4,X0)))),s('type.realax.real',X4)))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',sK17(X4,X0)),s('type.realax.real',sK18(X4,X0))))))))) | ! [X1,X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17,sK18])],[f431])).
fof(f431,plain,(
  ! [X0] : (sP0(X0) | ? [X3] : ! [X4] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | ? [X5,X6] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X3)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X6)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X6),s('type.realax.real',X5))),s('type.realax.real',X4)))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X6)))))))) | ! [X1,X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2))))))))),
  inference(rectify,[],[f413])).
fof(f413,plain,(
  ! [X0] : (sP0(X0) | ? [X1] : ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | ? [X3,X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X4)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X4),s('type.realax.real',X3))),s('type.realax.real',X2)))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X4)))))))) | ! [X8,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))))))))),
  inference(definition_folding,[],[f391,f412])).
fof(f412,plain,(
  ! [X0] : (~sP0(X0) | ? [X5,X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X6)))))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X7))))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f391,plain,(
  ! [X0] : (? [X5,X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X6)))))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7)))))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X7))))))) | ? [X1] : ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | ? [X3,X4] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X4)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X4),s('type.realax.real',X3))),s('type.realax.real',X2)))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X4)))))))) | ! [X8,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))))))))),
  inference(flattening,[],[f390])).
fof(f390,plain,(
  ! [X0] : ((? [X5,X6,X7] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X6)))))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7))))))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X7))))))) | ? [X1] : ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | ? [X3,X4] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X4)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X4),s('type.realax.real',X3))),s('type.realax.real',X2))))) & ~p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X4))))))))) | ! [X8,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))))))))),
  inference(ennf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0] : ((! [X5,X6,X7] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real',X6)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X6)))))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X6),s('type.realax.real',X7))))))) => p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X5),s('type.realax.real',X7))))))) & ! [X1] : ? [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) & ! [X3,X4] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X4)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X4),s('type.realax.real',X3))),s('type.realax.real',X2))))) => p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X3),s('type.realax.real',X4))))))))) => ! [X8,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X9)))) => p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X9))))))))),
  inference(rectify,[],[f80])).
fof(f80,axiom,(
  ! [X39] : ((! [X7,X8,X43] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X43)))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X39),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8)))))) & p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X39),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X8),s('type.realax.real',X43))))))) => p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X39),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X43))))))) & ! [X14] : ? [X31] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X31)))) & ! [X7,X8] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X14)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real',X8)))) & p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real',X7))),s('type.realax.real',X31))))) => p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X39),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))))))))) => ! [X7,X8] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X8)))) => p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X39),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8))))))))),
  file('thm.transc.DIVISION_LT','thm.seq.BOLZANO_LEMMA')).
fof(f15183,plain,(
  ( ! [X356,X354,X355] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK18(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X354)),s('type.realax.real',sK17(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X354)))),s('type.realax.real',sK2)))) | sP0(X354) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X355),s('type.realax.real',X356)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X354),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X355),s('type.realax.real',X356))))))) ) | $spl204),
  inference(cnf_transformation,[],[f15183_D])).
fof(f15183_D,plain,(
  ( ! [X356,X354,X355] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK18(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X354)),s('type.realax.real',sK17(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X354)))),s('type.realax.real',sK2)))) | sP0(X354) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X355),s('type.realax.real',X356)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X354),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X355),s('type.realax.real',X356))))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f15184,plain,(
  ( ! [X358,X357,X359] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK18(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X357)),s('type.realax.real',sK17(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X357)))),s('type.realax.real',sK2)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X357),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X358),s('type.realax.real',X359)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X358),s('type.realax.real',X359)))) | sP0(X357)) )),
  inference(forward_demodulation,[],[f14940,f704])).
fof(f14940,plain,(
  ( ! [X358,X357,X359] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X357),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X358),s('type.realax.real',X359)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X358),s('type.realax.real',X359)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK18(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X357)),s('type.realax.real',sK17(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X357)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')))))) | sP0(X357)) )),
  inference(superposition,[],[f518,f704])).
fof(f518,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK18(X4,X0)),s('type.realax.real',sK17(X4,X0)))),s('type.realax.real',X4)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | sP0(X0)) )),
  inference(cnf_transformation,[],[f432])).
fof(f15182,plain,(
  ( ! [X356,X354,X355] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK18(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X354)),s('type.realax.real',sK17(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X354)))),s('type.realax.real',sK2)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X354),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X355),s('type.realax.real',X356)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X355),s('type.realax.real',X356)))) | sP0(X354)) )),
  inference(forward_demodulation,[],[f14939,f704])).
fof(f14939,plain,(
  ( ! [X356,X354,X355] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',sK18(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X354)),s('type.realax.real',sK17(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X354)))),s('type.realax.real',sK2)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X354),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X355),s('type.realax.real',X356)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X355),s('type.realax.real',X356)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')))))) | sP0(X354)) )),
  inference(superposition,[],[f518,f704])).
fof(f15177,plain,(
  ( ! [X341,X339,X340] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X339),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X340),s('type.realax.real',X341)))))) | sP0(X339) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X339)),s('type.realax.real',sK18(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X339))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X340),s('type.realax.real',X341))))) ) | $spl202),
  inference(cnf_transformation,[],[f15177_D])).
fof(f15177_D,plain,(
  ( ! [X341,X339,X340] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X339),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X340),s('type.realax.real',X341)))))) | sP0(X339) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X339)),s('type.realax.real',sK18(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X339))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X340),s('type.realax.real',X341))))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f14934,plain,(
  ( ! [X341,X339,X340] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X339),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X340),s('type.realax.real',X341)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X340),s('type.realax.real',X341)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X339)),s('type.realax.real',sK18(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X339))))) | sP0(X339)) )),
  inference(superposition,[],[f517,f704])).
fof(f517,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK16(X0)),s('type.realax.real',sK18(X4,X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | sP0(X0)) )),
  inference(cnf_transformation,[],[f432])).
fof(f15172,plain,(
  ( ! [X325,X324,X326] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X324),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X325),s('type.realax.real',X326)))))) | sP0(X324) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X324)),s('type.realax.real',sK16(X324))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X325),s('type.realax.real',X326))))) ) | $spl200),
  inference(cnf_transformation,[],[f15172_D])).
fof(f15172_D,plain,(
  ( ! [X325,X324,X326] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X324),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X325),s('type.realax.real',X326)))))) | sP0(X324) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X324)),s('type.realax.real',sK16(X324))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X325),s('type.realax.real',X326))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f14929,plain,(
  ( ! [X325,X324,X326] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X324),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X325),s('type.realax.real',X326)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X325),s('type.realax.real',X326)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X324)),s('type.realax.real',sK16(X324))))) | sP0(X324)) )),
  inference(superposition,[],[f516,f704])).
fof(f516,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun('type.pair.prod'('type.realax.real','type.realax.real'),bool),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK17(X4,X0)),s('type.realax.real',sK16(X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4)))) | sP0(X0)) )),
  inference(cnf_transformation,[],[f432])).
fof(f15157,plain,(
  ( ! [X305] : (s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X305),s('type.realax.real',sK2))),s('type.realax.real',X305))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2))))) )),
  inference(forward_demodulation,[],[f15156,f704])).
fof(f15156,plain,(
  ( ! [X305] : (s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X305),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X305))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2))))) )),
  inference(forward_demodulation,[],[f14914,f622])).
fof(f14914,plain,(
  ( ! [X305] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X305),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X305),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f505,f704])).
fof(f505,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f387])).
fof(f387,plain,(
  ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real',X1) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))))),
  inference(ennf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) => s('type.realax.real',X1) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X14,X31] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X31)))) => s('type.realax.real',X31) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X14),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X14),s('type.realax.real',X31))))))))),
  file('thm.transc.DIVISION_LT','thm.topology.MR|49|_ADD_LT')).
fof(f15154,plain,(
  ( ! [X303] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X303),s('type.realax.real',sK2))),s('type.realax.real',X303)))))) )),
  inference(forward_demodulation,[],[f15153,f704])).
fof(f15153,plain,(
  ( ! [X303] : (s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X303),s('type.realax.real',sK2))),s('type.realax.real',X303))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f14912,f622])).
fof(f14912,plain,(
  ( ! [X303] : (s('type.realax.real',sK2) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X303),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X303),s('type.realax.real',sK2))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f505,f704])).
fof(f15152,plain,(
  ( ! [X302] : (s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X302),s('type.realax.real',sK2))),s('type.realax.real',X302))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2))))) )),
  inference(forward_demodulation,[],[f15151,f704])).
fof(f15151,plain,(
  ( ! [X302] : (s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X302),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X302))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2))))) )),
  inference(forward_demodulation,[],[f14911,f622])).
fof(f14911,plain,(
  ( ! [X302] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X302),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X302),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f504,f704])).
fof(f504,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f386])).
fof(f386,plain,(
  ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real',X1) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))))),
  inference(ennf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) => s('type.realax.real',X1) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))))),
  inference(rectify,[],[f81])).
fof(f81,axiom,(
  ! [X14,X31] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X31)))) => s('type.realax.real',X31) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X14),s('type.realax.real',X31))))))))),
  file('thm.transc.DIVISION_LT','thm.topology.MR|49|_SUB_LT')).
fof(f15149,plain,(
  ( ! [X300] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X300),s('type.realax.real',sK2))),s('type.realax.real',X300)))))) )),
  inference(forward_demodulation,[],[f15148,f704])).
fof(f15148,plain,(
  ( ! [X300] : (s('type.realax.real',sK2) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X300),s('type.realax.real',sK2))),s('type.realax.real',X300))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f14909,f622])).
fof(f14909,plain,(
  ( ! [X300] : (s('type.realax.real',sK2) = s('type.realax.real','const.topology.dist_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X300),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X300),s('type.realax.real',sK2))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f504,f704])).
fof(f15121,plain,(
  ( ! [X271,X270,X272] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X270),s('type.num.num',X271))),s('type.realax.real',X272))))),s('type.realax.real',sK2)))) | ~p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X270),s('type.realax.real',X272)))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X271),s('type.num.num',sK10(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X272,X270)))))) ) | $spl192),
  inference(cnf_transformation,[],[f15121_D])).
fof(f15121_D,plain,(
  ( ! [X271,X270,X272] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X270),s('type.num.num',X271))),s('type.realax.real',X272))))),s('type.realax.real',sK2)))) | ~p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X270),s('type.realax.real',X272)))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X271),s('type.num.num',sK10(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X272,X270)))))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f15124,plain,(
  ( ! [X273,X275,X274] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X273),s('type.num.num',X274))),s('type.realax.real',X275))))),s('type.realax.real',sK2)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X274),s('type.num.num',sK10(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X275,X273))))) | ~p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X273),s('type.realax.real',X275))))) )),
  inference(forward_demodulation,[],[f14889,f704])).
fof(f14889,plain,(
  ( ! [X273,X275,X274] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X273),s('type.num.num',X274))),s('type.realax.real',X275))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X274),s('type.num.num',sK10(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X275,X273))))) | ~p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X273),s('type.realax.real',X275))))) )),
  inference(superposition,[],[f497,f704])).
fof(f497,plain,(
  ( ! [X0,X7,X5,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X7))),s('type.realax.real',X1))))),s('type.realax.real',X5)))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X7),s('type.num.num',sK10(X5,X1,X0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) | ~p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f425])).
fof(f425,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X0),s('type.realax.real',X1)))) | ! [X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) | ! [X7] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X7),s('type.num.num',sK10(X5,X1,X0))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X7))),s('type.realax.real',X1))))),s('type.realax.real',X5))))))) & ((p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK11(X1,X0))))) & ! [X3] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK12(X3,X1,X0)),s('type.num.num',X3)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK12(X3,X1,X0)))),s('type.realax.real',X1))))),s('type.realax.real',sK11(X1,X0))))))) | p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X0),s('type.realax.real',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12])],[f424])).
fof(f424,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X0),s('type.realax.real',X1)))) | ! [X5] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) | ? [X6] : ! [X7] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X7),s('type.num.num',X6)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X7))),s('type.realax.real',X1))))),s('type.realax.real',X5))))))) & (? [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) & ! [X3] : ? [X4] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',X3)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',X1))))),s('type.realax.real',X2)))))) | p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X0),s('type.realax.real',X1))))))),
  inference(rectify,[],[f423])).
fof(f423,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X0),s('type.realax.real',X1)))) | ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | ? [X3] : ! [X4] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',X3)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',X1))))),s('type.realax.real',X2))))))) & (? [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) & ! [X3] : ? [X4] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',X3)))) & ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',X1))))),s('type.realax.real',X2)))))) | p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X0),s('type.realax.real',X1))))))),
  inference(nnf_transformation,[],[f385])).
fof(f385,plain,(
  ! [X0,X1] : (p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X0),s('type.realax.real',X1)))) <=> ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | ? [X3] : ! [X4] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',X3)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',X1))))),s('type.realax.real',X2)))))))),
  inference(ennf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1] : (p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X0),s('type.realax.real',X1)))) <=> ! [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) => ? [X3] : ! [X4] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',X3)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))),s('type.realax.real',X1))))),s('type.realax.real',X2)))))))),
  inference(rectify,[],[f180])).
fof(f180,axiom,(
  ! [X14,X47] : (p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X14),s('type.realax.real',X47)))) <=> ! [X42] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X42)))) => ? [X18] : ! [X9] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X9),s('type.num.num',X18)))) => p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X14),s('type.num.num',X9))),s('type.realax.real',X47))))),s('type.realax.real',X42)))))))),
  file('thm.transc.DIVISION_LT','thm.seq.SEQ')).
fof(f15120,plain,(
  ( ! [X271,X270,X272] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X270),s('type.num.num',X271))),s('type.realax.real',X272))))),s('type.realax.real',sK2)))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X271),s('type.num.num',sK10(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X272,X270))))) | ~p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X270),s('type.realax.real',X272))))) )),
  inference(forward_demodulation,[],[f14888,f704])).
fof(f14888,plain,(
  ( ! [X271,X270,X272] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X270),s('type.num.num',X271))),s('type.realax.real',X272))))),s('type.realax.real',sK2)))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X271),s('type.num.num',sK10(happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')),X272,X270))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.seq.-->_2'(s(fun('type.num.num','type.realax.real'),X270),s('type.realax.real',X272))))) )),
  inference(superposition,[],[f497,f704])).
fof(f987,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl39),
  inference(cnf_transformation,[],[f987_D])).
fof(f987_D,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f15021,plain,(
  ( ! [X6,X7] : (s('type.realax.real',sK2) = s('type.realax.real',X6) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl184),
  inference(cnf_transformation,[],[f15021_D])).
fof(f15021_D,plain,(
  ( ! [X6,X7] : (s('type.realax.real',sK2) = s('type.realax.real',X6) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s(fun('type.num.num','type.realax.real'),sK1))))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f14796,plain,(
  ( ! [X39,X38] : (s('type.realax.real',sK2) = s('type.realax.real',X38) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X39),s('type.realax.real',X38))),s(fun('type.num.num','type.realax.real'),sK1))))) )),
  inference(superposition,[],[f486,f704])).
fof(f15040,plain,(
  ( ! [X37,X36] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X36),s('type.realax.real',X37))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl10),
  inference(subsumption_resolution,[],[f14795,f10652])).
fof(f10652,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl10),
  inference(backward_demodulation,[],[f2713,f473])).
fof(f15037,plain,(
  ( ! [X33,X32] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5(sK1,X32,X33))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X33),s('type.realax.real',X32))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl188),
  inference(cnf_transformation,[],[f15037_D])).
fof(f15037_D,plain,(
  ( ! [X33,X32] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5(sK1,X32,X33))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X33),s('type.realax.real',X32))),s(fun('type.num.num','type.realax.real'),sK1))))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f14793,plain,(
  ( ! [X33,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5(sK1,X32,X33))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X33),s('type.realax.real',X32))),s(fun('type.num.num','type.realax.real'),sK1))))) )),
  inference(superposition,[],[f478,f704])).
fof(f478,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num',X8))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',sK5(X2,X1,X0))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),X2))))) )),
  inference(cnf_transformation,[],[f418])).
fof(f15036,plain,(
  ( ! [X26] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',X26)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X26)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X26),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))))))) ) | $spl186),
  inference(cnf_transformation,[],[f15036_D])).
fof(f15036_D,plain,(
  ( ! [X26] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',X26)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X26)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X26),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))))))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f14791,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',X26)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X26),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X26)))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X27),s('type.realax.real',X28))),s(fun('type.num.num','type.realax.real'),sK1))))) )),
  inference(superposition,[],[f475,f704])).
fof(f14766,plain,(
  ( ! [X6,X7] : (s('type.realax.real',sK2) = s('type.realax.real',X6) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s(fun('type.num.num','type.realax.real'),sK1))))) )),
  inference(superposition,[],[f704,f486])).
fof(f14648,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(resolution,[],[f698,f524])).
fof(f14587,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl176),
  inference(cnf_transformation,[],[f14587_D])).
fof(f14587_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f14589,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl179),
  inference(cnf_transformation,[],[f14589_D])).
fof(f14589_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl179),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl179])])).
fof(f14487,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) ) | $spl26),
  inference(equality_factoring,[],[f13594])).
fof(f14247,plain,(
  $false | ($spl71 | $spl117)),
  inference(subsumption_resolution,[],[f14246,f544])).
fof(f544,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f84])).
fof(f84,axiom,(
  ! [X19] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))))),
  file('thm.transc.DIVISION_LT','thm.arithmetic.LESS_EQ_SUC_REFL')).
fof(f14246,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl71 | $spl117)),
  inference(forward_demodulation,[],[f5575,f9133])).
fof(f9133,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK4)) | $spl71),
  inference(trivial_inequality_removal,[],[f9118])).
fof(f9118,plain,(
  s('type.num.num',sK20(sK4)) != s('type.num.num',sK20(sK4)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK4)) | $spl71),
  inference(superposition,[],[f1420,f539])).
fof(f1420,plain,(
  s('type.num.num',sK20(sK4)) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK4))))) | $spl71),
  inference(cnf_transformation,[],[f1420_D])).
fof(f1420_D,plain,(
  s('type.num.num',sK20(sK4)) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK4))))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f14244,plain,(
  $false | ($spl71 | $spl73)),
  inference(subsumption_resolution,[],[f14243,f542])).
fof(f542,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f260])).
fof(f260,axiom,(
  ! [X9] : p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))))),
  file('thm.transc.DIVISION_LT','thm.prim_rec.LESS_SUC_REFL')).
fof(f14243,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl71 | $spl73)),
  inference(forward_demodulation,[],[f1424,f9133])).
fof(f13462,plain,(
  ( ! [X35,X34] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X35)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X34),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl174),
  inference(cnf_transformation,[],[f13462_D])).
fof(f13462_D,plain,(
  ( ! [X35,X34] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X35)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X34),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f13453,plain,(
  ( ! [X8] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | $spl168),
  inference(cnf_transformation,[],[f13453_D])).
fof(f13453_D,plain,(
  ( ! [X8] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f13377,plain,(
  ( ! [X35,X33,X34] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X33))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X33),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X35)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X34),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl26),
  inference(superposition,[],[f13012,f562])).
fof(f13461,plain,(
  ( ! [X32] : (s('type.realax.real',X32) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real',X32)))) ) | $spl172),
  inference(cnf_transformation,[],[f13461_D])).
fof(f13461_D,plain,(
  ( ! [X32] : (s('type.realax.real',X32) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real',X32)))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f13376,plain,(
  ( ! [X31,X32] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X31))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X31),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real',X32) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real',X32)))) ) | $spl26),
  inference(superposition,[],[f13012,f550])).
fof(f13460,plain,(
  ( ! [X24,X23] : (s('type.realax.real',X23) = s('type.realax.real',X24) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real',X23))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real',X24)))) ) | $spl170),
  inference(cnf_transformation,[],[f13460_D])).
fof(f13460_D,plain,(
  ( ! [X24,X23] : (s('type.realax.real',X23) = s('type.realax.real',X24) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real',X23))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real',X24)))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f13373,plain,(
  ( ! [X24,X23,X22] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X22))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X22),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real',X23) = s('type.realax.real',X24) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real',X23))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real',X24)))) ) | $spl26),
  inference(superposition,[],[f13012,f549])).
fof(f13355,plain,(
  ( ! [X8] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl26),
  inference(superposition,[],[f13012,f539])).
fof(f13445,plain,(
  ( ! [X8] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) ) | $spl164),
  inference(cnf_transformation,[],[f13445_D])).
fof(f13445_D,plain,(
  ( ! [X8] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f13341,plain,(
  ( ! [X8] : (s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))) = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl26),
  inference(superposition,[],[f13012,f539])).
fof(f12762,plain,(
  $false | ($spl26 | $spl113)),
  inference(subsumption_resolution,[],[f12761,f544])).
fof(f12761,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl113)),
  inference(forward_demodulation,[],[f3782,f810])).
fof(f12758,plain,(
  $false | ($spl10 | $spl26 | $spl65)),
  inference(subsumption_resolution,[],[f12757,f544])).
fof(f12757,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl10 | $spl26 | $spl65)),
  inference(forward_demodulation,[],[f12756,f2713])).
fof(f12756,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl65)),
  inference(forward_demodulation,[],[f1134,f810])).
fof(f1007,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl59),
  inference(cnf_transformation,[],[f1007_D])).
fof(f1007_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f12608,plain,(
  ( ! [X24,X23] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X23),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X24)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X23),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl10),
  inference(superposition,[],[f10653,f562])).
fof(f12607,plain,(
  ( ! [X22] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X22) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X22)))) ) | $spl10),
  inference(superposition,[],[f10653,f550])).
fof(f12604,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X16) = s('type.realax.real',X17) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X16))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X17)))) ) | $spl10),
  inference(superposition,[],[f10653,f549])).
fof(f1004,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl54),
  inference(cnf_transformation,[],[f1004_D])).
fof(f1004_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s(fun('type.num.num','type.realax.real'),sK1))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f12596,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl10),
  inference(superposition,[],[f10653,f486])).
fof(f993,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X30)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X29),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl48),
  inference(cnf_transformation,[],[f993_D])).
fof(f993_D,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X30)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X29),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f12627,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl155),
  inference(cnf_transformation,[],[f12627_D])).
fof(f12627_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f12566,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X30)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X29),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl10),
  inference(superposition,[],[f10653,f562])).
fof(f992,plain,(
  ( ! [X28] : (s('type.realax.real',X28) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X28)))) ) | $spl46),
  inference(cnf_transformation,[],[f992_D])).
fof(f992_D,plain,(
  ( ! [X28] : (s('type.realax.real',X28) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X28)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f12565,plain,(
  ( ! [X28] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real',X28) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X28)))) ) | $spl10),
  inference(superposition,[],[f10653,f550])).
fof(f989,plain,(
  ( ! [X23,X22] : (s('type.realax.real',X22) = s('type.realax.real',X23) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X22))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23)))) ) | $spl42),
  inference(cnf_transformation,[],[f989_D])).
fof(f989_D,plain,(
  ( ! [X23,X22] : (s('type.realax.real',X22) = s('type.realax.real',X23) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X22))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f12562,plain,(
  ( ! [X23,X22] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real',X22) = s('type.realax.real',X23) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X22))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23)))) ) | $spl10),
  inference(superposition,[],[f10653,f549])).
fof(f12625,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl152),
  inference(cnf_transformation,[],[f12625_D])).
fof(f12625_D,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s(fun('type.num.num','type.realax.real'),sK1))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f12554,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl10),
  inference(superposition,[],[f10653,f486])).
fof(f12621,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl149),
  inference(cnf_transformation,[],[f12621_D])).
fof(f12621_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f12623,plain,(
  s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl150),
  inference(cnf_transformation,[],[f12623_D])).
fof(f12623_D,plain,(
  s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f12544,plain,(
  s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl10),
  inference(resolution,[],[f10653,f501])).
fof(f12619,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl147),
  inference(cnf_transformation,[],[f12619_D])).
fof(f12619_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f12617,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl10),
  inference(subsumption_resolution,[],[f12542,f542])).
fof(f12542,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X4),s('type.realax.real',X5))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl10),
  inference(resolution,[],[f10653,f475])).
fof(f12373,plain,(
  $false | ($spl10 | $spl26 | $spl80 | $spl89)),
  inference(subsumption_resolution,[],[f12372,f6958])).
fof(f6958,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl26 | $spl89)),
  inference(forward_demodulation,[],[f2407,f810])).
fof(f12372,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl10 | $spl80)),
  inference(forward_demodulation,[],[f2398,f2713])).
fof(f12265,plain,(
  $false | ($spl71 | $spl118)),
  inference(subsumption_resolution,[],[f12264,f536])).
fof(f536,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f300])).
fof(f300,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f299])).
fof(f299,plain,(
  ! [X0] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f106])).
fof(f106,axiom,(
  ! [X9] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))),
  file('thm.transc.DIVISION_LT','thm.num.NOT_SUC')).
fof(f12264,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl71 | $spl118)),
  inference(forward_demodulation,[],[f5577,f9133])).
fof(f5882,plain,(
  ( ! [X8] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))) ) | $spl120),
  inference(cnf_transformation,[],[f5882_D])).
fof(f5882_D,plain,(
  ( ! [X8] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f11875,plain,(
  ( ! [X8] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f4192,f539])).
fof(f11484,plain,(
  $false | ($spl70 | $spl74)),
  inference(subsumption_resolution,[],[f11483,f553])).
fof(f11483,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK20(sK4)),s('type.num.num','const.num.|48|_0')))) | ($spl70 | $spl74)),
  inference(subsumption_resolution,[],[f11444,f7607])).
fof(f7607,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK20(sK4)) | $spl70),
  inference(superposition,[],[f535,f1421])).
fof(f11444,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK4)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK20(sK4)),s('type.num.num','const.num.|48|_0')))) | $spl74),
  inference(resolution,[],[f1425,f574])).
fof(f1425,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20(sK4)),s('type.num.num','const.num.|48|_0')))) | $spl74),
  inference(cnf_transformation,[],[f1425_D])).
fof(f1425_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20(sK4)),s('type.num.num','const.num.|48|_0')))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f11482,plain,(
  $false | ($spl70 | $spl74)),
  inference(subsumption_resolution,[],[f11443,f7607])).
fof(f11443,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK20(sK4)) | $spl74),
  inference(resolution,[],[f1425,f556])).
fof(f11439,plain,(
  $false | ($spl10 | $spl26 | $spl66)),
  inference(subsumption_resolution,[],[f11438,f536])).
fof(f11438,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl10 | $spl26 | $spl66)),
  inference(forward_demodulation,[],[f9507,f2713])).
fof(f11435,plain,(
  $false | ($spl10 | $spl26 | $spl28 | $spl31)),
  inference(subsumption_resolution,[],[f11434,f570])).
fof(f11434,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl26 | $spl28 | $spl31)),
  inference(forward_demodulation,[],[f7561,f2713])).
fof(f7561,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | ($spl26 | $spl28 | $spl31)),
  inference(forward_demodulation,[],[f7560,f4199])).
fof(f4199,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) ) | $spl26),
  inference(backward_demodulation,[],[f810,f563])).
fof(f563,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f231])).
fof(f231,axiom,(
  ! [X9] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9))),s('type.num.num','const.num.|48|_0')))),
  file('thm.transc.DIVISION_LT','thm.numeral.numeral_distrib|split|27')).
fof(f7560,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))),s('type.num.num','const.num.|48|_0')))) | ($spl28 | $spl31)),
  inference(subsumption_resolution,[],[f7532,f852])).
fof(f7532,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))),s('type.num.num','const.num.|48|_0')))) | $spl31),
  inference(superposition,[],[f858,f556])).
fof(f994,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl50),
  inference(cnf_transformation,[],[f994_D])).
fof(f994_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f10985,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl10),
  inference(resolution,[],[f10652,f485])).
fof(f10247,plain,(
  $false | ($spl4 | $spl22 | $spl118)),
  inference(subsumption_resolution,[],[f6181,f9509])).
fof(f9509,plain,(
  s('type.num.num',sK4) = s('type.num.num',sK20(sK4)) | ($spl22 | $spl118)),
  inference(forward_demodulation,[],[f5577,f778])).
fof(f6181,plain,(
  s('type.num.num',sK4) != s('type.num.num',sK20(sK4)) | $spl4),
  inference(resolution,[],[f5460,f576])).
fof(f576,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f448])).
fof(f5460,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK20(sK4))))) | $spl4),
  inference(resolution,[],[f1244,f573])).
fof(f1244,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK20(sK4)),s('type.num.num',sK4)))) | $spl4),
  inference(superposition,[],[f542,f683])).
fof(f9542,plain,(
  ( ! [X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK4))))) ) | $spl22),
  inference(superposition,[],[f7971,f539])).
fof(f7971,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',sK4)))) ) | $spl22),
  inference(backward_demodulation,[],[f778,f4192])).
fof(f9508,plain,(
  $false | ($spl71 | $spl145)),
  inference(subsumption_resolution,[],[f9141,f9133])).
fof(f9317,plain,(
  $false | ($spl28 | $spl71 | $spl77)),
  inference(subsumption_resolution,[],[f9147,f852])).
fof(f9147,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl71 | $spl77)),
  inference(backward_demodulation,[],[f9133,f1428])).
fof(f9139,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK4))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK4)))))))) | $spl142),
  inference(cnf_transformation,[],[f9139_D])).
fof(f9139_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK4))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK4)))))))) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f9141,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK20(sK4)) | $spl145),
  inference(cnf_transformation,[],[f9141_D])).
fof(f9141_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK20(sK4)) <=> ~$spl145),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl145])])).
fof(f9123,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK20(sK4)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK4))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num',sK20(sK20(sK4)))))))) | $spl71),
  inference(superposition,[],[f1420,f539])).
fof(f9136,plain,(
  $false | ($spl71 | $spl75)),
  inference(subsumption_resolution,[],[f9133,f5533])).
fof(f8187,plain,(
  $false | ($spl4 | $spl21 | $spl22)),
  inference(subsumption_resolution,[],[f7970,f6217])).
fof(f6217,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK4)))) | $spl4),
  inference(forward_demodulation,[],[f6184,f683])).
fof(f6184,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK4)))),s('type.num.num',sK4)))) | $spl4),
  inference(resolution,[],[f5460,f522])).
fof(f7970,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK4)))) | ($spl21 | $spl22)),
  inference(backward_demodulation,[],[f778,f776])).
fof(f1421,plain,(
  s('type.num.num',sK20(sK4)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK4))))) | $spl70),
  inference(cnf_transformation,[],[f1421_D])).
fof(f1421_D,plain,(
  s('type.num.num',sK20(sK4)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK4))))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f5086,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK4)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK4))))) | $spl4),
  inference(superposition,[],[f683,f539])).
fof(f7963,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4)))) | $spl140),
  inference(cnf_transformation,[],[f7963_D])).
fof(f7963_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4)))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f6488,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4)))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl9),
  inference(resolution,[],[f688,f579])).
fof(f6489,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4)))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl9),
  inference(resolution,[],[f688,f579])).
fof(f7961,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl139),
  inference(cnf_transformation,[],[f7961_D])).
fof(f7961_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f6493,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl9),
  inference(resolution,[],[f688,f527])).
fof(f7779,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20(sK20(sK4))),s('type.num.num','const.num.|48|_0')))) | $spl137),
  inference(cnf_transformation,[],[f7779_D])).
fof(f7779_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20(sK20(sK4))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f7567,plain,(
  s('type.num.num',sK20(sK4)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20(sK20(sK4))),s('type.num.num','const.num.|48|_0')))) | $spl70),
  inference(superposition,[],[f1421,f556])).
fof(f7777,plain,(
  s('type.num.num',sK20(sK20(sK4))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK4)))))) | $spl134),
  inference(cnf_transformation,[],[f7777_D])).
fof(f7777_D,plain,(
  s('type.num.num',sK20(sK20(sK4))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK4)))))) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f7566,plain,(
  s('type.num.num',sK20(sK4)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK20(sK4))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK20(sK4)))))) | $spl70),
  inference(superposition,[],[f1421,f539])).
fof(f7446,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl21),
  inference(resolution,[],[f776,f524])).
fof(f7120,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl127),
  inference(cnf_transformation,[],[f7120_D])).
fof(f7120_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f7082,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl28)),
  inference(resolution,[],[f6782,f524])).
fof(f7002,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl123),
  inference(cnf_transformation,[],[f7002_D])).
fof(f7002_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl123),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl123])])).
fof(f6964,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl26 | $spl28)),
  inference(resolution,[],[f4742,f524])).
fof(f2676,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | $spl99),
  inference(cnf_transformation,[],[f2676_D])).
fof(f2676_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f5682,plain,(
  ( ! [X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(superposition,[],[f4192,f556])).
fof(f5681,plain,(
  ( ! [X8] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f4192,f539])).
fof(f5538,plain,(
  s('type.num.num',sK20(sK4)) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20(sK4)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl75),
  inference(resolution,[],[f1426,f524])).
fof(f5456,plain,(
  $false | ($spl26 | $spl78)),
  inference(subsumption_resolution,[],[f5455,f536])).
fof(f5455,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.ZERO_0')))) | ($spl26 | $spl78)),
  inference(forward_demodulation,[],[f2396,f810])).
fof(f5080,plain,(
  $false | ($spl26 | $spl68)),
  inference(subsumption_resolution,[],[f5079,f553])).
fof(f5079,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl26 | $spl68)),
  inference(forward_demodulation,[],[f1189,f810])).
fof(f4363,plain,(
  $false | ($spl25 | $spl26)),
  inference(subsumption_resolution,[],[f4201,f570])).
fof(f4201,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl25 | $spl26)),
  inference(backward_demodulation,[],[f810,f783])).
fof(f811,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | $spl27),
  inference(cnf_transformation,[],[f811_D])).
fof(f811_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f2683,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK4)))) | $spl95),
  inference(superposition,[],[f2672,f539])).
fof(f858,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))))) | $spl31),
  inference(cnf_transformation,[],[f858_D])).
fof(f858_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f3636,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))))) | $spl13),
  inference(superposition,[],[f692,f558])).
fof(f558,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f444])).
fof(f444,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f238])).
fof(f238,axiom,(
  ! [X9] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X9))),
  file('thm.transc.DIVISION_LT','thm.numeral.numeral_distrib|split|18')).
fof(f3637,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))))) | $spl13),
  inference(superposition,[],[f692,f560])).
fof(f560,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f445])).
fof(f445,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f244])).
fof(f244,axiom,(
  ! [X9] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X9))),
  file('thm.transc.DIVISION_LT','thm.numeral.numeral_distrib|split|17')).
fof(f2401,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl82),
  inference(cnf_transformation,[],[f2401_D])).
fof(f2401_D,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f3670,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(superposition,[],[f2996,f539])).
fof(f2996,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(subsumption_resolution,[],[f2796,f535])).
fof(f2796,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f546,f561])).
fof(f561,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f445])).
fof(f2405,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.num.|48|_0')))) | $spl87),
  inference(cnf_transformation,[],[f2405_D])).
fof(f2405_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f3671,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f2996,f556])).
fof(f4119,plain,(
  $false | ($spl27 | $spl28)),
  inference(subsumption_resolution,[],[f4118,f852])).
fof(f4118,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl27 | $spl28)),
  inference(forward_demodulation,[],[f4117,f852])).
fof(f4117,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl27 | $spl28)),
  inference(subsumption_resolution,[],[f3919,f811])).
fof(f3919,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | $spl28),
  inference(superposition,[],[f560,f852])).
fof(f4116,plain,(
  $false | ($spl27 | $spl28)),
  inference(subsumption_resolution,[],[f4115,f811])).
fof(f4115,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') | $spl28),
  inference(forward_demodulation,[],[f4114,f852])).
fof(f4114,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl28),
  inference(subsumption_resolution,[],[f3918,f852])).
fof(f3918,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl28),
  inference(superposition,[],[f560,f852])).
fof(f4113,plain,(
  $false | ($spl27 | $spl28)),
  inference(subsumption_resolution,[],[f4112,f852])).
fof(f4112,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl27 | $spl28)),
  inference(forward_demodulation,[],[f4111,f852])).
fof(f4111,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl27 | $spl28)),
  inference(subsumption_resolution,[],[f3915,f811])).
fof(f3915,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | $spl28),
  inference(superposition,[],[f558,f852])).
fof(f4110,plain,(
  $false | ($spl27 | $spl28)),
  inference(subsumption_resolution,[],[f4109,f811])).
fof(f4109,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') | $spl28),
  inference(forward_demodulation,[],[f4108,f852])).
fof(f4108,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl28),
  inference(subsumption_resolution,[],[f3914,f852])).
fof(f3914,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl28),
  inference(superposition,[],[f558,f852])).
fof(f4030,plain,(
  $false | ($spl27 | $spl28)),
  inference(subsumption_resolution,[],[f3992,f811])).
fof(f3992,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') | $spl28),
  inference(trivial_inequality_removal,[],[f3822])).
fof(f3822,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') | $spl28),
  inference(superposition,[],[f560,f852])).
fof(f4029,plain,(
  $false | ($spl27 | $spl28)),
  inference(subsumption_resolution,[],[f3993,f811])).
fof(f3993,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') | $spl28),
  inference(trivial_inequality_removal,[],[f3821])).
fof(f3821,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') | $spl28),
  inference(superposition,[],[f558,f852])).
fof(f4014,plain,(
  $false | ($spl27 | $spl28)),
  inference(subsumption_resolution,[],[f4013,f852])).
fof(f4013,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl27 | $spl28)),
  inference(forward_demodulation,[],[f4012,f852])).
fof(f4012,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl27 | $spl28)),
  inference(subsumption_resolution,[],[f3812,f811])).
fof(f3812,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | $spl28),
  inference(superposition,[],[f852,f560])).
fof(f4011,plain,(
  $false | ($spl27 | $spl28)),
  inference(subsumption_resolution,[],[f4010,f852])).
fof(f4010,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ($spl27 | $spl28)),
  inference(forward_demodulation,[],[f4009,f852])).
fof(f4009,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl27 | $spl28)),
  inference(subsumption_resolution,[],[f3811,f811])).
fof(f3811,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | $spl28),
  inference(superposition,[],[f852,f558])).
fof(f3782,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl113),
  inference(cnf_transformation,[],[f3782_D])).
fof(f3782_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f3784,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl114),
  inference(cnf_transformation,[],[f3784_D])).
fof(f3784_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f3741,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl25),
  inference(resolution,[],[f783,f524])).
fof(f783,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) | $spl25),
  inference(cnf_transformation,[],[f783_D])).
fof(f783_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f3685,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | $spl111),
  inference(cnf_transformation,[],[f3685_D])).
fof(f3685_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f3656,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f2996,f556])).
fof(f2396,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.ZERO_0')))) | $spl78),
  inference(cnf_transformation,[],[f2396_D])).
fof(f2396_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.ZERO_0')))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f3655,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.ZERO_0'))))),
  inference(superposition,[],[f2996,f539])).
fof(f853,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl29),
  inference(cnf_transformation,[],[f853_D])).
fof(f853_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f3615,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK4)))) | $spl13),
  inference(superposition,[],[f692,f539])).
fof(f3464,plain,(
  $false | ($spl12 | $spl95)),
  inference(subsumption_resolution,[],[f3432,f2672])).
fof(f3432,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',sK4) | $spl12),
  inference(trivial_inequality_removal,[],[f3262])).
fof(f3262,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',sK4) | $spl12),
  inference(superposition,[],[f560,f691])).
fof(f691,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))) | $spl12),
  inference(cnf_transformation,[],[f691_D])).
fof(f691_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3463,plain,(
  $false | ($spl12 | $spl95)),
  inference(subsumption_resolution,[],[f3433,f2672])).
fof(f3433,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',sK4) | $spl12),
  inference(trivial_inequality_removal,[],[f3261])).
fof(f3261,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',sK4) | $spl12),
  inference(superposition,[],[f558,f691])).
fof(f852,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl28),
  inference(cnf_transformation,[],[f852_D])).
fof(f852_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f3234,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK4)))) | $spl12),
  inference(superposition,[],[f691,f539])).
fof(f692,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))) | $spl13),
  inference(cnf_transformation,[],[f692_D])).
fof(f692_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f3195,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))) | $spl11),
  inference(superposition,[],[f690,f560])).
fof(f3194,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))) | $spl11),
  inference(superposition,[],[f690,f558])).
fof(f2407,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl89),
  inference(cnf_transformation,[],[f2407_D])).
fof(f2407_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f2993,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl108),
  inference(cnf_transformation,[],[f2993_D])).
fof(f2993_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f2785,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f546,f560])).
fof(f2784,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f546,f558])).
fof(f2991,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl106),
  inference(cnf_transformation,[],[f2991_D])).
fof(f2991_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f2783,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f546,f556])).
fof(f2782,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),
  inference(superposition,[],[f546,f539])).
fof(f2988,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | $spl104),
  inference(cnf_transformation,[],[f2988_D])).
fof(f2988_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f2768,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f546,f556])).
fof(f2767,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.ZERO_0'))))),
  inference(superposition,[],[f546,f539])).
fof(f2756,plain,(
  $false | ($spl10 | $spl101)),
  inference(subsumption_resolution,[],[f2755,f553])).
fof(f2755,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | ($spl10 | $spl101)),
  inference(subsumption_resolution,[],[f2714,f2678])).
fof(f2714,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK4) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(resolution,[],[f689,f574])).
fof(f2754,plain,(
  $false | ($spl10 | $spl101)),
  inference(subsumption_resolution,[],[f2713,f2678])).
fof(f995,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl51),
  inference(cnf_transformation,[],[f995_D])).
fof(f995_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f923,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f474,f556])).
fof(f2678,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) | $spl101),
  inference(cnf_transformation,[],[f2678_D])).
fof(f2678_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f2604,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl23),
  inference(superposition,[],[f777,f539])).
fof(f2605,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | $spl23),
  inference(superposition,[],[f777,f556])).
fof(f2674,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl97),
  inference(cnf_transformation,[],[f2674_D])).
fof(f2674_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f924,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4)))),
  inference(superposition,[],[f474,f558])).
fof(f925,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4)))),
  inference(superposition,[],[f474,f560])).
fof(f1189,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl68),
  inference(cnf_transformation,[],[f1189_D])).
fof(f1189_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f1060,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))) | $spl11),
  inference(superposition,[],[f733,f558])).
fof(f733,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4)))) | $spl11),
  inference(resolution,[],[f690,f572])).
fof(f1061,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))) | $spl11),
  inference(superposition,[],[f733,f560])).
fof(f2670,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl93),
  inference(cnf_transformation,[],[f2670_D])).
fof(f2670_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f2672,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',sK4) | $spl95),
  inference(cnf_transformation,[],[f2672_D])).
fof(f2672_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',sK4) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f2606,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',sK4) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl23),
  inference(superposition,[],[f777,f558])).
fof(f2607,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',sK4) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl23),
  inference(superposition,[],[f777,f560])).
fof(f2668,plain,(
  $false | $spl29),
  inference(subsumption_resolution,[],[f2667,f570])).
fof(f2667,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl29),
  inference(forward_demodulation,[],[f2657,f563])).
fof(f2657,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | $spl29),
  inference(trivial_inequality_removal,[],[f2646])).
fof(f2646,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) | $spl29),
  inference(superposition,[],[f853,f556])).
fof(f2409,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl90),
  inference(cnf_transformation,[],[f2409_D])).
fof(f2409_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f2193,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl22),
  inference(superposition,[],[f1619,f560])).
fof(f1619,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl22),
  inference(backward_demodulation,[],[f778,f547])).
fof(f547,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f241])).
fof(f241,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.transc.DIVISION_LT','thm.numeral.numeral_distrib|split|13')).
fof(f2192,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl22),
  inference(superposition,[],[f1619,f558])).
fof(f2403,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl84),
  inference(cnf_transformation,[],[f2403_D])).
fof(f2403_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f2191,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.num.|48|_0')))) | $spl22),
  inference(superposition,[],[f1619,f556])).
fof(f2190,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl22),
  inference(superposition,[],[f1619,f539])).
fof(f2398,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | $spl80),
  inference(cnf_transformation,[],[f2398_D])).
fof(f2398_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f2176,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) | $spl22),
  inference(superposition,[],[f1619,f556])).
fof(f2175,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20('const.arithmetic.ZERO_0')))) | $spl22),
  inference(superposition,[],[f1619,f539])).
fof(f1428,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK20(sK4)))) | $spl77),
  inference(cnf_transformation,[],[f1428_D])).
fof(f1428_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK20(sK4)))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f1136,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl66),
  inference(cnf_transformation,[],[f1136_D])).
fof(f1136_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f1198,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK20(sK4)))) | $spl4),
  inference(superposition,[],[f683,f560])).
fof(f1197,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK20(sK4)))) | $spl4),
  inference(superposition,[],[f683,f558])).
fof(f1426,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20(sK4)),s('type.num.num','const.num.|48|_0')))) | $spl75),
  inference(cnf_transformation,[],[f1426_D])).
fof(f1426_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20(sK4)),s('type.num.num','const.num.|48|_0')))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f1196,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK20(sK4)),s('type.num.num','const.num.|48|_0')))) | $spl4),
  inference(superposition,[],[f683,f556])).
fof(f1424,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK20(sK4)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl73),
  inference(cnf_transformation,[],[f1424_D])).
fof(f1424_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK20(sK4)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f1422,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK20(sK4)),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl4),
  inference(forward_demodulation,[],[f1195,f546])).
fof(f1195,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK20(sK4)),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl4),
  inference(superposition,[],[f683,f554])).
fof(f1194,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK20(sK4)) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK20(sK4))))) | $spl4),
  inference(superposition,[],[f683,f539])).
fof(f1190,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl69),
  inference(cnf_transformation,[],[f1190_D])).
fof(f1190_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f688,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl9),
  inference(cnf_transformation,[],[f688_D])).
fof(f688_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f1191,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl13),
  inference(forward_demodulation,[],[f1151,f546])).
fof(f1151,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl13),
  inference(superposition,[],[f1087,f554])).
fof(f1087,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl13),
  inference(resolution,[],[f856,f578])).
fof(f856,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl13),
  inference(forward_demodulation,[],[f847,f563])).
fof(f847,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))),s('type.num.num','const.num.|48|_0')))) | $spl13),
  inference(trivial_inequality_removal,[],[f836])).
fof(f836,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))),s('type.num.num','const.num.|48|_0')))) | $spl13),
  inference(superposition,[],[f692,f556])).
fof(f1150,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK4)))) | $spl13),
  inference(superposition,[],[f1087,f539])).
fof(f1134,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl65),
  inference(cnf_transformation,[],[f1134_D])).
fof(f1134_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f1091,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) | $spl13),
  inference(resolution,[],[f856,f524])).
fof(f1009,plain,(
  ( ! [X24,X23] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X23),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X24)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X23),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl62),
  inference(cnf_transformation,[],[f1009_D])).
fof(f1009_D,plain,(
  ( ! [X24,X23] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X23),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X24)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X23),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f963,plain,(
  ( ! [X24,X23] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X23),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X24)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X23),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f474,f562])).
fof(f1008,plain,(
  ( ! [X22] : (s('type.realax.real',X22) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))),s('type.realax.real',X22)))) ) | $spl60),
  inference(cnf_transformation,[],[f1008_D])).
fof(f1008_D,plain,(
  ( ! [X22] : (s('type.realax.real',X22) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))),s('type.realax.real',X22)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f962,plain,(
  ( ! [X22] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X22) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))),s('type.realax.real',X22)))) )),
  inference(superposition,[],[f474,f550])).
fof(f1005,plain,(
  ( ! [X17,X16] : (s('type.realax.real',X16) = s('type.realax.real',X17) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))),s('type.realax.real',X16))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))),s('type.realax.real',X17)))) ) | $spl56),
  inference(cnf_transformation,[],[f1005_D])).
fof(f1005_D,plain,(
  ( ! [X17,X16] : (s('type.realax.real',X16) = s('type.realax.real',X17) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))),s('type.realax.real',X16))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))),s('type.realax.real',X17)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f959,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X16) = s('type.realax.real',X17) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))),s('type.realax.real',X16))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))),s('type.realax.real',X17)))) )),
  inference(superposition,[],[f474,f549])).
fof(f951,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X1),s('type.realax.real',X0))),s(fun('type.num.num','type.realax.real'),sK1))))) )),
  inference(superposition,[],[f474,f486])).
fof(f996,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(forward_demodulation,[],[f922,f546])).
fof(f922,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(superposition,[],[f474,f554])).
fof(f921,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK4))))),
  inference(superposition,[],[f474,f539])).
fof(f991,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) | $spl45),
  inference(cnf_transformation,[],[f991_D])).
fof(f991_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f917,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real','const.real.abs_1'(s('type.realax.real',X30)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X29),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f474,f562])).
fof(f916,plain,(
  ( ! [X28] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) | s('type.realax.real',X28) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X28)))) )),
  inference(superposition,[],[f474,f550])).
fof(f913,plain,(
  ( ! [X23,X22] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) | s('type.realax.real',X22) = s('type.realax.real',X23) | s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X22))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X23)))) )),
  inference(superposition,[],[f474,f549])).
fof(f988,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s(fun('type.num.num','type.realax.real'),sK1))))) ) | $spl40),
  inference(cnf_transformation,[],[f988_D])).
fof(f988_D,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s(fun('type.num.num','type.realax.real'),sK1))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f905,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X6))),s(fun('type.num.num','type.realax.real'),sK1))))) )),
  inference(superposition,[],[f474,f486])).
fof(f975,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) | $spl35),
  inference(cnf_transformation,[],[f975_D])).
fof(f975_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f977,plain,(
  s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))) | $spl36),
  inference(cnf_transformation,[],[f977_D])).
fof(f977_D,plain,(
  s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f879,plain,(
  s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))) = s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.|48|_0'))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))))))))),
  inference(resolution,[],[f474,f501])).
fof(f973,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl33),
  inference(cnf_transformation,[],[f973_D])).
fof(f973_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f971,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),sK1))))) )),
  inference(subsumption_resolution,[],[f877,f543])).
fof(f877,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4))),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s(fun('type.num.num','type.realax.real'),sK1))))) )),
  inference(resolution,[],[f474,f475])).
fof(f876,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | $spl13),
  inference(forward_demodulation,[],[f862,f546])).
fof(f862,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl13),
  inference(superposition,[],[f850,f554])).
fof(f850,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',sK4) | $spl13),
  inference(trivial_inequality_removal,[],[f830])).
fof(f830,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',sK4) | $spl13),
  inference(superposition,[],[f692,f559])).
fof(f559,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f444])).
fof(f861,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK4)))) | $spl13),
  inference(superposition,[],[f850,f539])).
fof(f838,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))))) | $spl13),
  inference(superposition,[],[f692,f560])).
fof(f837,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))))) | $spl13),
  inference(superposition,[],[f692,f558])).
fof(f854,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | $spl13),
  inference(forward_demodulation,[],[f815,f546])).
fof(f815,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) | $spl13),
  inference(superposition,[],[f692,f554])).
fof(f814,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK4)))) | $spl13),
  inference(superposition,[],[f692,f539])).
fof(f797,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))) | $spl11),
  inference(superposition,[],[f729,f560])).
fof(f796,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ZERO_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))) | $spl11),
  inference(superposition,[],[f729,f558])).
fof(f744,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))) | $spl11),
  inference(superposition,[],[f690,f560])).
fof(f743,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4))) | $spl11),
  inference(superposition,[],[f690,f558])).
fof(f734,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl11),
  inference(resolution,[],[f690,f524])).
fof(f700,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))))) | $spl17),
  inference(cnf_transformation,[],[f700_D])).
fof(f700_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f702,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.ZERO_0')))) | $spl18),
  inference(cnf_transformation,[],[f702_D])).
fof(f702_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.ZERO_0')))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f663,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))),
  inference(superposition,[],[f473,f560])).
fof(f662,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))),
  inference(superposition,[],[f473,f558])).
fof(f694,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | $spl14),
  inference(cnf_transformation,[],[f694_D])).
fof(f694_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f647,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4)))),
  inference(superposition,[],[f473,f560])).
fof(f646,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',sK4)))),
  inference(superposition,[],[f473,f558])).
fof(f645,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f473,f556])).
fof(f686,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1))))))),
  inference(forward_demodulation,[],[f644,f546])).
fof(f644,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(superposition,[],[f473,f554])).
fof(f643,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.transc.dsize_1'(s(fun('type.num.num','type.realax.real'),sK1)))))) | s('type.num.num',sK4) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK20(sK4))))),
  inference(superposition,[],[f473,f539])).
fof(f632,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num',sK4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK4)))))))) | ~p(s(bool,'const.transc.division_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s(fun('type.num.num','type.realax.real'),sK1))))) )),
  inference(resolution,[],[f473,f485])).
% SZS output end Proof for thm.transc
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 392
% Passive clauses: 8887
% Generated clauses: 31347
% Final active clauses: 253
% Final passive clauses: 6697
% Input formulas: 261
% Initial clauses: 157
% 
% Selected by SInE selection: 95
% SInE iterations: 1
% 
% Duplicate literals: 243
% Trivial inequalities: 143
% Fw subsumption resolutions: 1830
% Fw demodulations: 8900
% Bw demodulations: 93
% 
% Simple tautologies: 312
% Equational tautologies: 23
% Forward subsumptions: 10663
% Fw demodulations to eq. taut.: 3
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 1224
% Forward superposition: 9854
% Backward superposition: 8479
% Self superposition: 11
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 172
% Split components: 345
% SAT solver clauses: 224
% SAT solver unit clauses: 8
% SAT solver binary clauses: 103
% 
% Sat splits: 172
% Sat splitting refutations: 51
% 
% Memory used [KB]: 15991
% Time elapsed: 0.580 s
% ------------------------------
% ------------------------------
% Success in time 12.306 s
