Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numeral
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numeral
% SZS output start Proof for thm.numeral
fof(f11167,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f11166,f369,f658,f656,f6589,f370,f659,f657,f381,f661,f457,f662,f660])).
fof(f660,plain,(
  p(s(bool,'F_0')) | $spl4),
  inference(cnf_transformation,[],[f660_D])).
fof(f660_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f662,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f662_D])).
fof(f662_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f457,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f281])).
fof(f281,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f207])).
fof(f207,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X9] : ((p(s(bool,'F_0')) => p(s(bool,X9))) <=> p(s(bool,'T_0')))),
  file('thm.numeral.numeral_distrib|split|30','thm.bool.IMP_CLAUSES|split|2')).
fof(f661,plain,(
  ~p(s(bool,'F_0')) | $spl5),
  inference(cnf_transformation,[],[f661_D])).
fof(f661_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f381,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.numeral.numeral_distrib|split|30',aHL_FALSITY)).
fof(f657,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | $spl1),
  inference(cnf_transformation,[],[f657_D])).
fof(f657_D,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f659,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK13) | $spl3),
  inference(cnf_transformation,[],[f659_D])).
fof(f659_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK13) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f370,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK13) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))))),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK13)) & (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK13))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f247])).
fof(f247,plain,(
  ? [X0] : ((p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)))),
  inference(nnf_transformation,[],[f196])).
fof(f196,plain,(
  ? [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <~> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(ennf_transformation,[],[f89])).
fof(f89,plain,(
  ~! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X6))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X6))),
  file('thm.numeral.numeral_distrib|split|30','thm.numeral.numeral_distrib|split|30')).
fof(f6589,plain,(
  $false | ($spl1 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f6587,f661])).
fof(f6587,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(superposition,[],[f6505,f715])).
fof(f715,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl2)),
  inference(resolution,[],[f414,f674])).
fof(f674,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f657,f658])).
fof(f414,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.numeral.numeral_distrib|split|30','thm.bool.EQ_CLAUSES|split|3')).
fof(f6505,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X0))))) )),
  inference(equality_resolution,[],[f504])).
fof(f504,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.>_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) & ((~p(s(bool,'const.arithmetic.>_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(flattening,[],[f304])).
fof(f304,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | (p(s(bool,'const.arithmetic.>_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1))) & ((~p(s(bool,'const.arithmetic.>_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> (p(s(bool,'const.arithmetic.>_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8,X6] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X8),s('type.num.num',X6)))) <=> (p(s(bool,'const.arithmetic.>_2'(s('type.num.num',X8),s('type.num.num',X6)))) | s('type.num.num',X6) = s('type.num.num',X8)))),
  file('thm.numeral.numeral_distrib|split|30','thm.arithmetic.GREATER_OR_EQ')).
fof(f656,plain,(
  p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | $spl0),
  inference(cnf_transformation,[],[f656_D])).
fof(f656_D,plain,(
  p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f658,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK13) | $spl2),
  inference(cnf_transformation,[],[f658_D])).
fof(f658_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK13) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f369,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK13) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))))),
  inference(cnf_transformation,[],[f248])).
fof(f11166,plain,(
  $false | ($spl0 | $spl3 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f11165,f661])).
fof(f11165,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3 | $spl6)),
  inference(forward_demodulation,[],[f11164,f7596])).
fof(f7596,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(superposition,[],[f483,f716])).
fof(f716,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) )),
  inference(resolution,[],[f414,f470])).
fof(f470,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X6] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0'))))),
  file('thm.numeral.numeral_distrib|split|30','thm.prim_rec.NOT_LESS_|48|')).
fof(f483,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.>_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.>_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8,X6] : s(bool,'const.arithmetic.>_2'(s('type.num.num',X8),s('type.num.num',X6))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X8)))),
  file('thm.numeral.numeral_distrib|split|30','thm.arithmetic.GREATER_DEF')).
fof(f11164,plain,(
  p(s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | ($spl0 | $spl3 | $spl6)),
  inference(subsumption_resolution,[],[f11163,f659])).
fof(f11163,plain,(
  p(s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK13) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f11059,f662])).
fof(f11059,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK13) | $spl0),
  inference(superposition,[],[f502,f6643])).
fof(f6643,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK13))) | $spl0),
  inference(resolution,[],[f656,f410])).
fof(f410,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.numeral.numeral_distrib|split|30','thm.bool.EQ_CLAUSES|split|0')).
fof(f502,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.>_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f305])).
% SZS output end Proof for thm.numeral
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 575
% Passive clauses: 2598
% Generated clauses: 10789
% Final active clauses: 545
% Final passive clauses: 1830
% Input formulas: 87
% Initial clauses: 278
% 
% Duplicate literals: 179
% Trivial inequalities: 37
% Fw subsumption resolutions: 804
% Fw demodulations: 609
% Bw demodulations: 11
% 
% Simple tautologies: 204
% Equational tautologies: 21
% Forward subsumptions: 6249
% Fw demodulations to eq. taut.: 6
% 
% Binary resolution: 3893
% Factoring: 822
% Forward superposition: 2207
% Backward superposition: 1850
% Self superposition: 20
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 154
% Split components: 310
% SAT solver clauses: 269
% SAT solver unit clauses: 18
% SAT solver binary clauses: 57
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 154
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 267
% TWLsolver calls for satisfiability: 223
% 
% Memory used [KB]: 4861
% Time elapsed: 0.197 s
% ------------------------------
% ------------------------------
% Success in time 0.201 s
