Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f859,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f853,f146,f148,f564,f562,f94,f151])).
fof(f151,plain,(
  ~p(s(bool,'F_0')) | $spl5),
  inference(cnf_transformation,[],[f151_D])).
fof(f151_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f94,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.relation.equivalence_inv_identity',aHL_FALSITY)).
fof(f562,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.relation.equivalence_1'(s(fun(X6,fun(X6,bool)),X7)))) | s(bool,'T_0') = s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7)))) ) | $spl26),
  inference(cnf_transformation,[],[f562_D])).
fof(f562_D,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.relation.equivalence_1'(s(fun(X6,fun(X6,bool)),X7)))) | s(bool,'T_0') = s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f564,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.relation.equivalence_1'(s(fun(X12,fun(X12,bool)),X13)))) | s(bool,'T_0') = s(bool,'const.relation.symmetric_1'(s(fun(X12,fun(X12,bool)),X13)))) ) | $spl5),
  inference(subsumption_resolution,[],[f560,f151])).
fof(f560,plain,(
  ( ! [X12,X13] : (p(s(bool,'F_0')) | ~p(s(bool,'const.relation.equivalence_1'(s(fun(X12,fun(X12,bool)),X13)))) | s(bool,'T_0') = s(bool,'const.relation.symmetric_1'(s(fun(X12,fun(X12,bool)),X13)))) )),
  inference(superposition,[],[f123,f95])).
fof(f95,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.equivalence_inv_identity',aHL_BOOL_CASES)).
fof(f123,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))))) & (~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) | (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))))) & ((~p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) | p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (p(s(bool,'const.relation.reflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X6,fun(X6,bool)),X7)))) <=> (p(s(bool,'const.relation.reflexive_1'(s(fun(X6,fun(X6,bool)),X7)))) & p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),X7))))))),
  file('thm.relation.equivalence_inv_identity','thm.relation.equivalence_def')).
fof(f148,plain,(
  p(s(bool,'T_0')) | $spl0),
  inference(cnf_transformation,[],[f148_D])).
fof(f148_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f146,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f120])).
fof(f120,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X9] : (s(X6,X9) = s(X6,X9) <=> p(s(bool,'T_0')))),
  file('thm.relation.equivalence_inv_identity','thm.bool.REFL_CLAUSE')).
fof(f853,plain,(
  $false | ($spl0 | $spl26)),
  inference(subsumption_resolution,[],[f852,f91])).
fof(f91,plain,(
  s(fun(sK0,fun(sK0,bool)),sK1) != s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)))),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  p(s(bool,'const.relation.equivalence_1'(s(fun(sK0,fun(sK0,bool)),sK1)))) & s(fun(sK0,fun(sK0,bool)),sK1) != s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f52])).
fof(f52,plain,(
  ? [X0,X1] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) & s(fun(X0,fun(X0,bool)),X1) != s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X6,fun(X6,bool)),X7)))) => s(fun(X6,fun(X6,bool)),X7) = s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.equivalence_1'(s(fun(X6,fun(X6,bool)),X7)))) => s(fun(X6,fun(X6,bool)),X7) = s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X7))))),
  file('thm.relation.equivalence_inv_identity','thm.relation.equivalence_inv_identity')).
fof(f852,plain,(
  s(fun(sK0,fun(sK0,bool)),sK1) = s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | ($spl0 | $spl26)),
  inference(subsumption_resolution,[],[f820,f148])).
fof(f820,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK0,fun(sK0,bool)),sK1) = s(fun(sK0,fun(sK0,bool)),'const.relation.inv_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | ($spl0 | $spl26)),
  inference(superposition,[],[f119,f814])).
fof(f814,plain,(
  s(bool,'T_0') = s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | ($spl0 | $spl26)),
  inference(subsumption_resolution,[],[f808,f148])).
fof(f808,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.relation.symmetric_1'(s(fun(sK0,fun(sK0,bool)),sK1))) | $spl26),
  inference(superposition,[],[f562,f158])).
fof(f158,plain,(
  s(bool,'T_0') = s(bool,'const.relation.equivalence_1'(s(fun(sK0,fun(sK0,bool)),sK1)))),
  inference(resolution,[],[f97,f90])).
fof(f90,plain,(
  p(s(bool,'const.relation.equivalence_1'(s(fun(sK0,fun(sK0,bool)),sK1))))),
  inference(cnf_transformation,[],[f64])).
fof(f97,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.relation.equivalence_inv_identity','thm.bool.EQ_CLAUSES|split|0')).
fof(f119,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1)))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : (~p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.inv_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.symmetric_1'(s(fun(X6,fun(X6,bool)),X7)))) => s(fun(X6,fun(X6,bool)),X7) = s(fun(X6,fun(X6,bool)),'const.relation.inv_1'(s(fun(X6,fun(X6,bool)),X7))))),
  file('thm.relation.equivalence_inv_identity','thm.relation.symmetric_inv_identity')).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 37
% Passive clauses: 148
% Generated clauses: 742
% Final active clauses: 36
% Final passive clauses: 84
% Input formulas: 22
% Initial clauses: 55
% 
% Duplicate literals: 65
% Trivial inequalities: 26
% Fw subsumption resolutions: 95
% Fw demodulations: 8
% Bw demodulations: 1
% 
% Simple tautologies: 67
% Equational tautologies: 16
% Forward subsumptions: 272
% 
% Binary resolution: 70
% Forward superposition: 201
% Backward superposition: 149
% Self superposition: 22
% Equality factoring: 26
% Equality resolution: 9
% 
% Split clauses: 37
% Split components: 76
% SAT solver clauses: 73
% SAT solver unit clauses: 11
% SAT solver binary clauses: 30
% SAT solver learnt clauses: 1
% 
% Sat splits: 37
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 72
% TWLsolver calls for satisfiability: 29
% 
% Memory used [KB]: 895
% Time elapsed: 0.025 s
% ------------------------------
% ------------------------------
% Success in time 0.027 s
