Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.rat
% (91060)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 171
% Passive clauses: 3142
% Generated clauses: 23323
% Final active clauses: 147
% Final passive clauses: 2637
% Input formulas: 18
% Initial clauses: 33
% 
% Duplicate literals: 843
% Trivial inequalities: 259
% Fw subsumption resolutions: 3467
% Fw demodulations: 3547
% Bw demodulations: 22
% 
% Simple tautologies: 392
% Equational tautologies: 44
% Forward subsumptions: 10344
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 1464
% Factoring: 4
% Forward superposition: 6972
% Backward superposition: 6271
% Self superposition: 165
% Equality factoring: 97
% Equality resolution: 22
% 
% Split clauses: 519
% Split components: 1051
% SAT solver clauses: 1481
% SAT solver unit clauses: 56
% SAT solver binary clauses: 307
% 
% Sat splits: 519
% 
% TWLsolver clauses: 1476
% TWLsolver calls for satisfiability: 187
% 
% Memory used [KB]: 9466
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (91079)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 103
% Passive clauses: 565
% Generated clauses: 15672
% Final active clauses: 98
% Final passive clauses: 450
% Input formulas: 18
% Initial clauses: 33
% 
% Duplicate literals: 745
% Trivial inequalities: 313
% Fw subsumption resolutions: 1933
% Fw demodulations: 2294
% Bw demodulations: 12
% Condensations: 1496
% 
% Simple tautologies: 250
% Equational tautologies: 46
% Forward subsumptions: 6810
% 
% Binary resolution: 216
% Factoring: 6
% Forward superposition: 3812
% Backward superposition: 3892
% Self superposition: 191
% Equality factoring: 122
% Equality resolution: 21
% 
% Split clauses: 101
% Split components: 202
% SAT solver clauses: 1145
% SAT solver unit clauses: 27
% SAT solver binary clauses: 34
% 
% Sat splits: 101
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1145
% TWLsolver calls for satisfiability: 198
% 
% Memory used [KB]: 9594
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f2655,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f62,f97,f95,f61,f96,f94,f318,f321,f706,f709,f707,f711,f715,f717,f721,f723,f726,f783,f785,f787,f788,f789,f790,f791,f792,f793,f918,f920,f919,f930,f932,f931,f968,f970,f969,f972,f973,f974,f975,f982,f976,f984,f981,f1220,f1215,f1216,f1227,f1225,f1217,f1218,f1228,f1219,f1224,f1222,f1266,f1263,f1264,f1272,f1265,f1270,f1268,f1396,f1532,f1529,f1530,f1538,f1531,f1536,f1534,f1610,f1617,f1618,f1619,f1620,f1621,f1622,f1623,f1624,f1625,f1626,f1627,f1628,f1629,f1630,f1631,f1632,f1633,f1634,f1635,f1636,f1637,f1638,f1639,f1640,f1641,f1642,f1643,f1644,f1645,f1646,f1647,f1648,f1689,f1693,f1695,f1696,f1698,f1700,f1701,f1702,f1750,f1752,f1755,f1756,f1757,f1759,f1761,f1656,f1655,f979,f977,f2159,f2161,f708,f2160,f2318,f2394,f2396,f2399,f2401,f2403,f2649])).
fof(f2649,plain,(
  $false | ($spl0 | $spl34 | $spl36)),
  inference(subsumption_resolution,[],[f2648,f2415])).
fof(f2415,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f373,f389,f453,f643,f644,f687,f689,f1188,f1190,f1194,f1196,f1201,f1203,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f1199,f1186,f79,f692,f694,f1205,f1198,f1192,f696,f691,f279,f192,f191,f185,f73,f72,f65,f2159,f2161,f2162,f2345,f2336])).
fof(f2336,plain,(
  ( ! [X17,X18] : (~p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X18),s('type.frac.frac',X17))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X17),s('type.frac.frac',X18)))) )),
  inference(superposition,[],[f191,f67])).
fof(f2345,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X6),s('type.frac.frac',X5))) | s(bool,X7) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X5),s('type.frac.frac',X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f2332,f63])).
fof(f2332,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X6),s('type.frac.frac',X5))) | s(bool,X7) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X5),s('type.frac.frac',X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f191,f132])).
fof(f2162,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X18),s('type.frac.frac',X17))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X17),s('type.frac.frac',X18)))) )),
  inference(subsumption_resolution,[],[f2152,f65])).
fof(f2152,plain,(
  ( ! [X17,X18] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X18),s('type.frac.frac',X17))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X17),s('type.frac.frac',X18)))) )),
  inference(superposition,[],[f189,f67])).
fof(f65,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f21])).
fof(f21,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.rat.RAT_EQUIV',aHL_FALSITY)).
fof(f72,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X16] : (s(bool,'F_0') = s(bool,X16) <=> ~p(s(bool,X16)))),
  file('thm.rat.RAT_EQUIV','thm.bool.EQ_CLAUSES|split|3')).
fof(f73,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f185,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X6),s('type.frac.frac',X5))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X5),s('type.frac.frac',X6)))) )),
  inference(superposition,[],[f78,f67])).
fof(f191,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X8),s('type.frac.frac',X7)))) | s(bool,'F_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X7),s('type.frac.frac',X8)))) )),
  inference(superposition,[],[f72,f78])).
fof(f192,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X10),s('type.frac.frac',X9)))) | s(bool,'F_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X9),s('type.frac.frac',X10)))) )),
  inference(superposition,[],[f73,f78])).
fof(f279,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1))) | s(bool,'T_0') = s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0))),s('type.frac.frac',X1)))) )),
  inference(superposition,[],[f67,f79])).
fof(f691,plain,(
  ( ! [X17,X18,X16] : (s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0'))) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(bool,'T_0') = s(bool,sK2(X16,X17,X18,bool))) )),
  inference(superposition,[],[f84,f67])).
fof(f696,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(X16,sK2(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'T_0') = s(bool,g(s(X16,sK2(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f84,f67])).
fof(f1192,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,happ(s(fun(X21,bool),X23),s(X21,sK3(X23,X22,X21,bool)))) | s(fun(X21,bool),X22) = s(fun(X21,bool),X23) | s(bool,'T_0') = s(bool,happ(s(fun(X21,bool),X22),s(X21,sK3(X23,X22,X21,bool))))) )),
  inference(superposition,[],[f86,f67])).
fof(f1198,plain,(
  ( ! [X21,X19,X20] : (s(X21,happ(s(fun(bool,X21),X19),s(bool,'F_0'))) != s(X21,happ(s(fun(bool,X21),X20),s(bool,'F_0'))) | s(fun(bool,X21),X19) = s(fun(bool,X21),X20) | s(bool,'T_0') = s(bool,sK3(X19,X20,bool,X21))) )),
  inference(superposition,[],[f86,f67])).
fof(f1205,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,happ(s(fun(X21,bool),X23),s(X21,sK3(X22,X23,X21,bool)))) | s(fun(X21,bool),X22) = s(fun(X21,bool),X23) | s(bool,'T_0') = s(bool,happ(s(fun(X21,bool),X22),s(X21,sK3(X22,X23,X21,bool))))) )),
  inference(superposition,[],[f86,f67])).
fof(f694,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK2(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK2(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f84,f132])).
fof(f692,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK2(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK2(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f84,f132])).
fof(f79,plain,(
  ( ! [X0,X1] : (s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1))) = s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0))),s('type.frac.frac',X1)))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1))) = s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0))),s('type.frac.frac',X1)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X21,X22] : s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X21),s('type.frac.frac',X22))) = s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X21))),s('type.frac.frac',X22)))),
  file('thm.rat.RAT_EQUIV','ARITY_1_2_const.rat.rat_equiv')).
fof(f1186,plain,(
  ( ! [X0,X1] : (s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',sK3(X1,'const.rat.rat_equiv_1'(s('type.frac.frac',X0)),'type.frac.frac',bool)))) != s(bool,happ(s(fun('type.frac.frac',bool),X1),s('type.frac.frac',sK3(X1,'const.rat.rat_equiv_1'(s('type.frac.frac',X0)),'type.frac.frac',bool)))) | s(fun('type.frac.frac',bool),X1) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0)))) )),
  inference(superposition,[],[f86,f79])).
fof(f1199,plain,(
  ( ! [X0,X1] : (s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',X0)),X1,'type.frac.frac',bool)))) != s(bool,happ(s(fun('type.frac.frac',bool),X1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',X0)),X1,'type.frac.frac',bool)))) | s(fun('type.frac.frac',bool),X1) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0)))) )),
  inference(superposition,[],[f86,f79])).
fof(f78,plain,(
  ( ! [X0,X1] : (s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1))) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X1),s('type.frac.frac',X0)))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1))) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X1),s('type.frac.frac',X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X18,X19] : s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X18),s('type.frac.frac',X19))) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X19),s('type.frac.frac',X18)))),
  file('thm.rat.RAT_EQUIV','thm.rat.RAT_EQUIV_SYM')).
fof(f87,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1)))) | ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X1),s('type.frac.frac',X2)))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X2))))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1)))) | ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X1),s('type.frac.frac',X2)))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X2)))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1)))) | ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X1),s('type.frac.frac',X2))))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X2)))))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1)))) & p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X1),s('type.frac.frac',X2))))) => p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X2)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X18,X19,X20] : ((p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X18),s('type.frac.frac',X19)))) & p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X19),s('type.frac.frac',X20))))) => p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X18),s('type.frac.frac',X20)))))),
  file('thm.rat.RAT_EQUIV','thm.rat.RAT_EQUIV_TRANS')).
fof(f189,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X3),s('type.frac.frac',X4)))) | s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X4),s('type.frac.frac',X3)))) )),
  inference(superposition,[],[f68,f78])).
fof(f190,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X6),s('type.frac.frac',X5)))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X5),s('type.frac.frac',X6)))) )),
  inference(superposition,[],[f69,f78])).
fof(f193,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X12),s('type.frac.frac',X11)))) | p(s(bool,X13)) | s(bool,X13) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',X12)))) )),
  inference(superposition,[],[f80,f78])).
fof(f195,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X18),s('type.frac.frac',X17)))) | ~p(s(bool,X19)) | s(bool,X19) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X17),s('type.frac.frac',X18)))) )),
  inference(superposition,[],[f83,f78])).
fof(f212,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X4),s('type.frac.frac',X3)))) | s(bool,X5) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X3),s('type.frac.frac',X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f167,f78])).
fof(f371,plain,(
  ( ! [X21,X22,X20] : (s(bool,X22) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X20),s('type.frac.frac',X21))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X20))),s('type.frac.frac',X21)))) )),
  inference(superposition,[],[f132,f79])).
fof(f399,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X14),s('type.frac.frac',X15))) | s(bool,X16) = s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X14))),s('type.frac.frac',X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f132,f79])).
fof(f650,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X3),s('type.frac.frac',X2)))) | ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X3),s('type.frac.frac',X4)))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X2),s('type.frac.frac',X4))))) )),
  inference(superposition,[],[f87,f78])).
fof(f659,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X10),s('type.frac.frac',X11)))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X9),s('type.frac.frac',X11)))) | s(bool,X12) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X9),s('type.frac.frac',X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f653,f63])).
fof(f653,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X10),s('type.frac.frac',X11)))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X9),s('type.frac.frac',X11)))) | s(bool,X12) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X9),s('type.frac.frac',X10))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f87,f132])).
fof(f1203,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,X16) != s(bool,happ(s(fun(X13,bool),X15),s(X13,sK3(X14,X15,X13,bool)))) | s(fun(X13,bool),X14) = s(fun(X13,bool),X15) | s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),X14),s(X13,sK3(X14,X15,X13,bool)))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f86,f132])).
fof(f1201,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') != s(bool,happ(s(fun(X5,bool),X7),s(X5,sK3(X6,X7,X5,bool)))) | s(fun(X5,bool),X6) = s(fun(X5,bool),X7) | s(bool,X8) = s(bool,happ(s(fun(X5,bool),X6),s(X5,sK3(X6,X7,X5,bool)))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f86,f132])).
fof(f1196,plain,(
  ( ! [X14,X12,X13,X11] : (s(X13,happ(s(fun(bool,X13),X11),s(bool,X14))) != s(X13,happ(s(fun(bool,X13),X12),s(bool,X14))) | s(fun(bool,X13),X11) = s(fun(bool,X13),X12) | s(bool,'T_0') = s(bool,sK3(X11,X12,bool,X13)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f86,f132])).
fof(f1194,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,happ(s(fun(bool,X5),X3),s(bool,'T_0'))) != s(X5,happ(s(fun(bool,X5),X4),s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK3(X3,X4,bool,X5)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f86,f132])).
fof(f1190,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,X16) != s(bool,happ(s(fun(X13,bool),X15),s(X13,sK3(X15,X14,X13,bool)))) | s(fun(X13,bool),X14) = s(fun(X13,bool),X15) | s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),X14),s(X13,sK3(X15,X14,X13,bool)))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f86,f132])).
fof(f1188,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') != s(bool,happ(s(fun(X5,bool),X7),s(X5,sK3(X7,X6,X5,bool)))) | s(fun(X5,bool),X6) = s(fun(X5,bool),X7) | s(bool,X8) = s(bool,happ(s(fun(X5,bool),X6),s(X5,sK3(X7,X6,X5,bool)))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f86,f132])).
fof(f689,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK2(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f84,f132])).
fof(f687,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK2(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f84,f132])).
fof(f644,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f470,f63])).
fof(f470,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f83,f132])).
fof(f643,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f466,f63])).
fof(f466,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f72,f132])).
fof(f453,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f68,f132])).
fof(f389,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f132,f132])).
fof(f373,plain,(
  ( ! [X26,X27,X25] : (s(bool,X27) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X26),s('type.frac.frac',X25))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X25),s('type.frac.frac',X26)))) )),
  inference(superposition,[],[f132,f78])).
fof(f167,plain,(
  ( ! [X6,X5] : (p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f160,f65])).
fof(f160,plain,(
  ( ! [X6,X5] : (p(s(bool,'F_0')) | p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f80,f67])).
fof(f132,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f67,f67])).
fof(f103,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X0)))) )),
  inference(unit_resulting_resolution,[],[f66,f69])).
fof(f99,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f65,f68])).
fof(f85,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK3(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK3(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X13,X10] : ! [X14,X15] : (s(fun(X10,X13),X14) = s(fun(X10,X13),X15) <=> ! [X11] : s(X13,happ(s(fun(X10,X13),X14),s(X10,X11))) = s(X13,happ(s(fun(X10,X13),X15),s(X10,X11))))),
  file('thm.rat.RAT_EQUIV','thm.bool.FUN_EQ_THM')).
fof(f86,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK3(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK3(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f58])).
fof(f84,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.rat.RAT_EQUIV',aHL_EXT)).
fof(f80,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.rat.RAT_EQUIV','thm.bool.IMP_ANTISYM_AX')).
fof(f83,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f42])).
fof(f68,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X16] : (s(bool,'T_0') = s(bool,X16) <=> p(s(bool,X16)))),
  file('thm.rat.RAT_EQUIV','thm.bool.EQ_CLAUSES|split|0')).
fof(f69,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f50])).
fof(f67,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.rat.RAT_EQUIV',aHL_BOOL_CASES)).
fof(f66,plain,(
  ( ! [X0] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X0))))) )),
  inference(cnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0] : p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X0))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X18] : p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X18),s('type.frac.frac',X18))))),
  file('thm.rat.RAT_EQUIV','thm.rat.RAT_EQUIV_REF')).
fof(f63,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rat.RAT_EQUIV',aHL_TRUTH)).
fof(f2648,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl34 | $spl36)),
  inference(forward_demodulation,[],[f2647,f1269])).
fof(f1269,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | $spl36),
  inference(cnf_transformation,[],[f1269_D])).
fof(f1269_D,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2647,plain,(
  p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool))))) | ($spl0 | $spl34)),
  inference(subsumption_resolution,[],[f2625,f63])).
fof(f2625,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool))))) | ($spl0 | $spl34)),
  inference(superposition,[],[f658,f1268])).
fof(f658,plain,(
  ( ! [X8] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X8)))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X8))))) ) | $spl0),
  inference(subsumption_resolution,[],[f652,f63])).
fof(f652,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X8)))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X8))))) ) | $spl0),
  inference(superposition,[],[f87,f338])).
fof(f338,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK1))) | $spl0),
  inference(unit_resulting_resolution,[],[f332,f69])).
fof(f332,plain,(
  p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK1)))) | $spl0),
  inference(forward_demodulation,[],[f94,f79])).
fof(f2403,plain,(
  $false | ($spl3 | $spl34 | $spl38)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f373,f389,f453,f643,f644,f687,f689,f686,f972,f973,f974,f975,f976,f1188,f1190,f1194,f1196,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1263,f1264,f1265,f1273,f1349,f1357,f1397,f1272,f1529,f1530,f1531,f1210,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f1199,f1186,f79,f692,f694,f1205,f1198,f1192,f696,f691,f279,f192,f191,f185,f73,f72,f65,f685,f1656,f1655,f2159,f2161,f2162,f1208,f1268,f2345,f2346,f2395,f2400,f2402])).
fof(f2402,plain,(
  ( ! [X0,X1] : (s(fun('type.frac.frac',bool),X1) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0)))) ) | ($spl3 | $spl34 | $spl38)),
  inference(subsumption_resolution,[],[f2366,f2346])).
fof(f2366,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',sK3(X1,'const.rat.rat_equiv_1'(s('type.frac.frac',X0)),'type.frac.frac',bool)))) | s(fun('type.frac.frac',bool),X1) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0)))) ) | ($spl3 | $spl34 | $spl38)),
  inference(backward_demodulation,[],[f2346,f1186])).
fof(f2400,plain,(
  ( ! [X0,X1] : (s(fun('type.frac.frac',bool),X1) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0)))) ) | ($spl3 | $spl34 | $spl38)),
  inference(subsumption_resolution,[],[f2365,f2346])).
fof(f2365,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',X0)),X1,'type.frac.frac',bool)))) | s(fun('type.frac.frac',bool),X1) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0)))) ) | ($spl3 | $spl34 | $spl38)),
  inference(backward_demodulation,[],[f2346,f1199])).
fof(f2395,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl3 | $spl34 | $spl38)),
  inference(subsumption_resolution,[],[f2348,f2346])).
fof(f2348,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl3 | $spl34 | $spl38)),
  inference(backward_demodulation,[],[f2346,f72])).
fof(f2346,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,X7)) ) | ($spl3 | $spl34 | $spl38)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f373,f389,f453,f643,f644,f687,f689,f686,f972,f973,f974,f975,f976,f1188,f1190,f1194,f1196,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1263,f1264,f1265,f1273,f1349,f1357,f1397,f1272,f1529,f1530,f1531,f1210,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f1199,f1186,f79,f692,f694,f1205,f1198,f1192,f696,f691,f279,f192,f191,f185,f73,f72,f65,f685,f1656,f1655,f2159,f2161,f2162,f1208,f1268,f2345])).
fof(f1208,plain,(
  s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | $spl3),
  inference(forward_demodulation,[],[f1207,f79])).
fof(f1207,plain,(
  s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) != s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | $spl3),
  inference(forward_demodulation,[],[f1185,f79])).
fof(f1185,plain,(
  s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) != s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1))),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | $spl3),
  inference(unit_resulting_resolution,[],[f97,f86])).
fof(f685,plain,(
  s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) != s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | $spl3),
  inference(unit_resulting_resolution,[],[f97,f84])).
fof(f1210,plain,(
  s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | $spl3),
  inference(forward_demodulation,[],[f1209,f79])).
fof(f1209,plain,(
  s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) != s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | $spl3),
  inference(forward_demodulation,[],[f1184,f79])).
fof(f1184,plain,(
  s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) != s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1))),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | $spl3),
  inference(unit_resulting_resolution,[],[f97,f86])).
fof(f1397,plain,(
  ( ! [X0] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)),s('type.frac.frac',X0)))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X0))))) ) | $spl34),
  inference(subsumption_resolution,[],[f1371,f63])).
fof(f1371,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)),s('type.frac.frac',X0)))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X0))))) ) | $spl34),
  inference(superposition,[],[f87,f1268])).
fof(f1357,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl34)),
  inference(trivial_inequality_removal,[],[f1351])).
fof(f1351,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl34)),
  inference(superposition,[],[f1273,f132])).
fof(f1349,plain,(
  ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool))))) | ($spl3 | $spl34)),
  inference(unit_resulting_resolution,[],[f63,f1273,f83])).
fof(f1273,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl34)),
  inference(backward_demodulation,[],[f1268,f1208])).
fof(f686,plain,(
  s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) != s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | $spl3),
  inference(unit_resulting_resolution,[],[f97,f84])).
fof(f2401,plain,(
  $false | ($spl3 | $spl34 | $spl38)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f373,f389,f453,f643,f644,f687,f689,f686,f972,f973,f974,f975,f976,f1188,f1190,f1194,f1196,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1263,f1264,f1265,f1273,f1349,f1357,f1397,f1272,f1529,f1530,f1531,f1210,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f1199,f1186,f79,f692,f694,f1205,f1198,f1192,f696,f691,f279,f192,f191,f185,f73,f72,f65,f685,f1656,f1655,f2159,f2161,f2162,f1208,f1268,f2345,f2346,f2395,f2400])).
fof(f2399,plain,(
  $false | ($spl3 | $spl34 | $spl38)),
  inference(subsumption_resolution,[],[f2398,f63])).
fof(f2398,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl34 | $spl38)),
  inference(forward_demodulation,[],[f2397,f2346])).
fof(f2397,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X8),s('type.frac.frac',X7))))) ) | ($spl3 | $spl34 | $spl38)),
  inference(subsumption_resolution,[],[f2351,f2346])).
fof(f2351,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X7),s('type.frac.frac',X8))) | ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X8),s('type.frac.frac',X7))))) ) | ($spl3 | $spl34 | $spl38)),
  inference(backward_demodulation,[],[f2346,f191])).
fof(f2396,plain,(
  $false | ($spl3 | $spl34 | $spl38)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f373,f389,f453,f643,f644,f687,f689,f686,f972,f973,f974,f975,f976,f1188,f1190,f1194,f1196,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1263,f1264,f1265,f1273,f1349,f1357,f1397,f1272,f1529,f1530,f1531,f1210,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f1199,f1186,f79,f692,f694,f1205,f1198,f1192,f696,f691,f279,f192,f191,f185,f73,f72,f65,f685,f1656,f1655,f2159,f2161,f2162,f1208,f1268,f2345,f2346,f2395])).
fof(f2394,plain,(
  $false | ($spl3 | $spl34 | $spl38)),
  inference(subsumption_resolution,[],[f2347,f63])).
fof(f2347,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl34 | $spl38)),
  inference(backward_demodulation,[],[f2346,f65])).
fof(f2318,plain,(
  $false | ($spl0 | $spl3 | $spl37)),
  inference(subsumption_resolution,[],[f2317,f65])).
fof(f2317,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3 | $spl37)),
  inference(forward_demodulation,[],[f2316,f1440])).
fof(f1440,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl37)),
  inference(unit_resulting_resolution,[],[f65,f1420,f167])).
fof(f1420,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl37)),
  inference(backward_demodulation,[],[f1409,f1208])).
fof(f1409,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | $spl37),
  inference(unit_resulting_resolution,[],[f65,f1270,f167])).
fof(f2316,plain,(
  p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool))))) | ($spl0 | $spl37)),
  inference(subsumption_resolution,[],[f2305,f63])).
fof(f2305,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool))))) | ($spl0 | $spl37)),
  inference(superposition,[],[f2296,f1409])).
fof(f2296,plain,(
  ( ! [X0] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X0)))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X0))))) ) | $spl0),
  inference(subsumption_resolution,[],[f2268,f63])).
fof(f2268,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X0)))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X0))))) ) | $spl0),
  inference(superposition,[],[f87,f2256])).
fof(f2256,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK0))) | $spl0),
  inference(subsumption_resolution,[],[f2242,f63])).
fof(f2242,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK0))) | $spl0),
  inference(superposition,[],[f190,f338])).
fof(f2160,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X12),s('type.frac.frac',X11))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',X12)))) ) | $spl46),
  inference(cnf_transformation,[],[f2160_D])).
fof(f2160_D,plain,(
  ( ! [X12,X11] : (s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X12),s('type.frac.frac',X11))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',X12)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f708,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) | $spl6),
  inference(cnf_transformation,[],[f708_D])).
fof(f708_D,plain,(
  ( ! [X13] : (p(s(bool,X13))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2161,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X16)) | s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X15),s('type.frac.frac',X14))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X14),s('type.frac.frac',X15)))) )),
  inference(subsumption_resolution,[],[f2151,f68])).
fof(f2151,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X16)) | s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X15),s('type.frac.frac',X14))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X14),s('type.frac.frac',X15)))) )),
  inference(superposition,[],[f189,f132])).
fof(f2159,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X13)) | s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X12),s('type.frac.frac',X11))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',X12)))) )),
  inference(subsumption_resolution,[],[f2150,f68])).
fof(f2150,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X13)) | s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X12),s('type.frac.frac',X11))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',X12))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f189,f132])).
fof(f977,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f977_D])).
fof(f977_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f979,plain,(
  s(bool,'T_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | $spl17),
  inference(cnf_transformation,[],[f979_D])).
fof(f979_D,plain,(
  s(bool,'T_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f1655,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,X0) = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f685,f132])).
fof(f1656,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,X1) = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(superposition,[],[f685,f132])).
fof(f1761,plain,(
  $false | ($spl17 | $spl20)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f373,f389,f453,f643,f644,f687,f689,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f696,f691,f279,f192,f191,f185,f73,f72,f65,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f982,f979,f1192,f1198,f1205,f1681,f1703,f1751,f1758,f1760])).
fof(f1760,plain,(
  ( ! [X0,X1] : (s(fun('type.frac.frac',bool),X1) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0)))) ) | ($spl17 | $spl20)),
  inference(subsumption_resolution,[],[f1725,f1703])).
fof(f1725,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',X0)),X1,'type.frac.frac',bool)))) | s(fun('type.frac.frac',bool),X1) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0)))) ) | ($spl17 | $spl20)),
  inference(backward_demodulation,[],[f1703,f1199])).
fof(f1758,plain,(
  ( ! [X0,X1] : (s(fun('type.frac.frac',bool),X1) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0)))) ) | ($spl17 | $spl20)),
  inference(subsumption_resolution,[],[f1724,f1703])).
fof(f1724,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',sK3(X1,'const.rat.rat_equiv_1'(s('type.frac.frac',X0)),'type.frac.frac',bool)))) | s(fun('type.frac.frac',bool),X1) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0)))) ) | ($spl17 | $spl20)),
  inference(backward_demodulation,[],[f1703,f1186])).
fof(f1751,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl17 | $spl20)),
  inference(subsumption_resolution,[],[f1705,f1703])).
fof(f1705,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl17 | $spl20)),
  inference(backward_demodulation,[],[f1703,f72])).
fof(f1703,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl17 | $spl20)),
  inference(subsumption_resolution,[],[f1688,f982])).
fof(f1688,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X0)) ) | $spl17),
  inference(trivial_inequality_removal,[],[f1682])).
fof(f1682,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X0)) ) | $spl17),
  inference(superposition,[],[f979,f132])).
fof(f1681,plain,(
  ~p(s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac'))))) | $spl17),
  inference(unit_resulting_resolution,[],[f63,f979,f83])).
fof(f1759,plain,(
  $false | ($spl17 | $spl20)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f373,f389,f453,f643,f644,f687,f689,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f696,f691,f279,f192,f191,f185,f73,f72,f65,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f982,f979,f1192,f1198,f1205,f1681,f1703,f1751,f1758])).
fof(f1757,plain,(
  $false | ($spl17 | $spl20)),
  inference(subsumption_resolution,[],[f1722,f63])).
fof(f1722,plain,(
  ~p(s(bool,'T_0')) | ($spl17 | $spl20)),
  inference(backward_demodulation,[],[f1703,f1681])).
fof(f1756,plain,(
  $false | ($spl17 | $spl20 | $spl43)),
  inference(subsumption_resolution,[],[f1719,f1703])).
fof(f1719,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | ($spl17 | $spl20 | $spl43)),
  inference(backward_demodulation,[],[f1703,f1536])).
fof(f1755,plain,(
  $false | ($spl17 | $spl20)),
  inference(subsumption_resolution,[],[f1754,f63])).
fof(f1754,plain,(
  ~p(s(bool,'T_0')) | ($spl17 | $spl20)),
  inference(forward_demodulation,[],[f1753,f1703])).
fof(f1753,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X8),s('type.frac.frac',X7))))) ) | ($spl17 | $spl20)),
  inference(subsumption_resolution,[],[f1708,f1703])).
fof(f1708,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X7),s('type.frac.frac',X8))) | ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X8),s('type.frac.frac',X7))))) ) | ($spl17 | $spl20)),
  inference(backward_demodulation,[],[f1703,f191])).
fof(f1752,plain,(
  $false | ($spl17 | $spl20)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f373,f389,f453,f643,f644,f687,f689,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f696,f691,f279,f192,f191,f185,f73,f72,f65,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f982,f979,f1192,f1198,f1205,f1681,f1703,f1751])).
fof(f1750,plain,(
  $false | ($spl17 | $spl20)),
  inference(subsumption_resolution,[],[f1704,f63])).
fof(f1704,plain,(
  ~p(s(bool,'T_0')) | ($spl17 | $spl20)),
  inference(backward_demodulation,[],[f1703,f65])).
fof(f1702,plain,(
  $false | ($spl17 | $spl23)),
  inference(subsumption_resolution,[],[f1660,f984])).
fof(f1660,plain,(
  s(bool,'F_0') = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | $spl17),
  inference(unit_resulting_resolution,[],[f979,f67])).
fof(f1701,plain,(
  $false | ($spl17 | $spl23)),
  inference(subsumption_resolution,[],[f1662,f984])).
fof(f1662,plain,(
  s(bool,'F_0') = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | $spl17),
  inference(unit_resulting_resolution,[],[f99,f979,f132])).
fof(f1700,plain,(
  $false | ($spl17 | $spl23 | $spl31)),
  inference(subsumption_resolution,[],[f1699,f984])).
fof(f1699,plain,(
  s(bool,'F_0') = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | ($spl17 | $spl31)),
  inference(forward_demodulation,[],[f1663,f1240])).
fof(f1240,plain,(
  s(bool,'F_0') = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | $spl31),
  inference(unit_resulting_resolution,[],[f65,f1227,f167])).
fof(f1663,plain,(
  s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | ($spl17 | $spl31)),
  inference(unit_resulting_resolution,[],[f1227,f979,f132])).
fof(f1698,plain,(
  $false | ($spl3 | $spl17 | $spl23 | $spl37)),
  inference(subsumption_resolution,[],[f1697,f984])).
fof(f1697,plain,(
  s(bool,'F_0') = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | ($spl3 | $spl17 | $spl37)),
  inference(forward_demodulation,[],[f1665,f1440])).
fof(f1665,plain,(
  s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl17 | $spl37)),
  inference(unit_resulting_resolution,[],[f1420,f979,f132])).
fof(f1696,plain,(
  $false | ($spl17 | $spl23)),
  inference(subsumption_resolution,[],[f1668,f984])).
fof(f1668,plain,(
  s(bool,'F_0') = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | $spl17),
  inference(unit_resulting_resolution,[],[f99,f979,f132])).
fof(f1695,plain,(
  $false | ($spl17 | $spl23 | $spl31)),
  inference(subsumption_resolution,[],[f1694,f984])).
fof(f1694,plain,(
  s(bool,'F_0') = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | ($spl17 | $spl31)),
  inference(forward_demodulation,[],[f1669,f1240])).
fof(f1669,plain,(
  s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | ($spl17 | $spl31)),
  inference(unit_resulting_resolution,[],[f1227,f979,f132])).
fof(f1693,plain,(
  $false | ($spl3 | $spl17 | $spl23 | $spl37)),
  inference(subsumption_resolution,[],[f1692,f984])).
fof(f1692,plain,(
  s(bool,'F_0') = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | ($spl3 | $spl17 | $spl37)),
  inference(forward_demodulation,[],[f1671,f1440])).
fof(f1671,plain,(
  s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl17 | $spl37)),
  inference(unit_resulting_resolution,[],[f1420,f979,f132])).
fof(f1689,plain,(
  $false | ($spl17 | $spl23)),
  inference(subsumption_resolution,[],[f1674,f984])).
fof(f1674,plain,(
  s(bool,'F_0') = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | $spl17),
  inference(unit_resulting_resolution,[],[f65,f979,f167])).
fof(f1648,plain,(
  $false | ($spl0 | $spl3 | $spl20 | $spl40)),
  inference(global_subsumption,[],[f982,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561,f1205,f1198,f1192,f685])).
fof(f1561,plain,(
  ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool))))) | ($spl3 | $spl40)),
  inference(unit_resulting_resolution,[],[f63,f1539,f83])).
fof(f1554,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | ($spl3 | $spl40)),
  inference(unit_resulting_resolution,[],[f65,f1539,f167])).
fof(f1548,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | ($spl3 | $spl40)),
  inference(unit_resulting_resolution,[],[f99,f1539,f132])).
fof(f1542,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | ($spl3 | $spl40)),
  inference(unit_resulting_resolution,[],[f99,f1539,f132])).
fof(f1541,plain,(
  ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool))))) | ($spl3 | $spl40)),
  inference(unit_resulting_resolution,[],[f1539,f69])).
fof(f1540,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | ($spl3 | $spl40)),
  inference(unit_resulting_resolution,[],[f1539,f67])).
fof(f1608,plain,(
  p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool))))) | ($spl0 | $spl40)),
  inference(subsumption_resolution,[],[f1584,f63])).
fof(f1584,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool))))) | ($spl0 | $spl40)),
  inference(superposition,[],[f658,f1534])).
fof(f1611,plain,(
  ( ! [X0] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)),s('type.frac.frac',X0)))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X0))))) ) | $spl40),
  inference(subsumption_resolution,[],[f1585,f63])).
fof(f1585,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)),s('type.frac.frac',X0)))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X0))))) ) | $spl40),
  inference(superposition,[],[f87,f1534])).
fof(f1569,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl40)),
  inference(trivial_inequality_removal,[],[f1563])).
fof(f1563,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl40)),
  inference(superposition,[],[f1539,f132])).
fof(f1539,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | ($spl3 | $spl40)),
  inference(backward_demodulation,[],[f1534,f1210])).
fof(f1537,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1527,f710])).
fof(f1527,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | s(bool,X0) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f1210,f132])).
fof(f1271,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1261,f710])).
fof(f1261,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | s(bool,X0) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f1208,f132])).
fof(f971,plain,(
  ( ! [X17,X18] : (s(bool,X18) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X17),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X17),s('type.frac.frac',sK0)))) ) | $spl0),
  inference(subsumption_resolution,[],[f967,f65])).
fof(f967,plain,(
  ( ! [X17,X18] : (p(s(bool,'F_0')) | s(bool,X18) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X17),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X17),s('type.frac.frac',sK0)))) ) | $spl0),
  inference(superposition,[],[f911,f67])).
fof(f933,plain,(
  ( ! [X17,X18] : (s(bool,X18) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X17))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X17),s('type.frac.frac',sK0)))) ) | $spl0),
  inference(subsumption_resolution,[],[f929,f65])).
fof(f929,plain,(
  ( ! [X17,X18] : (p(s(bool,'F_0')) | s(bool,X18) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X17))) | s(bool,'T_0') = s(bool,X18) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X17),s('type.frac.frac',sK0)))) ) | $spl0),
  inference(superposition,[],[f699,f67])).
fof(f921,plain,(
  ( ! [X19,X18] : (s(bool,X19) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X18),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X18)))) ) | $spl0),
  inference(subsumption_resolution,[],[f917,f65])).
fof(f917,plain,(
  ( ! [X19,X18] : (p(s(bool,'F_0')) | s(bool,X19) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X18),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X18)))) ) | $spl0),
  inference(superposition,[],[f684,f67])).
fof(f911,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X2),s('type.frac.frac',sK0)))) | s(bool,X3) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X2),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f684,f78])).
fof(f699,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X2),s('type.frac.frac',sK0)))) | s(bool,X3) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X2))) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f670,f78])).
fof(f684,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X2)))) | s(bool,X3) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X2),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(subsumption_resolution,[],[f675,f63])).
fof(f675,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X2)))) | s(bool,X3) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X2),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f661,f132])).
fof(f661,plain,(
  ( ! [X0] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',sK1)))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X0))))) ) | $spl0),
  inference(superposition,[],[f658,f78])).
fof(f710,plain,(
  ( ! [X19,X18] : (s(bool,X19) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X18))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X18)))) ) | $spl0),
  inference(subsumption_resolution,[],[f705,f65])).
fof(f705,plain,(
  ( ! [X19,X18] : (p(s(bool,'F_0')) | s(bool,X19) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X18))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X18)))) ) | $spl0),
  inference(superposition,[],[f670,f67])).
fof(f670,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X2)))) | s(bool,X3) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X2))) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(subsumption_resolution,[],[f663,f63])).
fof(f663,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X2)))) | s(bool,X3) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X2))) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f658,f132])).
fof(f1647,plain,(
  $false | ($spl0 | $spl3 | $spl23 | $spl40)),
  inference(global_subsumption,[],[f984,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561,f1205,f1198,f1192,f685])).
fof(f1646,plain,(
  $false | ($spl0 | $spl3 | $spl40)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561,f1205,f1198,f1192,f685])).
fof(f1645,plain,(
  $false | ($spl0 | $spl3 | $spl17 | $spl40)),
  inference(global_subsumption,[],[f979,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561,f1205,f1198,f1192])).
fof(f1644,plain,(
  $false | ($spl0 | $spl3 | $spl40)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561,f1205,f1198,f1192])).
fof(f1643,plain,(
  $false | ($spl0 | $spl3 | $spl40)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561,f1205,f1198])).
fof(f1642,plain,(
  $false | ($spl0 | $spl3 | $spl40)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561,f1205])).
fof(f1641,plain,(
  $false | ($spl0 | $spl3 | $spl31 | $spl40)),
  inference(global_subsumption,[],[f1247,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561])).
fof(f1247,plain,(
  ~p(s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac'))))) | $spl31),
  inference(unit_resulting_resolution,[],[f63,f1227,f83])).
fof(f1640,plain,(
  $false | ($spl0 | $spl3 | $spl31 | $spl40)),
  inference(global_subsumption,[],[f1240,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561])).
fof(f1639,plain,(
  $false | ($spl0 | $spl3 | $spl31 | $spl40)),
  inference(global_subsumption,[],[f1236,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561])).
fof(f1236,plain,(
  s(bool,'F_0') = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | $spl31),
  inference(unit_resulting_resolution,[],[f99,f1227,f132])).
fof(f1638,plain,(
  $false | ($spl0 | $spl3 | $spl31 | $spl40)),
  inference(global_subsumption,[],[f1232,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561])).
fof(f1232,plain,(
  s(bool,'F_0') = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | $spl31),
  inference(unit_resulting_resolution,[],[f99,f1227,f132])).
fof(f1637,plain,(
  $false | ($spl0 | $spl3 | $spl31 | $spl40)),
  inference(global_subsumption,[],[f1231,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561])).
fof(f1231,plain,(
  ~p(s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac'))))) | $spl31),
  inference(unit_resulting_resolution,[],[f1227,f69])).
fof(f1636,plain,(
  $false | ($spl0 | $spl3 | $spl31 | $spl40)),
  inference(global_subsumption,[],[f1230,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561])).
fof(f1230,plain,(
  s(bool,'F_0') = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | $spl31),
  inference(unit_resulting_resolution,[],[f1227,f67])).
fof(f1635,plain,(
  $false | ($spl0 | $spl3 | $spl37 | $spl40)),
  inference(global_subsumption,[],[f1447,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561])).
fof(f1447,plain,(
  ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool))))) | ($spl3 | $spl37)),
  inference(unit_resulting_resolution,[],[f63,f1420,f83])).
fof(f1634,plain,(
  $false | ($spl0 | $spl3 | $spl37 | $spl40)),
  inference(global_subsumption,[],[f1440,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561])).
fof(f1633,plain,(
  $false | ($spl0 | $spl3 | $spl31 | $spl37 | $spl40)),
  inference(global_subsumption,[],[f1437,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561])).
fof(f1437,plain,(
  s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl31 | $spl37)),
  inference(unit_resulting_resolution,[],[f1227,f1420,f132])).
fof(f1632,plain,(
  $false | ($spl0 | $spl3 | $spl37 | $spl40)),
  inference(global_subsumption,[],[f1435,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561])).
fof(f1435,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl37)),
  inference(unit_resulting_resolution,[],[f99,f1420,f132])).
fof(f1631,plain,(
  $false | ($spl0 | $spl3 | $spl31 | $spl37 | $spl40)),
  inference(global_subsumption,[],[f1432,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561])).
fof(f1432,plain,(
  s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl31 | $spl37)),
  inference(unit_resulting_resolution,[],[f1227,f1420,f132])).
fof(f1630,plain,(
  $false | ($spl0 | $spl3 | $spl37 | $spl40)),
  inference(global_subsumption,[],[f1430,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561])).
fof(f1430,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl37)),
  inference(unit_resulting_resolution,[],[f99,f1420,f132])).
fof(f1629,plain,(
  $false | ($spl0 | $spl3 | $spl37 | $spl40)),
  inference(global_subsumption,[],[f1429,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561])).
fof(f1429,plain,(
  ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool))))) | ($spl3 | $spl37)),
  inference(unit_resulting_resolution,[],[f1420,f69])).
fof(f1628,plain,(
  $false | ($spl0 | $spl3 | $spl37 | $spl40)),
  inference(global_subsumption,[],[f1428,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561])).
fof(f1428,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl37)),
  inference(unit_resulting_resolution,[],[f1420,f67])).
fof(f1627,plain,(
  $false | ($spl0 | $spl3 | $spl40)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554,f1561])).
fof(f1626,plain,(
  $false | ($spl0 | $spl3 | $spl40)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1561,f1569,f1611,f1608,f1540,f1541,f1542,f1548,f1554])).
fof(f1625,plain,(
  $false | ($spl0 | $spl3 | $spl37 | $spl40)),
  inference(global_subsumption,[],[f1551,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1561,f1569,f1611,f1608,f1540,f1541,f1542,f1548])).
fof(f1551,plain,(
  s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl37 | $spl40)),
  inference(unit_resulting_resolution,[],[f1420,f1539,f132])).
fof(f1624,plain,(
  $false | ($spl0 | $spl3 | $spl31 | $spl40)),
  inference(global_subsumption,[],[f1550,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1561,f1569,f1611,f1608,f1540,f1541,f1542,f1548])).
fof(f1550,plain,(
  s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | ($spl3 | $spl31 | $spl40)),
  inference(unit_resulting_resolution,[],[f1227,f1539,f132])).
fof(f1623,plain,(
  $false | ($spl0 | $spl3 | $spl40)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1561,f1569,f1611,f1608,f1540,f1541,f1542,f1548])).
fof(f1622,plain,(
  $false | ($spl0 | $spl3 | $spl37 | $spl40)),
  inference(global_subsumption,[],[f1545,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1561,f1569,f1611,f1608,f1540,f1541,f1542])).
fof(f1545,plain,(
  s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl37 | $spl40)),
  inference(unit_resulting_resolution,[],[f1420,f1539,f132])).
fof(f1621,plain,(
  $false | ($spl0 | $spl3 | $spl31 | $spl40)),
  inference(global_subsumption,[],[f1544,f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1561,f1569,f1611,f1608,f1540,f1541,f1542])).
fof(f1544,plain,(
  s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | ($spl3 | $spl31 | $spl40)),
  inference(unit_resulting_resolution,[],[f1227,f1539,f132])).
fof(f1620,plain,(
  $false | ($spl0 | $spl3 | $spl40)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1561,f1569,f1611,f1608,f1540,f1541,f1542])).
fof(f1619,plain,(
  $false | ($spl0 | $spl3 | $spl40)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1561,f1569,f1611,f1608,f1540,f1541])).
fof(f1618,plain,(
  $false | ($spl0 | $spl3 | $spl40)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1561,f1569,f1611,f1608,f1540])).
fof(f1617,plain,(
  $false | ($spl0 | $spl3 | $spl40)),
  inference(global_subsumption,[],[f62,f61,f63,f66,f67,f69,f68,f83,f80,f84,f86,f85,f99,f103,f132,f167,f97,f332,f338,f373,f389,f453,f643,f644,f687,f689,f706,f709,f658,f670,f710,f661,f684,f699,f659,f650,f399,f371,f212,f195,f193,f190,f189,f87,f78,f79,f692,f694,f685,f686,f696,f691,f279,f192,f191,f185,f73,f72,f65,f911,f918,f920,f921,f930,f932,f933,f968,f970,f971,f972,f973,f974,f975,f976,f1210,f1186,f1188,f1190,f1194,f1196,f1199,f1201,f1203,f1215,f1216,f1217,f1218,f1219,f1271,f1263,f1264,f1265,f1208,f1537,f1529,f1530,f1531,f1534,f1539,f1561,f1569,f1611,f1608])).
fof(f1610,plain,(
  $false | ($spl0 | $spl3 | $spl18 | $spl40)),
  inference(subsumption_resolution,[],[f1609,f65])).
fof(f1609,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3 | $spl18 | $spl40)),
  inference(forward_demodulation,[],[f1608,f1570])).
fof(f1570,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | ($spl3 | $spl18 | $spl40)),
  inference(forward_demodulation,[],[f1556,f1003])).
fof(f1003,plain,(
  s(bool,'F_0') = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | ($spl3 | $spl18)),
  inference(unit_resulting_resolution,[],[f65,f985,f167])).
fof(f985,plain,(
  s(bool,'T_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | ($spl3 | $spl18)),
  inference(backward_demodulation,[],[f981,f685])).
fof(f1556,plain,(
  s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | ($spl3 | $spl18 | $spl40)),
  inference(unit_resulting_resolution,[],[f1539,f1016])).
fof(f1016,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl18)),
  inference(trivial_inequality_removal,[],[f1010])).
fof(f1010,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl18)),
  inference(superposition,[],[f985,f132])).
fof(f1534,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | $spl40),
  inference(cnf_transformation,[],[f1534_D])).
fof(f1534_D,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1536,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | $spl43),
  inference(cnf_transformation,[],[f1536_D])).
fof(f1536_D,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f1531,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | $spl3),
  inference(superposition,[],[f1210,f67])).
fof(f1538,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl44),
  inference(cnf_transformation,[],[f1538_D])).
fof(f1538_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1530,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool))))) ) | $spl3),
  inference(superposition,[],[f1210,f132])).
fof(f1529,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f1210,f132])).
fof(f1532,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | ($spl3 | $spl18)),
  inference(forward_demodulation,[],[f1526,f1003])).
fof(f1526,plain,(
  s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'type.frac.frac',bool)))) | ($spl3 | $spl18)),
  inference(superposition,[],[f1210,f1016])).
fof(f1396,plain,(
  $false | ($spl0 | $spl3 | $spl18 | $spl34)),
  inference(subsumption_resolution,[],[f1395,f65])).
fof(f1395,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3 | $spl18 | $spl34)),
  inference(forward_demodulation,[],[f1394,f1358])).
fof(f1358,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl18 | $spl34)),
  inference(forward_demodulation,[],[f1344,f1003])).
fof(f1344,plain,(
  s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl18 | $spl34)),
  inference(unit_resulting_resolution,[],[f1273,f1016])).
fof(f1394,plain,(
  p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool))))) | ($spl0 | $spl34)),
  inference(subsumption_resolution,[],[f1370,f63])).
fof(f1370,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool))))) | ($spl0 | $spl34)),
  inference(superposition,[],[f658,f1268])).
fof(f1268,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | $spl34),
  inference(cnf_transformation,[],[f1268_D])).
fof(f1268_D,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1270,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | $spl37),
  inference(cnf_transformation,[],[f1270_D])).
fof(f1270_D,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f1265,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | $spl3),
  inference(superposition,[],[f1208,f67])).
fof(f1272,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl38),
  inference(cnf_transformation,[],[f1272_D])).
fof(f1272_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f1264,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool))))) ) | $spl3),
  inference(superposition,[],[f1208,f132])).
fof(f1263,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f1208,f132])).
fof(f1266,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl18)),
  inference(forward_demodulation,[],[f1260,f1003])).
fof(f1260,plain,(
  s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',sK3('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'type.frac.frac',bool)))) | ($spl3 | $spl18)),
  inference(superposition,[],[f1208,f1016])).
fof(f1222,plain,(
  s(bool,'T_0') = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | $spl24),
  inference(cnf_transformation,[],[f1222_D])).
fof(f1222_D,plain,(
  s(bool,'T_0') = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1224,plain,(
  s(bool,'F_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | $spl27),
  inference(cnf_transformation,[],[f1224_D])).
fof(f1224_D,plain,(
  s(bool,'F_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f1219,plain,(
  s(bool,'F_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | $spl3),
  inference(superposition,[],[f686,f67])).
fof(f1228,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X2)) ) | $spl32),
  inference(cnf_transformation,[],[f1228_D])).
fof(f1228_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1218,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac'))))) ) | $spl3),
  inference(superposition,[],[f686,f132])).
fof(f1217,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f686,f132])).
fof(f1225,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(cnf_transformation,[],[f1225_D])).
fof(f1225_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1227,plain,(
  s(bool,'T_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | $spl31),
  inference(cnf_transformation,[],[f1227_D])).
fof(f1227_D,plain,(
  s(bool,'T_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f1216,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | s(bool,X1) = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(superposition,[],[f686,f132])).
fof(f1215,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | s(bool,X0) = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f686,f132])).
fof(f1220,plain,(
  s(bool,'F_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | ($spl3 | $spl18)),
  inference(forward_demodulation,[],[f1214,f1003])).
fof(f1214,plain,(
  s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | ($spl3 | $spl18)),
  inference(superposition,[],[f686,f1016])).
fof(f981,plain,(
  s(bool,'T_0') = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | $spl18),
  inference(cnf_transformation,[],[f981_D])).
fof(f981_D,plain,(
  s(bool,'T_0') = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f984,plain,(
  s(bool,'F_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | $spl23),
  inference(cnf_transformation,[],[f984_D])).
fof(f984_D,plain,(
  s(bool,'F_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f976,plain,(
  s(bool,'F_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | $spl3),
  inference(superposition,[],[f685,f67])).
fof(f982,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X2)) ) | $spl20),
  inference(cnf_transformation,[],[f982_D])).
fof(f982_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f975,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac'))))) ) | $spl3),
  inference(superposition,[],[f685,f132])).
fof(f974,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f685,f132])).
fof(f973,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,X1) = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(superposition,[],[f685,f132])).
fof(f972,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,X0) = s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f685,f132])).
fof(f969,plain,(
  ( ! [X13,X11] : (s(bool,X13) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',sK0))) | s(bool,'T_0') = s(bool,X13)) ) | $spl12),
  inference(cnf_transformation,[],[f969_D])).
fof(f969_D,plain,(
  ( ! [X13,X11] : (s(bool,X13) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',sK0))) | s(bool,'T_0') = s(bool,X13)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f970,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X15)) | s(bool,X16) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X14),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X14),s('type.frac.frac',sK0)))) ) | $spl0),
  inference(subsumption_resolution,[],[f966,f68])).
fof(f966,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X15)) | s(bool,X16) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X14),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X14),s('type.frac.frac',sK0)))) ) | $spl0),
  inference(superposition,[],[f911,f132])).
fof(f968,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X12)) | s(bool,X13) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',sK0)))) ) | $spl0),
  inference(subsumption_resolution,[],[f965,f68])).
fof(f965,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X12)) | s(bool,X13) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',sK0))) | s(bool,'T_0') = s(bool,X12)) ) | $spl0),
  inference(superposition,[],[f911,f132])).
fof(f931,plain,(
  ( ! [X13,X11] : (s(bool,X13) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X11))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',sK0))) | s(bool,'T_0') = s(bool,X13)) ) | $spl10),
  inference(cnf_transformation,[],[f931_D])).
fof(f931_D,plain,(
  ( ! [X13,X11] : (s(bool,X13) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X11))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',sK0))) | s(bool,'T_0') = s(bool,X13)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f932,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X15)) | s(bool,X16) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X14))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X14),s('type.frac.frac',sK0)))) ) | $spl0),
  inference(subsumption_resolution,[],[f928,f68])).
fof(f928,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X15)) | s(bool,X16) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X14))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X14),s('type.frac.frac',sK0)))) ) | $spl0),
  inference(superposition,[],[f699,f132])).
fof(f930,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X12)) | s(bool,X13) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X11))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',sK0)))) ) | $spl0),
  inference(subsumption_resolution,[],[f927,f68])).
fof(f927,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X12)) | s(bool,X13) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X11))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X11),s('type.frac.frac',sK0))) | s(bool,'T_0') = s(bool,X12)) ) | $spl0),
  inference(superposition,[],[f699,f132])).
fof(f919,plain,(
  ( ! [X14,X12] : (s(bool,X14) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X12),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X12))) | s(bool,'T_0') = s(bool,X14)) ) | $spl8),
  inference(cnf_transformation,[],[f919_D])).
fof(f919_D,plain,(
  ( ! [X14,X12] : (s(bool,X14) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X12),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X12))) | s(bool,'T_0') = s(bool,X14)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f920,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X16)) | s(bool,X17) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X15),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X15)))) ) | $spl0),
  inference(subsumption_resolution,[],[f916,f68])).
fof(f916,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X16)) | s(bool,X17) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X15),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X15)))) ) | $spl0),
  inference(superposition,[],[f684,f132])).
fof(f918,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X13)) | s(bool,X14) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X12),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X12)))) ) | $spl0),
  inference(subsumption_resolution,[],[f915,f68])).
fof(f915,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X13)) | s(bool,X14) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X12),s('type.frac.frac',sK1))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X12))) | s(bool,'T_0') = s(bool,X13)) ) | $spl0),
  inference(superposition,[],[f684,f132])).
fof(f793,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f714,f63])).
fof(f714,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(unit_resulting_resolution,[],[f99,f708,f83])).
fof(f792,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f716,f708])).
fof(f716,plain,(
  ~p(s(bool,'F_0')) | $spl6),
  inference(unit_resulting_resolution,[],[f99,f708,f83])).
fof(f791,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f720,f63])).
fof(f720,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(unit_resulting_resolution,[],[f99,f708,f83])).
fof(f790,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f722,f708])).
fof(f722,plain,(
  ~p(s(bool,'F_0')) | $spl6),
  inference(unit_resulting_resolution,[],[f99,f708,f83])).
fof(f789,plain,(
  $false | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f758,f725])).
fof(f725,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f708,f69])).
fof(f758,plain,(
  s(bool,'T_0') != s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),bool,'type.frac.frac')))) | ($spl3 | $spl6)),
  inference(backward_demodulation,[],[f725,f685])).
fof(f788,plain,(
  $false | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f757,f725])).
fof(f757,plain,(
  s(bool,'T_0') != s(bool,g(s('type.frac.frac',sK2('const.rat.rat_equiv_1'(s('type.frac.frac',sK0)),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)),bool,'type.frac.frac')))) | ($spl3 | $spl6)),
  inference(backward_demodulation,[],[f725,f686])).
fof(f787,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f786,f708])).
fof(f786,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X8),s('type.frac.frac',X7))))) ) | $spl6),
  inference(subsumption_resolution,[],[f752,f725])).
fof(f752,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X7),s('type.frac.frac',X8))) | ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X8),s('type.frac.frac',X7))))) ) | $spl6),
  inference(backward_demodulation,[],[f725,f191])).
fof(f785,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f784,f708])).
fof(f784,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f749,f725])).
fof(f749,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl6),
  inference(backward_demodulation,[],[f725,f72])).
fof(f783,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f748,f63])).
fof(f748,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f725,f65])).
fof(f726,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f99,f708,f69])).
fof(f723,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f708,f99,f708,f83])).
fof(f721,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f63,f99,f708,f83])).
fof(f717,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f708,f99,f708,f83])).
fof(f715,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f63,f99,f708,f83])).
fof(f711,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f65,f708])).
fof(f707,plain,(
  ( ! [X14,X12] : (s(bool,X14) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X12))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X12))) | s(bool,'T_0') = s(bool,X14)) ) | $spl4),
  inference(cnf_transformation,[],[f707_D])).
fof(f707_D,plain,(
  ( ! [X14,X12] : (s(bool,X14) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X12))) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X12))) | s(bool,'T_0') = s(bool,X14)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f709,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X16)) | s(bool,X17) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X15))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X15)))) ) | $spl0),
  inference(subsumption_resolution,[],[f704,f68])).
fof(f704,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X16)) | s(bool,X17) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X15))) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X15)))) ) | $spl0),
  inference(superposition,[],[f670,f132])).
fof(f706,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X13)) | s(bool,X14) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X12))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X12)))) ) | $spl0),
  inference(subsumption_resolution,[],[f703,f68])).
fof(f703,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X13)) | s(bool,X14) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X12))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X12))) | s(bool,'T_0') = s(bool,X13)) ) | $spl0),
  inference(superposition,[],[f670,f132])).
fof(f321,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f305,f100])).
fof(f100,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK1))) | $spl1),
  inference(unit_resulting_resolution,[],[f98,f68])).
fof(f98,plain,(
  ~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK1)))) | $spl1),
  inference(forward_demodulation,[],[f95,f79])).
fof(f305,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK1))) | $spl2),
  inference(superposition,[],[f103,f289])).
fof(f289,plain,(
  ( ! [X0] : (s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',X0))) = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X0)))) ) | $spl2),
  inference(forward_demodulation,[],[f276,f79])).
fof(f276,plain,(
  ( ! [X0] : (s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK1),s('type.frac.frac',X0))) = s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))),s('type.frac.frac',X0)))) ) | $spl2),
  inference(superposition,[],[f79,f96])).
fof(f318,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f300,f100])).
fof(f300,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',sK0),s('type.frac.frac',sK1))) | $spl2),
  inference(superposition,[],[f289,f103])).
fof(f94,plain,(
  p(s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))),s('type.frac.frac',sK1)))) | $spl0),
  inference(cnf_transformation,[],[f94_D])).
fof(f94_D,plain,(
  p(s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))),s('type.frac.frac',sK1)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f96,plain,(
  s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1))) | $spl2),
  inference(cnf_transformation,[],[f96_D])).
fof(f96_D,plain,(
  s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f61,plain,(
  s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1))) | p(s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))),s('type.frac.frac',sK1))))),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  (p(s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))),s('type.frac.frac',sK1)))) | s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1)))) & (~p(s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))),s('type.frac.frac',sK1)))) | s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))) != s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f48])).
fof(f48,plain,(
  ? [X0,X1] : ((p(s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0))),s('type.frac.frac',X1)))) | s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0))) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X1)))) & (~p(s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0))),s('type.frac.frac',X1)))) | s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0))) != s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X1)))))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ? [X0,X1] : (p(s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0))),s('type.frac.frac',X1)))) <~> s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0))) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X1))))),
  inference(ennf_transformation,[],[f20])).
fof(f20,plain,(
  ~! [X0,X1] : (p(s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0))),s('type.frac.frac',X1)))) <=> s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X0))) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X6))),s('type.frac.frac',X7)))) <=> s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X6))) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X6))),s('type.frac.frac',X7)))) <=> s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X6))) = s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',X7))))),
  file('thm.rat.RAT_EQUIV','thm.rat.RAT_EQUIV')).
fof(f95,plain,(
  ~p(s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))),s('type.frac.frac',sK1)))) | $spl1),
  inference(cnf_transformation,[],[f95_D])).
fof(f95_D,plain,(
  ~p(s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))),s('type.frac.frac',sK1)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f97,plain,(
  s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))) != s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1))) | $spl3),
  inference(cnf_transformation,[],[f97_D])).
fof(f97_D,plain,(
  s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))) != s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f62,plain,(
  s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))) != s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK1))) | ~p(s(bool,happ(s(fun('type.frac.frac',bool),'const.rat.rat_equiv_1'(s('type.frac.frac',sK0))),s('type.frac.frac',sK1))))),
  inference(cnf_transformation,[],[f49])).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 100
% Passive clauses: 282
% Generated clauses: 2746
% Final active clauses: 32
% Final passive clauses: 59
% Input formulas: 18
% Initial clauses: 33
% 
% Duplicate literals: 108
% Trivial inequalities: 47
% Fw subsumption resolutions: 203
% Fw demodulations: 287
% Bw demodulations: 135
% Condensations: 110
% Global subsumptions: 40
% 
% Simple tautologies: 59
% Equational tautologies: 104
% Forward subsumptions: 1400
% Bw demodulations to eq. taut.: 11
% 
% Unit resulting resolution: 331
% Factoring: 1
% Forward superposition: 598
% Backward superposition: 588
% Self superposition: 26
% Equality factoring: 30
% Equality resolution: 3
% 
% Split clauses: 33
% Split components: 66
% SAT solver clauses: 463
% SAT solver unit clauses: 29
% SAT solver binary clauses: 132
% 
% Sat splits: 33
% Sat splitting refutations: 73
% 
% Memory used [KB]: 1663
% Time elapsed: 0.092 s
% ------------------------------
% ------------------------------
% Success in time 1.409 s
