Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.realax
% (96106)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 156
% Passive clauses: 2057
% Generated clauses: 23538
% Final active clauses: 111
% Final passive clauses: 1482
% Input formulas: 16
% Initial clauses: 32
% 
% Duplicate literals: 1136
% Trivial inequalities: 362
% Fw subsumption resolutions: 3818
% Fw demodulations: 3566
% Bw demodulations: 20
% 
% Simple tautologies: 423
% Equational tautologies: 63
% Forward subsumptions: 10811
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 220
% Factoring: 2
% Forward superposition: 7940
% Backward superposition: 6587
% Self superposition: 288
% Equality factoring: 157
% Equality resolution: 39
% 
% Split clauses: 812
% Split components: 1663
% SAT solver clauses: 1606
% SAT solver unit clauses: 12
% SAT solver binary clauses: 89
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 812
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1591
% TWLsolver calls for satisfiability: 232
% 
% Memory used [KB]: 10746
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.realax
% (96128)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 88
% Passive clauses: 363
% Generated clauses: 8549
% Final active clauses: 76
% Final passive clauses: 244
% Input formulas: 16
% Initial clauses: 32
% 
% Duplicate literals: 895
% Trivial inequalities: 89
% Fw subsumption resolutions: 651
% Fw demodulations: 500
% Bw demodulations: 32
% Condensations: 1569
% Global subsumptions: 4
% 
% Simple tautologies: 158
% Equational tautologies: 52
% Forward subsumptions: 3151
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 71
% Factoring: 2
% Forward superposition: 1855
% Backward superposition: 2154
% Self superposition: 201
% Equality factoring: 144
% Equality resolution: 19
% 
% Split clauses: 84
% Split components: 169
% SAT solver clauses: 833
% SAT solver unit clauses: 42
% SAT solver binary clauses: 60
% SAT solver learnt clauses: 8
% 
% Sat splits: 84
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 817
% TWLsolver calls for satisfiability: 159
% 
% Memory used [KB]: 7291
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.realax
% (96134)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 482
% Passive clauses: 8486
% Generated clauses: 117713
% Final active clauses: 78
% Final passive clauses: 984
% Input formulas: 16
% Initial clauses: 32
% Discarded non-redundant clauses: 30
% 
% Duplicate literals: 8870
% Trivial inequalities: 460
% Fw subsumption resolutions: 6136
% Fw demodulations: 22514
% Bw demodulations: 4851
% Condensations: 9917
% Global subsumptions: 405
% 
% Simple tautologies: 618
% Equational tautologies: 3157
% Forward subsumptions: 55209
% Fw demodulations to eq. taut.: 452
% Bw demodulations to eq. taut.: 90
% 
% Unit resulting resolution: 11878
% Factoring: 1
% Forward superposition: 23780
% Backward superposition: 18830
% Self superposition: 828
% Equality factoring: 403
% Equality resolution: 69
% 
% Split clauses: 958
% Split components: 1942
% SAT solver clauses: 11239
% SAT solver unit clauses: 52
% SAT solver binary clauses: 886
% 
% Sat splits: 958
% Sat splitting refutations: 430
% 
% Memory used [KB]: 65755
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.realax
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.realax
% SZS output start Proof for thm.realax
fof(f21749,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f101,f102,f103,f104,f105,f106,f125,f126,f132,f127,f128,f134,f233,f234,f238,f240,f308,f309,f310,f311,f312,f313,f359,f433,f435,f430,f431,f513,f531,f723,f707,f726,f728,f1795,f1816,f1796,f1877,f1904,f2626,f2648,f2818,f2822,f2824,f2823,f2826,f2827,f2829,f2843,f2845,f2844,f2846,f2848,f2847,f3017,f3019,f3082,f3084,f3083,f3137,f3139,f3203,f3205,f3264,f3266,f3318,f3319,f3138,f3384,f3386,f3438,f3409,f3441,f3443,f3442,f3450,f3451,f3452,f3454,f3439,f3559,f3561,f3631,f3632,f3685,f3686,f3787,f3789,f3878,f3849,f3881,f3883,f3890,f3891,f3892,f3894,f3879,f4043,f4044,f4389,f4392,f4390,f4585,f4587,f4616,f4640,f4675,f70,f4677,f4763,f131,f5175,f5176,f5693,f5694,f5799,f5800,f5853,f5854,f6015,f6016,f6070,f6071,f6187,f6188,f6647,f6649,f6648,f6703,f6705,f7482,f7470,f7466,f7572,f7582,f7461,f7632,f7793,f7795,f7800,f7808,f7811,f7815,f7871,f7875,f7880,f7887,f7889,f7891,f8145,f8146,f8148,f7573,f8151,f8152,f8203,f8204,f8205,f8291,f8293,f8298,f8300,f8502,f8826,f8828,f8833,f8835,f8995,f8997,f9193,f9194,f9197,f9198,f9286,f9287,f9289,f3560,f9290,f9296,f9371,f9372,f9374,f9375,f10906,f10905,f10904,f10900,f10899,f10898,f10891,f10888,f10887,f10886,f10882,f10881,f10880,f10876,f10870,f10867,f10864,f11034,f11041,f11053,f11055,f10856,f10854,f11068,f11081,f11101,f11103,f11104,f11105,f11110,f11112,f11118,f11119,f11120,f11122,f11124,f11128,f11129,f10850,f11133,f11134,f10837,f10834,f10833,f11137,f10832,f11146,f11147,f11149,f11160,f11165,f11176,f11177,f10827,f11184,f11185,f11195,f11196,f11200,f11206,f11212,f10809,f10806,f11220,f11241,f11242,f11243,f11248,f11250,f11251,f11252,f11253,f11255,f11257,f11258,f11259,f11260,f11261,f7583,f3385,f11262,f3788,f11263,f11264,f11266,f11267,f11270,f11272,f10791,f10790,f11288,f10789,f10788,f9291,f10787,f9297,f11290,f10785,f11295,f11312,f10782,f10781,f10777,f11331,f3882,f10776,f10775,f11340,f11363,f11372,f10770,f11381,f11769,f11807,f11949,f11291,f12133,f12134,f12138,f12139,f12142,f12144,f2627,f12417,f12418,f12424,f12425,f8996,f12501,f12579,f12592,f12593,f12595,f12596,f12597,f3204,f12598,f12599,f12600,f12602,f12693,f11273,f12692,f12691,f11271,f12724,f12690,f12689,f12729,f12795,f12796,f12802,f11268,f12803,f12898,f11113,f12914,f11123,f13017,f13032,f10941,f13115,f13125,f13132,f11244,f13211,f13220,f13228,f13230,f9376,f13231,f11161,f13232,f13234,f13235,f13310,f11102,f13326,f10917,f13328,f9199,f13331,f10916,f13406,f13415,f13423,f11111,f13425,f13426,f13429,f13430,f13603,f13604,f13608,f13609,f11148,f13612,f13614,f14010,f14019,f14027,f11125,f14029,f14032,f14033,f14163,f14171,f14174,f14175,f14178,f14179,f11135,f14180,f8292,f14181,f14182,f14184,f14185,f14263,f11136,f14280,f11256,f14561,f14562,f14566,f14567,f14571,f14573,f14679,f14680,f14683,f14684,f14687,f14568,f14689,f14691,f11313,f14694,f11373,f14706,f14710,f14712,f14714,f14715,f14716,f14718,f14711,f14719,f14720,f14721,f14722,f14724,f14726,f14707,f14816,f14817,f15257,f15259,f15263,f15264,f15265,f15268,f15270,f15645,f15627,f15706,f11296,f17350,f17352,f17356,f17357,f17358,f17361,f17363,f17617,f17647,f17648,f17649,f17652,f17654,f17656,f10981,f17994,f17989,f17979,f17978,f18106,f18110,f18111,f18112,f18149,f18150,f18168,f18169,f18170,f18171,f18172,f18174,f18175,f18176,f17972,f17653,f18180,f17632,f18188,f18189,f18216,f18215,f18266,f18267,f18268,f18288,f14799,f18292,f18289,f8478,f18295,f18296,f7504,f18298,f18300,f133,f18303,f18309,f11332,f7465,f7471,f7477,f7478,f7481,f5980,f5979,f5978,f5977,f5976,f5685,f18290,f18304,f3777,f3752,f3751,f3750,f3748,f3747,f18343,f18177,f18345,f18347,f18348,f18114,f18350,f18351,f18352,f18107,f18357,f18359,f18002,f18000,f18361,f18368,f18371,f18377,f8479,f18383,f8365,f18385,f18384,f8336,f8335,f18387,f8386,f8384,f18388,f7410,f18389,f11186,f18393,f14030,f7432,f7440,f7469,f7476,f18378,f6004,f18315,f18411,f18413,f18415,f5042,f18316,f18427,f18430,f18432,f18438,f18443,f18190,f18447,f18092,f18452,f18450,f18353,f18454,f18486,f18485,f18503,f18504,f18505,f278,f18449,f18529,f18530,f117,f116,f18531,f18532,f118,f18533,f110,f108,f18534,f114,f112,f18677,f18678,f18704,f18705,f115,f18783,f18784,f19015,f19016,f19461,f19537,f19535,f19666,f19682,f19684,f19686,f19688,f19690,f19815,f19667,f20277,f20278,f20279,f20286,f15258,f20287,f8827,f20288,f17351,f20963,f20968,f20969,f20981,f3265,f20986,f6704,f20987,f20970,f21092,f21094,f21093,f21095,f21096,f8147,f239,f21099,f21101,f21100,f21082,f21081,f18386,f21080,f724,f21183,f21184,f2819,f434,f21185,f21186,f436,f11035,f21189,f21190,f18313,f4391,f21297,f21298,f3018,f2828,f21627,f21628,f21629,f21748])).
fof(f21748,plain,(
  $false | ($spl2 | $spl9)),
  inference(trivial_inequality_removal,[],[f21737])).
fof(f21737,plain,(
  s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))) | ($spl2 | $spl9)),
  inference(superposition,[],[f21706,f18651])).
fof(f18651,plain,(
  s('type.hreal.hreal',sK2) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK3(sK2,sK1)))) | $spl2),
  inference(subsumption_resolution,[],[f18619,f85])).
fof(f85,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.realax.HREAL_LT_LADD','thm.bool.TRUTH')).
fof(f18619,plain,(
  ~p(s(bool,'T_0')) | s('type.hreal.hreal',sK2) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK3(sK2,sK1)))) | $spl2),
  inference(superposition,[],[f60,f109])).
fof(f109,plain,(
  s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | $spl2),
  inference(cnf_transformation,[],[f109_D])).
fof(f109_D,plain,(
  s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f60,plain,(
  ( ! [X0,X1] : (s('type.hreal.hreal',X1) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',sK3(X1,X0)))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | s('type.hreal.hreal',X1) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',sK3(X1,X0))))) & (! [X2] : s('type.hreal.hreal',X1) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2))) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f43])).
fof(f43,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | ? [X3] : s('type.hreal.hreal',X1) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X3)))) & (! [X2] : s('type.hreal.hreal',X1) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2))) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))))),
  inference(rectify,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | ? [X2] : s('type.hreal.hreal',X1) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2)))) & (! [X2] : s('type.hreal.hreal',X1) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2))) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))))),
  inference(nnf_transformation,[],[f21])).
fof(f21,plain,(
  ! [X0,X1] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) <=> ? [X2] : s('type.hreal.hreal',X1) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X13,X14] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X13),s('type.hreal.hreal',X14)))) <=> ? [X16] : s('type.hreal.hreal',X14) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X13),s('type.hreal.hreal',X16))))),
  file('thm.realax.HREAL_LT_LADD','thm.hreal.HREAL_LT')).
fof(f21706,plain,(
  ( ! [X0] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',X0)))))) ) | $spl9),
  inference(superposition,[],[f19022,f59])).
fof(f59,plain,(
  ( ! [X2,X0,X1] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X2))))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.hreal.hreal',X2)))) )),
  inference(cnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0,X1,X2] : s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X2))))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.hreal.hreal',X2)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X13,X14,X15] : s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X13),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X14),s('type.hreal.hreal',X15))))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X13),s('type.hreal.hreal',X14))),s('type.hreal.hreal',X15)))),
  file('thm.realax.HREAL_LT_LADD','thm.hreal.HREAL_ADD_ASSOC')).
fof(f19022,plain,(
  ( ! [X0] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal',X0)))) ) | $spl9),
  inference(resolution,[],[f115,f61])).
fof(f61,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))) | s('type.hreal.hreal',X1) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2)))) )),
  inference(cnf_transformation,[],[f44])).
fof(f21629,plain,(
  $false | ($spl2 | $spl16 | $spl266)),
  inference(trivial_inequality_removal,[],[f21624])).
fof(f21624,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl2 | $spl16 | $spl266)),
  inference(superposition,[],[f21587,f109])).
fof(f21587,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,X12)) ) | ($spl16 | $spl266)),
  inference(subsumption_resolution,[],[f21586,f133])).
fof(f21586,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,X12) | ~p(s(bool,X12))) ) | $spl266),
  inference(forward_demodulation,[],[f18387,f18715])).
fof(f18715,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X17),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X17),s('type.hreal.hreal',X18)))))) )),
  inference(resolution,[],[f4869,f4691])).
fof(f4691,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f84,f68])).
fof(f68,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.realax.HREAL_LT_LADD',aHL_BOOL_CASES)).
fof(f84,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.realax.HREAL_LT_LADD',aHL_FALSITY)).
fof(f4869,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))))) )),
  inference(equality_resolution,[],[f61])).
fof(f21628,plain,(
  $false | ($spl16 | $spl266)),
  inference(trivial_inequality_removal,[],[f21626])).
fof(f21626,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl16 | $spl266)),
  inference(superposition,[],[f21587,f18715])).
fof(f21627,plain,(
  $false | ($spl16 | $spl266)),
  inference(equality_resolution,[],[f21587])).
fof(f2828,plain,(
  ( ! [X44] : (s(bool,'F_0') = s(bool,X44)) ) | $spl36),
  inference(cnf_transformation,[],[f2828_D])).
fof(f2828_D,plain,(
  ( ! [X44] : (s(bool,'F_0') = s(bool,X44)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f3018,plain,(
  ( ! [X52,X51,X49] : (s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,X49) != s(bool,X51) | ~p(s(bool,X52))) ) | $spl42),
  inference(cnf_transformation,[],[f3018_D])).
fof(f3018_D,plain,(
  ( ! [X52,X51,X49] : (s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,X49) != s(bool,X51) | ~p(s(bool,X52))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f21298,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X32) = s(bool,X33) | ~p(s(bool,X33)) | s(bool,X30) != s(bool,X32) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X30)) ) | $spl16),
  inference(subsumption_resolution,[],[f21274,f4678])).
fof(f4678,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(resolution,[],[f84,f62])).
fof(f62,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.realax.HREAL_LT_LADD','thm.bool.IMP_ANTISYM_AX')).
fof(f21274,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,X31)) | s(bool,X32) = s(bool,X33) | ~p(s(bool,X33)) | s(bool,X30) != s(bool,X32) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X30)) ) | $spl16),
  inference(superposition,[],[f19130,f5008])).
fof(f5008,plain,(
  ( ! [X33,X32] : (s(bool,X32) = s(bool,X33) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,X33)) )),
  inference(resolution,[],[f4764,f4678])).
fof(f4764,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) )),
  inference(resolution,[],[f4678,f65])).
fof(f65,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f34])).
fof(f19130,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,X10) = s(bool,X11) | ~p(s(bool,X11)) | s(bool,X9) != s(bool,X10)) ) | $spl16),
  inference(subsumption_resolution,[],[f19112,f79])).
fof(f79,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) => p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.realax.HREAL_LT_LADD','thm.bool.IMP_CLAUSES|split|2')).
fof(f19112,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,X10) = s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'T_0')) | ~p(s(bool,X9))) ) | $spl16),
  inference(superposition,[],[f18584,f65])).
fof(f18584,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X0) | s(bool,X0) = s(bool,X1) | ~p(s(bool,X1))) ) | $spl16),
  inference(resolution,[],[f133,f65])).
fof(f21297,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,X28) = s(bool,X29) | ~p(s(bool,X29)) | s(bool,X26) != s(bool,X28) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X27)) ) | $spl16),
  inference(subsumption_resolution,[],[f21273,f4678])).
fof(f21273,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,X27)) | s(bool,X28) = s(bool,X29) | ~p(s(bool,X29)) | s(bool,X26) != s(bool,X28) | s(bool,'F_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X27)) ) | $spl16),
  inference(superposition,[],[f19130,f5008])).
fof(f4391,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,X46)) ) | $spl68),
  inference(cnf_transformation,[],[f4391_D])).
fof(f4391_D,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,X46)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f18313,plain,(
  ( ! [X76,X77] : (s(bool,X76) = s(bool,X77) | s(bool,'F_0') != s(bool,X76) | s(bool,'T_0') = s(bool,X77)) ) | $spl252),
  inference(cnf_transformation,[],[f18313_D])).
fof(f18313_D,plain,(
  ( ! [X76,X77] : (s(bool,X76) = s(bool,X77) | s(bool,'F_0') != s(bool,X76) | s(bool,'T_0') = s(bool,X77)) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f21190,plain,(
  ( ! [X39,X41,X40] : (s(bool,'F_0') != s(bool,X41) | s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X40)) ) | $spl62),
  inference(subsumption_resolution,[],[f21162,f68])).
fof(f21162,plain,(
  ( ! [X39,X41,X40] : (s(bool,'F_0') != s(bool,X40) | s(bool,'F_0') != s(bool,X41) | s(bool,X39) = s(bool,X41) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X40)) ) | $spl62),
  inference(superposition,[],[f3879,f19088])).
fof(f19088,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) )),
  inference(resolution,[],[f18556,f4691])).
fof(f18556,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(resolution,[],[f4691,f62])).
fof(f21189,plain,(
  ( ! [X37,X38,X36] : (s(bool,'F_0') != s(bool,X38) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,X36)) ) | $spl62),
  inference(subsumption_resolution,[],[f21161,f68])).
fof(f21161,plain,(
  ( ! [X37,X38,X36] : (s(bool,'F_0') != s(bool,X37) | s(bool,'F_0') != s(bool,X38) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,X36)) ) | $spl62),
  inference(superposition,[],[f3879,f19088])).
fof(f11035,plain,(
  ( ! [X114,X117] : (~p(s(bool,X114)) | s(bool,X114) = s(bool,X117)) ) | $spl110),
  inference(cnf_transformation,[],[f11035_D])).
fof(f11035_D,plain,(
  ( ! [X114,X117] : (~p(s(bool,X114)) | s(bool,X114) = s(bool,X117)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f436,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X2))) ) | $spl24),
  inference(cnf_transformation,[],[f436_D])).
fof(f436_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f21186,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X10) | s(bool,X9) = s(bool,X11) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | $spl62),
  inference(subsumption_resolution,[],[f21151,f4764])).
fof(f21151,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X10) | s(bool,'F_0') != s(bool,X11) | s(bool,X9) = s(bool,X11) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | $spl62),
  inference(superposition,[],[f3879,f65])).
fof(f21185,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,X6) = s(bool,X8) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | $spl62),
  inference(subsumption_resolution,[],[f21150,f4764])).
fof(f21150,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'F_0') != s(bool,X8) | s(bool,X6) = s(bool,X8) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | $spl62),
  inference(superposition,[],[f3879,f65])).
fof(f434,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl22),
  inference(cnf_transformation,[],[f434_D])).
fof(f434_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2819,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | s(bool,'F_0') != s(bool,X20) | s(bool,X20) = s(bool,X21)) ) | $spl32),
  inference(cnf_transformation,[],[f2819_D])).
fof(f2819_D,plain,(
  ( ! [X21,X20] : (p(s(bool,X21)) | s(bool,'F_0') != s(bool,X20) | s(bool,X20) = s(bool,X21)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f21184,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X5) | s(bool,X3) = s(bool,X5) | p(s(bool,X3)) | p(s(bool,X4))) ) | $spl62),
  inference(subsumption_resolution,[],[f21149,f4678])).
fof(f21149,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,'F_0') != s(bool,X5) | s(bool,X3) = s(bool,X5) | p(s(bool,X3)) | p(s(bool,X4))) ) | $spl62),
  inference(superposition,[],[f3879,f62])).
fof(f21183,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X2) | s(bool,X0) = s(bool,X2) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl62),
  inference(subsumption_resolution,[],[f21148,f4678])).
fof(f21148,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'F_0') != s(bool,X2) | s(bool,X0) = s(bool,X2) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl62),
  inference(superposition,[],[f3879,f62])).
fof(f724,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,X23) | s(bool,'T_0') != s(bool,X23)) ) | $spl26),
  inference(cnf_transformation,[],[f724_D])).
fof(f724_D,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,X23) | s(bool,'T_0') != s(bool,X23)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f21080,plain,(
  ( ! [X52,X51] : (s(bool,'F_0') != s(bool,X52) | s(bool,'T_0') != s(bool,X52) | s(bool,'T_0') = s(bool,X51)) ) | $spl54),
  inference(trivial_inequality_removal,[],[f21073])).
fof(f21073,plain,(
  ( ! [X52,X51] : (s(bool,X51) != s(bool,X51) | s(bool,'F_0') != s(bool,X52) | s(bool,'T_0') != s(bool,X52) | s(bool,'T_0') = s(bool,X51)) ) | $spl54),
  inference(superposition,[],[f3439,f68])).
fof(f18386,plain,(
  ( ! [X81,X84] : (s(bool,'F_0') = s(bool,X84) | s(bool,X81) = s(bool,X84)) ) | $spl264),
  inference(cnf_transformation,[],[f18386_D])).
fof(f18386_D,plain,(
  ( ! [X81,X84] : (s(bool,'F_0') = s(bool,X84) | s(bool,X81) = s(bool,X84)) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f21081,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') != s(bool,X49) | s(bool,'T_0') != s(bool,X49) | s(bool,X48) = s(bool,X50) | s(bool,'F_0') = s(bool,X50)) ) | $spl54),
  inference(trivial_inequality_removal,[],[f21072])).
fof(f21072,plain,(
  ( ! [X50,X48,X49] : (s(bool,X48) != s(bool,X48) | s(bool,'F_0') != s(bool,X49) | s(bool,'T_0') != s(bool,X49) | s(bool,X48) = s(bool,X50) | s(bool,'F_0') = s(bool,X50)) ) | $spl54),
  inference(superposition,[],[f3439,f5008])).
fof(f21082,plain,(
  ( ! [X47,X45,X46] : (s(bool,'F_0') != s(bool,X46) | s(bool,'T_0') != s(bool,X46) | s(bool,X45) = s(bool,X47) | s(bool,'F_0') = s(bool,X47)) ) | $spl54),
  inference(trivial_inequality_removal,[],[f21071])).
fof(f21071,plain,(
  ( ! [X47,X45,X46] : (s(bool,X45) != s(bool,X45) | s(bool,'F_0') != s(bool,X46) | s(bool,'T_0') != s(bool,X46) | s(bool,X45) = s(bool,X47) | s(bool,'F_0') = s(bool,X47)) ) | $spl54),
  inference(superposition,[],[f3439,f5008])).
fof(f21100,plain,(
  ( ! [X38,X36] : (s(bool,X36) != s(bool,X38) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') != s(bool,X38)) ) | $spl290),
  inference(cnf_transformation,[],[f21100_D])).
fof(f21100_D,plain,(
  ( ! [X38,X36] : (s(bool,X36) != s(bool,X38) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') != s(bool,X38)) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f21101,plain,(
  ( ! [X39,X41,X40] : (s(bool,X39) != s(bool,X41) | s(bool,'T_0') != s(bool,X41) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X40)) ) | $spl54),
  inference(subsumption_resolution,[],[f21070,f68])).
fof(f21070,plain,(
  ( ! [X39,X41,X40] : (s(bool,'F_0') != s(bool,X40) | s(bool,X39) != s(bool,X41) | s(bool,'T_0') != s(bool,X41) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X40)) ) | $spl54),
  inference(superposition,[],[f3439,f19088])).
fof(f21099,plain,(
  ( ! [X37,X38,X36] : (s(bool,X36) != s(bool,X38) | s(bool,'T_0') != s(bool,X38) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,X36)) ) | $spl54),
  inference(subsumption_resolution,[],[f21069,f68])).
fof(f21069,plain,(
  ( ! [X37,X38,X36] : (s(bool,'F_0') != s(bool,X37) | s(bool,X36) != s(bool,X38) | s(bool,'T_0') != s(bool,X38) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,X36)) ) | $spl54),
  inference(superposition,[],[f3439,f19088])).
fof(f239,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f239_D])).
fof(f239_D,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f8147,plain,(
  ( ! [X56] : (s(bool,'T_0') != s(bool,X56)) ) | $spl84),
  inference(cnf_transformation,[],[f8147_D])).
fof(f8147_D,plain,(
  ( ! [X56] : (s(bool,'T_0') != s(bool,X56)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f21096,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X10) | s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl16 | $spl54)),
  inference(subsumption_resolution,[],[f21059,f18584])).
fof(f21059,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X10) | s(bool,X9) != s(bool,X11) | s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | $spl54),
  inference(superposition,[],[f3439,f65])).
fof(f21095,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X8) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl16 | $spl54)),
  inference(subsumption_resolution,[],[f21058,f18584])).
fof(f21058,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,X6) != s(bool,X8) | s(bool,'T_0') != s(bool,X8) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | $spl54),
  inference(superposition,[],[f3439,f65])).
fof(f21093,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0)) | s(bool,'T_0') != s(bool,X2)) ) | $spl288),
  inference(cnf_transformation,[],[f21093_D])).
fof(f21093_D,plain,(
  ( ! [X2,X0] : (s(bool,X0) != s(bool,X2) | p(s(bool,X0)) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f21094,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X3)) | p(s(bool,X4))) ) | $spl54),
  inference(subsumption_resolution,[],[f21057,f4678])).
fof(f21057,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | s(bool,X3) != s(bool,X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X3)) | p(s(bool,X4))) ) | $spl54),
  inference(superposition,[],[f3439,f62])).
fof(f21092,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X2) | s(bool,'T_0') != s(bool,X2) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl54),
  inference(subsumption_resolution,[],[f21056,f4678])).
fof(f21056,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,'T_0') != s(bool,X2) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl54),
  inference(superposition,[],[f3439,f62])).
fof(f20970,plain,(
  ( ! [X70,X69,X67] : (s(bool,X67) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal',X70))))) | s(bool,'T_0') = s(bool,X67)) ) | $spl286),
  inference(cnf_transformation,[],[f20970_D])).
fof(f20970_D,plain,(
  ( ! [X70,X69,X67] : (s(bool,X67) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal',X70))))) | s(bool,'T_0') = s(bool,X67)) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f20987,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,X67) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal',X70))))) | s(bool,'T_0') = s(bool,X68) | s(bool,'T_0') = s(bool,X67)) ) | $spl86),
  inference(subsumption_resolution,[],[f20611,f68])).
fof(f20611,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,'F_0') != s(bool,X68) | s(bool,X67) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal',X70))))) | s(bool,'T_0') = s(bool,X68) | s(bool,'T_0') = s(bool,X67)) ) | $spl86),
  inference(superposition,[],[f8152,f19088])).
fof(f6704,plain,(
  ( ! [X30,X33,X32] : (s(bool,X30) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X32),s('type.hreal.hreal',X33))) | s(bool,'T_0') = s(bool,X30) | s('type.hreal.hreal',X33) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X32),s('type.hreal.hreal',sK3(X33,X32))))) ) | $spl74),
  inference(cnf_transformation,[],[f6704_D])).
fof(f6704_D,plain,(
  ( ! [X30,X33,X32] : (s(bool,X30) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X32),s('type.hreal.hreal',X33))) | s(bool,'T_0') = s(bool,X30) | s('type.hreal.hreal',X33) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X32),s('type.hreal.hreal',sK3(X33,X32))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f20986,plain,(
  ( ! [X66,X64,X65,X63] : (s(bool,X63) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X65),s('type.hreal.hreal',X66))) | s('type.hreal.hreal',X66) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X65),s('type.hreal.hreal',sK3(X66,X65)))) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f20610,f4691])).
fof(f20610,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | s(bool,X63) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X65),s('type.hreal.hreal',X66))) | s('type.hreal.hreal',X66) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X65),s('type.hreal.hreal',sK3(X66,X65)))) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f5061,f19088])).
fof(f5061,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X5)) | s(bool,X5) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X4),s('type.hreal.hreal',X3))) | s('type.hreal.hreal',X3) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X4),s('type.hreal.hreal',sK3(X3,X4))))) )),
  inference(resolution,[],[f60,f62])).
fof(f3265,plain,(
  ( ! [X52,X51,X49] : (s(bool,'T_0') = s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,X51) = s(bool,X52) | s(bool,X49) = s(bool,X52)) ) | $spl50),
  inference(cnf_transformation,[],[f3265_D])).
fof(f3265_D,plain,(
  ( ! [X52,X51,X49] : (s(bool,'T_0') = s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,X51) = s(bool,X52) | s(bool,X49) = s(bool,X52)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f20981,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,'F_0') = s(bool,X49) | s(bool,X47) = s(bool,X50) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f20605,f4691])).
fof(f20605,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X48)) | s(bool,'F_0') = s(bool,X49) | s(bool,X47) = s(bool,X50) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X48) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f4997,f19088])).
fof(f4997,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1)) )),
  inference(resolution,[],[f4764,f62])).
fof(f20969,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,X67) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal',X70))))) | s(bool,'T_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X68)) ) | $spl86),
  inference(subsumption_resolution,[],[f20555,f68])).
fof(f20555,plain,(
  ( ! [X70,X68,X69,X67] : (s(bool,'F_0') != s(bool,X68) | s(bool,X67) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal',X70))))) | s(bool,'T_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X68)) ) | $spl86),
  inference(superposition,[],[f8152,f19088])).
fof(f20968,plain,(
  ( ! [X66,X64,X65,X63] : (s(bool,X63) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X65),s('type.hreal.hreal',X66))) | s('type.hreal.hreal',X66) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X65),s('type.hreal.hreal',sK3(X66,X65)))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f20554,f4691])).
fof(f20554,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | s(bool,X63) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X65),s('type.hreal.hreal',X66))) | s('type.hreal.hreal',X66) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X65),s('type.hreal.hreal',sK3(X66,X65)))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f5061,f19088])).
fof(f20963,plain,(
  ( ! [X47,X50,X48,X49] : (s(bool,'F_0') = s(bool,X49) | s(bool,X47) = s(bool,X50) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f20549,f4691])).
fof(f20549,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X48)) | s(bool,'F_0') = s(bool,X49) | s(bool,X47) = s(bool,X50) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f4997,f19088])).
fof(f17351,plain,(
  ( ! [X61,X59,X60] : (s('type.hreal.hreal',X61) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X59),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X60),s('type.hreal.hreal',sK3(X61,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X59),s('type.hreal.hreal',X60)))))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X59),s('type.hreal.hreal',X60))),s('type.hreal.hreal',X61)))) ) | $spl218),
  inference(cnf_transformation,[],[f17351_D])).
fof(f17351_D,plain,(
  ( ! [X61,X59,X60] : (s('type.hreal.hreal',X61) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X59),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X60),s('type.hreal.hreal',sK3(X61,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X59),s('type.hreal.hreal',X60)))))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X59),s('type.hreal.hreal',X60))),s('type.hreal.hreal',X61)))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f20288,plain,(
  ( ! [X121,X118,X120,X119] : (s('type.hreal.hreal',X120) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X119),s('type.hreal.hreal',sK3(X120,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119)))))))) | s(bool,'F_0') = s(bool,X121) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119))),s('type.hreal.hreal',X120)))) )),
  inference(subsumption_resolution,[],[f20040,f4678])).
fof(f20040,plain,(
  ( ! [X121,X118,X120,X119] : (~p(s(bool,X121)) | s('type.hreal.hreal',X120) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X119),s('type.hreal.hreal',sK3(X120,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119)))))))) | s(bool,'F_0') = s(bool,X121) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119))),s('type.hreal.hreal',X120)))) )),
  inference(superposition,[],[f5083,f5008])).
fof(f5083,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X8),s('type.hreal.hreal',X9))),s('type.hreal.hreal',X10)))) | s('type.hreal.hreal',X10) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X8),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X9),s('type.hreal.hreal',sK3(X10,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X8),s('type.hreal.hreal',X9))))))))) )),
  inference(superposition,[],[f60,f59])).
fof(f8827,plain,(
  ( ! [X66,X64,X63] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X64),s('type.hreal.hreal',X66))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X64,X63)),s('type.hreal.hreal',X66))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal',X64)))) ) | $spl90),
  inference(cnf_transformation,[],[f8827_D])).
fof(f8827_D,plain,(
  ( ! [X66,X64,X63] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X64),s('type.hreal.hreal',X66))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X64,X63)),s('type.hreal.hreal',X66))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal',X64)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f20287,plain,(
  ( ! [X107,X105,X106,X104] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X105),s('type.hreal.hreal',X107))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X104),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X105,X104)),s('type.hreal.hreal',X107))))) | s(bool,'F_0') = s(bool,X106) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X104),s('type.hreal.hreal',X105)))) )),
  inference(subsumption_resolution,[],[f20036,f4678])).
fof(f20036,plain,(
  ( ! [X107,X105,X106,X104] : (~p(s(bool,X106)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X105),s('type.hreal.hreal',X107))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X104),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X105,X104)),s('type.hreal.hreal',X107))))) | s(bool,'F_0') = s(bool,X106) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X104),s('type.hreal.hreal',X105)))) )),
  inference(superposition,[],[f5121,f5008])).
fof(f5121,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X12),s('type.hreal.hreal',X13)))) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X13),s('type.hreal.hreal',X14))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X12),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X13,X12)),s('type.hreal.hreal',X14)))))) )),
  inference(superposition,[],[f59,f60])).
fof(f15258,plain,(
  ( ! [X66,X64,X63] : (s('type.hreal.hreal',sK3(X64,X63)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X66),s('type.hreal.hreal',X64)),X66),X63)) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal',X64)))) ) | $spl214),
  inference(cnf_transformation,[],[f15258_D])).
fof(f15258_D,plain,(
  ( ! [X66,X64,X63] : (s('type.hreal.hreal',sK3(X64,X63)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X66),s('type.hreal.hreal',X64)),X66),X63)) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal',X64)))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f20286,plain,(
  ( ! [X103,X101,X102,X100] : (s('type.hreal.hreal',sK3(X101,X100)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X103),s('type.hreal.hreal',X101)),X103),X100)) | s(bool,'F_0') = s(bool,X102) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X100),s('type.hreal.hreal',X101)))) )),
  inference(subsumption_resolution,[],[f20035,f4678])).
fof(f20035,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,X102)) | s('type.hreal.hreal',sK3(X101,X100)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X103),s('type.hreal.hreal',X101)),X103),X100)) | s(bool,'F_0') = s(bool,X102) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X100),s('type.hreal.hreal',X101)))) )),
  inference(superposition,[],[f14976,f5008])).
fof(f14976,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X11),s('type.hreal.hreal',X12)))) | s('type.hreal.hreal',sK3(X12,X11)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X13),s('type.hreal.hreal',X12)),X13),X11))) )),
  inference(superposition,[],[f14938,f60])).
fof(f14938,plain,(
  ( ! [X2,X0,X1] : (s('type.hreal.hreal',X0) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X0)))),X1),X2))) )),
  inference(equality_resolution,[],[f5355])).
fof(f5355,plain,(
  ( ! [X116,X114,X115,X113] : (s('type.hreal.hreal',X114) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal',X116))) | s('type.hreal.hreal',X116) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X113),s('type.hreal.hreal',X114)),X113),X115))) )),
  inference(superposition,[],[f5177,f5307])).
fof(f5307,plain,(
  ( ! [X0,X1] : (s('type.hreal.hreal',X0) = s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0)),X1))) )),
  inference(equality_resolution,[],[f5177])).
fof(f5177,plain,(
  ( ! [X2,X0,X1] : (s('type.hreal.hreal',X1) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2))) | s('type.hreal.hreal',X2) = s('type.hreal.hreal',sK3(X1,X0))) )),
  inference(subsumption_resolution,[],[f5117,f61])).
fof(f5117,plain,(
  ( ! [X2,X0,X1] : (s('type.hreal.hreal',X1) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2))) | s('type.hreal.hreal',X2) = s('type.hreal.hreal',sK3(X1,X0)) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))) )),
  inference(superposition,[],[f57,f60])).
fof(f57,plain,(
  ( ! [X2,X0,X1] : (s('type.hreal.hreal',X1) = s('type.hreal.hreal',X2) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2)))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1,X2] : ((s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2))) | s('type.hreal.hreal',X1) = s('type.hreal.hreal',X2)) & (s('type.hreal.hreal',X1) != s('type.hreal.hreal',X2) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2)))))),
  inference(nnf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1,X2] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2))) <=> s('type.hreal.hreal',X1) = s('type.hreal.hreal',X2))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6,X7,X8] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X8))) <=> s('type.hreal.hreal',X7) = s('type.hreal.hreal',X8))),
  file('thm.realax.HREAL_LT_LADD','thm.realax.HREAL_EQ_LADD')).
fof(f20279,plain,(
  ( ! [X121,X118,X120,X119] : (s('type.hreal.hreal',X120) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X119),s('type.hreal.hreal',sK3(X120,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119)))))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119))),s('type.hreal.hreal',X120))) | s(bool,'F_0') = s(bool,X121)) )),
  inference(subsumption_resolution,[],[f19990,f4678])).
fof(f19990,plain,(
  ( ! [X121,X118,X120,X119] : (~p(s(bool,X121)) | s('type.hreal.hreal',X120) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X119),s('type.hreal.hreal',sK3(X120,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119)))))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119))),s('type.hreal.hreal',X120))) | s(bool,'F_0') = s(bool,X121)) )),
  inference(superposition,[],[f5083,f5008])).
fof(f20278,plain,(
  ( ! [X107,X105,X106,X104] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X105),s('type.hreal.hreal',X107))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X104),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X105,X104)),s('type.hreal.hreal',X107))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X104),s('type.hreal.hreal',X105))) | s(bool,'F_0') = s(bool,X106)) )),
  inference(subsumption_resolution,[],[f19986,f4678])).
fof(f19986,plain,(
  ( ! [X107,X105,X106,X104] : (~p(s(bool,X106)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X105),s('type.hreal.hreal',X107))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X104),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X105,X104)),s('type.hreal.hreal',X107))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X104),s('type.hreal.hreal',X105))) | s(bool,'F_0') = s(bool,X106)) )),
  inference(superposition,[],[f5121,f5008])).
fof(f20277,plain,(
  ( ! [X103,X101,X102,X100] : (s('type.hreal.hreal',sK3(X101,X100)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X103),s('type.hreal.hreal',X101)),X103),X100)) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X100),s('type.hreal.hreal',X101))) | s(bool,'F_0') = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f19985,f4678])).
fof(f19985,plain,(
  ( ! [X103,X101,X102,X100] : (~p(s(bool,X102)) | s('type.hreal.hreal',sK3(X101,X100)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X103),s('type.hreal.hreal',X101)),X103),X100)) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X100),s('type.hreal.hreal',X101))) | s(bool,'F_0') = s(bool,X102)) )),
  inference(superposition,[],[f14976,f5008])).
fof(f19667,plain,(
  ( ! [X4] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK2),s('type.hreal.hreal',X4))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(sK2,sK1)),s('type.hreal.hreal',X4)))))) ) | $spl284),
  inference(cnf_transformation,[],[f19667_D])).
fof(f19667_D,plain,(
  ( ! [X4] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK2),s('type.hreal.hreal',X4))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(sK2,sK1)),s('type.hreal.hreal',X4)))))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f19815,plain,(
  ( ! [X4,X3] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK2),s('type.hreal.hreal',X4))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(sK2,sK1)),s('type.hreal.hreal',X4))))) | s(bool,'F_0') = s(bool,X3)) ) | $spl2),
  inference(subsumption_resolution,[],[f19748,f4678])).
fof(f19748,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK2),s('type.hreal.hreal',X4))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(sK2,sK1)),s('type.hreal.hreal',X4))))) | s(bool,'F_0') = s(bool,X3)) ) | $spl2),
  inference(superposition,[],[f5121,f18660])).
fof(f18660,plain,(
  ( ! [X11] : (s(bool,X11) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X11)) ) | $spl2),
  inference(subsumption_resolution,[],[f18638,f85])).
fof(f18638,plain,(
  ( ! [X11] : (~p(s(bool,'T_0')) | s(bool,X11) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X11)) ) | $spl2),
  inference(superposition,[],[f4764,f109])).
fof(f19690,plain,(
  $false | ($spl2 | $spl15)),
  inference(subsumption_resolution,[],[f19689,f86])).
fof(f86,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.realax.HREAL_LT_LADD',aHL_TRUTH)).
fof(f19689,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl15)),
  inference(forward_demodulation,[],[f131,f109])).
fof(f19688,plain,(
  $false | ($spl2 | $spl15)),
  inference(subsumption_resolution,[],[f19687,f86])).
fof(f19687,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl15)),
  inference(forward_demodulation,[],[f14484,f109])).
fof(f14484,plain,(
  ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) | $spl15),
  inference(forward_demodulation,[],[f14479,f5307])).
fof(f14479,plain,(
  ( ! [X0] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',sK2)),X0)))))) ) | $spl15),
  inference(equality_resolution,[],[f14451])).
fof(f14451,plain,(
  ( ! [X8,X9] : (s('type.hreal.hreal',X9) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X8),s('type.hreal.hreal',sK2))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK3(X9,X8)))))) ) | $spl15),
  inference(subsumption_resolution,[],[f14442,f61])).
fof(f14442,plain,(
  ( ! [X8,X9] : (s('type.hreal.hreal',X9) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X8),s('type.hreal.hreal',sK2))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK3(X9,X8))))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X8),s('type.hreal.hreal',X9))))) ) | $spl15),
  inference(superposition,[],[f259,f60])).
fof(f259,plain,(
  ( ! [X0,X1] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',sK2))) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',X0))))) ) | $spl15),
  inference(superposition,[],[f131,f57])).
fof(f19686,plain,(
  $false | ($spl2 | $spl15)),
  inference(subsumption_resolution,[],[f19685,f86])).
fof(f19685,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl15)),
  inference(forward_demodulation,[],[f14374,f109])).
fof(f14374,plain,(
  ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) | $spl15),
  inference(equality_resolution,[],[f257])).
fof(f257,plain,(
  ( ! [X0,X1] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',sK1))) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',sK2))))) ) | $spl15),
  inference(superposition,[],[f131,f57])).
fof(f19684,plain,(
  $false | ($spl2 | $spl15)),
  inference(subsumption_resolution,[],[f19683,f86])).
fof(f19683,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl15)),
  inference(forward_demodulation,[],[f14446,f109])).
fof(f14446,plain,(
  ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) | $spl15),
  inference(equality_resolution,[],[f259])).
fof(f19682,plain,(
  $false | ($spl2 | $spl15)),
  inference(subsumption_resolution,[],[f19681,f86])).
fof(f19681,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl15)),
  inference(forward_demodulation,[],[f14412,f109])).
fof(f14412,plain,(
  ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) | $spl15),
  inference(forward_demodulation,[],[f14407,f5307])).
fof(f14407,plain,(
  ( ! [X0] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',sK1)),X0)),s('type.hreal.hreal',sK2))))) ) | $spl15),
  inference(equality_resolution,[],[f14379])).
fof(f14379,plain,(
  ( ! [X8,X9] : (s('type.hreal.hreal',X9) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X8),s('type.hreal.hreal',sK1))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK3(X9,X8)),s('type.hreal.hreal',sK2))))) ) | $spl15),
  inference(subsumption_resolution,[],[f14370,f61])).
fof(f14370,plain,(
  ( ! [X8,X9] : (s('type.hreal.hreal',X9) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X8),s('type.hreal.hreal',sK1))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK3(X9,X8)),s('type.hreal.hreal',sK2)))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X8),s('type.hreal.hreal',X9))))) ) | $spl15),
  inference(superposition,[],[f257,f60])).
fof(f19666,plain,(
  ( ! [X4,X3] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK2),s('type.hreal.hreal',X4))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(sK2,sK1)),s('type.hreal.hreal',X4))))) | s(bool,'F_0') = s(bool,X3)) ) | $spl14),
  inference(subsumption_resolution,[],[f19599,f4678])).
fof(f19599,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK2),s('type.hreal.hreal',X4))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(sK2,sK1)),s('type.hreal.hreal',X4))))) | s(bool,'F_0') = s(bool,X3)) ) | $spl14),
  inference(superposition,[],[f5121,f18542])).
fof(f18542,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X4)) ) | $spl14),
  inference(resolution,[],[f132,f4764])).
fof(f19535,plain,(
  ( ! [X423] : (s('type.hreal.hreal',sK3(sK2,sK1)) = s('type.hreal.hreal',sK3(sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X423),s('type.hreal.hreal',sK2)),X423),sK1),sK1))) ) | $spl280),
  inference(cnf_transformation,[],[f19535_D])).
fof(f19535_D,plain,(
  ( ! [X423] : (s('type.hreal.hreal',sK3(sK2,sK1)) = s('type.hreal.hreal',sK3(sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X423),s('type.hreal.hreal',sK2)),X423),sK1),sK1))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f19537,plain,(
  s('type.hreal.hreal',sK2) != s('type.hreal.hreal',sK3(sK2,sK1)) | $spl283),
  inference(cnf_transformation,[],[f19537_D])).
fof(f19537_D,plain,(
  s('type.hreal.hreal',sK2) != s('type.hreal.hreal',sK3(sK2,sK1)) <=> ~$spl283),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl283])])).
fof(f19461,plain,(
  ( ! [X423] : (s('type.hreal.hreal',sK2) != s('type.hreal.hreal',sK3(sK2,sK1)) | s('type.hreal.hreal',sK3(sK2,sK1)) = s('type.hreal.hreal',sK3(sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X423),s('type.hreal.hreal',sK2)),X423),sK1),sK1))) ) | $spl2),
  inference(superposition,[],[f18810,f18814])).
fof(f18814,plain,(
  ( ! [X33] : (s('type.hreal.hreal',sK3(sK2,sK1)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X33),s('type.hreal.hreal',sK2)),X33),sK1))) ) | $spl2),
  inference(superposition,[],[f14938,f18651])).
fof(f18810,plain,(
  ( ! [X26] : (s('type.hreal.hreal',sK2) != s('type.hreal.hreal',X26) | s('type.hreal.hreal',sK3(sK2,sK1)) = s('type.hreal.hreal',sK3(X26,sK1))) ) | $spl2),
  inference(superposition,[],[f5177,f18651])).
fof(f19016,plain,(
  $false | ($spl2 | $spl12 | $spl16)),
  inference(trivial_inequality_removal,[],[f19014])).
fof(f19014,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl2 | $spl12 | $spl16)),
  inference(superposition,[],[f18999,f109])).
fof(f18999,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2)) ) | ($spl2 | $spl12 | $spl16)),
  inference(subsumption_resolution,[],[f18998,f133])).
fof(f18998,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2))) ) | ($spl2 | $spl12)),
  inference(forward_demodulation,[],[f118,f109])).
fof(f19015,plain,(
  $false | ($spl2 | $spl12 | $spl16)),
  inference(equality_resolution,[],[f18999])).
fof(f18784,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | ~p(s(bool,X7))) ) | ($spl16 | $spl26)),
  inference(subsumption_resolution,[],[f18774,f133])).
fof(f18774,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,'T_0') != s(bool,X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl26),
  inference(superposition,[],[f724,f65])).
fof(f18783,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,'T_0') != s(bool,X4) | ~p(s(bool,X5))) ) | ($spl16 | $spl26)),
  inference(subsumption_resolution,[],[f18773,f133])).
fof(f18773,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,'T_0') != s(bool,X4) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl26),
  inference(superposition,[],[f724,f65])).
fof(f115,plain,(
  ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) | $spl9),
  inference(cnf_transformation,[],[f115_D])).
fof(f115_D,plain,(
  ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f18705,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))))) ) | ($spl2 | $spl16)),
  inference(subsumption_resolution,[],[f18700,f133])).
fof(f18700,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))))) ) | $spl2),
  inference(superposition,[],[f18607,f65])).
fof(f18607,plain,(
  s(bool,'T_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | $spl2),
  inference(backward_demodulation,[],[f109,f56])).
fof(f56,plain,(
  s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f32])).
fof(f32,plain,(
  ? [X0,X1,X2] : s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X2))) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2)))))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ~! [X0,X1,X2] : s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X2))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X7),s('type.hreal.hreal',X8))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X7),s('type.hreal.hreal',X8))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X8)))))),
  file('thm.realax.HREAL_LT_LADD','thm.realax.HREAL_LT_LADD')).
fof(f18704,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))))) ) | ($spl2 | $spl16)),
  inference(subsumption_resolution,[],[f18699,f133])).
fof(f18699,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) | ~p(s(bool,X2))) ) | $spl2),
  inference(superposition,[],[f18607,f65])).
fof(f18678,plain,(
  $false | ($spl2 | $spl16 | $spl18)),
  inference(trivial_inequality_removal,[],[f18676])).
fof(f18676,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl2 | $spl16 | $spl18)),
  inference(superposition,[],[f18661,f109])).
fof(f18661,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2)) ) | ($spl16 | $spl18)),
  inference(subsumption_resolution,[],[f134,f133])).
fof(f18677,plain,(
  $false | ($spl16 | $spl18)),
  inference(equality_resolution,[],[f18661])).
fof(f112,plain,(
  s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | $spl4),
  inference(cnf_transformation,[],[f112_D])).
fof(f112_D,plain,(
  s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f114,plain,(
  s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | $spl7),
  inference(cnf_transformation,[],[f114_D])).
fof(f114_D,plain,(
  s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f18534,plain,(
  s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))),
  inference(superposition,[],[f56,f68])).
fof(f108,plain,(
  s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | $spl0),
  inference(cnf_transformation,[],[f108_D])).
fof(f108_D,plain,(
  s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f110,plain,(
  s(bool,'T_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | $spl3),
  inference(cnf_transformation,[],[f110_D])).
fof(f110_D,plain,(
  s(bool,'T_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f18533,plain,(
  s(bool,'T_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))),
  inference(superposition,[],[f56,f68])).
fof(f118,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | ~p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f118_D])).
fof(f118_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | ~p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f18532,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | ~p(s(bool,X3)) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))))) )),
  inference(superposition,[],[f56,f65])).
fof(f18531,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f56,f65])).
fof(f116,plain,(
  p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) | $spl8),
  inference(cnf_transformation,[],[f116_D])).
fof(f116_D,plain,(
  p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f117,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | p(s(bool,X0))) ) | $spl10),
  inference(cnf_transformation,[],[f117_D])).
fof(f117_D,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | p(s(bool,X0))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f18530,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) | p(s(bool,X1))) )),
  inference(superposition,[],[f56,f62])).
fof(f18529,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | p(s(bool,X0)) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))))) )),
  inference(superposition,[],[f56,f62])).
fof(f18449,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl275),
  inference(cnf_transformation,[],[f18449_D])).
fof(f18449_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl275),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl275])])).
fof(f278,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X9))) ) | $spl16),
  inference(resolution,[],[f133,f70])).
fof(f18505,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f18483,f239])).
fof(f18483,plain,(
  ( ! [X14] : (p(s(bool,X14))) ) | $spl20),
  inference(resolution,[],[f239,f81])).
fof(f81,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.realax.HREAL_LT_LADD','thm.bool.IMP_CLAUSES|split|3')).
fof(f18504,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f18481,f239])).
fof(f18481,plain,(
  ( ! [X12] : (p(s(bool,X12))) ) | $spl20),
  inference(resolution,[],[f239,f75])).
fof(f75,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11] : ((p(s(bool,'T_0')) => p(s(bool,X11))) <=> p(s(bool,X11)))),
  file('thm.realax.HREAL_LT_LADD','thm.bool.IMP_CLAUSES|split|0')).
fof(f18503,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f18480,f239])).
fof(f18480,plain,(
  ( ! [X11] : (p(s(bool,X11))) ) | $spl20),
  inference(resolution,[],[f239,f72])).
fof(f72,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.realax.HREAL_LT_LADD','thm.bool.IMP_CLAUSES|split|1')).
fof(f18485,plain,(
  $false | $spl20),
  inference(resolution,[],[f239,f85])).
fof(f18486,plain,(
  $false | $spl20),
  inference(resolution,[],[f239,f86])).
fof(f18454,plain,(
  s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))))))) | $spl278),
  inference(cnf_transformation,[],[f18454_D])).
fof(f18454_D,plain,(
  s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))))))) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f18353,plain,(
  ( ! [X3] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))))))) | s(bool,'F_0') = s(bool,X3)) ) | $spl4),
  inference(subsumption_resolution,[],[f17853,f4678])).
fof(f17853,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))))))) | s(bool,'F_0') = s(bool,X3)) ) | $spl4),
  inference(superposition,[],[f5083,f5058])).
fof(f5058,plain,(
  ( ! [X61] : (s(bool,X61) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X61)) ) | $spl4),
  inference(subsumption_resolution,[],[f5040,f85])).
fof(f5040,plain,(
  ( ! [X61] : (~p(s(bool,'T_0')) | s(bool,X61) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X61)) ) | $spl4),
  inference(superposition,[],[f4764,f112])).
fof(f18450,plain,(
  ( ! [X243,X244,X242] : (s(bool,X242) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X243),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X243),s('type.hreal.hreal',X244)))))) ) | $spl276),
  inference(cnf_transformation,[],[f18450_D])).
fof(f18450_D,plain,(
  ( ! [X243,X244,X242] : (s(bool,X242) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X243),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X243),s('type.hreal.hreal',X244)))))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f18452,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X16) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X17),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X17),s('type.hreal.hreal',X18)))))) ) | $spl86),
  inference(subsumption_resolution,[],[f8337,f8152])).
fof(f8337,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X16) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X17),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X17),s('type.hreal.hreal',X18))))) | s(bool,'F_0') = s(bool,X16)) ) | $spl86),
  inference(superposition,[],[f8152,f68])).
fof(f18092,plain,(
  ( ! [X243,X244,X242] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X242) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X243),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X243),s('type.hreal.hreal',X244)))))) ) | ($spl4 | $spl86)),
  inference(forward_demodulation,[],[f18091,f112])).
fof(f18091,plain,(
  ( ! [X243,X244,X242] : (s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X242) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X243),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X243),s('type.hreal.hreal',X244)))))) ) | ($spl4 | $spl86)),
  inference(subsumption_resolution,[],[f17835,f8152])).
fof(f17835,plain,(
  ( ! [X243,X244,X242] : (s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X242) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X243),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X243),s('type.hreal.hreal',X244))))) | s(bool,'F_0') = s(bool,X242)) ) | ($spl4 | $spl86)),
  inference(superposition,[],[f8152,f5058])).
fof(f18447,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) ) | $spl272),
  inference(cnf_transformation,[],[f18447_D])).
fof(f18447_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f18190,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) ) | $spl4),
  inference(forward_demodulation,[],[f17969,f112])).
fof(f17969,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X0) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) ) | $spl4),
  inference(equality_factoring,[],[f5058])).
fof(f18443,plain,(
  $false | ($spl4 | $spl9)),
  inference(subsumption_resolution,[],[f18442,f86])).
fof(f18442,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl9)),
  inference(forward_demodulation,[],[f115,f112])).
fof(f18438,plain,(
  $false | ($spl6 | $spl14)),
  inference(subsumption_resolution,[],[f18437,f84])).
fof(f18437,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl14)),
  inference(forward_demodulation,[],[f132,f113])).
fof(f113,plain,(
  s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | $spl6),
  inference(cnf_transformation,[],[f113_D])).
fof(f113_D,plain,(
  s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f18432,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,X4) | s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X5))) ) | ($spl6 | $spl32)),
  inference(subsumption_resolution,[],[f18431,f4678])).
fof(f18431,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,X4) | s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | ($spl6 | $spl32)),
  inference(forward_demodulation,[],[f2920,f113])).
fof(f2920,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,X4) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | ($spl6 | $spl32)),
  inference(superposition,[],[f2900,f65])).
fof(f2900,plain,(
  ( ! [X60] : (s(bool,'F_0') != s(bool,X60) | s(bool,X60) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl32)),
  inference(subsumption_resolution,[],[f2889,f84])).
fof(f2889,plain,(
  ( ! [X60] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X60) | s(bool,X60) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl32)),
  inference(superposition,[],[f2819,f113])).
fof(f18430,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | s(bool,'F_0') != s(bool,X7) | ~p(s(bool,X7))) ) | ($spl6 | $spl32)),
  inference(subsumption_resolution,[],[f18429,f4678])).
fof(f18429,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | s(bool,'F_0') != s(bool,X7) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | ($spl6 | $spl32)),
  inference(forward_demodulation,[],[f2921,f113])).
fof(f2921,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,X6) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | ($spl6 | $spl32)),
  inference(superposition,[],[f2900,f65])).
fof(f18427,plain,(
  ( ! [X66,X67,X65] : (s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67) | s(bool,'F_0') != s(bool,X66) | s(bool,'F_0') = s(bool,X65)) ) | $spl32),
  inference(subsumption_resolution,[],[f3774,f4678])).
fof(f3774,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,X65)) | s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67) | s(bool,'F_0') != s(bool,X66) | s(bool,'F_0') = s(bool,X65)) ) | $spl32),
  inference(superposition,[],[f2853,f68])).
fof(f2853,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') != s(bool,X0)) ) | $spl32),
  inference(resolution,[],[f2819,f65])).
fof(f18316,plain,(
  ( ! [X76] : (~p(s(bool,X76)) | s(bool,'T_0') = s(bool,X76)) ) | $spl256),
  inference(cnf_transformation,[],[f18316_D])).
fof(f18316_D,plain,(
  ( ! [X76] : (~p(s(bool,X76)) | s(bool,'T_0') = s(bool,X76)) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f5042,plain,(
  ( ! [X57,X56] : (~p(s(bool,X56)) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X56)) )),
  inference(duplicate_literal_removal,[],[f5037])).
fof(f5037,plain,(
  ( ! [X57,X56] : (~p(s(bool,X56)) | s(bool,'F_0') = s(bool,X57) | s(bool,'F_0') = s(bool,X57) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f4764,f68])).
fof(f18415,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,X4) | s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X5))) ) | ($spl6 | $spl32)),
  inference(subsumption_resolution,[],[f18414,f4678])).
fof(f18414,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,X4) | s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | ($spl6 | $spl32)),
  inference(forward_demodulation,[],[f5508,f113])).
fof(f5508,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,X4) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | ($spl6 | $spl32)),
  inference(superposition,[],[f2900,f65])).
fof(f18413,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | s(bool,'F_0') != s(bool,X7) | ~p(s(bool,X7))) ) | ($spl6 | $spl32)),
  inference(subsumption_resolution,[],[f18412,f4678])).
fof(f18412,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,X6) | s(bool,'F_0') != s(bool,X7) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | ($spl6 | $spl32)),
  inference(forward_demodulation,[],[f5509,f113])).
fof(f5509,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,X6) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | ($spl6 | $spl32)),
  inference(superposition,[],[f2900,f65])).
fof(f18411,plain,(
  ( ! [X66,X67,X65] : (s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67) | s(bool,'F_0') != s(bool,X66) | s(bool,'F_0') = s(bool,X65)) ) | $spl32),
  inference(subsumption_resolution,[],[f6001,f4678])).
fof(f6001,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,X65)) | s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67) | s(bool,'F_0') != s(bool,X66) | s(bool,'F_0') = s(bool,X65)) ) | $spl32),
  inference(superposition,[],[f2853,f68])).
fof(f18315,plain,(
  ( ! [X78,X77] : (s(bool,X77) = s(bool,X78) | s(bool,'F_0') != s(bool,X77) | s(bool,'F_0') = s(bool,X78)) ) | $spl254),
  inference(cnf_transformation,[],[f18315_D])).
fof(f18315_D,plain,(
  ( ! [X78,X77] : (s(bool,X77) = s(bool,X78) | s(bool,'F_0') != s(bool,X77) | s(bool,'F_0') = s(bool,X78)) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f6004,plain,(
  ( ! [X78,X76,X77] : (~p(s(bool,X76)) | s(bool,X77) = s(bool,X78) | s(bool,'F_0') = s(bool,X78) | s(bool,'F_0') != s(bool,X77) | s(bool,'T_0') = s(bool,X76)) ) | $spl32),
  inference(superposition,[],[f2853,f68])).
fof(f18378,plain,(
  ( ! [X65] : (s(bool,X65) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X65)) ) | $spl260),
  inference(cnf_transformation,[],[f18378_D])).
fof(f18378_D,plain,(
  ( ! [X65] : (s(bool,X65) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X65)) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f7476,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X12) | s(bool,X13) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f6811])).
fof(f6811,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X12) | s(bool,X13) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X12)) ) | $spl6),
  inference(superposition,[],[f5704,f68])).
fof(f5704,plain,(
  ( ! [X80,X79] : (s(bool,'F_0') = s(bool,X79) | s(bool,X80) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X79) = s(bool,X80)) ) | $spl6),
  inference(subsumption_resolution,[],[f5686,f84])).
fof(f5686,plain,(
  ( ! [X80,X79] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,X79) | s(bool,X80) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X79) = s(bool,X80)) ) | $spl6),
  inference(superposition,[],[f4997,f113])).
fof(f7469,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') = s(bool,X15) | s(bool,X14) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X14)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f6845])).
fof(f6845,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') = s(bool,X15) | s(bool,'F_0') = s(bool,X15) | s(bool,X14) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X14)) ) | $spl6),
  inference(superposition,[],[f5704,f68])).
fof(f7440,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,'F_0') = s(bool,X24) | s(bool,X25) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f7196])).
fof(f7196,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,'F_0') = s(bool,X24) | s(bool,'F_0') = s(bool,X24) | s(bool,X25) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(superposition,[],[f68,f5704])).
fof(f7432,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X26) | s(bool,X26) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f7304])).
fof(f7304,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X26) | s(bool,'F_0') = s(bool,X27) | s(bool,X26) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(superposition,[],[f68,f5704])).
fof(f14030,plain,(
  ( ! [X107,X106,X104] : (s(bool,X104) = s(bool,X106) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X106) = s(bool,X107) | s(bool,'F_0') != s(bool,X107)) ) | $spl190),
  inference(cnf_transformation,[],[f14030_D])).
fof(f14030_D,plain,(
  ( ! [X107,X106,X104] : (s(bool,X104) = s(bool,X106) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X106) = s(bool,X107) | s(bool,'F_0') != s(bool,X107)) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f18393,plain,(
  ( ! [X173,X171,X172,X170] : (s(bool,X172) = s(bool,X173) | s(bool,X170) = s(bool,X173) | s(bool,'F_0') != s(bool,X172) | s(bool,'F_0') = s(bool,X171) | s(bool,X170) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl32)),
  inference(subsumption_resolution,[],[f7352,f4678])).
fof(f7352,plain,(
  ( ! [X173,X171,X172,X170] : (~p(s(bool,X171)) | s(bool,X172) = s(bool,X173) | s(bool,X170) = s(bool,X173) | s(bool,'F_0') != s(bool,X172) | s(bool,'F_0') = s(bool,X171) | s(bool,X170) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl32)),
  inference(superposition,[],[f2853,f5704])).
fof(f11186,plain,(
  ( ! [X375] : (s(bool,X375) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl140),
  inference(cnf_transformation,[],[f11186_D])).
fof(f11186_D,plain,(
  ( ! [X375] : (s(bool,X375) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f18389,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X1)) ) | $spl270),
  inference(cnf_transformation,[],[f18389_D])).
fof(f18389_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f7410,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(equality_factoring,[],[f5704])).
fof(f18388,plain,(
  ( ! [X6,X4,X7] : (s(bool,X4) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7))))) | p(s(bool,X4))) ) | $spl268),
  inference(cnf_transformation,[],[f18388_D])).
fof(f18388_D,plain,(
  ( ! [X6,X4,X7] : (s(bool,X4) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7))))) | p(s(bool,X4))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f8384,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X3))))) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl86),
  inference(subsumption_resolution,[],[f8333,f4678])).
fof(f8333,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X0) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X3))))) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl86),
  inference(superposition,[],[f8152,f62])).
fof(f8386,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7))))) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl86),
  inference(subsumption_resolution,[],[f8334,f4678])).
fof(f8334,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,X4) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X6),s('type.hreal.hreal',X7))))) | p(s(bool,X4)) | p(s(bool,X5))) ) | $spl86),
  inference(superposition,[],[f8152,f62])).
fof(f18387,plain,(
  ( ! [X14,X12,X15] : (s(bool,X12) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X14),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X14),s('type.hreal.hreal',X15))))) | ~p(s(bool,X12))) ) | $spl266),
  inference(cnf_transformation,[],[f18387_D])).
fof(f18387_D,plain,(
  ( ! [X14,X12,X15] : (s(bool,X12) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X14),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X14),s('type.hreal.hreal',X15))))) | ~p(s(bool,X12))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f8335,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s(bool,X8) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X10),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X10),s('type.hreal.hreal',X11))))) | ~p(s(bool,X8)) | ~p(s(bool,X9))) ) | $spl86),
  inference(superposition,[],[f8152,f65])).
fof(f8336,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X12) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X14),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X14),s('type.hreal.hreal',X15))))) | ~p(s(bool,X13)) | ~p(s(bool,X12))) ) | $spl86),
  inference(superposition,[],[f8152,f65])).
fof(f18384,plain,(
  ( ! [X97,X98] : (s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X97),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X97),s('type.hreal.hreal',X98)))))) ) | $spl262),
  inference(cnf_transformation,[],[f18384_D])).
fof(f18384_D,plain,(
  ( ! [X97,X98] : (s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X97),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X97),s('type.hreal.hreal',X98)))))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f18385,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'F_0') = s(bool,X84) | s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X82),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X82),s('type.hreal.hreal',X83))))) | s(bool,X81) = s(bool,X84)) ) | ($spl6 | $spl86)),
  inference(forward_demodulation,[],[f8368,f113])).
fof(f8368,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X82),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X82),s('type.hreal.hreal',X83))))) | s(bool,X84) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X81) = s(bool,X84)) ) | ($spl6 | $spl86)),
  inference(trivial_inequality_removal,[],[f8353])).
fof(f8353,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,X81) != s(bool,X81) | s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X82),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X82),s('type.hreal.hreal',X83))))) | s(bool,X84) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X81) = s(bool,X84)) ) | ($spl6 | $spl86)),
  inference(superposition,[],[f8152,f5704])).
fof(f8365,plain,(
  ( ! [X94,X95,X93] : (s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X94),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X94),s('type.hreal.hreal',X95))))) | s(bool,'T_0') = s(bool,X93)) ) | $spl86),
  inference(trivial_inequality_removal,[],[f8356])).
fof(f8356,plain,(
  ( ! [X94,X95,X93] : (s(bool,X93) != s(bool,X93) | s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X94),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X94),s('type.hreal.hreal',X95))))) | s(bool,'T_0') = s(bool,X93)) ) | $spl86),
  inference(superposition,[],[f8152,f68])).
fof(f18383,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X97),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X97),s('type.hreal.hreal',X98))))) | s(bool,'F_0') != s(bool,X96) | s(bool,'F_0') = s(bool,X99) | s(bool,X96) = s(bool,X99)) ) | ($spl6 | $spl86)),
  inference(forward_demodulation,[],[f8357,f113])).
fof(f8357,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'F_0') != s(bool,X96) | s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X97),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X97),s('type.hreal.hreal',X98))))) | s(bool,'F_0') = s(bool,X99) | s(bool,X96) = s(bool,X99)) ) | ($spl6 | $spl86)),
  inference(superposition,[],[f8152,f5704])).
fof(f8479,plain,(
  ( ! [X64,X65] : (s(bool,X65) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X64)) ) | ($spl3 | $spl6)),
  inference(trivial_inequality_removal,[],[f8462])).
fof(f8462,plain,(
  ( ! [X64,X65] : (s(bool,X64) != s(bool,X64) | s(bool,X65) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') = s(bool,X64)) ) | ($spl3 | $spl6)),
  inference(superposition,[],[f6875,f68])).
fof(f6875,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,X5) | s(bool,X6) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X5) = s(bool,X6)) ) | ($spl3 | $spl6)),
  inference(superposition,[],[f130,f5704])).
fof(f130,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(subsumption_resolution,[],[f124,f110])).
fof(f124,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | $spl3),
  inference(superposition,[],[f110,f68])).
fof(f18377,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal',X89))) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X89),s('type.hreal.hreal',X91))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X89,X88)),s('type.hreal.hreal',X91))))) | s(bool,'F_0') = s(bool,X90)) ) | $spl6),
  inference(forward_demodulation,[],[f18376,f113])).
fof(f18376,plain,(
  ( ! [X90,X88,X91,X89] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X89),s('type.hreal.hreal',X91))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X89,X88)),s('type.hreal.hreal',X91))))) | s(bool,'F_0') = s(bool,X90) | s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal',X89)))) ) | $spl6),
  inference(subsumption_resolution,[],[f8797,f4678])).
fof(f8797,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X90)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X89),s('type.hreal.hreal',X91))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X89,X88)),s('type.hreal.hreal',X91))))) | s(bool,'F_0') = s(bool,X90) | s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal',X89)))) ) | $spl6),
  inference(superposition,[],[f5121,f5704])).
fof(f18371,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119))) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X119),s('type.hreal.hreal',X121))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X119,X118)),s('type.hreal.hreal',X121))))) | s(bool,'F_0') = s(bool,X120)) ) | $spl6),
  inference(forward_demodulation,[],[f18370,f113])).
fof(f18370,plain,(
  ( ! [X121,X118,X120,X119] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X119),s('type.hreal.hreal',X121))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X119,X118)),s('type.hreal.hreal',X121))))) | s(bool,'F_0') = s(bool,X120) | s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119)))) ) | $spl6),
  inference(subsumption_resolution,[],[f13572,f4678])).
fof(f13572,plain,(
  ( ! [X121,X118,X120,X119] : (~p(s(bool,X120)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X119),s('type.hreal.hreal',X121))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X119,X118)),s('type.hreal.hreal',X121))))) | s(bool,'F_0') = s(bool,X120) | s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119)))) ) | $spl6),
  inference(superposition,[],[f5121,f5704])).
fof(f18368,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119))) | s('type.hreal.hreal',sK3(X119,X118)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X121),s('type.hreal.hreal',X119)),X121),X118)) | s(bool,'F_0') = s(bool,X120)) ) | $spl6),
  inference(forward_demodulation,[],[f18367,f113])).
fof(f18367,plain,(
  ( ! [X121,X118,X120,X119] : (s('type.hreal.hreal',sK3(X119,X118)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X121),s('type.hreal.hreal',X119)),X121),X118)) | s(bool,'F_0') = s(bool,X120) | s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119)))) ) | $spl6),
  inference(subsumption_resolution,[],[f15241,f4678])).
fof(f15241,plain,(
  ( ! [X121,X118,X120,X119] : (~p(s(bool,X120)) | s('type.hreal.hreal',sK3(X119,X118)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X121),s('type.hreal.hreal',X119)),X121),X118)) | s(bool,'F_0') = s(bool,X120) | s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119)))) ) | $spl6),
  inference(superposition,[],[f14976,f5704])).
fof(f18361,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115))),s('type.hreal.hreal',X116))) | s('type.hreal.hreal',X116) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal',sK3(X116,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115)))))))) | s(bool,'F_0') = s(bool,X117)) ) | $spl6),
  inference(forward_demodulation,[],[f18360,f113])).
fof(f18360,plain,(
  ( ! [X116,X114,X117,X115] : (s('type.hreal.hreal',X116) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal',sK3(X116,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115)))))))) | s(bool,'F_0') = s(bool,X117) | s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115))),s('type.hreal.hreal',X116)))) ) | $spl6),
  inference(subsumption_resolution,[],[f17296,f4678])).
fof(f17296,plain,(
  ( ! [X116,X114,X117,X115] : (~p(s(bool,X117)) | s('type.hreal.hreal',X116) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal',sK3(X116,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115)))))))) | s(bool,'F_0') = s(bool,X117) | s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115))),s('type.hreal.hreal',X116)))) ) | $spl6),
  inference(superposition,[],[f5083,f5704])).
fof(f18000,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,X4) | ~p(s(bool,X5))) ) | $spl4),
  inference(forward_demodulation,[],[f17999,f112])).
fof(f17999,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X4) | ~p(s(bool,X5))) ) | $spl4),
  inference(subsumption_resolution,[],[f17659,f4678])).
fof(f17659,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X4) | ~p(s(bool,X4)) | ~p(s(bool,X5))) ) | $spl4),
  inference(superposition,[],[f5058,f65])).
fof(f18002,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X6) | ~p(s(bool,X7))) ) | $spl4),
  inference(forward_demodulation,[],[f18001,f112])).
fof(f18001,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X6) | ~p(s(bool,X7))) ) | $spl4),
  inference(subsumption_resolution,[],[f17660,f4678])).
fof(f17660,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl4),
  inference(superposition,[],[f5058,f65])).
fof(f18359,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,X9) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,X8)) ) | $spl4),
  inference(forward_demodulation,[],[f18358,f112])).
fof(f18358,plain,(
  ( ! [X8,X9] : (s(bool,X9) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,X8)) ) | $spl4),
  inference(subsumption_resolution,[],[f17733,f4678])).
fof(f17733,plain,(
  ( ! [X8,X9] : (s(bool,X9) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | ~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'F_0') = s(bool,X8)) ) | $spl4),
  inference(superposition,[],[f65,f5058])).
fof(f18357,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,X11) | ~p(s(bool,X11)) | s(bool,'F_0') = s(bool,X10)) ) | $spl4),
  inference(forward_demodulation,[],[f18356,f112])).
fof(f18356,plain,(
  ( ! [X10,X11] : (s(bool,X11) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | ~p(s(bool,X11)) | s(bool,'F_0') = s(bool,X10)) ) | $spl4),
  inference(subsumption_resolution,[],[f17734,f4678])).
fof(f17734,plain,(
  ( ! [X10,X11] : (s(bool,X11) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | ~p(s(bool,X11)) | ~p(s(bool,X10)) | s(bool,'F_0') = s(bool,X10)) ) | $spl4),
  inference(superposition,[],[f65,f5058])).
fof(f18107,plain,(
  ( ! [X5] : (s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X5),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))),X5),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))))) ) | $spl226),
  inference(cnf_transformation,[],[f18107_D])).
fof(f18107_D,plain,(
  ( ! [X5] : (s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X5),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))),X5),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f18352,plain,(
  ( ! [X4,X5] : (s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X5),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))),X5),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))) | s(bool,'F_0') = s(bool,X4)) ) | $spl4),
  inference(subsumption_resolution,[],[f17854,f4678])).
fof(f17854,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X5),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))),X5),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))) | s(bool,'F_0') = s(bool,X4)) ) | $spl4),
  inference(superposition,[],[f14976,f5058])).
fof(f18351,plain,(
  ( ! [X7] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK2),s('type.hreal.hreal',X7))))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))),s('type.hreal.hreal',X7)))))) ) | $spl258),
  inference(cnf_transformation,[],[f18351_D])).
fof(f18351_D,plain,(
  ( ! [X7] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK2),s('type.hreal.hreal',X7))))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))),s('type.hreal.hreal',X7)))))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f18350,plain,(
  ( ! [X6,X7] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK2),s('type.hreal.hreal',X7))))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))),s('type.hreal.hreal',X7))))) | s(bool,'F_0') = s(bool,X6)) ) | $spl4),
  inference(forward_demodulation,[],[f18349,f59])).
fof(f18349,plain,(
  ( ! [X6,X7] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))),s('type.hreal.hreal',X7))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))),s('type.hreal.hreal',X7))))) | s(bool,'F_0') = s(bool,X6)) ) | $spl4),
  inference(subsumption_resolution,[],[f17855,f4678])).
fof(f17855,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))),s('type.hreal.hreal',X7))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))),s('type.hreal.hreal',X7))))) | s(bool,'F_0') = s(bool,X6)) ) | $spl4),
  inference(superposition,[],[f5121,f5058])).
fof(f18114,plain,(
  s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))))) | $spl230),
  inference(cnf_transformation,[],[f18114_D])).
fof(f18114_D,plain,(
  s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))))) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f18348,plain,(
  ( ! [X9] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))))) | s(bool,'F_0') = s(bool,X9)) ) | $spl4),
  inference(subsumption_resolution,[],[f17857,f4678])).
fof(f17857,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))))) | s(bool,'F_0') = s(bool,X9)) ) | $spl4),
  inference(superposition,[],[f60,f5058])).
fof(f18347,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,X25) | ~p(s(bool,X25)) | s(bool,'F_0') = s(bool,X24)) ) | $spl4),
  inference(forward_demodulation,[],[f18346,f112])).
fof(f18346,plain,(
  ( ! [X24,X25] : (s(bool,X25) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | ~p(s(bool,X25)) | s(bool,'F_0') = s(bool,X24)) ) | $spl4),
  inference(subsumption_resolution,[],[f17865,f4678])).
fof(f17865,plain,(
  ( ! [X24,X25] : (~p(s(bool,X24)) | s(bool,X25) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | ~p(s(bool,X25)) | s(bool,'F_0') = s(bool,X24)) ) | $spl4),
  inference(superposition,[],[f65,f5058])).
fof(f18345,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,X27) | ~p(s(bool,X27)) | s(bool,'F_0') = s(bool,X26)) ) | $spl4),
  inference(forward_demodulation,[],[f18344,f112])).
fof(f18344,plain,(
  ( ! [X26,X27] : (s(bool,X27) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | ~p(s(bool,X27)) | s(bool,'F_0') = s(bool,X26)) ) | $spl4),
  inference(subsumption_resolution,[],[f17866,f4678])).
fof(f17866,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | s(bool,X27) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | ~p(s(bool,X27)) | s(bool,'F_0') = s(bool,X26)) ) | $spl4),
  inference(superposition,[],[f65,f5058])).
fof(f18177,plain,(
  ( ! [X191,X192] : (s(bool,X191) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X191) = s(bool,X192) | s(bool,'F_0') != s(bool,X192)) ) | $spl240),
  inference(cnf_transformation,[],[f18177_D])).
fof(f18177_D,plain,(
  ( ! [X191,X192] : (s(bool,X191) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X191) = s(bool,X192) | s(bool,'F_0') != s(bool,X192)) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f18343,plain,(
  ( ! [X149,X150,X148] : (s(bool,X149) = s(bool,X150) | s(bool,X150) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') != s(bool,X149) | s(bool,'F_0') = s(bool,X148)) ) | ($spl4 | $spl32)),
  inference(subsumption_resolution,[],[f17925,f4678])).
fof(f17925,plain,(
  ( ! [X149,X150,X148] : (~p(s(bool,X148)) | s(bool,X149) = s(bool,X150) | s(bool,X150) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') != s(bool,X149) | s(bool,'F_0') = s(bool,X148)) ) | ($spl4 | $spl32)),
  inference(superposition,[],[f2853,f5058])).
fof(f3747,plain,(
  ( ! [X68,X69,X67] : (s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') != s(bool,X67) | p(s(bool,X69))) ) | $spl32),
  inference(resolution,[],[f2853,f72])).
fof(f3748,plain,(
  ( ! [X70,X72,X71] : (s(bool,X70) = s(bool,X71) | s(bool,'T_0') = s(bool,X71) | s(bool,'F_0') != s(bool,X70) | p(s(bool,X72))) ) | $spl32),
  inference(resolution,[],[f2853,f75])).
fof(f3750,plain,(
  ( ! [X76,X77,X75] : (s(bool,X75) = s(bool,X76) | s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') != s(bool,X75) | ~p(s(bool,X77))) ) | $spl32),
  inference(resolution,[],[f2853,f79])).
fof(f3751,plain,(
  ( ! [X80,X78,X79] : (s(bool,X78) = s(bool,X79) | s(bool,'T_0') = s(bool,X79) | s(bool,'F_0') != s(bool,X78) | p(s(bool,X80))) ) | $spl32),
  inference(resolution,[],[f2853,f81])).
fof(f3752,plain,(
  ( ! [X83,X81,X82] : (s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'F_0') != s(bool,X81) | ~p(s(bool,X83))) ) | $spl32),
  inference(resolution,[],[f2853,f82])).
fof(f82,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f55])).
fof(f3777,plain,(
  ( ! [X78,X76,X77] : (~p(s(bool,X76)) | s(bool,X77) = s(bool,X78) | s(bool,'F_0') = s(bool,X78) | s(bool,'F_0') != s(bool,X77) | s(bool,'T_0') = s(bool,X76)) ) | $spl32),
  inference(superposition,[],[f2853,f68])).
fof(f18304,plain,(
  ( ! [X204,X202] : (s(bool,X202) = s(bool,X204) | s(bool,'F_0') = s(bool,X202) | s(bool,'F_0') = s(bool,X204)) ) | $spl250),
  inference(cnf_transformation,[],[f18304_D])).
fof(f18304_D,plain,(
  ( ! [X204,X202] : (s(bool,X202) = s(bool,X204) | s(bool,'F_0') = s(bool,X202) | s(bool,'F_0') = s(bool,X204)) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f18290,plain,(
  ( ! [X141] : (p(s(bool,X141)) | s(bool,'T_0') = s(bool,X141)) ) | $spl246),
  inference(cnf_transformation,[],[f18290_D])).
fof(f18290_D,plain,(
  ( ! [X141] : (p(s(bool,X141)) | s(bool,'T_0') = s(bool,X141)) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f5685,plain,(
  ( ! [X78,X76,X77] : (p(s(bool,X76)) | s(bool,'F_0') = s(bool,X77) | s(bool,'F_0') = s(bool,X78) | s(bool,X77) = s(bool,X78) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f4997,f68])).
fof(f5976,plain,(
  ( ! [X66,X64,X65] : (s(bool,X64) = s(bool,X65) | s(bool,'T_0') = s(bool,X65) | s(bool,'F_0') != s(bool,X64) | p(s(bool,X66))) ) | $spl32),
  inference(resolution,[],[f2853,f72])).
fof(f5977,plain,(
  ( ! [X68,X69,X67] : (s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') != s(bool,X67) | p(s(bool,X69))) ) | $spl32),
  inference(resolution,[],[f2853,f75])).
fof(f5978,plain,(
  ( ! [X70,X72,X71] : (s(bool,X70) = s(bool,X71) | s(bool,'T_0') = s(bool,X71) | s(bool,'F_0') != s(bool,X70) | ~p(s(bool,X72))) ) | $spl32),
  inference(resolution,[],[f2853,f79])).
fof(f5979,plain,(
  ( ! [X74,X75,X73] : (s(bool,X73) = s(bool,X74) | s(bool,'T_0') = s(bool,X74) | s(bool,'F_0') != s(bool,X73) | p(s(bool,X75))) ) | $spl32),
  inference(resolution,[],[f2853,f81])).
fof(f5980,plain,(
  ( ! [X78,X76,X77] : (s(bool,X76) = s(bool,X77) | s(bool,'T_0') = s(bool,X77) | s(bool,'F_0') != s(bool,X76) | ~p(s(bool,X78))) ) | $spl32),
  inference(resolution,[],[f2853,f82])).
fof(f7481,plain,(
  ( ! [X59,X60,X58] : (s(bool,'F_0') = s(bool,X58) | s(bool,X59) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X60) | s(bool,X58) = s(bool,X60)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f6755])).
fof(f6755,plain,(
  ( ! [X59,X60,X58] : (s(bool,'F_0') = s(bool,X58) | s(bool,X59) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X59) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X60) | s(bool,X58) = s(bool,X60)) ) | $spl6),
  inference(superposition,[],[f5704,f5704])).
fof(f7478,plain,(
  ( ! [X50,X51,X49] : (s(bool,X50) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X51) | s(bool,X49) = s(bool,X51) | s(bool,'F_0') = s(bool,X49)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f6776])).
fof(f6776,plain,(
  ( ! [X50,X51,X49] : (s(bool,X50) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X51) | s(bool,X49) = s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,X50) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(superposition,[],[f5704,f5704])).
fof(f7477,plain,(
  ( ! [X57,X58,X56] : (s(bool,'F_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X57) | s(bool,X56) = s(bool,X57) | s(bool,X58) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f6804])).
fof(f6804,plain,(
  ( ! [X57,X58,X56] : (s(bool,'F_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X57) | s(bool,X56) = s(bool,X57) | s(bool,X58) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X58) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(superposition,[],[f5704,f5704])).
fof(f7471,plain,(
  ( ! [X47,X48,X46] : (s(bool,X47) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X46) | s(bool,'F_0') = s(bool,X48) | s(bool,X46) = s(bool,X48)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f6823])).
fof(f6823,plain,(
  ( ! [X47,X48,X46] : (s(bool,X47) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X46) | s(bool,X47) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X48) | s(bool,X46) = s(bool,X48)) ) | $spl6),
  inference(superposition,[],[f5704,f5704])).
fof(f7465,plain,(
  ( ! [X57,X56,X55] : (s(bool,X56) = s(bool,X57) | s(bool,'F_0') = s(bool,X57) | s(bool,X55) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X56)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f6859])).
fof(f6859,plain,(
  ( ! [X57,X56,X55] : (s(bool,X56) = s(bool,X57) | s(bool,'F_0') = s(bool,X57) | s(bool,X55) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X56) | s(bool,X55) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(superposition,[],[f5704,f5704])).
fof(f11332,plain,(
  ( ! [X269,X266] : (p(s(bool,X266)) | s(bool,X266) = s(bool,X269)) ) | $spl168),
  inference(cnf_transformation,[],[f11332_D])).
fof(f11332_D,plain,(
  ( ! [X269,X266] : (p(s(bool,X266)) | s(bool,X266) = s(bool,X269)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f18309,plain,(
  ( ! [X235,X233,X234,X232] : (p(s(bool,X232)) | s(bool,'F_0') = s(bool,X233) | s(bool,'F_0') = s(bool,X234) | s(bool,X233) = s(bool,X234) | s(bool,X232) = s(bool,X235)) ) | ($spl6 | $spl32)),
  inference(subsumption_resolution,[],[f18308,f2819])).
fof(f18308,plain,(
  ( ! [X235,X233,X234,X232] : (s(bool,'F_0') = s(bool,X235) | p(s(bool,X232)) | s(bool,'F_0') = s(bool,X233) | s(bool,'F_0') = s(bool,X234) | s(bool,X233) = s(bool,X234) | s(bool,X232) = s(bool,X235)) ) | $spl6),
  inference(forward_demodulation,[],[f6946,f113])).
fof(f6946,plain,(
  ( ! [X235,X233,X234,X232] : (p(s(bool,X232)) | s(bool,'F_0') = s(bool,X233) | s(bool,'F_0') = s(bool,X234) | s(bool,X233) = s(bool,X234) | s(bool,X235) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X232) = s(bool,X235)) ) | $spl6),
  inference(superposition,[],[f4997,f5704])).
fof(f18303,plain,(
  ( ! [X204,X202,X203] : (s(bool,'F_0') = s(bool,X203) | s(bool,X202) = s(bool,X204) | s(bool,'F_0') = s(bool,X204) | s(bool,'F_0') = s(bool,X202)) ) | $spl6),
  inference(subsumption_resolution,[],[f18302,f4678])).
fof(f18302,plain,(
  ( ! [X204,X202,X203] : (s(bool,'F_0') = s(bool,X203) | ~p(s(bool,X203)) | s(bool,X202) = s(bool,X204) | s(bool,'F_0') = s(bool,X204) | s(bool,'F_0') = s(bool,X202)) ) | $spl6),
  inference(forward_demodulation,[],[f7255,f113])).
fof(f7255,plain,(
  ( ! [X204,X202,X203] : (~p(s(bool,X203)) | s(bool,X202) = s(bool,X204) | s(bool,'F_0') = s(bool,X204) | s(bool,'F_0') = s(bool,X202) | s(bool,X203) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(superposition,[],[f4764,f5704])).
fof(f133,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f133_D])).
fof(f133_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f18300,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X0) = s(bool,X2) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl3 | $spl6 | $spl32)),
  inference(subsumption_resolution,[],[f18299,f2819])).
fof(f18299,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') != s(bool,X1) | s(bool,X0) = s(bool,X2) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl3 | $spl6)),
  inference(forward_demodulation,[],[f8440,f113])).
fof(f8440,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X2) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X0) = s(bool,X2) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl3 | $spl6)),
  inference(superposition,[],[f6875,f62])).
fof(f18298,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X4) | s(bool,X3) = s(bool,X5) | p(s(bool,X3)) | p(s(bool,X4))) ) | ($spl3 | $spl6 | $spl32)),
  inference(subsumption_resolution,[],[f18297,f2819])).
fof(f18297,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') = s(bool,X5) | s(bool,'T_0') != s(bool,X4) | s(bool,X3) = s(bool,X5) | p(s(bool,X3)) | p(s(bool,X4))) ) | ($spl3 | $spl6)),
  inference(forward_demodulation,[],[f8441,f113])).
fof(f8441,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X4) | s(bool,X5) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X3) = s(bool,X5) | p(s(bool,X3)) | p(s(bool,X4))) ) | ($spl3 | $spl6)),
  inference(superposition,[],[f6875,f62])).
fof(f7504,plain,(
  ( ! [X50,X48] : (s(bool,X50) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X48) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X48) = s(bool,X50)) ) | $spl76),
  inference(cnf_transformation,[],[f7504_D])).
fof(f7504_D,plain,(
  ( ! [X50,X48] : (s(bool,X50) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X48) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X48) = s(bool,X50)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f18296,plain,(
  ( ! [X57,X56,X55] : (s(bool,X57) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X55) = s(bool,X57) | s(bool,'F_0') = s(bool,X56) | s(bool,X55) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f8459,f68])).
fof(f8459,plain,(
  ( ! [X57,X56,X55] : (s(bool,'T_0') != s(bool,X56) | s(bool,X57) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X55) = s(bool,X57) | s(bool,'F_0') = s(bool,X56) | s(bool,X55) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl3 | $spl6)),
  inference(superposition,[],[f6875,f5704])).
fof(f18295,plain,(
  ( ! [X79,X77] : (s(bool,'T_0') != s(bool,X77) | s(bool,X77) = s(bool,X79) | s(bool,'F_0') = s(bool,X79)) ) | $spl248),
  inference(cnf_transformation,[],[f18295_D])).
fof(f18295_D,plain,(
  ( ! [X79,X77] : (s(bool,'T_0') != s(bool,X77) | s(bool,X77) = s(bool,X79) | s(bool,'F_0') = s(bool,X79)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f8478,plain,(
  ( ! [X78,X79,X77] : (s(bool,'T_0') != s(bool,X77) | s(bool,X78) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X79) | s(bool,X77) = s(bool,X79)) ) | ($spl3 | $spl6)),
  inference(duplicate_literal_removal,[],[f8467])).
fof(f8467,plain,(
  ( ! [X78,X79,X77] : (s(bool,'T_0') != s(bool,X77) | s(bool,X78) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X78) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X79) | s(bool,X77) = s(bool,X79)) ) | ($spl3 | $spl6)),
  inference(superposition,[],[f6875,f5704])).
fof(f18289,plain,(
  ( ! [X146,X145] : (s(bool,'F_0') = s(bool,X146) | s(bool,'F_0') != s(bool,X145) | s(bool,X145) != s(bool,X146)) ) | $spl244),
  inference(cnf_transformation,[],[f18289_D])).
fof(f18289_D,plain,(
  ( ! [X146,X145] : (s(bool,'F_0') = s(bool,X146) | s(bool,'F_0') != s(bool,X145) | s(bool,X145) != s(bool,X146)) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f18292,plain,(
  ( ! [X132,X130,X131,X129] : (p(s(bool,X129)) | s(bool,'F_0') != s(bool,X130) | s(bool,'F_0') = s(bool,X131) | s(bool,X130) != s(bool,X131) | s(bool,X129) = s(bool,X132)) ) | ($spl6 | $spl32 | $spl204)),
  inference(subsumption_resolution,[],[f18291,f2819])).
fof(f18291,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,'F_0') = s(bool,X132) | p(s(bool,X129)) | s(bool,'F_0') != s(bool,X130) | s(bool,'F_0') = s(bool,X131) | s(bool,X130) != s(bool,X131) | s(bool,X129) = s(bool,X132)) ) | ($spl6 | $spl204)),
  inference(forward_demodulation,[],[f14796,f113])).
fof(f14796,plain,(
  ( ! [X132,X130,X131,X129] : (p(s(bool,X129)) | s(bool,'F_0') != s(bool,X130) | s(bool,'F_0') = s(bool,X131) | s(bool,X130) != s(bool,X131) | s(bool,X132) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X129) = s(bool,X132)) ) | ($spl6 | $spl204)),
  inference(superposition,[],[f14707,f5704])).
fof(f14799,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,X141)) | s(bool,'F_0') != s(bool,X142) | s(bool,'F_0') = s(bool,X143) | s(bool,X142) != s(bool,X143) | s(bool,'T_0') = s(bool,X141)) ) | $spl204),
  inference(superposition,[],[f14707,f68])).
fof(f18288,plain,(
  ( ! [X146,X144,X147,X145] : (s(bool,'F_0') = s(bool,X146) | p(s(bool,X144)) | s(bool,'F_0') != s(bool,X145) | s(bool,X145) != s(bool,X146) | s(bool,X144) = s(bool,X147)) ) | ($spl6 | $spl32 | $spl204)),
  inference(forward_demodulation,[],[f14831,f113])).
fof(f14831,plain,(
  ( ! [X146,X144,X147,X145] : (p(s(bool,X144)) | s(bool,'F_0') != s(bool,X145) | s(bool,X146) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X145) != s(bool,X146) | s(bool,X144) = s(bool,X147)) ) | ($spl6 | $spl32 | $spl204)),
  inference(subsumption_resolution,[],[f14800,f2819])).
fof(f14800,plain,(
  ( ! [X146,X144,X147,X145] : (p(s(bool,X144)) | s(bool,'F_0') != s(bool,X145) | s(bool,X146) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X145) != s(bool,X146) | s(bool,'F_0') = s(bool,X147) | s(bool,X144) = s(bool,X147)) ) | ($spl6 | $spl204)),
  inference(superposition,[],[f14707,f5704])).
fof(f18268,plain,(
  $false | ($spl4 | $spl6 | $spl24)),
  inference(subsumption_resolution,[],[f18213,f18105])).
fof(f18105,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | ($spl4 | $spl6 | $spl24)),
  inference(subsumption_resolution,[],[f18104,f17493])).
fof(f17493,plain,(
  ( ! [X0,X1] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',sK2))) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',X1)))))) ) | $spl6),
  inference(equality_resolution,[],[f502])).
fof(f502,plain,(
  ( ! [X6,X7,X5] : (s('type.hreal.hreal',sK2) != s('type.hreal.hreal',X6) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X7),s('type.hreal.hreal',X6))) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X7),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',X5)))))) ) | $spl6),
  inference(superposition,[],[f401,f57])).
fof(f401,plain,(
  ( ! [X0] : (s('type.hreal.hreal',sK2) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',X0)))) ) | $spl6),
  inference(subsumption_resolution,[],[f377,f84])).
fof(f377,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.hreal.hreal',sK2) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',X0)))) ) | $spl6),
  inference(superposition,[],[f61,f113])).
fof(f18104,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))))))))) ) | ($spl4 | $spl24)),
  inference(subsumption_resolution,[],[f17853,f436])).
fof(f18213,plain,(
  ( ! [X41] : (p(s(bool,X41))) ) | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f18105,f81])).
fof(f18267,plain,(
  $false | ($spl4 | $spl6 | $spl24)),
  inference(subsumption_resolution,[],[f18211,f18105])).
fof(f18211,plain,(
  ( ! [X39] : (p(s(bool,X39))) ) | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f18105,f75])).
fof(f18266,plain,(
  $false | ($spl4 | $spl6 | $spl24)),
  inference(subsumption_resolution,[],[f18210,f18105])).
fof(f18210,plain,(
  ( ! [X38] : (p(s(bool,X38))) ) | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f18105,f72])).
fof(f18215,plain,(
  $false | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f18105,f85])).
fof(f18216,plain,(
  $false | ($spl4 | $spl6 | $spl24)),
  inference(resolution,[],[f18105,f86])).
fof(f18189,plain,(
  ( ! [X259,X258] : (s(bool,X258) != s(bool,X259) | s(bool,X258) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') != s(bool,X259)) ) | $spl242),
  inference(cnf_transformation,[],[f18189_D])).
fof(f18189_D,plain,(
  ( ! [X259,X258] : (s(bool,X258) != s(bool,X259) | s(bool,X258) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') != s(bool,X259)) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f18188,plain,(
  ( ! [X257,X259,X258] : (s(bool,'T_0') != s(bool,X257) | s(bool,X258) != s(bool,X259) | s(bool,'F_0') != s(bool,X259) | s(bool,X258) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) ) | ($spl4 | $spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f17966,f705])).
fof(f705,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,X22) | s(bool,'F_0') != s(bool,X22)) ) | ($spl16 | $spl24)),
  inference(resolution,[],[f538,f84])).
fof(f538,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,X14) != s(bool,X15) | s(bool,'T_0') != s(bool,X15)) ) | ($spl16 | $spl24)),
  inference(resolution,[],[f458,f133])).
fof(f458,plain,(
  ( ! [X4,X5] : (~p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | p(s(bool,X4))) ) | $spl24),
  inference(subsumption_resolution,[],[f440,f70])).
fof(f440,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | ~p(s(bool,X5)) | p(s(bool,'F_0')) | p(s(bool,X4))) ) | $spl24),
  inference(superposition,[],[f436,f62])).
fof(f17966,plain,(
  ( ! [X257,X259,X258] : (s(bool,'T_0') != s(bool,X257) | s(bool,X258) != s(bool,X259) | s(bool,'F_0') != s(bool,X259) | s(bool,X258) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X257)) ) | ($spl4 | $spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8231,f5058])).
fof(f8231,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') != s(bool,X17) | s(bool,X18) != s(bool,X19) | s(bool,'F_0') != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | ($spl6 | $spl16 | $spl24)),
  inference(resolution,[],[f8131,f133])).
fof(f8131,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | s(bool,X17) != s(bool,X18) | s(bool,X18) != s(bool,X19) | s(bool,'F_0') != s(bool,X19)) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f8130,f113])).
fof(f8130,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,X18) | s(bool,X19) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X18) != s(bool,X19) | ~p(s(bool,X17))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f8072,f79])).
fof(f8072,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,X18) | s(bool,X19) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X18) != s(bool,X19) | ~p(s(bool,'T_0')) | ~p(s(bool,X17))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3455,f65])).
fof(f3455,plain,(
  ( ! [X83,X82] : (s(bool,'T_0') != s(bool,X82) | s(bool,X83) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X82) != s(bool,X83)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3433,f84])).
fof(f3433,plain,(
  ( ! [X83,X82] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,X82) | s(bool,X83) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X82) != s(bool,X83)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f703,f113])).
fof(f703,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X19)) | s(bool,'T_0') != s(bool,X18) | s(bool,X17) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | ($spl16 | $spl24)),
  inference(resolution,[],[f538,f458])).
fof(f17632,plain,(
  ( ! [X37,X38] : (s(bool,X37) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X38) | s(bool,X37) = s(bool,X38)) ) | $spl222),
  inference(cnf_transformation,[],[f17632_D])).
fof(f17632_D,plain,(
  ( ! [X37,X38] : (s(bool,X37) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X38) | s(bool,X37) = s(bool,X38)) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f18180,plain,(
  ( ! [X229,X227,X228] : (s(bool,'T_0') != s(bool,X227) | s(bool,'F_0') = s(bool,X228) | s(bool,X229) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X228) = s(bool,X229)) ) | ($spl4 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f17954,f705])).
fof(f17954,plain,(
  ( ! [X229,X227,X228] : (s(bool,'T_0') != s(bool,X227) | s(bool,'F_0') = s(bool,X228) | s(bool,X229) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X228) = s(bool,X229) | s(bool,'F_0') = s(bool,X227)) ) | ($spl4 | $spl16 | $spl24)),
  inference(superposition,[],[f5004,f5058])).
fof(f5004,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') = s(bool,X17) | s(bool,X18) != s(bool,X19) | s(bool,X17) = s(bool,X18)) ) | ($spl16 | $spl24)),
  inference(resolution,[],[f4764,f538])).
fof(f17653,plain,(
  ( ! [X47,X48] : (s(bool,X47) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'T_0') = s(bool,X48) | s(bool,X47) = s(bool,X48)) ) | $spl224),
  inference(cnf_transformation,[],[f17653_D])).
fof(f17653_D,plain,(
  ( ! [X47,X48] : (s(bool,X47) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'T_0') = s(bool,X48) | s(bool,X47) = s(bool,X48)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f17972,plain,(
  ( ! [X225,X226,X224] : (s(bool,'F_0') = s(bool,X224) | s(bool,X225) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X225) = s(bool,X226) | s(bool,'T_0') = s(bool,X226)) ) | ($spl4 | $spl24)),
  inference(duplicate_literal_removal,[],[f17953])).
fof(f17953,plain,(
  ( ! [X225,X226,X224] : (s(bool,'F_0') = s(bool,X224) | s(bool,X225) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X225) = s(bool,X226) | s(bool,'T_0') = s(bool,X226) | s(bool,'F_0') = s(bool,X224)) ) | ($spl4 | $spl24)),
  inference(superposition,[],[f5003,f5058])).
fof(f5003,plain,(
  ( ! [X14,X15,X16] : (s(bool,X14) = s(bool,X15) | s(bool,'F_0') = s(bool,X14) | s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl24),
  inference(resolution,[],[f4764,f461])).
fof(f461,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl24),
  inference(resolution,[],[f457,f62])).
fof(f457,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl24),
  inference(trivial_inequality_removal,[],[f445])).
fof(f445,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl24),
  inference(superposition,[],[f436,f68])).
fof(f18176,plain,(
  ( ! [X191,X192,X190] : (s(bool,'T_0') != s(bool,X190) | s(bool,X191) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') != s(bool,X192) | s(bool,X191) = s(bool,X192)) ) | ($spl4 | $spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f17940,f705])).
fof(f17940,plain,(
  ( ! [X191,X192,X190] : (s(bool,'T_0') != s(bool,X190) | s(bool,X191) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') != s(bool,X192) | s(bool,X191) = s(bool,X192) | s(bool,'F_0') = s(bool,X190)) ) | ($spl4 | $spl16 | $spl32)),
  inference(superposition,[],[f3736,f5058])).
fof(f3736,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X22) | s(bool,X21) = s(bool,X22) | s(bool,'F_0') != s(bool,X20) | s(bool,X20) = s(bool,X21)) ) | ($spl16 | $spl32)),
  inference(resolution,[],[f2853,f133])).
fof(f18175,plain,(
  ( ! [X189] : (s(bool,'F_0') != s(bool,X189) | s(bool,X189) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) ) | $spl238),
  inference(cnf_transformation,[],[f18175_D])).
fof(f18175_D,plain,(
  ( ! [X189] : (s(bool,'F_0') != s(bool,X189) | s(bool,X189) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f18174,plain,(
  ( ! [X189,X188] : (s(bool,'F_0') != s(bool,X189) | s(bool,'T_0') != s(bool,X188) | s(bool,X189) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) ) | ($spl4 | $spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f18173,f113])).
fof(f18173,plain,(
  ( ! [X189,X188] : (s(bool,'T_0') != s(bool,X188) | s(bool,X189) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X189) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) ) | ($spl4 | $spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f17939,f705])).
fof(f17939,plain,(
  ( ! [X189,X188] : (s(bool,'T_0') != s(bool,X188) | s(bool,X189) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X189) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X188)) ) | ($spl4 | $spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3455,f5058])).
fof(f18172,plain,(
  ( ! [X185,X184] : (s(bool,X184) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'T_0') != s(bool,X185) | s(bool,X184) != s(bool,X185)) ) | $spl236),
  inference(cnf_transformation,[],[f18172_D])).
fof(f18172_D,plain,(
  ( ! [X185,X184] : (s(bool,X184) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'T_0') != s(bool,X185) | s(bool,X184) != s(bool,X185)) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f18171,plain,(
  ( ! [X185,X184,X183] : (s(bool,'T_0') != s(bool,X183) | s(bool,X184) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X184) != s(bool,X185) | s(bool,'T_0') != s(bool,X185)) ) | ($spl4 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f17937,f705])).
fof(f17937,plain,(
  ( ! [X185,X184,X183] : (s(bool,'T_0') != s(bool,X183) | s(bool,X184) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X184) != s(bool,X185) | s(bool,'T_0') != s(bool,X185) | s(bool,'F_0') = s(bool,X183)) ) | ($spl4 | $spl16 | $spl24)),
  inference(superposition,[],[f3336,f5058])).
fof(f3336,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X22) | s(bool,X21) = s(bool,X22) | s(bool,X20) != s(bool,X21) | s(bool,'T_0') != s(bool,X20)) ) | ($spl16 | $spl24)),
  inference(resolution,[],[f696,f133])).
fof(f696,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | s(bool,'T_0') != s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,X0) != s(bool,X1)) ) | ($spl16 | $spl24)),
  inference(resolution,[],[f538,f65])).
fof(f18170,plain,(
  ( ! [X182,X180,X181] : (s(bool,'T_0') != s(bool,X180) | s(bool,X181) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X181) = s(bool,X182) | s(bool,'T_0') = s(bool,X182)) ) | ($spl4 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f17936,f705])).
fof(f17936,plain,(
  ( ! [X182,X180,X181] : (s(bool,'T_0') != s(bool,X180) | s(bool,X181) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X181) = s(bool,X182) | s(bool,'T_0') = s(bool,X182) | s(bool,'F_0') = s(bool,X180)) ) | ($spl4 | $spl16 | $spl24)),
  inference(superposition,[],[f3217,f5058])).
fof(f3217,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X22) | s(bool,X21) = s(bool,X22) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl16 | $spl24)),
  inference(resolution,[],[f666,f133])).
fof(f666,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1)) ) | $spl24),
  inference(resolution,[],[f461,f65])).
fof(f18169,plain,(
  ( ! [X178,X179] : (s(bool,X178) != s(bool,X179) | s(bool,X178) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X179)) ) | $spl234),
  inference(cnf_transformation,[],[f18169_D])).
fof(f18169_D,plain,(
  ( ! [X178,X179] : (s(bool,X178) != s(bool,X179) | s(bool,X178) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X179)) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f18168,plain,(
  ( ! [X177,X178,X179] : (s(bool,'T_0') != s(bool,X177) | s(bool,X178) != s(bool,X179) | s(bool,'F_0') = s(bool,X179) | s(bool,X178) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) ) | ($spl4 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f17935,f705])).
fof(f17935,plain,(
  ( ! [X177,X178,X179] : (s(bool,'T_0') != s(bool,X177) | s(bool,X178) != s(bool,X179) | s(bool,'F_0') = s(bool,X179) | s(bool,X178) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X177)) ) | ($spl4 | $spl16)),
  inference(superposition,[],[f3156,f5058])).
fof(f3156,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X21) | s(bool,X20) != s(bool,X22) | s(bool,'F_0') = s(bool,X22) | s(bool,X20) = s(bool,X21)) ) | $spl16),
  inference(resolution,[],[f642,f133])).
fof(f642,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | s(bool,X1) = s(bool,X2) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(superposition,[],[f273,f68])).
fof(f273,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,X3) = s(bool,X4) | ~p(s(bool,X4))) ) | $spl16),
  inference(resolution,[],[f133,f65])).
fof(f18150,plain,(
  ( ! [X97,X98] : (s(bool,X97) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'T_0') = s(bool,X98) | s(bool,X97) != s(bool,X98)) ) | $spl232),
  inference(cnf_transformation,[],[f18150_D])).
fof(f18150_D,plain,(
  ( ! [X97,X98] : (s(bool,X97) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'T_0') = s(bool,X98) | s(bool,X97) != s(bool,X98)) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f18149,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,X96) | s(bool,X97) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X97) != s(bool,X98) | s(bool,'T_0') = s(bool,X98)) ) | ($spl4 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f17902,f705])).
fof(f17902,plain,(
  ( ! [X97,X98,X96] : (s(bool,'T_0') != s(bool,X96) | s(bool,X97) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X97) != s(bool,X98) | s(bool,'T_0') = s(bool,X98) | s(bool,'F_0') = s(bool,X96)) ) | ($spl4 | $spl16 | $spl24)),
  inference(superposition,[],[f701,f5058])).
fof(f701,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,X13) | s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'T_0') = s(bool,X14)) ) | ($spl16 | $spl24)),
  inference(resolution,[],[f538,f437])).
fof(f437,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f436,f68])).
fof(f18112,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))))))) ) | ($spl4 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f17856,f705])).
fof(f17856,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,X8) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))))) | s(bool,'F_0') = s(bool,X8)) ) | ($spl4 | $spl16)),
  inference(superposition,[],[f279,f5058])).
fof(f279,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X10),s('type.hreal.hreal',X11))) | s('type.hreal.hreal',X11) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X10),s('type.hreal.hreal',sK3(X11,X10))))) ) | $spl16),
  inference(resolution,[],[f133,f60])).
fof(f18111,plain,(
  ( ! [X7] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK2),s('type.hreal.hreal',X7))))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))),s('type.hreal.hreal',X7)))))))) ) | $spl228),
  inference(cnf_transformation,[],[f18111_D])).
fof(f18111_D,plain,(
  ( ! [X7] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK2),s('type.hreal.hreal',X7))))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))),s('type.hreal.hreal',X7)))))))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f18110,plain,(
  ( ! [X6,X7] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK2),s('type.hreal.hreal',X7))))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))),s('type.hreal.hreal',X7))))))) | ~p(s(bool,X6))) ) | ($spl4 | $spl24)),
  inference(forward_demodulation,[],[f18109,f59])).
fof(f18109,plain,(
  ( ! [X6,X7] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))),s('type.hreal.hreal',X7))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))),s('type.hreal.hreal',X7))))))) | ~p(s(bool,X6))) ) | ($spl4 | $spl24)),
  inference(forward_demodulation,[],[f18108,f59])).
fof(f18108,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))),s('type.hreal.hreal',X7))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))),s('type.hreal.hreal',X7)))))) ) | ($spl4 | $spl24)),
  inference(subsumption_resolution,[],[f17855,f436])).
fof(f18106,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s('type.hreal.hreal',sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1)))) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X5),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))),X5),'const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))))) ) | ($spl4 | $spl24)),
  inference(subsumption_resolution,[],[f17854,f436])).
fof(f17978,plain,(
  ( ! [X222,X223,X221] : (s(bool,X222) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X221) | s(bool,X222) = s(bool,X223) | s(bool,'F_0') = s(bool,X223)) ) | ($spl4 | $spl24)),
  inference(duplicate_literal_removal,[],[f17826])).
fof(f17826,plain,(
  ( ! [X222,X223,X221] : (s(bool,X222) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X221) | s(bool,X222) = s(bool,X223) | s(bool,'F_0') = s(bool,X223) | s(bool,'F_0') = s(bool,X221)) ) | ($spl4 | $spl24)),
  inference(superposition,[],[f5192,f5058])).
fof(f5192,plain,(
  ( ! [X39,X38,X40] : (s(bool,X38) != s(bool,X39) | s(bool,'F_0') = s(bool,X38) | s(bool,X39) = s(bool,X40) | s(bool,'F_0') = s(bool,X40)) ) | $spl24),
  inference(resolution,[],[f561,f4764])).
fof(f561,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(subsumption_resolution,[],[f550,f85])).
fof(f550,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f458,f68])).
fof(f17979,plain,(
  ( ! [X208,X206,X207] : (s(bool,X207) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X206) | s(bool,X207) = s(bool,X208) | s(bool,'T_0') = s(bool,X208)) ) | ($spl4 | $spl24)),
  inference(duplicate_literal_removal,[],[f17821])).
fof(f17821,plain,(
  ( ! [X208,X206,X207] : (s(bool,X207) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X206) | s(bool,X207) = s(bool,X208) | s(bool,'T_0') = s(bool,X208) | s(bool,'F_0') = s(bool,X206)) ) | ($spl4 | $spl24)),
  inference(superposition,[],[f5003,f5058])).
fof(f17989,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') = s(bool,X48) | s(bool,X49) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | ($spl4 | $spl24)),
  inference(duplicate_literal_removal,[],[f17723])).
fof(f17723,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') = s(bool,X48) | s(bool,'F_0') = s(bool,X48) | s(bool,X49) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | ($spl4 | $spl24)),
  inference(superposition,[],[f5058,f5003])).
fof(f17994,plain,(
  ( ! [X39,X41,X40] : (s(bool,X40) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X39) | s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) ) | ($spl4 | $spl24)),
  inference(duplicate_literal_removal,[],[f17676])).
fof(f17676,plain,(
  ( ! [X39,X41,X40] : (s(bool,X40) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X39) | s(bool,'F_0') = s(bool,X39) | s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) ) | ($spl4 | $spl24)),
  inference(superposition,[],[f5058,f5003])).
fof(f10981,plain,(
  ( ! [X154] : (s(bool,X154) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X154)) ) | $spl108),
  inference(cnf_transformation,[],[f10981_D])).
fof(f10981_D,plain,(
  ( ! [X154] : (s(bool,X154) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X154)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f17656,plain,(
  ( ! [X52,X51] : (s(bool,'F_0') = s(bool,X52) | s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) ) | ($spl6 | $spl8 | $spl24)),
  inference(forward_demodulation,[],[f17655,f113])).
fof(f17655,plain,(
  ( ! [X52,X51] : (s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X52) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) ) | ($spl6 | $spl8 | $spl24)),
  inference(subsumption_resolution,[],[f17612,f68])).
fof(f17612,plain,(
  ( ! [X52,X51] : (s(bool,'F_0') != s(bool,X51) | s(bool,'T_0') = s(bool,X51) | s(bool,X52) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X52) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) ) | ($spl6 | $spl8 | $spl24)),
  inference(superposition,[],[f616,f5704])).
fof(f616,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl8 | $spl24)),
  inference(resolution,[],[f116,f437])).
fof(f17654,plain,(
  ( ! [X50,X49] : (s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50) | s(bool,X50) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) ) | ($spl6 | $spl8 | $spl24)),
  inference(subsumption_resolution,[],[f17611,f695])).
fof(f695,plain,(
  ( ! [X20] : (s(bool,X20) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f687,f84])).
fof(f687,plain,(
  ( ! [X20] : (p(s(bool,'F_0')) | s(bool,X20) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X20)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f461,f113])).
fof(f17611,plain,(
  ( ! [X50,X49] : (s(bool,X49) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50) | s(bool,X50) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) ) | ($spl6 | $spl8 | $spl24)),
  inference(superposition,[],[f616,f5704])).
fof(f17652,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') = s(bool,X46) | s(bool,X47) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) ) | ($spl8 | $spl24)),
  inference(subsumption_resolution,[],[f17610,f68])).
fof(f17610,plain,(
  ( ! [X47,X48,X46] : (s(bool,'F_0') != s(bool,X46) | s(bool,'T_0') = s(bool,X46) | s(bool,X47) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,X47) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) ) | ($spl8 | $spl24)),
  inference(superposition,[],[f616,f5003])).
fof(f17649,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X34)) ) | ($spl8 | $spl24)),
  inference(subsumption_resolution,[],[f17606,f672])).
fof(f672,plain,(
  ( ! [X15,X16] : (s(bool,X15) = s(bool,X16) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X16)) ) | $spl24),
  inference(resolution,[],[f461,f457])).
fof(f17606,plain,(
  ( ! [X35,X36,X34] : (s(bool,X34) != s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X34)) ) | ($spl8 | $spl24)),
  inference(superposition,[],[f616,f5003])).
fof(f17648,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X23)) ) | ($spl8 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f17601,f68])).
fof(f17601,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X23)) ) | ($spl8 | $spl16 | $spl24)),
  inference(superposition,[],[f616,f782])).
fof(f782,plain,(
  ( ! [X17,X16] : (s(bool,X16) = s(bool,X17) | s(bool,'F_0') = s(bool,X16) | s(bool,'F_0') = s(bool,X17)) ) | $spl16),
  inference(resolution,[],[f317,f289])).
fof(f289,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') = s(bool,X10)) ) | $spl16),
  inference(trivial_inequality_removal,[],[f287])).
fof(f287,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,X10)) ) | $spl16),
  inference(superposition,[],[f133,f68])).
fof(f317,plain,(
  ( ! [X4,X3] : (~p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'F_0') = s(bool,X3)) ) | $spl16),
  inference(resolution,[],[f289,f65])).
fof(f17647,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X21)) ) | ($spl8 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f17600,f68])).
fof(f17600,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X21)) ) | ($spl8 | $spl16 | $spl24)),
  inference(superposition,[],[f616,f782])).
fof(f17617,plain,(
  ( ! [X37,X38,X36] : (s(bool,X37) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,X38) | s(bool,'F_0') = s(bool,X38)) ) | ($spl8 | $spl24)),
  inference(duplicate_literal_removal,[],[f17562])).
fof(f17562,plain,(
  ( ! [X37,X38,X36] : (s(bool,X37) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'T_0') = s(bool,X36) | s(bool,X37) = s(bool,X38) | s(bool,'F_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X36)) ) | ($spl8 | $spl24)),
  inference(superposition,[],[f616,f5003])).
fof(f17363,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115))),s('type.hreal.hreal',X116))) | ~p(s(bool,X117)) | s('type.hreal.hreal',X116) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal',sK3(X116,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115))))))))) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f17362,f113])).
fof(f17362,plain,(
  ( ! [X116,X114,X117,X115] : (~p(s(bool,X117)) | s('type.hreal.hreal',X116) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal',sK3(X116,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115)))))))) | s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115))),s('type.hreal.hreal',X116)))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f17296,f436])).
fof(f17361,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'F_0') = s(bool,X113) | s('type.hreal.hreal',X112) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X110),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X111),s('type.hreal.hreal',sK3(X112,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X110),s('type.hreal.hreal',X111)))))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X110),s('type.hreal.hreal',X111))),s('type.hreal.hreal',X112)))) ) | $spl6),
  inference(subsumption_resolution,[],[f17360,f4678])).
fof(f17360,plain,(
  ( ! [X111,X112,X110,X113] : (s(bool,'F_0') = s(bool,X113) | ~p(s(bool,X113)) | s('type.hreal.hreal',X112) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X110),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X111),s('type.hreal.hreal',sK3(X112,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X110),s('type.hreal.hreal',X111)))))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X110),s('type.hreal.hreal',X111))),s('type.hreal.hreal',X112)))) ) | $spl6),
  inference(forward_demodulation,[],[f17295,f113])).
fof(f17295,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,X113)) | s('type.hreal.hreal',X112) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X110),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X111),s('type.hreal.hreal',sK3(X112,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X110),s('type.hreal.hreal',X111)))))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X110),s('type.hreal.hreal',X111))),s('type.hreal.hreal',X112))) | s(bool,X113) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(superposition,[],[f5083,f5704])).
fof(f17358,plain,(
  ( ! [X94,X92,X93,X96] : (s('type.hreal.hreal',X94) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X92),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X93),s('type.hreal.hreal',sK3(X94,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X92),s('type.hreal.hreal',X93)))))))) | s(bool,'T_0') = s(bool,X96) | s(bool,X96) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X92),s('type.hreal.hreal',X93))),s('type.hreal.hreal',X94)))) ) | $spl220),
  inference(cnf_transformation,[],[f17358_D])).
fof(f17358_D,plain,(
  ( ! [X94,X92,X93,X96] : (s('type.hreal.hreal',X94) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X92),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X93),s('type.hreal.hreal',sK3(X94,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X92),s('type.hreal.hreal',X93)))))))) | s(bool,'T_0') = s(bool,X96) | s(bool,X96) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X92),s('type.hreal.hreal',X93))),s('type.hreal.hreal',X94)))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f17357,plain,(
  ( ! [X94,X92,X95,X93,X96] : (~p(s(bool,X95)) | s('type.hreal.hreal',X94) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X92),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X93),s('type.hreal.hreal',sK3(X94,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X92),s('type.hreal.hreal',X93)))))))) | s(bool,X96) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X92),s('type.hreal.hreal',X93))),s('type.hreal.hreal',X94))) | s(bool,'T_0') = s(bool,X96)) ) | $spl24),
  inference(subsumption_resolution,[],[f17291,f436])).
fof(f17291,plain,(
  ( ! [X94,X92,X95,X93,X96] : (~p(s(bool,X95)) | s('type.hreal.hreal',X94) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X92),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X93),s('type.hreal.hreal',sK3(X94,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X92),s('type.hreal.hreal',X93)))))))) | s(bool,'F_0') = s(bool,X95) | s(bool,X96) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X92),s('type.hreal.hreal',X93))),s('type.hreal.hreal',X94))) | s(bool,'T_0') = s(bool,X96)) ) | $spl24),
  inference(superposition,[],[f5083,f5003])).
fof(f17356,plain,(
  ( ! [X90,X88,X87,X91,X89] : (~p(s(bool,X90)) | s('type.hreal.hreal',X89) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X87),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal',sK3(X89,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X87),s('type.hreal.hreal',X88)))))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X87),s('type.hreal.hreal',X88))),s('type.hreal.hreal',X89))) | s(bool,X90) = s(bool,X91)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f17290,f273])).
fof(f17290,plain,(
  ( ! [X90,X88,X87,X91,X89] : (~p(s(bool,X90)) | s('type.hreal.hreal',X89) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X87),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal',sK3(X89,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X87),s('type.hreal.hreal',X88)))))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X87),s('type.hreal.hreal',X88))),s('type.hreal.hreal',X89))) | s(bool,X90) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) ) | $spl24),
  inference(superposition,[],[f5083,f5003])).
fof(f17352,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X66)) | s('type.hreal.hreal',X65) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X64),s('type.hreal.hreal',sK3(X65,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal',X64)))))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal',X64))),s('type.hreal.hreal',X65)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f17285,f436])).
fof(f17285,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X66)) | s('type.hreal.hreal',X65) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X64),s('type.hreal.hreal',sK3(X65,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal',X64)))))))) | s(bool,'F_0') = s(bool,X66) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal',X64))),s('type.hreal.hreal',X65)))) ) | $spl16),
  inference(superposition,[],[f5083,f782])).
fof(f17350,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X62)) | s('type.hreal.hreal',X61) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X59),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X60),s('type.hreal.hreal',sK3(X61,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X59),s('type.hreal.hreal',X60)))))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X59),s('type.hreal.hreal',X60))),s('type.hreal.hreal',X61)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f17284,f436])).
fof(f17284,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X62)) | s('type.hreal.hreal',X61) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X59),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X60),s('type.hreal.hreal',sK3(X61,'const.hreal.hreal_add_2'(s('type.hreal.hreal',X59),s('type.hreal.hreal',X60)))))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X59),s('type.hreal.hreal',X60))),s('type.hreal.hreal',X61))) | s(bool,'F_0') = s(bool,X62)) ) | $spl16),
  inference(superposition,[],[f5083,f782])).
fof(f11296,plain,(
  ( ! [X540,X542,X538,X539] : (s(bool,X542) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X538),s('type.hreal.hreal',X539))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X538),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X539),s('type.hreal.hreal',X540))))))) | s(bool,'F_0') = s(bool,X542)) ) | $spl164),
  inference(cnf_transformation,[],[f11296_D])).
fof(f11296_D,plain,(
  ( ! [X540,X542,X538,X539] : (s(bool,X542) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X538),s('type.hreal.hreal',X539))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X538),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X539),s('type.hreal.hreal',X540))))))) | s(bool,'F_0') = s(bool,X542)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f15706,plain,(
  ( ! [X432,X434,X431,X433,X435] : (s(bool,'T_0') != s(bool,X434) | s(bool,X435) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X431),s('type.hreal.hreal',X432))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X431),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X432),s('type.hreal.hreal',X433))))))) | s(bool,'F_0') = s(bool,X435)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f15612,f705])).
fof(f15612,plain,(
  ( ! [X432,X434,X431,X433,X435] : (s(bool,'T_0') != s(bool,X434) | s(bool,'F_0') = s(bool,X434) | s(bool,X435) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X431),s('type.hreal.hreal',X432))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X431),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X432),s('type.hreal.hreal',X433))))))) | s(bool,'F_0') = s(bool,X435)) ) | $spl24),
  inference(superposition,[],[f5192,f8910])).
fof(f8910,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X19),s('type.hreal.hreal',X20))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X19),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X20),s('type.hreal.hreal',X21)))))))) ) | $spl24),
  inference(resolution,[],[f8633,f457])).
fof(f8633,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X2))))))))) )),
  inference(equality_resolution,[],[f4866])).
fof(f4866,plain,(
  ( ! [X10,X8,X11,X9] : (s('type.hreal.hreal',X11) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X8),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X9),s('type.hreal.hreal',X10))))) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X8),s('type.hreal.hreal',X9))),s('type.hreal.hreal',X11))))) )),
  inference(superposition,[],[f61,f59])).
fof(f15627,plain,(
  ( ! [X391,X393,X395,X392,X394] : (s(bool,'T_0') = s(bool,X394) | s(bool,X395) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X391),s('type.hreal.hreal',X392))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X391),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X392),s('type.hreal.hreal',X393))))))) | s(bool,'F_0') = s(bool,X395)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f15604])).
fof(f15604,plain,(
  ( ! [X391,X393,X395,X392,X394] : (s(bool,'T_0') = s(bool,X394) | s(bool,X395) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X391),s('type.hreal.hreal',X392))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X391),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X392),s('type.hreal.hreal',X393))))))) | s(bool,'F_0') = s(bool,X395) | s(bool,'T_0') = s(bool,X394)) ) | $spl24),
  inference(superposition,[],[f5003,f8910])).
fof(f15645,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,'T_0') = s(bool,X69) | s(bool,X70) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X66),s('type.hreal.hreal',X67))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X66),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal',X68))))))) | s(bool,'F_0') = s(bool,X70)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f15505])).
fof(f15505,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,'T_0') = s(bool,X69) | s(bool,X70) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X66),s('type.hreal.hreal',X67))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X66),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal',X68))))))) | s(bool,'F_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X69)) ) | $spl24),
  inference(superposition,[],[f8910,f5003])).
fof(f15270,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119))) | ~p(s(bool,X120)) | s('type.hreal.hreal',sK3(X119,X118)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X121),s('type.hreal.hreal',X119)),X121),X118))) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f15269,f113])).
fof(f15269,plain,(
  ( ! [X121,X118,X120,X119] : (~p(s(bool,X120)) | s('type.hreal.hreal',sK3(X119,X118)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X121),s('type.hreal.hreal',X119)),X121),X118)) | s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119)))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f15241,f436])).
fof(f15268,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,'F_0') = s(bool,X116) | s('type.hreal.hreal',sK3(X115,X114)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X117),s('type.hreal.hreal',X115)),X117),X114)) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115)))) ) | $spl6),
  inference(subsumption_resolution,[],[f15267,f4678])).
fof(f15267,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,'F_0') = s(bool,X116) | ~p(s(bool,X116)) | s('type.hreal.hreal',sK3(X115,X114)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X117),s('type.hreal.hreal',X115)),X117),X114)) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115)))) ) | $spl6),
  inference(forward_demodulation,[],[f15240,f113])).
fof(f15240,plain,(
  ( ! [X116,X114,X117,X115] : (~p(s(bool,X116)) | s('type.hreal.hreal',sK3(X115,X114)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X117),s('type.hreal.hreal',X115)),X117),X114)) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115))) | s(bool,X116) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(superposition,[],[f14976,f5704])).
fof(f15265,plain,(
  ( ! [X99,X97,X100,X96] : (s('type.hreal.hreal',sK3(X97,X96)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X99),s('type.hreal.hreal',X97)),X99),X96)) | s(bool,'T_0') = s(bool,X100) | s(bool,X100) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X96),s('type.hreal.hreal',X97)))) ) | $spl216),
  inference(cnf_transformation,[],[f15265_D])).
fof(f15265_D,plain,(
  ( ! [X99,X97,X100,X96] : (s('type.hreal.hreal',sK3(X97,X96)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X99),s('type.hreal.hreal',X97)),X99),X96)) | s(bool,'T_0') = s(bool,X100) | s(bool,X100) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X96),s('type.hreal.hreal',X97)))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f15264,plain,(
  ( ! [X99,X97,X100,X98,X96] : (~p(s(bool,X98)) | s('type.hreal.hreal',sK3(X97,X96)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X99),s('type.hreal.hreal',X97)),X99),X96)) | s(bool,X100) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X96),s('type.hreal.hreal',X97))) | s(bool,'T_0') = s(bool,X100)) ) | $spl24),
  inference(subsumption_resolution,[],[f15236,f436])).
fof(f15236,plain,(
  ( ! [X99,X97,X100,X98,X96] : (~p(s(bool,X98)) | s('type.hreal.hreal',sK3(X97,X96)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X99),s('type.hreal.hreal',X97)),X99),X96)) | s(bool,'F_0') = s(bool,X98) | s(bool,X100) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X96),s('type.hreal.hreal',X97))) | s(bool,'T_0') = s(bool,X100)) ) | $spl24),
  inference(superposition,[],[f14976,f5003])).
fof(f15263,plain,(
  ( ! [X94,X92,X95,X93,X91] : (~p(s(bool,X93)) | s('type.hreal.hreal',sK3(X92,X91)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X94),s('type.hreal.hreal',X92)),X94),X91)) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X91),s('type.hreal.hreal',X92))) | s(bool,X93) = s(bool,X95)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f15235,f273])).
fof(f15235,plain,(
  ( ! [X94,X92,X95,X93,X91] : (~p(s(bool,X93)) | s('type.hreal.hreal',sK3(X92,X91)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X94),s('type.hreal.hreal',X92)),X94),X91)) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X91),s('type.hreal.hreal',X92))) | s(bool,X93) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) ) | $spl24),
  inference(superposition,[],[f14976,f5003])).
fof(f15259,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X69)) | s('type.hreal.hreal',sK3(X68,X67)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X70),s('type.hreal.hreal',X68)),X70),X67)) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal',X68)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f15230,f436])).
fof(f15230,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X69)) | s('type.hreal.hreal',sK3(X68,X67)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X70),s('type.hreal.hreal',X68)),X70),X67)) | s(bool,'F_0') = s(bool,X69) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal',X68)))) ) | $spl16),
  inference(superposition,[],[f14976,f782])).
fof(f15257,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | s('type.hreal.hreal',sK3(X64,X63)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X66),s('type.hreal.hreal',X64)),X66),X63)) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal',X64)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f15229,f436])).
fof(f15229,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | s('type.hreal.hreal',sK3(X64,X63)) = s('type.hreal.hreal',sK3(sK3('const.hreal.hreal_add_2'(s('type.hreal.hreal',X66),s('type.hreal.hreal',X64)),X66),X63)) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal',X64))) | s(bool,'F_0') = s(bool,X65)) ) | $spl16),
  inference(superposition,[],[f14976,f782])).
fof(f14817,plain,(
  ( ! [X70,X71,X69,X67] : (s(bool,'F_0') != s(bool,X69) | s(bool,'F_0') = s(bool,X71) | s(bool,X67) = s(bool,X71) | s(bool,X69) != s(bool,X70) | s(bool,X67) = s(bool,X70)) ) | $spl212),
  inference(cnf_transformation,[],[f14817_D])).
fof(f14817_D,plain,(
  ( ! [X70,X71,X69,X67] : (s(bool,'F_0') != s(bool,X69) | s(bool,'F_0') = s(bool,X71) | s(bool,X67) = s(bool,X71) | s(bool,X69) != s(bool,X70) | s(bool,X67) = s(bool,X70)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f14816,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,X68)) | s(bool,'F_0') != s(bool,X69) | s(bool,X67) = s(bool,X70) | s(bool,X69) != s(bool,X70) | s(bool,X67) = s(bool,X71) | s(bool,'F_0') = s(bool,X71)) ) | ($spl16 | $spl24 | $spl204)),
  inference(subsumption_resolution,[],[f14782,f133])).
fof(f14782,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,X68)) | s(bool,'F_0') != s(bool,X69) | s(bool,X67) = s(bool,X70) | s(bool,X69) != s(bool,X70) | s(bool,X67) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X68)) ) | ($spl24 | $spl204)),
  inference(superposition,[],[f14707,f5003])).
fof(f14707,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X27)) | s(bool,'F_0') != s(bool,X25) | s(bool,X26) = s(bool,X27) | s(bool,X25) != s(bool,X26)) ) | $spl204),
  inference(cnf_transformation,[],[f14707_D])).
fof(f14707_D,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X27)) | s(bool,'F_0') != s(bool,X25) | s(bool,X26) = s(bool,X27) | s(bool,X25) != s(bool,X26)) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f14726,plain,(
  ( ! [X123,X121,X122,X120,X119] : (s(bool,'F_0') != s(bool,X120) | s(bool,X120) != s(bool,X121) | p(s(bool,X119)) | s(bool,X121) = s(bool,X122) | p(s(bool,X122)) | s(bool,X119) = s(bool,X123)) ) | ($spl6 | $spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f14725,f2819])).
fof(f14725,plain,(
  ( ! [X123,X121,X122,X120,X119] : (s(bool,'F_0') = s(bool,X123) | s(bool,'F_0') != s(bool,X120) | s(bool,X120) != s(bool,X121) | p(s(bool,X119)) | s(bool,X121) = s(bool,X122) | p(s(bool,X122)) | s(bool,X119) = s(bool,X123)) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f14652,f113])).
fof(f14652,plain,(
  ( ! [X123,X121,X122,X120,X119] : (s(bool,'F_0') != s(bool,X120) | s(bool,X120) != s(bool,X121) | p(s(bool,X119)) | s(bool,X121) = s(bool,X122) | p(s(bool,X122)) | s(bool,X123) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X119) = s(bool,X123)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3026,f5704])).
fof(f3026,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X2) | s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | s(bool,X1) = s(bool,X3) | p(s(bool,X3))) ) | ($spl16 | $spl24)),
  inference(resolution,[],[f907,f62])).
fof(f907,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,X9)) | s(bool,X8) != s(bool,X9) | s(bool,X7) != s(bool,X8) | p(s(bool,X7))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f890,f70])).
fof(f890,plain,(
  ( ! [X8,X7,X9] : (s(bool,X7) != s(bool,X8) | s(bool,X8) != s(bool,X9) | ~p(s(bool,X9)) | p(s(bool,'F_0')) | p(s(bool,X7))) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f746,f62])).
fof(f746,plain,(
  ( ! [X8,X7] : (s(bool,'F_0') != s(bool,X8) | s(bool,X7) != s(bool,X8) | ~p(s(bool,X7))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f733,f79])).
fof(f733,plain,(
  ( ! [X8,X7] : (s(bool,X7) != s(bool,X8) | s(bool,'F_0') != s(bool,X8) | ~p(s(bool,X7)) | ~p(s(bool,'T_0'))) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f705,f65])).
fof(f14724,plain,(
  ( ! [X118,X116,X114,X117,X115] : (s(bool,'F_0') != s(bool,X115) | s(bool,X115) != s(bool,X116) | p(s(bool,X114)) | s(bool,X116) = s(bool,X117) | p(s(bool,X117)) | s(bool,X114) = s(bool,X118)) ) | ($spl6 | $spl16 | $spl24 | $spl32)),
  inference(forward_demodulation,[],[f14723,f113])).
fof(f14723,plain,(
  ( ! [X118,X116,X114,X117,X115] : (s(bool,X115) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X115) != s(bool,X116) | p(s(bool,X114)) | s(bool,X116) = s(bool,X117) | p(s(bool,X117)) | s(bool,X114) = s(bool,X118)) ) | ($spl6 | $spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f14651,f2819])).
fof(f14651,plain,(
  ( ! [X118,X116,X114,X117,X115] : (s(bool,X115) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X115) != s(bool,X116) | p(s(bool,X114)) | s(bool,X116) = s(bool,X117) | p(s(bool,X117)) | s(bool,'F_0') = s(bool,X118) | s(bool,X114) = s(bool,X118)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3026,f5704])).
fof(f14722,plain,(
  ( ! [X112,X113] : (s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) ) | $spl210),
  inference(cnf_transformation,[],[f14722_D])).
fof(f14722_D,plain,(
  ( ! [X112,X113] : (s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f14721,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (s(bool,'F_0') != s(bool,X109) | s(bool,X109) != s(bool,X110) | p(s(bool,X108)) | s(bool,X110) = s(bool,X111) | p(s(bool,X111)) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14650,f673])).
fof(f673,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) != s(bool,X19) | s(bool,X17) = s(bool,X18)) ) | $spl24),
  inference(resolution,[],[f461,f458])).
fof(f14650,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (s(bool,'F_0') != s(bool,X109) | s(bool,X109) != s(bool,X110) | p(s(bool,X108)) | s(bool,X110) = s(bool,X111) | p(s(bool,X111)) | s(bool,X108) = s(bool,X112) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3026,f5003])).
fof(f14720,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | s(bool,'F_0') = s(bool,X95) | s(bool,X91) = s(bool,X95) | p(s(bool,X94)) | s(bool,X93) = s(bool,X94) | s(bool,X92) != s(bool,X93)) ) | $spl208),
  inference(cnf_transformation,[],[f14720_D])).
fof(f14720_D,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | s(bool,'F_0') = s(bool,X95) | s(bool,X91) = s(bool,X95) | p(s(bool,X94)) | s(bool,X93) = s(bool,X94) | s(bool,X92) != s(bool,X93)) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f14719,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | s(bool,X92) != s(bool,X93) | p(s(bool,X90)) | s(bool,X93) = s(bool,X94) | p(s(bool,X94)) | s(bool,X91) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14647,f133])).
fof(f14647,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s(bool,X91) != s(bool,X92) | s(bool,X92) != s(bool,X93) | p(s(bool,X90)) | s(bool,X93) = s(bool,X94) | p(s(bool,X94)) | s(bool,X91) = s(bool,X95) | s(bool,'F_0') = s(bool,X95) | s(bool,'T_0') = s(bool,X90)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3026,f5003])).
fof(f14711,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X39) != s(bool,X40) | s(bool,'T_0') = s(bool,X39) | p(s(bool,X42)) | s(bool,X41) = s(bool,X42) | s(bool,X40) != s(bool,X41)) ) | $spl206),
  inference(cnf_transformation,[],[f14711_D])).
fof(f14711_D,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X39) != s(bool,X40) | s(bool,'T_0') = s(bool,X39) | p(s(bool,X42)) | s(bool,X41) = s(bool,X42) | s(bool,X40) != s(bool,X41)) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f14718,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (s(bool,X85) != s(bool,X86) | s(bool,X86) != s(bool,X87) | p(s(bool,X84)) | s(bool,X87) = s(bool,X88) | p(s(bool,X88)) | s(bool,X84) = s(bool,X89) | s(bool,'T_0') = s(bool,X85)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f14646,f2819])).
fof(f14646,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (s(bool,X85) != s(bool,X86) | s(bool,X86) != s(bool,X87) | p(s(bool,X84)) | s(bool,X87) = s(bool,X88) | p(s(bool,X88)) | s(bool,X84) = s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X85)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3026,f5003])).
fof(f14716,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s(bool,'F_0') != s(bool,X58) | s(bool,X58) != s(bool,X59) | p(s(bool,X57)) | s(bool,X59) = s(bool,X60) | p(s(bool,X60)) | s(bool,X57) = s(bool,X61)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f14641,f2819])).
fof(f14641,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s(bool,'F_0') != s(bool,X58) | s(bool,X58) != s(bool,X59) | p(s(bool,X57)) | s(bool,X59) = s(bool,X60) | p(s(bool,X60)) | s(bool,X57) = s(bool,X61) | s(bool,'F_0') = s(bool,X61)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3026,f782])).
fof(f14715,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,'F_0') != s(bool,X53) | s(bool,X53) != s(bool,X54) | p(s(bool,X52)) | s(bool,X54) = s(bool,X55) | p(s(bool,X55)) | s(bool,X52) = s(bool,X56)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f14640,f2819])).
fof(f14640,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,'F_0') != s(bool,X53) | s(bool,X53) != s(bool,X54) | p(s(bool,X52)) | s(bool,X54) = s(bool,X55) | p(s(bool,X55)) | s(bool,X52) = s(bool,X56) | s(bool,'F_0') = s(bool,X56)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3026,f782])).
fof(f14714,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X49) | s(bool,X49) != s(bool,X50) | p(s(bool,X48)) | s(bool,X50) = s(bool,X51) | p(s(bool,X51))) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f14713,f113])).
fof(f14713,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X49) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X49) != s(bool,X50) | p(s(bool,X48)) | s(bool,X50) = s(bool,X51) | p(s(bool,X51))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14639,f133])).
fof(f14639,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X49) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X49) != s(bool,X50) | p(s(bool,X48)) | s(bool,X50) = s(bool,X51) | p(s(bool,X51)) | s(bool,'T_0') = s(bool,X48)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3026,f695])).
fof(f14712,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(bool,X44) != s(bool,X45) | s(bool,X45) != s(bool,X46) | p(s(bool,X43)) | s(bool,X46) = s(bool,X47) | p(s(bool,X47)) | s(bool,'T_0') = s(bool,X44)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14638,f133])).
fof(f14638,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(bool,X44) != s(bool,X45) | s(bool,X45) != s(bool,X46) | p(s(bool,X43)) | s(bool,X46) = s(bool,X47) | p(s(bool,X47)) | s(bool,'T_0') = s(bool,X44) | s(bool,'T_0') = s(bool,X43)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3026,f672])).
fof(f14710,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,X39) != s(bool,X40) | s(bool,X40) != s(bool,X41) | p(s(bool,X38)) | s(bool,X41) = s(bool,X42) | p(s(bool,X42)) | s(bool,'T_0') = s(bool,X39)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14637,f133])).
fof(f14637,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,X39) != s(bool,X40) | s(bool,X40) != s(bool,X41) | p(s(bool,X38)) | s(bool,X41) = s(bool,X42) | p(s(bool,X42)) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,X39)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3026,f672])).
fof(f14706,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,X25) | s(bool,X25) != s(bool,X26) | p(s(bool,X24)) | s(bool,X26) = s(bool,X27) | p(s(bool,X27))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14634,f133])).
fof(f14634,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'F_0') != s(bool,X25) | s(bool,X25) != s(bool,X26) | p(s(bool,X24)) | s(bool,X26) = s(bool,X27) | p(s(bool,X27)) | s(bool,'T_0') = s(bool,X24)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3026,f68])).
fof(f11373,plain,(
  ( ! [X163,X159,X162,X160] : (s(bool,X159) != s(bool,X160) | s(bool,X159) = s(bool,X163) | s(bool,X162) = s(bool,X163) | s(bool,'F_0') = s(bool,X162) | p(s(bool,X160))) ) | $spl170),
  inference(cnf_transformation,[],[f11373_D])).
fof(f11373_D,plain,(
  ( ! [X163,X159,X162,X160] : (s(bool,X159) != s(bool,X160) | s(bool,X159) = s(bool,X163) | s(bool,X162) = s(bool,X163) | s(bool,'F_0') = s(bool,X162) | p(s(bool,X160))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f14694,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (s(bool,X134) != s(bool,X135) | s(bool,'T_0') != s(bool,X136) | p(s(bool,X135)) | p(s(bool,X137)) | s(bool,X138) = s(bool,X139) | s(bool,'F_0') = s(bool,X138) | s(bool,X134) = s(bool,X139)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14614,f538])).
fof(f14614,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (s(bool,X134) != s(bool,X135) | s(bool,'T_0') != s(bool,X136) | p(s(bool,X135)) | s(bool,X136) = s(bool,X137) | p(s(bool,X137)) | s(bool,X138) = s(bool,X139) | s(bool,'F_0') = s(bool,X138) | s(bool,X134) = s(bool,X139)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3026,f5003])).
fof(f11313,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,X79) != s(bool,X80) | s(bool,'F_0') = s(bool,X82) | s(bool,X79) = s(bool,X82) | p(s(bool,X81)) | s(bool,X80) = s(bool,X81)) ) | $spl166),
  inference(cnf_transformation,[],[f11313_D])).
fof(f11313_D,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,X79) != s(bool,X80) | s(bool,'F_0') = s(bool,X82) | s(bool,X79) = s(bool,X82) | p(s(bool,X81)) | s(bool,X80) = s(bool,X81)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f14691,plain,(
  ( ! [X123,X121,X122,X120,X119] : (s(bool,'F_0') = s(bool,X123) | s(bool,X119) != s(bool,X121) | p(s(bool,X120)) | s(bool,X121) = s(bool,X122) | p(s(bool,X122)) | s(bool,X119) = s(bool,X123)) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f14690,f113])).
fof(f14690,plain,(
  ( ! [X123,X121,X122,X120,X119] : (s(bool,X119) != s(bool,X121) | p(s(bool,X120)) | s(bool,X121) = s(bool,X122) | p(s(bool,X122)) | s(bool,X123) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X119) = s(bool,X123)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14611,f4678])).
fof(f14611,plain,(
  ( ! [X123,X121,X122,X120,X119] : (s(bool,'F_0') != s(bool,X120) | s(bool,X119) != s(bool,X121) | p(s(bool,X120)) | s(bool,X121) = s(bool,X122) | p(s(bool,X122)) | s(bool,X123) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X119) = s(bool,X123)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3026,f5704])).
fof(f14689,plain,(
  ( ! [X118,X116,X114,X117,X115] : (s(bool,X114) != s(bool,X116) | p(s(bool,X115)) | s(bool,X116) = s(bool,X117) | p(s(bool,X117)) | s(bool,'F_0') = s(bool,X118) | s(bool,X114) = s(bool,X118)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14688,f4678])).
fof(f14688,plain,(
  ( ! [X118,X116,X114,X117,X115] : (s(bool,'F_0') != s(bool,X115) | s(bool,X114) != s(bool,X116) | p(s(bool,X115)) | s(bool,X116) = s(bool,X117) | p(s(bool,X117)) | s(bool,'F_0') = s(bool,X118) | s(bool,X114) = s(bool,X118)) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f14610,f113])).
fof(f14610,plain,(
  ( ! [X118,X116,X114,X117,X115] : (s(bool,X115) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X114) != s(bool,X116) | p(s(bool,X115)) | s(bool,X116) = s(bool,X117) | p(s(bool,X117)) | s(bool,'F_0') = s(bool,X118) | s(bool,X114) = s(bool,X118)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3026,f5704])).
fof(f14568,plain,(
  ( ! [X107,X105,X102,X106,X104] : (s(bool,X102) = s(bool,X104) | s(bool,'T_0') = s(bool,X107) | s(bool,X102) = s(bool,X107) | p(s(bool,X106)) | s(bool,X105) = s(bool,X106) | s(bool,X104) != s(bool,X105)) ) | $spl200),
  inference(cnf_transformation,[],[f14568_D])).
fof(f14568_D,plain,(
  ( ! [X107,X105,X102,X106,X104] : (s(bool,X102) = s(bool,X104) | s(bool,'T_0') = s(bool,X107) | s(bool,X102) = s(bool,X107) | p(s(bool,X106)) | s(bool,X105) = s(bool,X106) | s(bool,X104) != s(bool,X105)) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f14687,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (s(bool,X108) != s(bool,X110) | p(s(bool,X109)) | s(bool,X110) = s(bool,X111) | p(s(bool,X111)) | s(bool,X108) = s(bool,X112) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14609,f4678])).
fof(f14609,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (s(bool,'F_0') != s(bool,X109) | s(bool,X108) != s(bool,X110) | p(s(bool,X109)) | s(bool,X110) = s(bool,X111) | p(s(bool,X111)) | s(bool,X108) = s(bool,X112) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3026,f5003])).
fof(f14684,plain,(
  ( ! [X85,X86] : (s(bool,X85) != s(bool,X86) | p(s(bool,X86))) ) | $spl202),
  inference(cnf_transformation,[],[f14684_D])).
fof(f14684_D,plain,(
  ( ! [X85,X86] : (s(bool,X85) != s(bool,X86) | p(s(bool,X86))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f14683,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (s(bool,X85) != s(bool,X86) | s(bool,X84) != s(bool,X87) | p(s(bool,X86)) | s(bool,X87) = s(bool,X88) | p(s(bool,X88)) | s(bool,X84) = s(bool,X89) | s(bool,'F_0') = s(bool,X89)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14605,f538])).
fof(f14605,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (s(bool,X85) != s(bool,X86) | s(bool,X84) != s(bool,X87) | p(s(bool,X86)) | s(bool,X87) = s(bool,X88) | p(s(bool,X88)) | s(bool,X84) = s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X85)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3026,f5003])).
fof(f14680,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s(bool,X57) != s(bool,X59) | p(s(bool,X58)) | s(bool,X59) = s(bool,X60) | p(s(bool,X60)) | s(bool,X57) = s(bool,X61) | s(bool,'F_0') = s(bool,X61)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14600,f4678])).
fof(f14600,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s(bool,'F_0') != s(bool,X58) | s(bool,X57) != s(bool,X59) | p(s(bool,X58)) | s(bool,X59) = s(bool,X60) | p(s(bool,X60)) | s(bool,X57) = s(bool,X61) | s(bool,'F_0') = s(bool,X61)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3026,f782])).
fof(f14679,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,X52) != s(bool,X54) | p(s(bool,X53)) | s(bool,X54) = s(bool,X55) | p(s(bool,X55)) | s(bool,X52) = s(bool,X56) | s(bool,'F_0') = s(bool,X56)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14599,f4678])).
fof(f14599,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,'F_0') != s(bool,X53) | s(bool,X52) != s(bool,X54) | p(s(bool,X53)) | s(bool,X54) = s(bool,X55) | p(s(bool,X55)) | s(bool,X52) = s(bool,X56) | s(bool,'F_0') = s(bool,X56)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3026,f782])).
fof(f14573,plain,(
  ( ! [X132,X130,X133,X131,X129] : (s(bool,'F_0') = s(bool,X129) | ~p(s(bool,X130)) | s(bool,X129) = s(bool,X131) | s(bool,X131) != s(bool,X132) | s(bool,X132) = s(bool,X133) | p(s(bool,X133))) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f14572,f113])).
fof(f14572,plain,(
  ( ! [X132,X130,X133,X131,X129] : (~p(s(bool,X130)) | s(bool,X129) = s(bool,X131) | s(bool,X131) != s(bool,X132) | s(bool,X132) = s(bool,X133) | p(s(bool,X133)) | s(bool,X129) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14539,f436])).
fof(f14539,plain,(
  ( ! [X132,X130,X133,X131,X129] : (~p(s(bool,X130)) | s(bool,X129) = s(bool,X131) | s(bool,X131) != s(bool,X132) | s(bool,X132) = s(bool,X133) | p(s(bool,X133)) | s(bool,'F_0') = s(bool,X130) | s(bool,X129) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16)),
  inference(superposition,[],[f2961,f5704])).
fof(f2961,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,X0) != s(bool,X2) | s(bool,X2) = s(bool,X3) | p(s(bool,X3))) ) | $spl16),
  inference(resolution,[],[f662,f62])).
fof(f662,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,X11)) | s(bool,X12) = s(bool,X13) | ~p(s(bool,X13)) | s(bool,X11) != s(bool,X12)) ) | $spl16),
  inference(subsumption_resolution,[],[f646,f79])).
fof(f646,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | s(bool,X12) = s(bool,X13) | ~p(s(bool,X13)) | ~p(s(bool,X11)) | ~p(s(bool,'T_0'))) ) | $spl16),
  inference(superposition,[],[f273,f65])).
fof(f14571,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,'F_0') = s(bool,X125) | s(bool,X124) = s(bool,X126) | s(bool,X126) != s(bool,X127) | s(bool,X127) = s(bool,X128) | p(s(bool,X128)) | s(bool,'F_0') = s(bool,X124)) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f14570,f4678])).
fof(f14570,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,'F_0') = s(bool,X125) | ~p(s(bool,X125)) | s(bool,X124) = s(bool,X126) | s(bool,X126) != s(bool,X127) | s(bool,X127) = s(bool,X128) | p(s(bool,X128)) | s(bool,'F_0') = s(bool,X124)) ) | ($spl6 | $spl16)),
  inference(forward_demodulation,[],[f14538,f113])).
fof(f14538,plain,(
  ( ! [X127,X125,X128,X126,X124] : (~p(s(bool,X125)) | s(bool,X124) = s(bool,X126) | s(bool,X126) != s(bool,X127) | s(bool,X127) = s(bool,X128) | p(s(bool,X128)) | s(bool,'F_0') = s(bool,X124) | s(bool,X125) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16)),
  inference(superposition,[],[f2961,f5704])).
fof(f14567,plain,(
  ( ! [X103,X107,X105,X102,X106,X104] : (~p(s(bool,X103)) | s(bool,X102) = s(bool,X104) | s(bool,X104) != s(bool,X105) | s(bool,X105) = s(bool,X106) | p(s(bool,X106)) | s(bool,X102) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14534,f436])).
fof(f14534,plain,(
  ( ! [X103,X107,X105,X102,X106,X104] : (~p(s(bool,X103)) | s(bool,X102) = s(bool,X104) | s(bool,X104) != s(bool,X105) | s(bool,X105) = s(bool,X106) | p(s(bool,X106)) | s(bool,'F_0') = s(bool,X103) | s(bool,X102) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f2961,f5003])).
fof(f14566,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (~p(s(bool,X97)) | s(bool,X96) = s(bool,X98) | s(bool,X98) != s(bool,X99) | s(bool,X99) = s(bool,X100) | p(s(bool,X100)) | s(bool,'F_0') = s(bool,X96) | s(bool,X97) = s(bool,X101)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14533,f273])).
fof(f14533,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (~p(s(bool,X97)) | s(bool,X96) = s(bool,X98) | s(bool,X98) != s(bool,X99) | s(bool,X99) = s(bool,X100) | p(s(bool,X100)) | s(bool,'F_0') = s(bool,X96) | s(bool,X97) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f2961,f5003])).
fof(f14562,plain,(
  ( ! [X70,X68,X71,X69,X67] : (~p(s(bool,X68)) | s(bool,X67) = s(bool,X69) | s(bool,X69) != s(bool,X70) | s(bool,X70) = s(bool,X71) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X67)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14528,f436])).
fof(f14528,plain,(
  ( ! [X70,X68,X71,X69,X67] : (~p(s(bool,X68)) | s(bool,X67) = s(bool,X69) | s(bool,X69) != s(bool,X70) | s(bool,X70) = s(bool,X71) | p(s(bool,X71)) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X67)) ) | $spl16),
  inference(superposition,[],[f2961,f782])).
fof(f14561,plain,(
  ( ! [X66,X64,X62,X65,X63] : (~p(s(bool,X63)) | s(bool,X62) = s(bool,X64) | s(bool,X64) != s(bool,X65) | s(bool,X65) = s(bool,X66) | p(s(bool,X66)) | s(bool,'F_0') = s(bool,X62)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14527,f436])).
fof(f14527,plain,(
  ( ! [X66,X64,X62,X65,X63] : (~p(s(bool,X63)) | s(bool,X62) = s(bool,X64) | s(bool,X64) != s(bool,X65) | s(bool,X65) = s(bool,X66) | p(s(bool,X66)) | s(bool,'F_0') = s(bool,X62) | s(bool,'F_0') = s(bool,X63)) ) | $spl16),
  inference(superposition,[],[f2961,f782])).
fof(f11256,plain,(
  ( ! [X204,X205,X203,X201] : (s(bool,'F_0') = s(bool,X203) | s(bool,'F_0') = s(bool,X205) | s(bool,X201) = s(bool,X205) | s(bool,X203) != s(bool,X204) | s(bool,X201) != s(bool,X204)) ) | $spl148),
  inference(cnf_transformation,[],[f11256_D])).
fof(f11256_D,plain,(
  ( ! [X204,X205,X203,X201] : (s(bool,'F_0') = s(bool,X203) | s(bool,'F_0') = s(bool,X205) | s(bool,X201) = s(bool,X205) | s(bool,X203) != s(bool,X204) | s(bool,X201) != s(bool,X204)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f14280,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,X67) != s(bool,X69) | s(bool,X69) != s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,X67) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X68)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14231,f68])).
fof(f14231,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,'F_0') != s(bool,X68) | s(bool,X67) != s(bool,X69) | s(bool,X69) != s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,X67) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X68)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8241,f5003])).
fof(f8241,plain,(
  ( ! [X61,X62,X63] : (s(bool,'F_0') != s(bool,X63) | s(bool,X62) != s(bool,X63) | s(bool,X61) != s(bool,X62) | s(bool,'F_0') = s(bool,X61)) ) | ($spl6 | $spl16 | $spl24)),
  inference(resolution,[],[f8131,f4678])).
fof(f11136,plain,(
  ( ! [X401,X403,X404,X402] : (s(bool,X401) != s(bool,X402) | s(bool,'F_0') = s(bool,X401) | s(bool,'T_0') = s(bool,X404) | s(bool,X403) != s(bool,X404) | s(bool,X402) != s(bool,X403)) ) | $spl130),
  inference(cnf_transformation,[],[f11136_D])).
fof(f11136_D,plain,(
  ( ! [X401,X403,X404,X402] : (s(bool,X401) != s(bool,X402) | s(bool,'F_0') = s(bool,X401) | s(bool,'T_0') = s(bool,X404) | s(bool,X403) != s(bool,X404) | s(bool,X402) != s(bool,X403)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f14263,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X86) != s(bool,X87) | s(bool,X87) != s(bool,X88) | s(bool,X88) != s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X86)) ) | ($spl6 | $spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f14206])).
fof(f14206,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X86) != s(bool,X87) | s(bool,X87) != s(bool,X88) | s(bool,X88) != s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X90) | s(bool,'F_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X86)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8241,f5003])).
fof(f14185,plain,(
  ( ! [X158,X156,X159,X155] : (s(bool,X158) != s(bool,X159) | s(bool,X158) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,'F_0') != s(bool,X159)) ) | $spl198),
  inference(cnf_transformation,[],[f14185_D])).
fof(f14185_D,plain,(
  ( ! [X158,X156,X159,X155] : (s(bool,X158) != s(bool,X159) | s(bool,X158) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,'F_0') != s(bool,X159)) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f14184,plain,(
  ( ! [X158,X156,X159,X157,X155] : (s(bool,'T_0') != s(bool,X157) | s(bool,X158) != s(bool,X159) | s(bool,'F_0') != s(bool,X159) | s(bool,X158) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156)))))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14150,f705])).
fof(f14150,plain,(
  ( ! [X158,X156,X159,X157,X155] : (s(bool,'T_0') != s(bool,X157) | s(bool,X158) != s(bool,X159) | s(bool,'F_0') != s(bool,X159) | s(bool,X158) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,'F_0') = s(bool,X157)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8231,f4294])).
fof(f4294,plain,(
  ( ! [X54,X56,X55] : (s(bool,X56) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X54),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X54),s('type.hreal.hreal',X55))))) | s(bool,'F_0') = s(bool,X56)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f4221,f85])).
fof(f4221,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'T_0')) | s(bool,X56) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X54),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X54),s('type.hreal.hreal',X55))))) | s(bool,'F_0') = s(bool,X56)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f317,f4167])).
fof(f4167,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))))) ) | $spl24),
  inference(equality_resolution,[],[f478])).
fof(f478,plain,(
  ( ! [X17,X15,X16] : (s('type.hreal.hreal',X16) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X15),s('type.hreal.hreal',X17))) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X15),s('type.hreal.hreal',X16)))) ) | $spl24),
  inference(resolution,[],[f457,f61])).
fof(f14182,plain,(
  ( ! [X107,X106,X104] : (s(bool,X106) != s(bool,X107) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X104) != s(bool,X106) | s(bool,'F_0') != s(bool,X107)) ) | $spl196),
  inference(cnf_transformation,[],[f14182_D])).
fof(f14182_D,plain,(
  ( ! [X107,X106,X104] : (s(bool,X106) != s(bool,X107) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X104) != s(bool,X106) | s(bool,'F_0') != s(bool,X107)) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f14181,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,X105) | s(bool,X106) != s(bool,X107) | s(bool,'F_0') != s(bool,X107) | s(bool,X104) != s(bool,X106) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14136,f705])).
fof(f14136,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,X105) | s(bool,X106) != s(bool,X107) | s(bool,'F_0') != s(bool,X107) | s(bool,X104) != s(bool,X106) | s(bool,'F_0') = s(bool,X105) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8231,f5704])).
fof(f8292,plain,(
  ( ! [X52,X51,X49] : (s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') != s(bool,X52) | s(bool,X51) != s(bool,X52)) ) | $spl88),
  inference(cnf_transformation,[],[f8292_D])).
fof(f8292_D,plain,(
  ( ! [X52,X51,X49] : (s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') != s(bool,X52) | s(bool,X51) != s(bool,X52)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f14180,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,'T_0') != s(bool,X101) | s(bool,X102) != s(bool,X103) | s(bool,'F_0') != s(bool,X103) | s(bool,X100) != s(bool,X102) | s(bool,'F_0') = s(bool,X100)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14135,f727])).
fof(f727,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,X20) | s(bool,X20) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f717,f84])).
fof(f717,plain,(
  ( ! [X20] : (p(s(bool,'F_0')) | s(bool,X20) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') != s(bool,X20)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f538,f113])).
fof(f14135,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,'T_0') != s(bool,X101) | s(bool,X102) != s(bool,X103) | s(bool,'F_0') != s(bool,X103) | s(bool,X100) != s(bool,X102) | s(bool,'F_0') = s(bool,X100) | s(bool,X101) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8231,f5704])).
fof(f11135,plain,(
  ( ! [X350,X352,X351,X348] : (s(bool,X348) != s(bool,X350) | s(bool,'T_0') = s(bool,X352) | s(bool,X348) = s(bool,X352) | s(bool,'F_0') != s(bool,X351) | s(bool,X350) != s(bool,X351)) ) | $spl128),
  inference(cnf_transformation,[],[f11135_D])).
fof(f11135_D,plain,(
  ( ! [X350,X352,X351,X348] : (s(bool,X348) != s(bool,X350) | s(bool,'T_0') = s(bool,X352) | s(bool,X348) = s(bool,X352) | s(bool,'F_0') != s(bool,X351) | s(bool,X350) != s(bool,X351)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f14179,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,X84) != s(bool,X85) | s(bool,'F_0') != s(bool,X85) | s(bool,X82) != s(bool,X84) | s(bool,X82) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14131,f705])).
fof(f14131,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,X84) != s(bool,X85) | s(bool,'F_0') != s(bool,X85) | s(bool,X82) != s(bool,X84) | s(bool,'F_0') = s(bool,X83) | s(bool,X82) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8231,f5003])).
fof(f14178,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,'T_0') != s(bool,X78) | s(bool,X79) != s(bool,X80) | s(bool,'F_0') != s(bool,X80) | s(bool,X77) != s(bool,X79) | s(bool,'F_0') = s(bool,X77) | s(bool,'T_0') = s(bool,X81)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14130,f586])).
fof(f586,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | s(bool,'T_0') = s(bool,X14) | s(bool,X14) != s(bool,X15)) ) | ($spl16 | $spl24)),
  inference(resolution,[],[f437,f133])).
fof(f14130,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,'T_0') != s(bool,X78) | s(bool,X79) != s(bool,X80) | s(bool,'F_0') != s(bool,X80) | s(bool,X77) != s(bool,X79) | s(bool,'F_0') = s(bool,X77) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8231,f5003])).
fof(f14175,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,'T_0') != s(bool,X54) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') != s(bool,X56) | s(bool,X53) != s(bool,X55) | s(bool,'F_0') = s(bool,X53)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14125,f705])).
fof(f14125,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,'T_0') != s(bool,X54) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') != s(bool,X56) | s(bool,X53) != s(bool,X55) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8231,f782])).
fof(f14174,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'T_0') != s(bool,X50) | s(bool,X51) != s(bool,X52) | s(bool,'F_0') != s(bool,X52) | s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X49)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f14124,f705])).
fof(f14124,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'T_0') != s(bool,X50) | s(bool,X51) != s(bool,X52) | s(bool,'F_0') != s(bool,X52) | s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8231,f782])).
fof(f14171,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X85) | s(bool,X86) != s(bool,X87) | s(bool,X87) != s(bool,X88) | s(bool,'F_0') != s(bool,X88)) ) | $spl194),
  inference(cnf_transformation,[],[f14171_D])).
fof(f14171_D,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X85) | s(bool,X86) != s(bool,X87) | s(bool,X87) != s(bool,X88) | s(bool,'F_0') != s(bool,X88)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f14163,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X85) != s(bool,X86) | s(bool,X87) != s(bool,X88) | s(bool,'F_0') != s(bool,X88) | s(bool,X86) != s(bool,X87) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X89)) ) | ($spl6 | $spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f14108])).
fof(f14108,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X85) != s(bool,X86) | s(bool,X87) != s(bool,X88) | s(bool,'F_0') != s(bool,X88) | s(bool,X86) != s(bool,X87) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8231,f5003])).
fof(f14033,plain,(
  ( ! [X158,X156,X159,X155] : (s(bool,X158) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,X158) = s(bool,X159) | s(bool,'F_0') != s(bool,X159)) ) | $spl192),
  inference(cnf_transformation,[],[f14033_D])).
fof(f14033_D,plain,(
  ( ! [X158,X156,X159,X155] : (s(bool,X158) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,X158) = s(bool,X159) | s(bool,'F_0') != s(bool,X159)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f14032,plain,(
  ( ! [X158,X156,X159,X157,X155] : (s(bool,'T_0') != s(bool,X157) | s(bool,X158) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,'F_0') != s(bool,X159) | s(bool,X158) = s(bool,X159)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f13997,f705])).
fof(f13997,plain,(
  ( ! [X158,X156,X159,X157,X155] : (s(bool,'T_0') != s(bool,X157) | s(bool,X158) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,'F_0') != s(bool,X159) | s(bool,X158) = s(bool,X159) | s(bool,'F_0') = s(bool,X157)) ) | ($spl16 | $spl24 | $spl32)),
  inference(superposition,[],[f3736,f4294])).
fof(f14029,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,X105) | s(bool,X104) = s(bool,X106) | s(bool,'F_0') != s(bool,X107) | s(bool,X106) = s(bool,X107) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f13983,f705])).
fof(f13983,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,X105) | s(bool,X104) = s(bool,X106) | s(bool,'F_0') != s(bool,X107) | s(bool,X106) = s(bool,X107) | s(bool,'F_0') = s(bool,X105) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl32)),
  inference(superposition,[],[f3736,f5704])).
fof(f11125,plain,(
  ( ! [X263,X260,X262,X264] : (s(bool,X262) = s(bool,X263) | s(bool,'T_0') = s(bool,X264) | s(bool,X260) = s(bool,X264) | s(bool,'F_0') != s(bool,X262) | s(bool,X260) = s(bool,X263)) ) | $spl124),
  inference(cnf_transformation,[],[f11125_D])).
fof(f11125_D,plain,(
  ( ! [X263,X260,X262,X264] : (s(bool,X262) = s(bool,X263) | s(bool,'T_0') = s(bool,X264) | s(bool,X260) = s(bool,X264) | s(bool,'F_0') != s(bool,X262) | s(bool,X260) = s(bool,X263)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f14027,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,'F_0') != s(bool,X85) | s(bool,X84) = s(bool,X85) | s(bool,X82) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f13978,f705])).
fof(f13978,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,'F_0') != s(bool,X85) | s(bool,X84) = s(bool,X85) | s(bool,'F_0') = s(bool,X83) | s(bool,X82) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) ) | ($spl16 | $spl24 | $spl32)),
  inference(superposition,[],[f3736,f5003])).
fof(f14019,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X85) | s(bool,X87) = s(bool,X88) | s(bool,'F_0') != s(bool,X88) | s(bool,X86) = s(bool,X87)) ) | $spl188),
  inference(cnf_transformation,[],[f14019_D])).
fof(f14019_D,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X85) | s(bool,X87) = s(bool,X88) | s(bool,'F_0') != s(bool,X88) | s(bool,X86) = s(bool,X87)) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f14010,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X85) != s(bool,X86) | s(bool,X86) = s(bool,X87) | s(bool,'F_0') != s(bool,X88) | s(bool,X87) = s(bool,X88) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X89)) ) | ($spl16 | $spl24 | $spl32)),
  inference(duplicate_literal_removal,[],[f13955])).
fof(f13955,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X85) != s(bool,X86) | s(bool,X86) = s(bool,X87) | s(bool,'F_0') != s(bool,X88) | s(bool,X87) = s(bool,X88) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) ) | ($spl16 | $spl24 | $spl32)),
  inference(superposition,[],[f3736,f5003])).
fof(f13614,plain,(
  ( ! [X121,X118,X120,X119] : (s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119))) | ~p(s(bool,X120)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X119),s('type.hreal.hreal',X121))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X119,X118)),s('type.hreal.hreal',X121)))))) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f13613,f113])).
fof(f13613,plain,(
  ( ! [X121,X118,X120,X119] : (~p(s(bool,X120)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X119),s('type.hreal.hreal',X121))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X119,X118)),s('type.hreal.hreal',X121))))) | s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X118),s('type.hreal.hreal',X119)))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f13572,f436])).
fof(f13612,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,'F_0') = s(bool,X116) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal',X117))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X115,X114)),s('type.hreal.hreal',X117))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115)))) ) | $spl6),
  inference(subsumption_resolution,[],[f13611,f4678])).
fof(f13611,plain,(
  ( ! [X116,X114,X117,X115] : (s(bool,'F_0') = s(bool,X116) | ~p(s(bool,X116)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal',X117))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X115,X114)),s('type.hreal.hreal',X117))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115)))) ) | $spl6),
  inference(forward_demodulation,[],[f13571,f113])).
fof(f13571,plain,(
  ( ! [X116,X114,X117,X115] : (~p(s(bool,X116)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal',X117))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X115,X114)),s('type.hreal.hreal',X117))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X114),s('type.hreal.hreal',X115))) | s(bool,X116) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(superposition,[],[f5121,f5704])).
fof(f11148,plain,(
  ( ! [X500,X504,X501,X503] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X501),s('type.hreal.hreal',X503))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X500),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X501,X500)),s('type.hreal.hreal',X503))))) | s(bool,'T_0') = s(bool,X504) | s(bool,X504) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X500),s('type.hreal.hreal',X501)))) ) | $spl134),
  inference(cnf_transformation,[],[f11148_D])).
fof(f11148_D,plain,(
  ( ! [X500,X504,X501,X503] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X501),s('type.hreal.hreal',X503))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X500),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X501,X500)),s('type.hreal.hreal',X503))))) | s(bool,'T_0') = s(bool,X504) | s(bool,X504) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X500),s('type.hreal.hreal',X501)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f13609,plain,(
  ( ! [X99,X97,X100,X98,X96] : (~p(s(bool,X98)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X97),s('type.hreal.hreal',X99))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X96),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X97,X96)),s('type.hreal.hreal',X99))))) | s(bool,X100) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X96),s('type.hreal.hreal',X97))) | s(bool,'T_0') = s(bool,X100)) ) | $spl24),
  inference(subsumption_resolution,[],[f13567,f436])).
fof(f13567,plain,(
  ( ! [X99,X97,X100,X98,X96] : (~p(s(bool,X98)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X97),s('type.hreal.hreal',X99))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X96),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X97,X96)),s('type.hreal.hreal',X99))))) | s(bool,'F_0') = s(bool,X98) | s(bool,X100) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X96),s('type.hreal.hreal',X97))) | s(bool,'T_0') = s(bool,X100)) ) | $spl24),
  inference(superposition,[],[f5121,f5003])).
fof(f13608,plain,(
  ( ! [X94,X92,X95,X93,X91] : (~p(s(bool,X93)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X92),s('type.hreal.hreal',X94))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X91),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X92,X91)),s('type.hreal.hreal',X94))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X91),s('type.hreal.hreal',X92))) | s(bool,X93) = s(bool,X95)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f13566,f273])).
fof(f13566,plain,(
  ( ! [X94,X92,X95,X93,X91] : (~p(s(bool,X93)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X92),s('type.hreal.hreal',X94))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X91),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X92,X91)),s('type.hreal.hreal',X94))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X91),s('type.hreal.hreal',X92))) | s(bool,X93) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) ) | $spl24),
  inference(superposition,[],[f5121,f5003])).
fof(f13604,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X69)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X68),s('type.hreal.hreal',X70))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X68,X67)),s('type.hreal.hreal',X70))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal',X68)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f13561,f436])).
fof(f13561,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X69)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X68),s('type.hreal.hreal',X70))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X68,X67)),s('type.hreal.hreal',X70))))) | s(bool,'F_0') = s(bool,X69) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal',X68)))) ) | $spl16),
  inference(superposition,[],[f5121,f782])).
fof(f13603,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X64),s('type.hreal.hreal',X66))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X64,X63)),s('type.hreal.hreal',X66))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal',X64)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f13560,f436])).
fof(f13560,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X64),s('type.hreal.hreal',X66))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X64,X63)),s('type.hreal.hreal',X66))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal',X64))) | s(bool,'F_0') = s(bool,X65)) ) | $spl16),
  inference(superposition,[],[f5121,f782])).
fof(f13430,plain,(
  ( ! [X158,X156,X159,X155] : (s(bool,X158) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,'T_0') != s(bool,X159) | s(bool,X158) != s(bool,X159)) ) | $spl186),
  inference(cnf_transformation,[],[f13430_D])).
fof(f13430_D,plain,(
  ( ! [X158,X156,X159,X155] : (s(bool,X158) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,'T_0') != s(bool,X159) | s(bool,X158) != s(bool,X159)) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f13429,plain,(
  ( ! [X158,X156,X159,X157,X155] : (s(bool,'T_0') != s(bool,X157) | s(bool,X158) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,X158) != s(bool,X159) | s(bool,'T_0') != s(bool,X159)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f13393,f705])).
fof(f13393,plain,(
  ( ! [X158,X156,X159,X157,X155] : (s(bool,'T_0') != s(bool,X157) | s(bool,X158) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,X158) != s(bool,X159) | s(bool,'T_0') != s(bool,X159) | s(bool,'F_0') = s(bool,X157)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3336,f4294])).
fof(f13426,plain,(
  ( ! [X107,X106,X104] : (s(bool,X104) = s(bool,X106) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') != s(bool,X107) | s(bool,X106) != s(bool,X107)) ) | $spl184),
  inference(cnf_transformation,[],[f13426_D])).
fof(f13426_D,plain,(
  ( ! [X107,X106,X104] : (s(bool,X104) = s(bool,X106) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') != s(bool,X107) | s(bool,X106) != s(bool,X107)) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f13425,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,X105) | s(bool,X104) = s(bool,X106) | s(bool,X106) != s(bool,X107) | s(bool,'T_0') != s(bool,X107) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f13379,f705])).
fof(f13379,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,X105) | s(bool,X104) = s(bool,X106) | s(bool,X106) != s(bool,X107) | s(bool,'T_0') != s(bool,X107) | s(bool,'F_0') = s(bool,X105) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3336,f5704])).
fof(f11111,plain,(
  ( ! [X165,X166,X164,X162] : (s(bool,'T_0') != s(bool,X164) | s(bool,'T_0') = s(bool,X166) | s(bool,X162) = s(bool,X166) | s(bool,X164) != s(bool,X165) | s(bool,X162) = s(bool,X165)) ) | $spl116),
  inference(cnf_transformation,[],[f11111_D])).
fof(f11111_D,plain,(
  ( ! [X165,X166,X164,X162] : (s(bool,'T_0') != s(bool,X164) | s(bool,'T_0') = s(bool,X166) | s(bool,X162) = s(bool,X166) | s(bool,X164) != s(bool,X165) | s(bool,X162) = s(bool,X165)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f13423,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,X84) != s(bool,X85) | s(bool,'T_0') != s(bool,X85) | s(bool,X82) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f13374,f705])).
fof(f13374,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,X84) != s(bool,X85) | s(bool,'T_0') != s(bool,X85) | s(bool,'F_0') = s(bool,X83) | s(bool,X82) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3336,f5003])).
fof(f13415,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') != s(bool,X88) | s(bool,X87) != s(bool,X88) | s(bool,X86) = s(bool,X87)) ) | $spl182),
  inference(cnf_transformation,[],[f13415_D])).
fof(f13415_D,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') != s(bool,X88) | s(bool,X87) != s(bool,X88) | s(bool,X86) = s(bool,X87)) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f13406,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X85) != s(bool,X86) | s(bool,X86) = s(bool,X87) | s(bool,X87) != s(bool,X88) | s(bool,'T_0') != s(bool,X88) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X89)) ) | ($spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f13351])).
fof(f13351,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X85) != s(bool,X86) | s(bool,X86) = s(bool,X87) | s(bool,X87) != s(bool,X88) | s(bool,'T_0') != s(bool,X88) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3336,f5003])).
fof(f10916,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,X60) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X58),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X58),s('type.hreal.hreal',X59))))) | s(bool,'T_0') = s(bool,X61) | s(bool,X60) = s(bool,X61)) ) | $spl102),
  inference(cnf_transformation,[],[f10916_D])).
fof(f10916_D,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,X60) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X58),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X58),s('type.hreal.hreal',X59))))) | s(bool,'T_0') = s(bool,X61) | s(bool,X60) = s(bool,X61)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f13331,plain,(
  ( ! [X158,X156,X159,X157,X155] : (s(bool,'T_0') != s(bool,X157) | s(bool,X158) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,X158) = s(bool,X159) | s(bool,'T_0') = s(bool,X159)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f13297,f705])).
fof(f13297,plain,(
  ( ! [X158,X156,X159,X157,X155] : (s(bool,'T_0') != s(bool,X157) | s(bool,X158) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,X158) = s(bool,X159) | s(bool,'T_0') = s(bool,X159) | s(bool,'F_0') = s(bool,X157)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3217,f4294])).
fof(f9199,plain,(
  ( ! [X68,X66,X69] : (s(bool,'T_0') = s(bool,X68) | s(bool,X66) = s(bool,X69) | s(bool,X69) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X66) = s(bool,X68)) ) | $spl94),
  inference(cnf_transformation,[],[f9199_D])).
fof(f9199_D,plain,(
  ( ! [X68,X66,X69] : (s(bool,'T_0') = s(bool,X68) | s(bool,X66) = s(bool,X69) | s(bool,X69) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X66) = s(bool,X68)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f13328,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,X105) | s(bool,X104) = s(bool,X106) | s(bool,X106) = s(bool,X107) | s(bool,'T_0') = s(bool,X107) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f13283,f705])).
fof(f13283,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,X105) | s(bool,X104) = s(bool,X106) | s(bool,X106) = s(bool,X107) | s(bool,'T_0') = s(bool,X107) | s(bool,'F_0') = s(bool,X105) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3217,f5704])).
fof(f10917,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,X78) = s(bool,X79) | s(bool,'T_0') = s(bool,X81) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X80) | s(bool,X79) = s(bool,X80)) ) | $spl104),
  inference(cnf_transformation,[],[f10917_D])).
fof(f10917_D,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,X78) = s(bool,X79) | s(bool,'T_0') = s(bool,X81) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X80) | s(bool,X79) = s(bool,X80)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f13326,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,X84) = s(bool,X85) | s(bool,'T_0') = s(bool,X85) | s(bool,X82) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f13278,f705])).
fof(f13278,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,X82) = s(bool,X84) | s(bool,X84) = s(bool,X85) | s(bool,'T_0') = s(bool,X85) | s(bool,'F_0') = s(bool,X83) | s(bool,X82) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3217,f5003])).
fof(f11102,plain,(
  ( ! [X116,X117,X115,X113] : (s(bool,X113) = s(bool,X115) | s(bool,'T_0') = s(bool,X117) | s(bool,X113) = s(bool,X117) | s(bool,'F_0') = s(bool,X116) | s(bool,X115) != s(bool,X116)) ) | $spl112),
  inference(cnf_transformation,[],[f11102_D])).
fof(f11102_D,plain,(
  ( ! [X116,X117,X115,X113] : (s(bool,X113) = s(bool,X115) | s(bool,'T_0') = s(bool,X117) | s(bool,X113) = s(bool,X117) | s(bool,'F_0') = s(bool,X116) | s(bool,X115) != s(bool,X116)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f13310,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X85) != s(bool,X86) | s(bool,X86) = s(bool,X87) | s(bool,X87) = s(bool,X88) | s(bool,'T_0') = s(bool,X88) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X89)) ) | ($spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f13255])).
fof(f13255,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X85) != s(bool,X86) | s(bool,X86) = s(bool,X87) | s(bool,X87) = s(bool,X88) | s(bool,'T_0') = s(bool,X88) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3217,f5003])).
fof(f13235,plain,(
  ( ! [X158,X156,X159,X155] : (s(bool,X158) != s(bool,X159) | s(bool,X158) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,'F_0') = s(bool,X159)) ) | $spl180),
  inference(cnf_transformation,[],[f13235_D])).
fof(f13235_D,plain,(
  ( ! [X158,X156,X159,X155] : (s(bool,X158) != s(bool,X159) | s(bool,X158) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,'F_0') = s(bool,X159)) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f13234,plain,(
  ( ! [X158,X156,X159,X157,X155] : (s(bool,'T_0') != s(bool,X157) | s(bool,X158) != s(bool,X159) | s(bool,'F_0') = s(bool,X159) | s(bool,X158) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156)))))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f13198,f705])).
fof(f13198,plain,(
  ( ! [X158,X156,X159,X157,X155] : (s(bool,'T_0') != s(bool,X157) | s(bool,X158) != s(bool,X159) | s(bool,'F_0') = s(bool,X159) | s(bool,X158) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,'F_0') = s(bool,X157)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3156,f4294])).
fof(f13232,plain,(
  ( ! [X125,X126] : (s(bool,X125) != s(bool,X126) | s(bool,'F_0') = s(bool,X126)) ) | $spl178),
  inference(cnf_transformation,[],[f13232_D])).
fof(f13232_D,plain,(
  ( ! [X125,X126] : (s(bool,X125) != s(bool,X126) | s(bool,'F_0') = s(bool,X126)) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f11161,plain,(
  ( ! [X128,X126,X129] : (s(bool,'T_0') != s(bool,X126) | s(bool,'T_0') = s(bool,X129) | s(bool,X128) = s(bool,X129) | s(bool,X126) = s(bool,X128)) ) | $spl136),
  inference(cnf_transformation,[],[f11161_D])).
fof(f11161_D,plain,(
  ( ! [X128,X126,X129] : (s(bool,'T_0') != s(bool,X126) | s(bool,'T_0') = s(bool,X129) | s(bool,X128) = s(bool,X129) | s(bool,X126) = s(bool,X128)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f13231,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,'T_0') != s(bool,X124) | s(bool,X125) != s(bool,X126) | s(bool,'F_0') = s(bool,X126) | s(bool,X124) = s(bool,X127) | s(bool,X127) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f13189,f729])).
fof(f729,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f705,f68])).
fof(f13189,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,'T_0') != s(bool,X124) | s(bool,X125) != s(bool,X126) | s(bool,'F_0') = s(bool,X126) | s(bool,'F_0') = s(bool,X125) | s(bool,X124) = s(bool,X127) | s(bool,X127) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3156,f5003])).
fof(f9376,plain,(
  ( ! [X68,X66,X69] : (s(bool,'F_0') = s(bool,X68) | s(bool,X66) = s(bool,X69) | s(bool,X69) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X66) != s(bool,X68)) ) | $spl100),
  inference(cnf_transformation,[],[f9376_D])).
fof(f9376_D,plain,(
  ( ! [X68,X66,X69] : (s(bool,'F_0') = s(bool,X68) | s(bool,X66) = s(bool,X69) | s(bool,X69) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X66) != s(bool,X68)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f13230,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,X105) | s(bool,X106) != s(bool,X107) | s(bool,'F_0') = s(bool,X107) | s(bool,X104) = s(bool,X106) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f13184,f705])).
fof(f13184,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,X105) | s(bool,X106) != s(bool,X107) | s(bool,'F_0') = s(bool,X107) | s(bool,X104) = s(bool,X106) | s(bool,'F_0') = s(bool,X105) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16)),
  inference(superposition,[],[f3156,f5704])).
fof(f13228,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,X84) != s(bool,X85) | s(bool,'F_0') = s(bool,X85) | s(bool,X82) = s(bool,X84) | s(bool,X82) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f13179,f705])).
fof(f13179,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,X84) != s(bool,X85) | s(bool,'F_0') = s(bool,X85) | s(bool,X82) = s(bool,X84) | s(bool,'F_0') = s(bool,X83) | s(bool,X82) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3156,f5003])).
fof(f13220,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X85) | s(bool,X86) = s(bool,X87) | s(bool,X87) != s(bool,X88) | s(bool,'F_0') = s(bool,X88)) ) | $spl176),
  inference(cnf_transformation,[],[f13220_D])).
fof(f13220_D,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X85) | s(bool,X86) = s(bool,X87) | s(bool,X87) != s(bool,X88) | s(bool,'F_0') = s(bool,X88)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f13211,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X85) != s(bool,X86) | s(bool,X87) != s(bool,X88) | s(bool,'F_0') = s(bool,X88) | s(bool,X86) = s(bool,X87) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X89)) ) | ($spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f13156])).
fof(f13156,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X85) != s(bool,X86) | s(bool,X87) != s(bool,X88) | s(bool,'F_0') = s(bool,X88) | s(bool,X86) = s(bool,X87) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3156,f5003])).
fof(f11244,plain,(
  ( ! [X99,X100,X98,X96] : (s(bool,'T_0') = s(bool,X98) | s(bool,'F_0') = s(bool,X100) | s(bool,X96) = s(bool,X100) | s(bool,X98) != s(bool,X99) | s(bool,X96) = s(bool,X99)) ) | $spl144),
  inference(cnf_transformation,[],[f11244_D])).
fof(f11244_D,plain,(
  ( ! [X99,X100,X98,X96] : (s(bool,'T_0') = s(bool,X98) | s(bool,'F_0') = s(bool,X100) | s(bool,X96) = s(bool,X100) | s(bool,X98) != s(bool,X99) | s(bool,X96) = s(bool,X99)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f13132,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,X67) = s(bool,X69) | s(bool,X69) != s(bool,X70) | s(bool,'T_0') = s(bool,X70) | s(bool,X67) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X68)) ) | ($spl24 | $spl32)),
  inference(subsumption_resolution,[],[f13083,f68])).
fof(f13083,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,'F_0') != s(bool,X68) | s(bool,X67) = s(bool,X69) | s(bool,X69) != s(bool,X70) | s(bool,'T_0') = s(bool,X70) | s(bool,X67) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X68)) ) | ($spl24 | $spl32)),
  inference(superposition,[],[f2859,f5003])).
fof(f2859,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X15) | s(bool,X15) = s(bool,X16) | s(bool,X16) != s(bool,X17) | s(bool,'T_0') = s(bool,X17)) ) | ($spl24 | $spl32)),
  inference(resolution,[],[f2819,f437])).
fof(f13125,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,X86) != s(bool,X87) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X89) | s(bool,X88) != s(bool,X89) | s(bool,X87) = s(bool,X88)) ) | $spl174),
  inference(cnf_transformation,[],[f13125_D])).
fof(f13125_D,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,X86) != s(bool,X87) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X89) | s(bool,X88) != s(bool,X89) | s(bool,X87) = s(bool,X88)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f13115,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X86) != s(bool,X87) | s(bool,X87) = s(bool,X88) | s(bool,X88) != s(bool,X89) | s(bool,'T_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X86)) ) | ($spl24 | $spl32)),
  inference(duplicate_literal_removal,[],[f13058])).
fof(f13058,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X86) != s(bool,X87) | s(bool,X87) = s(bool,X88) | s(bool,X88) != s(bool,X89) | s(bool,'T_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X90) | s(bool,'F_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X86)) ) | ($spl24 | $spl32)),
  inference(superposition,[],[f2859,f5003])).
fof(f10941,plain,(
  ( ! [X111,X109,X112,X108] : (s(bool,X108) = s(bool,X109) | s(bool,X108) = s(bool,X112) | s(bool,X111) = s(bool,X112) | s(bool,'F_0') = s(bool,X111) | s(bool,'F_0') = s(bool,X109)) ) | $spl106),
  inference(cnf_transformation,[],[f10941_D])).
fof(f10941_D,plain,(
  ( ! [X111,X109,X112,X108] : (s(bool,X108) = s(bool,X109) | s(bool,X108) = s(bool,X112) | s(bool,X111) = s(bool,X112) | s(bool,'F_0') = s(bool,X111) | s(bool,'F_0') = s(bool,X109)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f13032,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,X67) = s(bool,X69) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,X67) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X68)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f12985,f68])).
fof(f12985,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,'F_0') != s(bool,X68) | s(bool,X67) = s(bool,X69) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,X70) | s(bool,X67) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X68)) ) | ($spl16 | $spl24 | $spl32)),
  inference(superposition,[],[f2858,f5003])).
fof(f2858,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,X12) | s(bool,X12) = s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'F_0') = s(bool,X14)) ) | ($spl16 | $spl32)),
  inference(resolution,[],[f2819,f317])).
fof(f13017,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X86) != s(bool,X87) | s(bool,X87) = s(bool,X88) | s(bool,X88) = s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X86)) ) | ($spl16 | $spl24 | $spl32)),
  inference(duplicate_literal_removal,[],[f12960])).
fof(f12960,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X86) != s(bool,X87) | s(bool,X87) = s(bool,X88) | s(bool,X88) = s(bool,X89) | s(bool,'F_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X90) | s(bool,'F_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X86)) ) | ($spl16 | $spl24 | $spl32)),
  inference(superposition,[],[f2858,f5003])).
fof(f11123,plain,(
  ( ! [X255,X253,X251,X254] : (s(bool,'F_0') = s(bool,X253) | s(bool,'T_0') = s(bool,X255) | s(bool,X251) = s(bool,X255) | s(bool,X251) != s(bool,X254) | s(bool,X253) = s(bool,X254)) ) | $spl122),
  inference(cnf_transformation,[],[f11123_D])).
fof(f11123_D,plain,(
  ( ! [X255,X253,X251,X254] : (s(bool,'F_0') = s(bool,X253) | s(bool,'T_0') = s(bool,X255) | s(bool,X251) = s(bool,X255) | s(bool,X251) != s(bool,X254) | s(bool,X253) = s(bool,X254)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f12914,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,X67) != s(bool,X69) | s(bool,X69) = s(bool,X70) | s(bool,'T_0') = s(bool,X70) | s(bool,X67) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X68)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12866,f68])).
fof(f12866,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,'F_0') != s(bool,X68) | s(bool,X67) != s(bool,X69) | s(bool,X69) = s(bool,X70) | s(bool,'T_0') = s(bool,X70) | s(bool,X67) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X68)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f1426,f5003])).
fof(f1426,plain,(
  ( ! [X39,X37,X38] : (s(bool,'F_0') != s(bool,X38) | s(bool,X37) != s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f705,f672])).
fof(f11113,plain,(
  ( ! [X167,X171,X169,X170] : (s(bool,X167) != s(bool,X169) | s(bool,'T_0') = s(bool,X171) | s(bool,X167) = s(bool,X171) | s(bool,'T_0') = s(bool,X170) | s(bool,X169) != s(bool,X170)) ) | $spl118),
  inference(cnf_transformation,[],[f11113_D])).
fof(f11113_D,plain,(
  ( ! [X167,X171,X169,X170] : (s(bool,X167) != s(bool,X169) | s(bool,'T_0') = s(bool,X171) | s(bool,X167) = s(bool,X171) | s(bool,'T_0') = s(bool,X170) | s(bool,X169) != s(bool,X170)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f12898,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X86) != s(bool,X87) | s(bool,X87) != s(bool,X88) | s(bool,X88) = s(bool,X89) | s(bool,'T_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X86)) ) | ($spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f12841])).
fof(f12841,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X86) != s(bool,X87) | s(bool,X87) != s(bool,X88) | s(bool,X88) = s(bool,X89) | s(bool,'T_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X90) | s(bool,'F_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X86)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f1426,f5003])).
fof(f12803,plain,(
  ( ! [X76,X74,X72,X75,X73] : (p(s(bool,X73)) | s(bool,X72) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X74),s('type.hreal.hreal',X75))) | s('type.hreal.hreal',X75) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X74),s('type.hreal.hreal',sK3(X75,X74)))) | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X72)) ) | ($spl24 | $spl32)),
  inference(subsumption_resolution,[],[f12769,f2819])).
fof(f12769,plain,(
  ( ! [X76,X74,X72,X75,X73] : (p(s(bool,X73)) | s(bool,X72) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X74),s('type.hreal.hreal',X75))) | s('type.hreal.hreal',X75) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X74),s('type.hreal.hreal',sK3(X75,X74)))) | s(bool,X73) = s(bool,X76) | s(bool,'F_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X72)) ) | $spl24),
  inference(superposition,[],[f5061,f5003])).
fof(f11268,plain,(
  ( ! [X321,X323,X322,X319] : (s(bool,X319) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X321),s('type.hreal.hreal',X322))) | s(bool,'F_0') = s(bool,X323) | s(bool,X319) = s(bool,X323) | s('type.hreal.hreal',X322) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X321),s('type.hreal.hreal',sK3(X322,X321))))) ) | $spl154),
  inference(cnf_transformation,[],[f11268_D])).
fof(f11268_D,plain,(
  ( ! [X321,X323,X322,X319] : (s(bool,X319) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X321),s('type.hreal.hreal',X322))) | s(bool,'F_0') = s(bool,X323) | s(bool,X319) = s(bool,X323) | s('type.hreal.hreal',X322) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X321),s('type.hreal.hreal',sK3(X322,X321))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f12802,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,X68)) | s(bool,X67) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal',X70))) | s('type.hreal.hreal',X70) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal',sK3(X70,X69)))) | s(bool,X67) = s(bool,X71) | s(bool,'F_0') = s(bool,X71)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12768,f133])).
fof(f12768,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,X68)) | s(bool,X67) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal',X70))) | s('type.hreal.hreal',X70) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X69),s('type.hreal.hreal',sK3(X70,X69)))) | s(bool,X67) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X68)) ) | $spl24),
  inference(superposition,[],[f5061,f5003])).
fof(f12796,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,X34) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X36),s('type.hreal.hreal',X37))) | s('type.hreal.hreal',X37) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X36),s('type.hreal.hreal',sK3(X37,X36)))) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12760,f133])).
fof(f12760,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,X34) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X36),s('type.hreal.hreal',X37))) | s('type.hreal.hreal',X37) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X36),s('type.hreal.hreal',sK3(X37,X36)))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl24),
  inference(superposition,[],[f5061,f672])).
fof(f12795,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,X30) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X32),s('type.hreal.hreal',X33))) | s('type.hreal.hreal',X33) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X32),s('type.hreal.hreal',sK3(X33,X32)))) | s(bool,'T_0') = s(bool,X30)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12759,f133])).
fof(f12759,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,X30) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X32),s('type.hreal.hreal',X33))) | s('type.hreal.hreal',X33) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X32),s('type.hreal.hreal',sK3(X33,X32)))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl24),
  inference(superposition,[],[f5061,f672])).
fof(f12729,plain,(
  ( ! [X154,X153] : (s(bool,'T_0') != s(bool,X153) | s(bool,X154) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X154)) ) | ($spl4 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12682,f705])).
fof(f12682,plain,(
  ( ! [X154,X153] : (s(bool,'T_0') != s(bool,X153) | s(bool,'F_0') = s(bool,X153) | s(bool,X154) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X154)) ) | ($spl4 | $spl24)),
  inference(superposition,[],[f5192,f112])).
fof(f12689,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,X129) != s(bool,X130) | s(bool,'F_0') = s(bool,X130) | s(bool,'F_0') = s(bool,X131) | s(bool,X132) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X129) = s(bool,X132)) ) | ($spl6 | $spl24)),
  inference(duplicate_literal_removal,[],[f12675])).
fof(f12675,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,X129) != s(bool,X130) | s(bool,'F_0') = s(bool,X130) | s(bool,'F_0') = s(bool,X131) | s(bool,'F_0') = s(bool,X131) | s(bool,X132) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X129) = s(bool,X132)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f5192,f5704])).
fof(f12690,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,X124) != s(bool,X125) | s(bool,'F_0') = s(bool,X125) | s(bool,'F_0') = s(bool,X126) | s(bool,X124) = s(bool,X127) | s(bool,X127) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f12674])).
fof(f12674,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,X124) != s(bool,X125) | s(bool,'F_0') = s(bool,X125) | s(bool,'F_0') = s(bool,X126) | s(bool,'F_0') = s(bool,X126) | s(bool,X124) = s(bool,X127) | s(bool,X127) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) ) | $spl24),
  inference(superposition,[],[f5192,f5003])).
fof(f12724,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s(bool,'T_0') != s(bool,X63) | s(bool,X62) = s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,X65) = s(bool,X66) | s(bool,'F_0') = s(bool,X65) | s(bool,X62) = s(bool,X66)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12660,f705])).
fof(f12660,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s(bool,'T_0') != s(bool,X63) | s(bool,'F_0') = s(bool,X63) | s(bool,X62) = s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,X65) = s(bool,X66) | s(bool,'F_0') = s(bool,X65) | s(bool,X62) = s(bool,X66)) ) | $spl24),
  inference(superposition,[],[f5192,f5003])).
fof(f11271,plain,(
  ( ! [X344,X346,X347] : (s(bool,X344) != s(bool,X346) | s(bool,'F_0') = s(bool,X347) | s(bool,X344) = s(bool,X347) | s(bool,X346) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl156),
  inference(cnf_transformation,[],[f11271_D])).
fof(f11271_D,plain,(
  ( ! [X344,X346,X347] : (s(bool,X344) != s(bool,X346) | s(bool,'F_0') = s(bool,X347) | s(bool,X344) = s(bool,X347) | s(bool,X346) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f12691,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,X101) != s(bool,X102) | s(bool,'F_0') = s(bool,X100) | s(bool,X102) = s(bool,X103) | s(bool,'F_0') = s(bool,X103) | s(bool,X101) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl24)),
  inference(duplicate_literal_removal,[],[f12627])).
fof(f12627,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,X101) != s(bool,X102) | s(bool,'F_0') = s(bool,X100) | s(bool,X102) = s(bool,X103) | s(bool,'F_0') = s(bool,X103) | s(bool,'F_0') = s(bool,X100) | s(bool,X101) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f5192,f5704])).
fof(f12692,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,X78) != s(bool,X79) | s(bool,'F_0') = s(bool,X77) | s(bool,X79) = s(bool,X80) | s(bool,'F_0') = s(bool,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f12622])).
fof(f12622,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,X78) != s(bool,X79) | s(bool,'F_0') = s(bool,X77) | s(bool,X79) = s(bool,X80) | s(bool,'F_0') = s(bool,X80) | s(bool,'F_0') = s(bool,X77) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) ) | $spl24),
  inference(superposition,[],[f5192,f5003])).
fof(f11273,plain,(
  ( ! [X356,X357,X353,X355] : (s(bool,X353) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X355),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X355),s('type.hreal.hreal',X356))))) | s(bool,'F_0') = s(bool,X357) | s(bool,X353) = s(bool,X357)) ) | $spl158),
  inference(cnf_transformation,[],[f11273_D])).
fof(f11273_D,plain,(
  ( ! [X356,X357,X353,X355] : (s(bool,X353) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X355),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X355),s('type.hreal.hreal',X356))))) | s(bool,'F_0') = s(bool,X357) | s(bool,X353) = s(bool,X357)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f12693,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s(bool,X60) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X58),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X58),s('type.hreal.hreal',X59))))) | s(bool,'F_0') = s(bool,X57) | s(bool,X60) = s(bool,X61) | s(bool,'F_0') = s(bool,X61)) ) | ($spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f12618])).
fof(f12618,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s(bool,X60) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X58),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X58),s('type.hreal.hreal',X59))))) | s(bool,'F_0') = s(bool,X57) | s(bool,X60) = s(bool,X61) | s(bool,'F_0') = s(bool,X61) | s(bool,'F_0') = s(bool,X57)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f5192,f4294])).
fof(f12602,plain,(
  ( ! [X158,X156,X159,X157,X155] : (s(bool,'T_0') != s(bool,X157) | s(bool,'F_0') = s(bool,X158) | s(bool,X159) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,X158) = s(bool,X159)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12566,f705])).
fof(f12566,plain,(
  ( ! [X158,X156,X159,X157,X155] : (s(bool,'T_0') != s(bool,X157) | s(bool,'F_0') = s(bool,X158) | s(bool,X159) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X155),s('type.hreal.hreal',X156))))) | s(bool,X158) = s(bool,X159) | s(bool,'F_0') = s(bool,X157)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f5004,f4294])).
fof(f12600,plain,(
  ( ! [X125,X126] : (s(bool,'F_0') != s(bool,X126) | s(bool,X125) = s(bool,X126)) ) | $spl172),
  inference(cnf_transformation,[],[f12600_D])).
fof(f12600_D,plain,(
  ( ! [X125,X126] : (s(bool,'F_0') != s(bool,X126) | s(bool,X125) = s(bool,X126)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f12599,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,'T_0') != s(bool,X124) | s(bool,'F_0') != s(bool,X126) | s(bool,X125) = s(bool,X126) | s(bool,X124) = s(bool,X127) | s(bool,X127) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f12557,f3847])).
fof(f3847,plain,(
  ( ! [X70,X71] : (s(bool,'F_0') != s(bool,X70) | s(bool,'F_0') != s(bool,X71) | s(bool,X70) = s(bool,X71)) ) | ($spl24 | $spl32)),
  inference(resolution,[],[f2861,f84])).
fof(f2861,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X22)) | s(bool,X20) = s(bool,X21) | s(bool,X21) != s(bool,X22) | s(bool,'F_0') != s(bool,X20)) ) | ($spl24 | $spl32)),
  inference(resolution,[],[f2819,f458])).
fof(f12557,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,'T_0') != s(bool,X124) | s(bool,'F_0') = s(bool,X125) | s(bool,'F_0') != s(bool,X126) | s(bool,X125) = s(bool,X126) | s(bool,X124) = s(bool,X127) | s(bool,X127) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f5004,f5003])).
fof(f12598,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,X105) | s(bool,'F_0') = s(bool,X106) | s(bool,X104) != s(bool,X107) | s(bool,X106) = s(bool,X107) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12552,f705])).
fof(f12552,plain,(
  ( ! [X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,X105) | s(bool,'F_0') = s(bool,X106) | s(bool,X104) != s(bool,X107) | s(bool,X106) = s(bool,X107) | s(bool,'F_0') = s(bool,X105) | s(bool,X104) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f5004,f5704])).
fof(f3204,plain,(
  ( ! [X52,X51,X49] : (s(bool,X49) = s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X52) | s(bool,X51) != s(bool,X52)) ) | $spl48),
  inference(cnf_transformation,[],[f3204_D])).
fof(f3204_D,plain,(
  ( ! [X52,X51,X49] : (s(bool,X49) = s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X52) | s(bool,X51) != s(bool,X52)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f12597,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,'T_0') != s(bool,X101) | s(bool,'F_0') = s(bool,X102) | s(bool,X100) != s(bool,X103) | s(bool,X102) = s(bool,X103) | s(bool,'F_0') = s(bool,X100)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12551,f727])).
fof(f12551,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,'T_0') != s(bool,X101) | s(bool,'F_0') = s(bool,X102) | s(bool,X100) != s(bool,X103) | s(bool,X102) = s(bool,X103) | s(bool,'F_0') = s(bool,X100) | s(bool,X101) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f5004,f5704])).
fof(f12596,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,'F_0') = s(bool,X84) | s(bool,X82) != s(bool,X85) | s(bool,X84) = s(bool,X85) | s(bool,X82) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12547,f705])).
fof(f12547,plain,(
  ( ! [X85,X83,X86,X84,X82] : (s(bool,'T_0') != s(bool,X83) | s(bool,'F_0') = s(bool,X84) | s(bool,X82) != s(bool,X85) | s(bool,X84) = s(bool,X85) | s(bool,'F_0') = s(bool,X83) | s(bool,X82) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f5004,f5003])).
fof(f12595,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,'T_0') != s(bool,X78) | s(bool,'F_0') = s(bool,X79) | s(bool,X77) != s(bool,X80) | s(bool,X79) = s(bool,X80) | s(bool,'F_0') = s(bool,X77) | s(bool,'T_0') = s(bool,X81)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12546,f586])).
fof(f12546,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,'T_0') != s(bool,X78) | s(bool,'F_0') = s(bool,X79) | s(bool,X77) != s(bool,X80) | s(bool,X79) = s(bool,X80) | s(bool,'F_0') = s(bool,X77) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f5004,f5003])).
fof(f12593,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,'T_0') != s(bool,X54) | s(bool,'F_0') = s(bool,X55) | s(bool,X53) != s(bool,X56) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12541,f705])).
fof(f12541,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,'T_0') != s(bool,X54) | s(bool,'F_0') = s(bool,X55) | s(bool,X53) != s(bool,X56) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f5004,f782])).
fof(f12592,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'T_0') != s(bool,X50) | s(bool,'F_0') = s(bool,X51) | s(bool,X49) != s(bool,X52) | s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X49)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12540,f705])).
fof(f12540,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'T_0') != s(bool,X50) | s(bool,'F_0') = s(bool,X51) | s(bool,X49) != s(bool,X52) | s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f5004,f782])).
fof(f12579,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X87) | s(bool,X86) != s(bool,X88) | s(bool,X87) = s(bool,X88) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X89)) ) | ($spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f12524])).
fof(f12524,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X87) | s(bool,X86) != s(bool,X88) | s(bool,X87) = s(bool,X88) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f5004,f5003])).
fof(f12501,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') != s(bool,X72) | s('type.hreal.hreal',X71) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X70),s('type.hreal.hreal',sK3(X71,X70)))) | s(bool,X73) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X70),s('type.hreal.hreal',X71))) | s(bool,'T_0') = s(bool,X73)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12476,f705])).
fof(f12476,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') != s(bool,X72) | s('type.hreal.hreal',X71) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X70),s('type.hreal.hreal',sK3(X71,X70)))) | s(bool,'F_0') = s(bool,X72) | s(bool,X73) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X70),s('type.hreal.hreal',X71))) | s(bool,'T_0') = s(bool,X73)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f279,f5003])).
fof(f8996,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X30),s('type.hreal.hreal',X31))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X30),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X31),s('type.hreal.hreal',X32)))))))) ) | $spl92),
  inference(cnf_transformation,[],[f8996_D])).
fof(f8996_D,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X30),s('type.hreal.hreal',X31))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X30),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X31),s('type.hreal.hreal',X32)))))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f12425,plain,(
  ( ! [X76,X74,X72,X75,X73] : (p(s(bool,X75)) | s(bool,X75) = s(bool,X76) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X72),s('type.hreal.hreal',X73))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X72),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X73),s('type.hreal.hreal',X74)))))))) ) | ($spl24 | $spl32)),
  inference(subsumption_resolution,[],[f12404,f2819])).
fof(f12404,plain,(
  ( ! [X76,X74,X72,X75,X73] : (p(s(bool,X75)) | s(bool,X75) = s(bool,X76) | s(bool,'F_0') = s(bool,X76) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X72),s('type.hreal.hreal',X73))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X72),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X73),s('type.hreal.hreal',X74)))))))) ) | $spl24),
  inference(superposition,[],[f8633,f5003])).
fof(f12424,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,X70)) | s(bool,X71) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal',X68))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X68),s('type.hreal.hreal',X69))))))) | s(bool,'F_0') = s(bool,X71)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12403,f133])).
fof(f12403,plain,(
  ( ! [X70,X68,X71,X69,X67] : (p(s(bool,X70)) | s(bool,X71) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal',X68))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X68),s('type.hreal.hreal',X69))))))) | s(bool,'F_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X70)) ) | $spl24),
  inference(superposition,[],[f8633,f5003])).
fof(f12418,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X37)) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X34),s('type.hreal.hreal',X35))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X34),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X35),s('type.hreal.hreal',X36)))))))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12395,f133])).
fof(f12395,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X37)) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X34),s('type.hreal.hreal',X35))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X34),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X35),s('type.hreal.hreal',X36)))))))) ) | $spl24),
  inference(superposition,[],[f8633,f672])).
fof(f12417,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X33)) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X30),s('type.hreal.hreal',X31))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X30),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X31),s('type.hreal.hreal',X32)))))))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12394,f133])).
fof(f12394,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X33)) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X30),s('type.hreal.hreal',X31))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X30),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X31),s('type.hreal.hreal',X32))))))) | s(bool,'T_0') = s(bool,X33)) ) | $spl24),
  inference(superposition,[],[f8633,f672])).
fof(f2627,plain,(
  ( ! [X165,X164] : (s('type.hreal.hreal',X165) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X164),s('type.hreal.hreal',sK3(X165,X164)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X164),s('type.hreal.hreal',X165)))) ) | $spl30),
  inference(cnf_transformation,[],[f2627_D])).
fof(f2627_D,plain,(
  ( ! [X165,X164] : (s('type.hreal.hreal',X165) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X164),s('type.hreal.hreal',sK3(X165,X164)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X164),s('type.hreal.hreal',X165)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f12144,plain,(
  ( ! [X88,X87,X89] : (s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X87),s('type.hreal.hreal',X88))) | ~p(s(bool,X89)) | s('type.hreal.hreal',X88) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X87),s('type.hreal.hreal',sK3(X88,X87))))) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f12143,f113])).
fof(f12143,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,X89)) | s('type.hreal.hreal',X88) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X87),s('type.hreal.hreal',sK3(X88,X87)))) | s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X87),s('type.hreal.hreal',X88)))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f12048,f436])).
fof(f12048,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,X89)) | s('type.hreal.hreal',X88) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X87),s('type.hreal.hreal',sK3(X88,X87)))) | s(bool,'F_0') = s(bool,X89) | s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X87),s('type.hreal.hreal',X88)))) ) | $spl6),
  inference(superposition,[],[f60,f5704])).
fof(f12142,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') = s(bool,X86) | s('type.hreal.hreal',X85) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal',sK3(X85,X84)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal',X85)))) ) | $spl6),
  inference(subsumption_resolution,[],[f12141,f4678])).
fof(f12141,plain,(
  ( ! [X85,X86,X84] : (s(bool,'F_0') = s(bool,X86) | ~p(s(bool,X86)) | s('type.hreal.hreal',X85) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal',sK3(X85,X84)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal',X85)))) ) | $spl6),
  inference(forward_demodulation,[],[f12047,f113])).
fof(f12047,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X86)) | s('type.hreal.hreal',X85) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal',sK3(X85,X84)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal',X85))) | s(bool,X86) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(superposition,[],[f60,f5704])).
fof(f12139,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X72)) | s('type.hreal.hreal',X71) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X70),s('type.hreal.hreal',sK3(X71,X70)))) | s(bool,X73) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X70),s('type.hreal.hreal',X71))) | s(bool,'T_0') = s(bool,X73)) ) | $spl24),
  inference(subsumption_resolution,[],[f12043,f436])).
fof(f12043,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X72)) | s('type.hreal.hreal',X71) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X70),s('type.hreal.hreal',sK3(X71,X70)))) | s(bool,'F_0') = s(bool,X72) | s(bool,X73) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X70),s('type.hreal.hreal',X71))) | s(bool,'T_0') = s(bool,X73)) ) | $spl24),
  inference(superposition,[],[f60,f5003])).
fof(f12138,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X68)) | s('type.hreal.hreal',X67) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X66),s('type.hreal.hreal',sK3(X67,X66)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X66),s('type.hreal.hreal',X67))) | s(bool,X68) = s(bool,X69)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12042,f273])).
fof(f12042,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X68)) | s('type.hreal.hreal',X67) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X66),s('type.hreal.hreal',sK3(X67,X66)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X66),s('type.hreal.hreal',X67))) | s(bool,X68) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) ) | $spl24),
  inference(superposition,[],[f60,f5003])).
fof(f12134,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,X49)) | s('type.hreal.hreal',X48) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X47),s('type.hreal.hreal',sK3(X48,X47)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X47),s('type.hreal.hreal',X48)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12037,f436])).
fof(f12037,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,X49)) | s('type.hreal.hreal',X48) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X47),s('type.hreal.hreal',sK3(X48,X47)))) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X47),s('type.hreal.hreal',X48)))) ) | $spl16),
  inference(superposition,[],[f60,f782])).
fof(f12133,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,X46)) | s('type.hreal.hreal',X45) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X44),s('type.hreal.hreal',sK3(X45,X44)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X44),s('type.hreal.hreal',X45)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f12036,f436])).
fof(f12036,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,X46)) | s('type.hreal.hreal',X45) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X44),s('type.hreal.hreal',sK3(X45,X44)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X44),s('type.hreal.hreal',X45))) | s(bool,'F_0') = s(bool,X46)) ) | $spl16),
  inference(superposition,[],[f60,f782])).
fof(f11291,plain,(
  ( ! [X509,X512,X510,X513] : (s('type.hreal.hreal',X510) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X509),s('type.hreal.hreal',X512))) | s(bool,'F_0') = s(bool,X513) | s(bool,X513) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X509),s('type.hreal.hreal',X510)))) ) | $spl162),
  inference(cnf_transformation,[],[f11291_D])).
fof(f11291_D,plain,(
  ( ! [X509,X512,X510,X513] : (s('type.hreal.hreal',X510) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X509),s('type.hreal.hreal',X512))) | s(bool,'F_0') = s(bool,X513) | s(bool,X513) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X509),s('type.hreal.hreal',X510)))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f11949,plain,(
  ( ! [X85,X83,X81,X84,X82] : (p(s(bool,X83)) | s('type.hreal.hreal',X82) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X81),s('type.hreal.hreal',X84))) | s(bool,X85) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X81),s('type.hreal.hreal',X82))) | s(bool,'F_0') = s(bool,X85)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f11896,f133])).
fof(f11896,plain,(
  ( ! [X85,X83,X81,X84,X82] : (p(s(bool,X83)) | s('type.hreal.hreal',X82) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X81),s('type.hreal.hreal',X84))) | s(bool,X85) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X81),s('type.hreal.hreal',X82))) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X83)) ) | $spl24),
  inference(superposition,[],[f61,f5003])).
fof(f11807,plain,(
  $false | ($spl36 | $spl86)),
  inference(subsumption_resolution,[],[f11744,f2828])).
fof(f11744,plain,(
  ( ! [X235,X234] : (s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X234),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X234),s('type.hreal.hreal',X235)))))) ) | ($spl36 | $spl86)),
  inference(trivial_inequality_removal,[],[f11614])).
fof(f11614,plain,(
  ( ! [X235,X233,X234] : (s(bool,X233) != s(bool,X233) | s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X234),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X234),s('type.hreal.hreal',X235)))))) ) | ($spl36 | $spl86)),
  inference(superposition,[],[f8152,f2828])).
fof(f11769,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f11460,f84])).
fof(f11460,plain,(
  p(s(bool,'F_0')) | $spl36),
  inference(backward_demodulation,[],[f2828,f8633])).
fof(f11381,plain,(
  ( ! [X249,X246,X248,X247,X245] : (s(bool,X245) != s(bool,X246) | s(bool,'T_0') != s(bool,X247) | s(bool,'T_0') = s(bool,X246) | s(bool,X248) = s(bool,X249) | s(bool,'F_0') = s(bool,X248) | s(bool,X245) = s(bool,X249)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10536,f705])).
fof(f10536,plain,(
  ( ! [X249,X246,X248,X247,X245] : (s(bool,X245) != s(bool,X246) | s(bool,'F_0') = s(bool,X247) | s(bool,'T_0') != s(bool,X247) | s(bool,'T_0') = s(bool,X246) | s(bool,X248) = s(bool,X249) | s(bool,'F_0') = s(bool,X248) | s(bool,X245) = s(bool,X249)) ) | $spl24),
  inference(superposition,[],[f753,f5003])).
fof(f753,plain,(
  ( ! [X14,X12,X13] : (s(bool,X13) != s(bool,X14) | s(bool,'F_0') = s(bool,X12) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | $spl24),
  inference(resolution,[],[f561,f437])).
fof(f10770,plain,(
  ( ! [X206,X204,X202,X205,X203] : (s(bool,X202) != s(bool,X203) | s(bool,'T_0') = s(bool,X204) | s(bool,'T_0') = s(bool,X203) | s(bool,X205) = s(bool,X206) | s(bool,'F_0') = s(bool,X205) | s(bool,X202) = s(bool,X206)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f10526])).
fof(f10526,plain,(
  ( ! [X206,X204,X202,X205,X203] : (s(bool,X202) != s(bool,X203) | s(bool,'T_0') = s(bool,X204) | s(bool,'T_0') = s(bool,X204) | s(bool,'T_0') = s(bool,X203) | s(bool,X205) = s(bool,X206) | s(bool,'F_0') = s(bool,X205) | s(bool,X202) = s(bool,X206)) ) | $spl24),
  inference(superposition,[],[f671,f5003])).
fof(f671,plain,(
  ( ! [X14,X12,X13] : (s(bool,X13) != s(bool,X14) | s(bool,'T_0') = s(bool,X12) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X14)) ) | $spl24),
  inference(resolution,[],[f461,f437])).
fof(f11372,plain,(
  ( ! [X163,X161,X159,X162,X160] : (s(bool,X159) != s(bool,X160) | p(s(bool,X160)) | p(s(bool,X161)) | s(bool,X162) = s(bool,X163) | s(bool,'F_0') = s(bool,X162) | s(bool,X159) = s(bool,X163)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10516,f133])).
fof(f10516,plain,(
  ( ! [X163,X161,X159,X162,X160] : (s(bool,X159) != s(bool,X160) | p(s(bool,X160)) | s(bool,'T_0') = s(bool,X161) | p(s(bool,X161)) | s(bool,X162) = s(bool,X163) | s(bool,'F_0') = s(bool,X162) | s(bool,X159) = s(bool,X163)) ) | $spl24),
  inference(superposition,[],[f532,f5003])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X0)) | s(bool,X1) = s(bool,X2) | p(s(bool,X2))) ) | $spl24),
  inference(resolution,[],[f458,f62])).
fof(f11363,plain,(
  ( ! [X540,X542,X538,X541,X539] : (p(s(bool,X541)) | s(bool,X541) = s(bool,X542) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X538),s('type.hreal.hreal',X539))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X538),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X539),s('type.hreal.hreal',X540)))))))) ) | ($spl24 | $spl32)),
  inference(subsumption_resolution,[],[f10477,f2819])).
fof(f10477,plain,(
  ( ! [X540,X542,X538,X541,X539] : (p(s(bool,X541)) | s(bool,X541) = s(bool,X542) | s(bool,'F_0') = s(bool,X542) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X538),s('type.hreal.hreal',X539))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X538),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X539),s('type.hreal.hreal',X540)))))))) ) | $spl24),
  inference(superposition,[],[f8633,f5003])).
fof(f11340,plain,(
  ( ! [X321,X323,X320,X322,X319] : (p(s(bool,X320)) | s(bool,X319) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X321),s('type.hreal.hreal',X322))) | s('type.hreal.hreal',X322) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X321),s('type.hreal.hreal',sK3(X322,X321)))) | s(bool,X320) = s(bool,X323) | s(bool,'T_0') = s(bool,X319)) ) | ($spl24 | $spl32)),
  inference(subsumption_resolution,[],[f10419,f2819])).
fof(f10419,plain,(
  ( ! [X321,X323,X320,X322,X319] : (p(s(bool,X320)) | s(bool,X319) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X321),s('type.hreal.hreal',X322))) | s('type.hreal.hreal',X322) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X321),s('type.hreal.hreal',sK3(X322,X321)))) | s(bool,X320) = s(bool,X323) | s(bool,'F_0') = s(bool,X323) | s(bool,'T_0') = s(bool,X319)) ) | $spl24),
  inference(superposition,[],[f5061,f5003])).
fof(f10775,plain,(
  ( ! [X300,X298,X301,X297,X299] : (s(bool,X298) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X299),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X299),s('type.hreal.hreal',X300))))) | s(bool,'T_0') = s(bool,X297) | s(bool,X298) = s(bool,X301) | s(bool,'F_0') = s(bool,X301)) ) | ($spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f10414])).
fof(f10414,plain,(
  ( ! [X300,X298,X301,X297,X299] : (s(bool,X298) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X299),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X299),s('type.hreal.hreal',X300))))) | s(bool,'T_0') = s(bool,X297) | s(bool,X298) = s(bool,X301) | s(bool,'F_0') = s(bool,X301) | s(bool,'T_0') = s(bool,X297)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f4271,f5003])).
fof(f4271,plain,(
  ( ! [X90,X88,X89] : (s(bool,X90) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal',X89))))) | s(bool,'T_0') = s(bool,X90)) ) | ($spl16 | $spl24)),
  inference(trivial_inequality_removal,[],[f4232])).
fof(f4232,plain,(
  ( ! [X90,X88,X89] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,X90) | s(bool,X90) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal',X89)))))) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f586,f4167])).
fof(f10776,plain,(
  ( ! [X273,X275,X276,X274] : (s(bool,X274) != s(bool,X275) | s(bool,X275) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X273) | s(bool,X274) = s(bool,X276) | s(bool,'F_0') = s(bool,X276)) ) | ($spl6 | $spl24)),
  inference(duplicate_literal_removal,[],[f10408])).
fof(f10408,plain,(
  ( ! [X273,X275,X276,X274] : (s(bool,X274) != s(bool,X275) | s(bool,X275) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X273) | s(bool,X274) = s(bool,X276) | s(bool,'F_0') = s(bool,X276) | s(bool,'T_0') = s(bool,X273)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f3149,f5003])).
fof(f3149,plain,(
  ( ! [X83,X82] : (s(bool,X82) != s(bool,X83) | s(bool,X83) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X82)) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f3131,f84])).
fof(f3131,plain,(
  ( ! [X83,X82] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X82) | s(bool,X83) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X82) != s(bool,X83)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f580,f113])).
fof(f580,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X0) != s(bool,X1)) ) | $spl24),
  inference(resolution,[],[f437,f62])).
fof(f3882,plain,(
  ( ! [X30,X33,X32] : (s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') != s(bool,X32) | s(bool,X30) != s(bool,X33)) ) | $spl64),
  inference(cnf_transformation,[],[f3882_D])).
fof(f3882_D,plain,(
  ( ! [X30,X33,X32] : (s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') != s(bool,X32) | s(bool,X30) != s(bool,X33)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f11331,plain,(
  ( ! [X269,X265,X267,X268,X266] : (p(s(bool,X266)) | s(bool,X267) = s(bool,X268) | s(bool,X265) != s(bool,X268) | s(bool,'F_0') != s(bool,X267) | s(bool,X266) = s(bool,X269) | s(bool,'T_0') = s(bool,X265)) ) | ($spl24 | $spl32)),
  inference(subsumption_resolution,[],[f10406,f2819])).
fof(f10406,plain,(
  ( ! [X269,X265,X267,X268,X266] : (p(s(bool,X266)) | s(bool,X267) = s(bool,X268) | s(bool,X265) != s(bool,X268) | s(bool,'F_0') != s(bool,X267) | s(bool,X266) = s(bool,X269) | s(bool,'F_0') = s(bool,X269) | s(bool,'T_0') = s(bool,X265)) ) | ($spl24 | $spl32)),
  inference(superposition,[],[f2861,f5003])).
fof(f10777,plain,(
  ( ! [X198,X196,X200,X199,X197] : (s(bool,X197) != s(bool,X198) | s(bool,'F_0') = s(bool,X199) | s(bool,X198) != s(bool,X199) | s(bool,'T_0') = s(bool,X196) | s(bool,X197) = s(bool,X200) | s(bool,'F_0') = s(bool,X200)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f10390])).
fof(f10390,plain,(
  ( ! [X198,X196,X200,X199,X197] : (s(bool,X197) != s(bool,X198) | s(bool,'F_0') = s(bool,X199) | s(bool,X198) != s(bool,X199) | s(bool,'T_0') = s(bool,X196) | s(bool,X197) = s(bool,X200) | s(bool,'F_0') = s(bool,X200) | s(bool,'T_0') = s(bool,X196)) ) | $spl24),
  inference(superposition,[],[f753,f5003])).
fof(f10781,plain,(
  ( ! [X134,X136,X135,X133,X137] : (s(bool,X134) != s(bool,X135) | s(bool,'T_0') = s(bool,X136) | s(bool,X135) = s(bool,X136) | s(bool,'T_0') = s(bool,X133) | s(bool,X134) = s(bool,X137) | s(bool,'F_0') = s(bool,X137)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f10375])).
fof(f10375,plain,(
  ( ! [X134,X136,X135,X133,X137] : (s(bool,X134) != s(bool,X135) | s(bool,'T_0') = s(bool,X136) | s(bool,X135) = s(bool,X136) | s(bool,'T_0') = s(bool,X133) | s(bool,X134) = s(bool,X137) | s(bool,'F_0') = s(bool,X137) | s(bool,'T_0') = s(bool,X133)) ) | $spl24),
  inference(superposition,[],[f671,f5003])).
fof(f10782,plain,(
  ( ! [X111,X109,X112,X110] : (s(bool,X110) != s(bool,X111) | s(bool,'T_0') = s(bool,X109) | s(bool,'F_0') = s(bool,X111) | s(bool,X110) = s(bool,X112) | s(bool,'F_0') = s(bool,X112)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f10370])).
fof(f10370,plain,(
  ( ! [X111,X109,X112,X110] : (s(bool,X110) != s(bool,X111) | s(bool,'T_0') = s(bool,X109) | s(bool,'F_0') = s(bool,X111) | s(bool,X110) = s(bool,X112) | s(bool,'F_0') = s(bool,X112) | s(bool,'T_0') = s(bool,X109)) ) | $spl24),
  inference(superposition,[],[f609,f5003])).
fof(f609,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(subsumption_resolution,[],[f598,f85])).
fof(f598,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl24),
  inference(superposition,[],[f437,f68])).
fof(f11312,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,X79) != s(bool,X80) | p(s(bool,X78)) | s(bool,X80) = s(bool,X81) | p(s(bool,X81)) | s(bool,X79) = s(bool,X82) | s(bool,'F_0') = s(bool,X82)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10363,f133])).
fof(f10363,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,X79) != s(bool,X80) | p(s(bool,X78)) | s(bool,X80) = s(bool,X81) | p(s(bool,X81)) | s(bool,X79) = s(bool,X82) | s(bool,'F_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X78)) ) | $spl24),
  inference(superposition,[],[f532,f5003])).
fof(f11295,plain,(
  ( ! [X540,X542,X538,X541,X539] : (p(s(bool,X541)) | s(bool,X542) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X538),s('type.hreal.hreal',X539))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X538),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X539),s('type.hreal.hreal',X540))))))) | s(bool,'F_0') = s(bool,X542)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10340,f133])).
fof(f10340,plain,(
  ( ! [X540,X542,X538,X541,X539] : (p(s(bool,X541)) | s(bool,X542) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X538),s('type.hreal.hreal',X539))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X538),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X539),s('type.hreal.hreal',X540))))))) | s(bool,'F_0') = s(bool,X542) | s(bool,'T_0') = s(bool,X541)) ) | $spl24),
  inference(superposition,[],[f8633,f5003])).
fof(f10785,plain,(
  ( ! [X536,X537] : (s(bool,'T_0') = s(bool,X536) | s(bool,X537) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X537)) ) | ($spl4 | $spl24)),
  inference(duplicate_literal_removal,[],[f10339])).
fof(f10339,plain,(
  ( ! [X536,X537] : (s(bool,'T_0') = s(bool,X536) | s(bool,X537) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X537) | s(bool,'T_0') = s(bool,X536)) ) | ($spl4 | $spl24)),
  inference(superposition,[],[f112,f5003])).
fof(f11290,plain,(
  ( ! [X509,X511,X512,X510,X513] : (p(s(bool,X511)) | s('type.hreal.hreal',X510) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X509),s('type.hreal.hreal',X512))) | s(bool,X513) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X509),s('type.hreal.hreal',X510))) | s(bool,'F_0') = s(bool,X513)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10331,f133])).
fof(f10331,plain,(
  ( ! [X509,X511,X512,X510,X513] : (p(s(bool,X511)) | s('type.hreal.hreal',X510) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X509),s('type.hreal.hreal',X512))) | s(bool,X513) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X509),s('type.hreal.hreal',X510))) | s(bool,'F_0') = s(bool,X513) | s(bool,'T_0') = s(bool,X511)) ) | $spl24),
  inference(superposition,[],[f61,f5003])).
fof(f9297,plain,(
  ( ! [X118,X116,X119,X115] : (s(bool,X118) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal',X116))))) | s(bool,'T_0') = s(bool,X119) | s(bool,X118) != s(bool,X119)) ) | $spl98),
  inference(cnf_transformation,[],[f9297_D])).
fof(f9297_D,plain,(
  ( ! [X118,X116,X119,X115] : (s(bool,X118) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal',X116))))) | s(bool,'T_0') = s(bool,X119) | s(bool,X118) != s(bool,X119)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f10787,plain,(
  ( ! [X478,X480,X477,X479,X481] : (s(bool,X477) != s(bool,X478) | s(bool,X478) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X479),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X479),s('type.hreal.hreal',X480))))) | s(bool,'F_0') = s(bool,X481) | s(bool,'T_0') = s(bool,X477)) ) | ($spl24 | $spl86)),
  inference(duplicate_literal_removal,[],[f10324])).
fof(f10324,plain,(
  ( ! [X478,X480,X477,X479,X481] : (s(bool,X477) != s(bool,X478) | s(bool,X478) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X479),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X479),s('type.hreal.hreal',X480))))) | s(bool,'F_0') = s(bool,X481) | s(bool,'F_0') = s(bool,X481) | s(bool,'T_0') = s(bool,X477)) ) | ($spl24 | $spl86)),
  inference(superposition,[],[f8152,f5003])).
fof(f9291,plain,(
  ( ! [X76,X74,X77] : (s(bool,X74) != s(bool,X76) | s(bool,X74) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X77) | s(bool,X76) != s(bool,X77)) ) | $spl96),
  inference(cnf_transformation,[],[f9291_D])).
fof(f9291_D,plain,(
  ( ! [X76,X74,X77] : (s(bool,X74) != s(bool,X76) | s(bool,X74) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X77) | s(bool,X76) != s(bool,X77)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f10788,plain,(
  ( ! [X476,X474,X475,X473] : (s(bool,X473) != s(bool,X474) | s(bool,X474) != s(bool,X475) | s(bool,X475) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X476) | s(bool,'T_0') = s(bool,X473)) ) | ($spl3 | $spl6 | $spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f10323])).
fof(f10323,plain,(
  ( ! [X476,X474,X475,X473] : (s(bool,X473) != s(bool,X474) | s(bool,X474) != s(bool,X475) | s(bool,X475) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X476) | s(bool,'F_0') = s(bool,X476) | s(bool,'T_0') = s(bool,X473)) ) | ($spl3 | $spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f7843,f5003])).
fof(f7843,plain,(
  ( ! [X241,X242] : (s(bool,'F_0') != s(bool,X242) | s(bool,X241) != s(bool,X242) | s(bool,X241) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl3 | $spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f7270,f130])).
fof(f7270,plain,(
  ( ! [X241,X242] : (s(bool,X241) != s(bool,X242) | s(bool,'F_0') != s(bool,X242) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X241) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f705,f5704])).
fof(f10789,plain,(
  ( ! [X472,X469,X471,X470] : (s(bool,X469) = s(bool,X470) | s(bool,X471) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X470) = s(bool,X471) | s(bool,'F_0') = s(bool,X472) | s(bool,'T_0') = s(bool,X469)) ) | ($spl6 | $spl24)),
  inference(duplicate_literal_removal,[],[f10322])).
fof(f10322,plain,(
  ( ! [X472,X469,X471,X470] : (s(bool,X469) = s(bool,X470) | s(bool,X471) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X470) = s(bool,X471) | s(bool,'F_0') = s(bool,X472) | s(bool,'F_0') = s(bool,X472) | s(bool,'T_0') = s(bool,X469)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f5704,f5003])).
fof(f11288,plain,(
  ( ! [X466,X465,X467] : (s(bool,X465) != s(bool,X466) | s(bool,'T_0') = s(bool,X465) | s(bool,X466) = s(bool,X467) | s(bool,X467) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl160),
  inference(cnf_transformation,[],[f11288_D])).
fof(f11288_D,plain,(
  ( ! [X466,X465,X467] : (s(bool,X465) != s(bool,X466) | s(bool,'T_0') = s(bool,X465) | s(bool,X466) = s(bool,X467) | s(bool,X467) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f10790,plain,(
  ( ! [X466,X465,X467,X468] : (s(bool,X465) != s(bool,X466) | s(bool,X467) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X466) = s(bool,X467) | s(bool,'F_0') = s(bool,X468) | s(bool,'T_0') = s(bool,X465)) ) | ($spl6 | $spl24 | $spl32)),
  inference(duplicate_literal_removal,[],[f10321])).
fof(f10321,plain,(
  ( ! [X466,X465,X467,X468] : (s(bool,X465) != s(bool,X466) | s(bool,X467) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X466) = s(bool,X467) | s(bool,'F_0') = s(bool,X468) | s(bool,'F_0') = s(bool,X468) | s(bool,'T_0') = s(bool,X465)) ) | ($spl6 | $spl24 | $spl32)),
  inference(superposition,[],[f3895,f5003])).
fof(f3895,plain,(
  ( ! [X83,X82] : (s(bool,'F_0') != s(bool,X82) | s(bool,X83) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X82) = s(bool,X83)) ) | ($spl6 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f3873,f84])).
fof(f3873,plain,(
  ( ! [X83,X82] : (p(s(bool,'F_0')) | s(bool,X82) = s(bool,X83) | s(bool,X83) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') != s(bool,X82)) ) | ($spl6 | $spl24 | $spl32)),
  inference(superposition,[],[f2861,f113])).
fof(f10791,plain,(
  ( ! [X461,X463,X462,X464] : (s(bool,X461) != s(bool,X462) | s(bool,'F_0') != s(bool,X463) | s(bool,X462) = s(bool,X463) | s(bool,'F_0') = s(bool,X464) | s(bool,'T_0') = s(bool,X461)) ) | ($spl24 | $spl32)),
  inference(duplicate_literal_removal,[],[f10320])).
fof(f10320,plain,(
  ( ! [X461,X463,X462,X464] : (s(bool,X461) != s(bool,X462) | s(bool,'F_0') != s(bool,X463) | s(bool,X462) = s(bool,X463) | s(bool,'F_0') = s(bool,X464) | s(bool,'F_0') = s(bool,X464) | s(bool,'T_0') = s(bool,X461)) ) | ($spl24 | $spl32)),
  inference(superposition,[],[f3847,f5003])).
fof(f11272,plain,(
  ( ! [X356,X354,X357,X353,X355] : (s(bool,X353) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X355),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X355),s('type.hreal.hreal',X356))))) | s(bool,X353) = s(bool,X357) | s(bool,'F_0') = s(bool,X357) | s(bool,'T_0') = s(bool,X354)) ) | ($spl24 | $spl86)),
  inference(subsumption_resolution,[],[f10290,f68])).
fof(f10290,plain,(
  ( ! [X356,X354,X357,X353,X355] : (s(bool,'F_0') != s(bool,X354) | s(bool,X353) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X355),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X355),s('type.hreal.hreal',X356))))) | s(bool,X353) = s(bool,X357) | s(bool,'F_0') = s(bool,X357) | s(bool,'T_0') = s(bool,X354)) ) | ($spl24 | $spl86)),
  inference(superposition,[],[f8152,f5003])).
fof(f11270,plain,(
  ( ! [X344,X346,X345,X347] : (s(bool,X344) != s(bool,X346) | s(bool,X346) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X344) = s(bool,X347) | s(bool,'F_0') = s(bool,X347) | s(bool,'T_0') = s(bool,X345)) ) | ($spl3 | $spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10288,f68])).
fof(f10288,plain,(
  ( ! [X344,X346,X345,X347] : (s(bool,'F_0') != s(bool,X345) | s(bool,X344) != s(bool,X346) | s(bool,X346) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X344) = s(bool,X347) | s(bool,'F_0') = s(bool,X347) | s(bool,'T_0') = s(bool,X345)) ) | ($spl3 | $spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f7843,f5003])).
fof(f11267,plain,(
  ( ! [X321,X323,X320,X322,X319] : (p(s(bool,X320)) | s(bool,X319) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X321),s('type.hreal.hreal',X322))) | s('type.hreal.hreal',X322) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X321),s('type.hreal.hreal',sK3(X322,X321)))) | s(bool,X319) = s(bool,X323) | s(bool,'F_0') = s(bool,X323)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10282,f133])).
fof(f10282,plain,(
  ( ! [X321,X323,X320,X322,X319] : (p(s(bool,X320)) | s(bool,X319) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X321),s('type.hreal.hreal',X322))) | s('type.hreal.hreal',X322) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X321),s('type.hreal.hreal',sK3(X322,X321)))) | s(bool,X319) = s(bool,X323) | s(bool,'F_0') = s(bool,X323) | s(bool,'T_0') = s(bool,X320)) ) | $spl24),
  inference(superposition,[],[f5061,f5003])).
fof(f11266,plain,(
  ( ! [X315,X316,X318,X314,X317] : (p(s(bool,X315)) | s(bool,'F_0') = s(bool,X316) | s(bool,X314) = s(bool,X317) | s(bool,X316) = s(bool,X317) | s(bool,X314) = s(bool,X318) | s(bool,'F_0') = s(bool,X318)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10281,f133])).
fof(f10281,plain,(
  ( ! [X315,X316,X318,X314,X317] : (p(s(bool,X315)) | s(bool,'F_0') = s(bool,X316) | s(bool,X314) = s(bool,X317) | s(bool,X316) = s(bool,X317) | s(bool,X314) = s(bool,X318) | s(bool,'F_0') = s(bool,X318) | s(bool,'T_0') = s(bool,X315)) ) | $spl24),
  inference(superposition,[],[f4997,f5003])).
fof(f11264,plain,(
  ( ! [X296,X293,X295] : (s(bool,X295) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X296) | s(bool,X293) = s(bool,X296) | s(bool,X293) = s(bool,X295)) ) | $spl152),
  inference(cnf_transformation,[],[f11264_D])).
fof(f11264_D,plain,(
  ( ! [X296,X293,X295] : (s(bool,X295) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X296) | s(bool,X293) = s(bool,X296) | s(bool,X293) = s(bool,X295)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f11263,plain,(
  ( ! [X294,X296,X293,X295] : (s(bool,X295) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X293) = s(bool,X295) | s(bool,X293) = s(bool,X296) | s(bool,'F_0') = s(bool,X296) | s(bool,'T_0') = s(bool,X294)) ) | ($spl6 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f10276,f68])).
fof(f10276,plain,(
  ( ! [X294,X296,X293,X295] : (s(bool,'F_0') != s(bool,X294) | s(bool,X295) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X293) = s(bool,X295) | s(bool,X293) = s(bool,X296) | s(bool,'F_0') = s(bool,X296) | s(bool,'T_0') = s(bool,X294)) ) | ($spl6 | $spl24 | $spl32)),
  inference(superposition,[],[f3895,f5003])).
fof(f3788,plain,(
  ( ! [X52,X51,X49] : (s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') != s(bool,X51) | s(bool,X49) = s(bool,X52)) ) | $spl60),
  inference(cnf_transformation,[],[f3788_D])).
fof(f3788_D,plain,(
  ( ! [X52,X51,X49] : (s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') != s(bool,X51) | s(bool,X49) = s(bool,X52)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f11262,plain,(
  ( ! [X292,X290,X289,X291] : (s(bool,'F_0') != s(bool,X291) | s(bool,X289) = s(bool,X291) | s(bool,X289) = s(bool,X292) | s(bool,'F_0') = s(bool,X292) | s(bool,'T_0') = s(bool,X290)) ) | ($spl24 | $spl32)),
  inference(subsumption_resolution,[],[f10275,f68])).
fof(f10275,plain,(
  ( ! [X292,X290,X289,X291] : (s(bool,'F_0') != s(bool,X290) | s(bool,'F_0') != s(bool,X291) | s(bool,X289) = s(bool,X291) | s(bool,X289) = s(bool,X292) | s(bool,'F_0') = s(bool,X292) | s(bool,'T_0') = s(bool,X290)) ) | ($spl24 | $spl32)),
  inference(superposition,[],[f3847,f5003])).
fof(f3385,plain,(
  ( ! [X52,X51,X49] : (s(bool,'T_0') != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,X51) != s(bool,X52) | s(bool,X49) = s(bool,X52)) ) | $spl52),
  inference(cnf_transformation,[],[f3385_D])).
fof(f3385_D,plain,(
  ( ! [X52,X51,X49] : (s(bool,'T_0') != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,X51) != s(bool,X52) | s(bool,X49) = s(bool,X52)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f7583,plain,(
  ( ! [X114] : (s(bool,'F_0') != s(bool,X114)) ) | $spl80),
  inference(cnf_transformation,[],[f7583_D])).
fof(f7583_D,plain,(
  ( ! [X114] : (s(bool,'F_0') != s(bool,X114)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f11261,plain,(
  ( ! [X282,X281,X283,X284] : (s(bool,'F_0') != s(bool,X282) | s(bool,'T_0') != s(bool,X283) | s(bool,X281) != s(bool,X283) | s(bool,X281) = s(bool,X284) | s(bool,'F_0') = s(bool,X284)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10273,f705])).
fof(f10273,plain,(
  ( ! [X282,X281,X283,X284] : (s(bool,'F_0') != s(bool,X282) | s(bool,'T_0') != s(bool,X283) | s(bool,X281) != s(bool,X283) | s(bool,X281) = s(bool,X284) | s(bool,'F_0') = s(bool,X284) | s(bool,'T_0') = s(bool,X282)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f3407,f5003])).
fof(f3407,plain,(
  ( ! [X61,X60] : (s(bool,'F_0') != s(bool,X61) | s(bool,'T_0') != s(bool,X60) | s(bool,X60) != s(bool,X61)) ) | ($spl16 | $spl24)),
  inference(resolution,[],[f703,f84])).
fof(f11260,plain,(
  ( ! [X269,X265,X267,X268] : (s(bool,X267) = s(bool,X268) | s(bool,'F_0') = s(bool,X269) | s(bool,X265) = s(bool,X269) | s(bool,'F_0') != s(bool,X267) | s(bool,X265) != s(bool,X268)) ) | $spl150),
  inference(cnf_transformation,[],[f11260_D])).
fof(f11260_D,plain,(
  ( ! [X269,X265,X267,X268] : (s(bool,X267) = s(bool,X268) | s(bool,'F_0') = s(bool,X269) | s(bool,X265) = s(bool,X269) | s(bool,'F_0') != s(bool,X267) | s(bool,X265) != s(bool,X268)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f11259,plain,(
  ( ! [X269,X265,X267,X268,X266] : (p(s(bool,X266)) | s(bool,X267) = s(bool,X268) | s(bool,X265) != s(bool,X268) | s(bool,'F_0') != s(bool,X267) | s(bool,X265) = s(bool,X269) | s(bool,'F_0') = s(bool,X269)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f10269,f133])).
fof(f10269,plain,(
  ( ! [X269,X265,X267,X268,X266] : (p(s(bool,X266)) | s(bool,X267) = s(bool,X268) | s(bool,X265) != s(bool,X268) | s(bool,'F_0') != s(bool,X267) | s(bool,X265) = s(bool,X269) | s(bool,'F_0') = s(bool,X269) | s(bool,'T_0') = s(bool,X266)) ) | ($spl24 | $spl32)),
  inference(superposition,[],[f2861,f5003])).
fof(f11258,plain,(
  ( ! [X257,X259,X256,X258] : (p(s(bool,X257)) | s(bool,'F_0') != s(bool,X258) | s(bool,X256) = s(bool,X258) | s(bool,X256) = s(bool,X259) | s(bool,'F_0') = s(bool,X259)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f10267,f133])).
fof(f10267,plain,(
  ( ! [X257,X259,X256,X258] : (p(s(bool,X257)) | s(bool,'F_0') != s(bool,X258) | s(bool,X256) = s(bool,X258) | s(bool,X256) = s(bool,X259) | s(bool,'F_0') = s(bool,X259) | s(bool,'T_0') = s(bool,X257)) ) | ($spl24 | $spl32)),
  inference(superposition,[],[f2819,f5003])).
fof(f11257,plain,(
  ( ! [X237,X235,X236,X234] : (s(bool,X234) != s(bool,X236) | s(bool,'F_0') = s(bool,X236) | s(bool,X234) = s(bool,X237) | s(bool,'F_0') = s(bool,X237) | s(bool,'T_0') = s(bool,X235)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10262,f695])).
fof(f10262,plain,(
  ( ! [X237,X235,X236,X234] : (s(bool,X235) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X234) != s(bool,X236) | s(bool,'F_0') = s(bool,X236) | s(bool,X234) = s(bool,X237) | s(bool,'F_0') = s(bool,X237) | s(bool,'T_0') = s(bool,X235)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f1059,f5003])).
fof(f1059,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f727,f68])).
fof(f11255,plain,(
  ( ! [X204,X202,X205,X203,X201] : (p(s(bool,X202)) | s(bool,'F_0') = s(bool,X203) | s(bool,X201) != s(bool,X204) | s(bool,X203) != s(bool,X204) | s(bool,X201) = s(bool,X205) | s(bool,'F_0') = s(bool,X205)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10254,f133])).
fof(f10254,plain,(
  ( ! [X204,X202,X205,X203,X201] : (p(s(bool,X202)) | s(bool,'F_0') = s(bool,X203) | s(bool,X201) != s(bool,X204) | s(bool,X203) != s(bool,X204) | s(bool,X201) = s(bool,X205) | s(bool,'F_0') = s(bool,X205) | s(bool,'T_0') = s(bool,X202)) ) | $spl24),
  inference(superposition,[],[f755,f5003])).
fof(f755,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X19)) | s(bool,'F_0') = s(bool,X17) | s(bool,X18) != s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl24),
  inference(resolution,[],[f561,f458])).
fof(f11253,plain,(
  ( ! [X191,X194,X192,X195,X193] : (s(bool,'F_0') = s(bool,X194) | s(bool,X191) != s(bool,X194) | s(bool,'T_0') = s(bool,X193) | s(bool,X191) = s(bool,X195) | s(bool,'F_0') = s(bool,X195) | s(bool,'T_0') = s(bool,X192)) ) | $spl24),
  inference(subsumption_resolution,[],[f10252,f672])).
fof(f10252,plain,(
  ( ! [X191,X194,X192,X195,X193] : (s(bool,X192) != s(bool,X193) | s(bool,'F_0') = s(bool,X194) | s(bool,X191) != s(bool,X194) | s(bool,'T_0') = s(bool,X193) | s(bool,X191) = s(bool,X195) | s(bool,'F_0') = s(bool,X195) | s(bool,'T_0') = s(bool,X192)) ) | $spl24),
  inference(superposition,[],[f753,f5003])).
fof(f11252,plain,(
  ( ! [X185,X186,X184,X183] : (s(bool,X183) != s(bool,X185) | s(bool,'F_0') = s(bool,X185) | s(bool,X183) = s(bool,X186) | s(bool,'F_0') = s(bool,X186) | s(bool,'T_0') = s(bool,X184)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10250,f68])).
fof(f10250,plain,(
  ( ! [X185,X186,X184,X183] : (s(bool,'F_0') != s(bool,X184) | s(bool,X183) != s(bool,X185) | s(bool,'F_0') = s(bool,X185) | s(bool,X183) = s(bool,X186) | s(bool,'F_0') = s(bool,X186) | s(bool,'T_0') = s(bool,X184)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f729,f5003])).
fof(f11251,plain,(
  ( ! [X175,X176,X174,X172] : (s(bool,'T_0') != s(bool,X174) | s(bool,'F_0') = s(bool,X176) | s(bool,X172) = s(bool,X176) | s(bool,X174) != s(bool,X175) | s(bool,X172) != s(bool,X175)) ) | $spl146),
  inference(cnf_transformation,[],[f11251_D])).
fof(f11251_D,plain,(
  ( ! [X175,X176,X174,X172] : (s(bool,'T_0') != s(bool,X174) | s(bool,'F_0') = s(bool,X176) | s(bool,X172) = s(bool,X176) | s(bool,X174) != s(bool,X175) | s(bool,X172) != s(bool,X175)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f11250,plain,(
  ( ! [X175,X173,X176,X174,X172] : (p(s(bool,X173)) | s(bool,'T_0') != s(bool,X174) | s(bool,X172) != s(bool,X175) | s(bool,X174) != s(bool,X175) | s(bool,X172) = s(bool,X176) | s(bool,'F_0') = s(bool,X176)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10247,f133])).
fof(f10247,plain,(
  ( ! [X175,X173,X176,X174,X172] : (p(s(bool,X173)) | s(bool,'T_0') != s(bool,X174) | s(bool,X172) != s(bool,X175) | s(bool,X174) != s(bool,X175) | s(bool,X172) = s(bool,X176) | s(bool,'F_0') = s(bool,X176) | s(bool,'T_0') = s(bool,X173)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f703,f5003])).
fof(f11248,plain,(
  ( ! [X158,X156,X154,X157,X155] : (p(s(bool,X155)) | s(bool,'T_0') = s(bool,X156) | s(bool,X154) != s(bool,X157) | s(bool,X156) = s(bool,X157) | s(bool,X154) = s(bool,X158) | s(bool,'F_0') = s(bool,X158)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10243,f133])).
fof(f10243,plain,(
  ( ! [X158,X156,X154,X157,X155] : (p(s(bool,X155)) | s(bool,'T_0') = s(bool,X156) | s(bool,X154) != s(bool,X157) | s(bool,X156) = s(bool,X157) | s(bool,X154) = s(bool,X158) | s(bool,'F_0') = s(bool,X158) | s(bool,'T_0') = s(bool,X155)) ) | $spl24),
  inference(superposition,[],[f673,f5003])).
fof(f11243,plain,(
  ( ! [X99,X97,X100,X98,X96] : (p(s(bool,X97)) | s(bool,'T_0') = s(bool,X98) | s(bool,X96) = s(bool,X99) | s(bool,X98) != s(bool,X99) | s(bool,X96) = s(bool,X100) | s(bool,'F_0') = s(bool,X100)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10230,f133])).
fof(f10230,plain,(
  ( ! [X99,X97,X100,X98,X96] : (p(s(bool,X97)) | s(bool,'T_0') = s(bool,X98) | s(bool,X96) = s(bool,X99) | s(bool,X98) != s(bool,X99) | s(bool,X96) = s(bool,X100) | s(bool,'F_0') = s(bool,X100) | s(bool,'T_0') = s(bool,X97)) ) | $spl24),
  inference(superposition,[],[f580,f5003])).
fof(f11242,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,X93)) | s(bool,X92) != s(bool,X94) | s(bool,'F_0') = s(bool,X94) | s(bool,X92) = s(bool,X95) | s(bool,'F_0') = s(bool,X95)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10229,f133])).
fof(f10229,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,X93)) | s(bool,X92) != s(bool,X94) | s(bool,'F_0') = s(bool,X94) | s(bool,X92) = s(bool,X95) | s(bool,'F_0') = s(bool,X95) | s(bool,'T_0') = s(bool,X93)) ) | $spl24),
  inference(superposition,[],[f561,f5003])).
fof(f11241,plain,(
  ( ! [X90,X88,X91,X89] : (p(s(bool,X89)) | s(bool,X88) != s(bool,X90) | s(bool,'T_0') != s(bool,X90) | s(bool,X88) = s(bool,X91) | s(bool,'F_0') = s(bool,X91)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10228,f133])).
fof(f10228,plain,(
  ( ! [X90,X88,X91,X89] : (p(s(bool,X89)) | s(bool,X88) != s(bool,X90) | s(bool,'T_0') != s(bool,X90) | s(bool,X88) = s(bool,X91) | s(bool,'F_0') = s(bool,X91) | s(bool,'T_0') = s(bool,X89)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f538,f5003])).
fof(f11220,plain,(
  ( ! [X412,X414,X410,X413,X411] : (s(bool,'T_0') = s(bool,X412) | s(bool,X413) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X410),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X410),s('type.hreal.hreal',X411))))) | s(bool,X413) = s(bool,X414) | s(bool,'T_0') = s(bool,X414)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10190,f68])).
fof(f10190,plain,(
  ( ! [X412,X414,X410,X413,X411] : (s(bool,'F_0') != s(bool,X412) | s(bool,'T_0') = s(bool,X412) | s(bool,X413) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X410),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X410),s('type.hreal.hreal',X411))))) | s(bool,X413) = s(bool,X414) | s(bool,'T_0') = s(bool,X414)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f4271,f5003])).
fof(f10806,plain,(
  ( ! [X408,X405,X407,X409,X406] : (s(bool,'F_0') = s(bool,X407) | s(bool,X408) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X405),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X405),s('type.hreal.hreal',X406))))) | s(bool,X408) = s(bool,X409) | s(bool,'T_0') = s(bool,X409)) ) | ($spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f10189])).
fof(f10189,plain,(
  ( ! [X408,X405,X407,X409,X406] : (s(bool,'F_0') = s(bool,X407) | s(bool,'F_0') = s(bool,X407) | s(bool,X408) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X405),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X405),s('type.hreal.hreal',X406))))) | s(bool,X408) = s(bool,X409) | s(bool,'T_0') = s(bool,X409)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f4294,f5003])).
fof(f10809,plain,(
  ( ! [X325,X327,X326,X328] : (s(bool,'F_0') = s(bool,X326) | s(bool,X325) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X325) = s(bool,X327) | s(bool,X327) = s(bool,X328) | s(bool,'T_0') = s(bool,X328)) ) | ($spl6 | $spl24)),
  inference(duplicate_literal_removal,[],[f10168])).
fof(f10168,plain,(
  ( ! [X325,X327,X326,X328] : (s(bool,'F_0') = s(bool,X326) | s(bool,'F_0') = s(bool,X326) | s(bool,X325) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X325) = s(bool,X327) | s(bool,X327) = s(bool,X328) | s(bool,'T_0') = s(bool,X328)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f5704,f5003])).
fof(f11212,plain,(
  ( ! [X269,X271,X270,X272] : (s(bool,X269) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X270) | s(bool,X269) = s(bool,X271) | s(bool,X271) = s(bool,X272) | s(bool,'T_0') = s(bool,X272)) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f10155,f68])).
fof(f10155,plain,(
  ( ! [X269,X271,X270,X272] : (s(bool,'F_0') != s(bool,X270) | s(bool,X269) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X270) | s(bool,X269) = s(bool,X271) | s(bool,X271) = s(bool,X272) | s(bool,'T_0') = s(bool,X272)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f3149,f5003])).
fof(f11206,plain,(
  ( ! [X191,X189,X192,X190,X193] : (s(bool,'F_0') = s(bool,X191) | s(bool,X189) != s(bool,X191) | s(bool,'T_0') = s(bool,X190) | s(bool,X189) = s(bool,X192) | s(bool,X192) = s(bool,X193) | s(bool,'T_0') = s(bool,X193)) ) | $spl24),
  inference(subsumption_resolution,[],[f10137,f68])).
fof(f10137,plain,(
  ( ! [X191,X189,X192,X190,X193] : (s(bool,'F_0') != s(bool,X190) | s(bool,'F_0') = s(bool,X191) | s(bool,X189) != s(bool,X191) | s(bool,'T_0') = s(bool,X190) | s(bool,X189) = s(bool,X192) | s(bool,X192) = s(bool,X193) | s(bool,'T_0') = s(bool,X193)) ) | $spl24),
  inference(superposition,[],[f753,f5003])).
fof(f11200,plain,(
  ( ! [X127,X130,X128,X126,X129] : (s(bool,'T_0') = s(bool,X128) | s(bool,X126) = s(bool,X128) | s(bool,'T_0') = s(bool,X127) | s(bool,X126) = s(bool,X129) | s(bool,X129) = s(bool,X130) | s(bool,'T_0') = s(bool,X130)) ) | $spl24),
  inference(subsumption_resolution,[],[f10122,f68])).
fof(f10122,plain,(
  ( ! [X127,X130,X128,X126,X129] : (s(bool,'F_0') != s(bool,X127) | s(bool,'T_0') = s(bool,X128) | s(bool,X126) = s(bool,X128) | s(bool,'T_0') = s(bool,X127) | s(bool,X126) = s(bool,X129) | s(bool,X129) = s(bool,X130) | s(bool,'T_0') = s(bool,X130)) ) | $spl24),
  inference(superposition,[],[f671,f5003])).
fof(f11196,plain,(
  ( ! [X85,X83,X81,X84] : (s(bool,X81) = s(bool,X83) | s(bool,'T_0') = s(bool,X85) | s(bool,X84) = s(bool,X85) | s(bool,X81) = s(bool,X84) | p(s(bool,X83))) ) | $spl142),
  inference(cnf_transformation,[],[f11196_D])).
fof(f11196_D,plain,(
  ( ! [X85,X83,X81,X84] : (s(bool,X81) = s(bool,X83) | s(bool,'T_0') = s(bool,X85) | s(bool,X84) = s(bool,X85) | s(bool,X81) = s(bool,X84) | p(s(bool,X83))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f11195,plain,(
  ( ! [X85,X83,X81,X84,X82] : (p(s(bool,X82)) | s(bool,X81) = s(bool,X83) | p(s(bool,X83)) | s(bool,X81) = s(bool,X84) | s(bool,X84) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) ) | $spl24),
  inference(subsumption_resolution,[],[f10112,f4678])).
fof(f10112,plain,(
  ( ! [X85,X83,X81,X84,X82] : (s(bool,'F_0') != s(bool,X82) | p(s(bool,X82)) | s(bool,X81) = s(bool,X83) | p(s(bool,X83)) | s(bool,X81) = s(bool,X84) | s(bool,X84) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) ) | $spl24),
  inference(superposition,[],[f532,f5003])).
fof(f11185,plain,(
  ( ! [X377,X374,X376,X375] : (s(bool,'T_0') != s(bool,X374) | s(bool,X375) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X374) = s(bool,X376) | s(bool,X376) = s(bool,X377) | s(bool,'T_0') = s(bool,X377)) ) | ($spl3 | $spl6 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f10086,f2900])).
fof(f10086,plain,(
  ( ! [X377,X374,X376,X375] : (s(bool,'T_0') != s(bool,X374) | s(bool,X375) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X375) | s(bool,X374) = s(bool,X376) | s(bool,X376) = s(bool,X377) | s(bool,'T_0') = s(bool,X377)) ) | ($spl3 | $spl6 | $spl24)),
  inference(superposition,[],[f6875,f5003])).
fof(f11184,plain,(
  ( ! [X333,X335,X334,X336] : (s(bool,'T_0') != s(bool,X333) | s(bool,'F_0') != s(bool,X334) | s(bool,X333) = s(bool,X335) | s(bool,X335) = s(bool,X336) | s(bool,'T_0') = s(bool,X336)) ) | ($spl6 | $spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f10077,f2900])).
fof(f10077,plain,(
  ( ! [X333,X335,X334,X336] : (s(bool,'T_0') != s(bool,X333) | s(bool,X334) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') != s(bool,X334) | s(bool,X333) = s(bool,X335) | s(bool,X335) = s(bool,X336) | s(bool,'T_0') = s(bool,X336)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3455,f5003])).
fof(f10827,plain,(
  ( ! [X273,X275,X276,X274] : (s(bool,'T_0') != s(bool,X273) | s(bool,'F_0') = s(bool,X274) | s(bool,X273) = s(bool,X275) | s(bool,X275) = s(bool,X276) | s(bool,'T_0') = s(bool,X276)) ) | ($spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f10063])).
fof(f10063,plain,(
  ( ! [X273,X275,X276,X274] : (s(bool,'T_0') != s(bool,X273) | s(bool,'F_0') = s(bool,X274) | s(bool,'F_0') = s(bool,X274) | s(bool,X273) = s(bool,X275) | s(bool,X275) = s(bool,X276) | s(bool,'T_0') = s(bool,X276)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f781,f5003])).
fof(f781,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X15) | s(bool,'F_0') = s(bool,X14) | s(bool,X14) = s(bool,X15)) ) | $spl16),
  inference(resolution,[],[f317,f133])).
fof(f11177,plain,(
  ( ! [X235,X234] : (s(bool,'F_0') != s(bool,X234) | s(bool,X234) != s(bool,X235)) ) | $spl138),
  inference(cnf_transformation,[],[f11177_D])).
fof(f11177_D,plain,(
  ( ! [X235,X234] : (s(bool,'F_0') != s(bool,X234) | s(bool,X234) != s(bool,X235)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f11176,plain,(
  ( ! [X237,X235,X233,X236,X234] : (s(bool,'T_0') != s(bool,X233) | s(bool,'F_0') != s(bool,X234) | s(bool,X234) != s(bool,X235) | s(bool,X233) = s(bool,X236) | s(bool,X236) = s(bool,X237) | s(bool,'T_0') = s(bool,X237)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10054,f3407])).
fof(f10054,plain,(
  ( ! [X237,X235,X233,X236,X234] : (s(bool,'T_0') != s(bool,X233) | s(bool,'F_0') != s(bool,X234) | s(bool,X234) != s(bool,X235) | s(bool,'T_0') = s(bool,X235) | s(bool,X233) = s(bool,X236) | s(bool,X236) = s(bool,X237) | s(bool,'T_0') = s(bool,X237)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f701,f5003])).
fof(f11165,plain,(
  ( ! [X173,X171,X174,X172] : (s(bool,'T_0') != s(bool,X171) | s(bool,'F_0') != s(bool,X172) | s(bool,X171) = s(bool,X173) | s(bool,X173) = s(bool,X174) | s(bool,'T_0') = s(bool,X174)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10040,f705])).
fof(f10040,plain,(
  ( ! [X173,X171,X174,X172] : (s(bool,'T_0') != s(bool,X171) | s(bool,'T_0') = s(bool,X172) | s(bool,'F_0') != s(bool,X172) | s(bool,X171) = s(bool,X173) | s(bool,X173) = s(bool,X174) | s(bool,'T_0') = s(bool,X174)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f586,f5003])).
fof(f11160,plain,(
  ( ! [X127,X128,X126,X129] : (s(bool,'T_0') != s(bool,X126) | s(bool,'F_0') = s(bool,X127) | s(bool,X126) = s(bool,X128) | s(bool,X128) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f10029,f4678])).
fof(f10029,plain,(
  ( ! [X127,X128,X126,X129] : (s(bool,'T_0') != s(bool,X126) | s(bool,'F_0') = s(bool,X127) | ~p(s(bool,X127)) | s(bool,X126) = s(bool,X128) | s(bool,X128) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f273,f5003])).
fof(f11149,plain,(
  ( ! [X505,X507,X508,X506] : (~p(s(bool,X507)) | s('type.hreal.hreal',X506) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X505),s('type.hreal.hreal',sK3(X506,X505)))) | s(bool,X508) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X505),s('type.hreal.hreal',X506))) | s(bool,'T_0') = s(bool,X508)) ) | $spl24),
  inference(subsumption_resolution,[],[f9987,f436])).
fof(f9987,plain,(
  ( ! [X505,X507,X508,X506] : (~p(s(bool,X507)) | s('type.hreal.hreal',X506) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X505),s('type.hreal.hreal',sK3(X506,X505)))) | s(bool,'F_0') = s(bool,X507) | s(bool,X508) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X505),s('type.hreal.hreal',X506))) | s(bool,'T_0') = s(bool,X508)) ) | $spl24),
  inference(superposition,[],[f60,f5003])).
fof(f11147,plain,(
  ( ! [X500,X502,X504,X501,X503] : (~p(s(bool,X502)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X501),s('type.hreal.hreal',X503))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X500),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X501,X500)),s('type.hreal.hreal',X503))))) | s(bool,X504) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X500),s('type.hreal.hreal',X501))) | s(bool,'T_0') = s(bool,X504)) ) | $spl24),
  inference(subsumption_resolution,[],[f9986,f436])).
fof(f9986,plain,(
  ( ! [X500,X502,X504,X501,X503] : (~p(s(bool,X502)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X501),s('type.hreal.hreal',X503))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X500),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X501,X500)),s('type.hreal.hreal',X503))))) | s(bool,'F_0') = s(bool,X502) | s(bool,X504) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X500),s('type.hreal.hreal',X501))) | s(bool,'T_0') = s(bool,X504)) ) | $spl24),
  inference(superposition,[],[f5121,f5003])).
fof(f11146,plain,(
  ( ! [X499,X496,X498,X497] : (s(bool,'T_0') != s(bool,X498) | s('type.hreal.hreal',X497) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X496),s('type.hreal.hreal',sK3(X497,X496)))) | s(bool,X499) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X496),s('type.hreal.hreal',X497))) | s(bool,'T_0') = s(bool,X499)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9985,f705])).
fof(f9985,plain,(
  ( ! [X499,X496,X498,X497] : (s(bool,'T_0') != s(bool,X498) | s('type.hreal.hreal',X497) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X496),s('type.hreal.hreal',sK3(X497,X496)))) | s(bool,'F_0') = s(bool,X498) | s(bool,X499) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X496),s('type.hreal.hreal',X497))) | s(bool,'T_0') = s(bool,X499)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f279,f5003])).
fof(f10832,plain,(
  ( ! [X420,X422,X421,X423] : (s(bool,X420) != s(bool,X421) | s(bool,X422) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X421) = s(bool,X422) | s(bool,'F_0') = s(bool,X420) | s(bool,'T_0') = s(bool,X423)) ) | ($spl3 | $spl6 | $spl24)),
  inference(duplicate_literal_removal,[],[f9965])).
fof(f9965,plain,(
  ( ! [X420,X422,X421,X423] : (s(bool,X420) != s(bool,X421) | s(bool,X422) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X421) = s(bool,X422) | s(bool,'F_0') = s(bool,X420) | s(bool,'T_0') = s(bool,X423) | s(bool,'T_0') = s(bool,X423)) ) | ($spl3 | $spl6 | $spl24)),
  inference(superposition,[],[f6875,f5003])).
fof(f11137,plain,(
  ( ! [X416,X418,X417] : (s(bool,X416) != s(bool,X417) | s(bool,'F_0') = s(bool,X416) | s(bool,X417) != s(bool,X418) | s(bool,X418) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl132),
  inference(cnf_transformation,[],[f11137_D])).
fof(f11137_D,plain,(
  ( ! [X416,X418,X417] : (s(bool,X416) != s(bool,X417) | s(bool,'F_0') = s(bool,X416) | s(bool,X417) != s(bool,X418) | s(bool,X418) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f10833,plain,(
  ( ! [X416,X418,X417,X419] : (s(bool,X416) != s(bool,X417) | s(bool,X418) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X417) != s(bool,X418) | s(bool,'F_0') = s(bool,X416) | s(bool,'T_0') = s(bool,X419)) ) | ($spl6 | $spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f9964])).
fof(f9964,plain,(
  ( ! [X416,X418,X417,X419] : (s(bool,X416) != s(bool,X417) | s(bool,X418) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X417) != s(bool,X418) | s(bool,'F_0') = s(bool,X416) | s(bool,'T_0') = s(bool,X419) | s(bool,'T_0') = s(bool,X419)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3455,f5003])).
fof(f10834,plain,(
  ( ! [X412,X414,X413,X415] : (s(bool,X412) != s(bool,X413) | s(bool,'F_0') = s(bool,X414) | s(bool,X413) = s(bool,X414) | s(bool,'F_0') = s(bool,X412) | s(bool,'T_0') = s(bool,X415)) ) | ($spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f9963])).
fof(f9963,plain,(
  ( ! [X412,X414,X413,X415] : (s(bool,X412) != s(bool,X413) | s(bool,'F_0') = s(bool,X414) | s(bool,X413) = s(bool,X414) | s(bool,'F_0') = s(bool,X412) | s(bool,'T_0') = s(bool,X415) | s(bool,'T_0') = s(bool,X415)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f781,f5003])).
fof(f10837,plain,(
  ( ! [X405,X401,X403,X404,X402] : (s(bool,X401) != s(bool,X402) | s(bool,X402) != s(bool,X403) | s(bool,X403) != s(bool,X404) | s(bool,'T_0') = s(bool,X404) | s(bool,'F_0') = s(bool,X401) | s(bool,'T_0') = s(bool,X405)) ) | ($spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f9960])).
fof(f9960,plain,(
  ( ! [X405,X401,X403,X404,X402] : (s(bool,X401) != s(bool,X402) | s(bool,X402) != s(bool,X403) | s(bool,X403) != s(bool,X404) | s(bool,'T_0') = s(bool,X404) | s(bool,'F_0') = s(bool,X401) | s(bool,'T_0') = s(bool,X405) | s(bool,'T_0') = s(bool,X405)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f701,f5003])).
fof(f11134,plain,(
  ( ! [X350,X352,X349,X351,X348] : (~p(s(bool,X349)) | s(bool,X348) != s(bool,X350) | s(bool,X350) != s(bool,X351) | s(bool,'F_0') != s(bool,X351) | s(bool,X348) = s(bool,X352) | s(bool,'T_0') = s(bool,X352)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9946,f436])).
fof(f9946,plain,(
  ( ! [X350,X352,X349,X351,X348] : (~p(s(bool,X349)) | s(bool,X348) != s(bool,X350) | s(bool,X350) != s(bool,X351) | s(bool,'F_0') != s(bool,X351) | s(bool,'F_0') = s(bool,X349) | s(bool,X348) = s(bool,X352) | s(bool,'T_0') = s(bool,X352)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8131,f5003])).
fof(f11133,plain,(
  ( ! [X341,X343,X340,X342] : (s(bool,'T_0') != s(bool,X341) | s(bool,X342) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X340) = s(bool,X342) | s(bool,X340) = s(bool,X343) | s(bool,'T_0') = s(bool,X343)) ) | ($spl3 | $spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9944,f705])).
fof(f9944,plain,(
  ( ! [X341,X343,X340,X342] : (s(bool,'T_0') != s(bool,X341) | s(bool,X342) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X340) = s(bool,X342) | s(bool,'F_0') = s(bool,X341) | s(bool,X340) = s(bool,X343) | s(bool,'T_0') = s(bool,X343)) ) | ($spl3 | $spl6 | $spl24)),
  inference(superposition,[],[f6875,f5003])).
fof(f10850,plain,(
  ( ! [X329,X331,X328,X330] : (s(bool,'F_0') = s(bool,X329) | s(bool,X330) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X328) = s(bool,X330) | s(bool,X328) = s(bool,X331) | s(bool,'T_0') = s(bool,X331)) ) | ($spl6 | $spl24)),
  inference(duplicate_literal_removal,[],[f9941])).
fof(f9941,plain,(
  ( ! [X329,X331,X328,X330] : (s(bool,'F_0') = s(bool,X329) | s(bool,X330) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X328) = s(bool,X330) | s(bool,'F_0') = s(bool,X329) | s(bool,X328) = s(bool,X331) | s(bool,'T_0') = s(bool,X331)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f5704,f5003])).
fof(f11129,plain,(
  ( ! [X288,X285,X287] : (s(bool,X287) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X288) | s(bool,X285) = s(bool,X288) | s(bool,X285) != s(bool,X287)) ) | $spl126),
  inference(cnf_transformation,[],[f11129_D])).
fof(f11129_D,plain,(
  ( ! [X288,X285,X287] : (s(bool,X287) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X288) | s(bool,X285) = s(bool,X288) | s(bool,X285) != s(bool,X287)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f11128,plain,(
  ( ! [X288,X285,X287,X286] : (s(bool,'T_0') != s(bool,X286) | s(bool,X287) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X285) != s(bool,X287) | s(bool,X285) = s(bool,X288) | s(bool,'T_0') = s(bool,X288)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9931,f705])).
fof(f9931,plain,(
  ( ! [X288,X285,X287,X286] : (s(bool,'T_0') != s(bool,X286) | s(bool,X287) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X285) != s(bool,X287) | s(bool,'F_0') = s(bool,X286) | s(bool,X285) = s(bool,X288) | s(bool,'T_0') = s(bool,X288)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3455,f5003])).
fof(f11124,plain,(
  ( ! [X261,X263,X260,X262,X264] : (~p(s(bool,X261)) | s(bool,X262) = s(bool,X263) | s(bool,X260) = s(bool,X263) | s(bool,'F_0') != s(bool,X262) | s(bool,X260) = s(bool,X264) | s(bool,'T_0') = s(bool,X264)) ) | ($spl24 | $spl32)),
  inference(subsumption_resolution,[],[f9925,f436])).
fof(f9925,plain,(
  ( ! [X261,X263,X260,X262,X264] : (~p(s(bool,X261)) | s(bool,X262) = s(bool,X263) | s(bool,X260) = s(bool,X263) | s(bool,'F_0') != s(bool,X262) | s(bool,'F_0') = s(bool,X261) | s(bool,X260) = s(bool,X264) | s(bool,'T_0') = s(bool,X264)) ) | ($spl24 | $spl32)),
  inference(superposition,[],[f2853,f5003])).
fof(f11122,plain,(
  ( ! [X255,X253,X251,X254,X252] : (~p(s(bool,X252)) | s(bool,'F_0') = s(bool,X253) | s(bool,X253) = s(bool,X254) | s(bool,X251) != s(bool,X254) | s(bool,X251) = s(bool,X255) | s(bool,'T_0') = s(bool,X255)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9923,f436])).
fof(f9923,plain,(
  ( ! [X255,X253,X251,X254,X252] : (~p(s(bool,X252)) | s(bool,'F_0') = s(bool,X253) | s(bool,X253) = s(bool,X254) | s(bool,X251) != s(bool,X254) | s(bool,'F_0') = s(bool,X252) | s(bool,X251) = s(bool,X255) | s(bool,'T_0') = s(bool,X255)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f1996,f5003])).
fof(f1996,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | s(bool,'F_0') = s(bool,X19) | s(bool,X18) = s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl16),
  inference(subsumption_resolution,[],[f1960,f79])).
fof(f1960,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,X18) | s(bool,'F_0') = s(bool,X19) | s(bool,X18) = s(bool,X19) | ~p(s(bool,'T_0')) | ~p(s(bool,X17))) ) | $spl16),
  inference(superposition,[],[f781,f65])).
fof(f11120,plain,(
  ( ! [X241,X239,X242,X240,X238] : (~p(s(bool,X239)) | s(bool,'T_0') = s(bool,X240) | s(bool,X240) != s(bool,X241) | s(bool,X238) != s(bool,X241) | s(bool,X238) = s(bool,X242) | s(bool,'T_0') = s(bool,X242)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9920,f436])).
fof(f9920,plain,(
  ( ! [X241,X239,X242,X240,X238] : (~p(s(bool,X239)) | s(bool,'T_0') = s(bool,X240) | s(bool,X240) != s(bool,X241) | s(bool,X238) != s(bool,X241) | s(bool,'F_0') = s(bool,X239) | s(bool,X238) = s(bool,X242) | s(bool,'T_0') = s(bool,X242)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f1148,f5003])).
fof(f1148,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) != s(bool,X11) | s(bool,X9) != s(bool,X10)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f1130,f79])).
fof(f1130,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) != s(bool,X11) | ~p(s(bool,'T_0')) | ~p(s(bool,X9))) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f586,f65])).
fof(f11119,plain,(
  ( ! [X231,X229,X233,X232] : (s(bool,X229) != s(bool,X231) | s(bool,'T_0') = s(bool,X233) | s(bool,X229) = s(bool,X233) | p(s(bool,X232)) | s(bool,X231) != s(bool,X232)) ) | $spl120),
  inference(cnf_transformation,[],[f11119_D])).
fof(f11119_D,plain,(
  ( ! [X231,X229,X233,X232] : (s(bool,X229) != s(bool,X231) | s(bool,'T_0') = s(bool,X233) | s(bool,X229) = s(bool,X233) | p(s(bool,X232)) | s(bool,X231) != s(bool,X232)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f11118,plain,(
  ( ! [X231,X229,X233,X230,X232] : (~p(s(bool,X230)) | s(bool,X229) != s(bool,X231) | s(bool,X231) != s(bool,X232) | p(s(bool,X232)) | s(bool,X229) = s(bool,X233) | s(bool,'T_0') = s(bool,X233)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9918,f436])).
fof(f9918,plain,(
  ( ! [X231,X229,X233,X230,X232] : (~p(s(bool,X230)) | s(bool,X229) != s(bool,X231) | s(bool,X231) != s(bool,X232) | p(s(bool,X232)) | s(bool,'F_0') = s(bool,X230) | s(bool,X229) = s(bool,X233) | s(bool,'T_0') = s(bool,X233)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f907,f5003])).
fof(f11112,plain,(
  ( ! [X167,X171,X169,X170,X168] : (s(bool,'T_0') != s(bool,X168) | s(bool,X167) != s(bool,X169) | s(bool,X169) != s(bool,X170) | s(bool,'T_0') = s(bool,X170) | s(bool,X167) = s(bool,X171) | s(bool,'T_0') = s(bool,X171)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9903,f705])).
fof(f9903,plain,(
  ( ! [X167,X171,X169,X170,X168] : (s(bool,'T_0') != s(bool,X168) | s(bool,X167) != s(bool,X169) | s(bool,X169) != s(bool,X170) | s(bool,'T_0') = s(bool,X170) | s(bool,'F_0') = s(bool,X168) | s(bool,X167) = s(bool,X171) | s(bool,'T_0') = s(bool,X171)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f701,f5003])).
fof(f11110,plain,(
  ( ! [X165,X163,X166,X164,X162] : (~p(s(bool,X163)) | s(bool,'T_0') != s(bool,X164) | s(bool,X162) = s(bool,X165) | s(bool,X164) != s(bool,X165) | s(bool,X162) = s(bool,X166) | s(bool,'T_0') = s(bool,X166)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9902,f436])).
fof(f9902,plain,(
  ( ! [X165,X163,X166,X164,X162] : (~p(s(bool,X163)) | s(bool,'T_0') != s(bool,X164) | s(bool,X162) = s(bool,X165) | s(bool,X164) != s(bool,X165) | s(bool,'F_0') = s(bool,X163) | s(bool,X162) = s(bool,X166) | s(bool,'T_0') = s(bool,X166)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f696,f5003])).
fof(f11105,plain,(
  ( ! [X127,X125,X123,X126,X124] : (~p(s(bool,X124)) | s(bool,'T_0') = s(bool,X125) | s(bool,X123) = s(bool,X126) | s(bool,X125) = s(bool,X126) | s(bool,X123) = s(bool,X127) | s(bool,'T_0') = s(bool,X127)) ) | $spl24),
  inference(subsumption_resolution,[],[f9893,f436])).
fof(f9893,plain,(
  ( ! [X127,X125,X123,X126,X124] : (~p(s(bool,X124)) | s(bool,'T_0') = s(bool,X125) | s(bool,X123) = s(bool,X126) | s(bool,X125) = s(bool,X126) | s(bool,'F_0') = s(bool,X124) | s(bool,X123) = s(bool,X127) | s(bool,'T_0') = s(bool,X127)) ) | $spl24),
  inference(superposition,[],[f666,f5003])).
fof(f11104,plain,(
  ( ! [X121,X118,X122,X120] : (s(bool,X120) = s(bool,X121) | s(bool,'T_0') = s(bool,X122) | s(bool,X118) = s(bool,X122) | s(bool,X118) != s(bool,X120) | ~p(s(bool,X121))) ) | $spl114),
  inference(cnf_transformation,[],[f11104_D])).
fof(f11104_D,plain,(
  ( ! [X121,X118,X122,X120] : (s(bool,X120) = s(bool,X121) | s(bool,'T_0') = s(bool,X122) | s(bool,X118) = s(bool,X122) | s(bool,X118) != s(bool,X120) | ~p(s(bool,X121))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f11103,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,X119)) | s(bool,X120) = s(bool,X121) | ~p(s(bool,X121)) | s(bool,X118) != s(bool,X120) | s(bool,X118) = s(bool,X122) | s(bool,'T_0') = s(bool,X122)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9892,f436])).
fof(f9892,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,X119)) | s(bool,X120) = s(bool,X121) | ~p(s(bool,X121)) | s(bool,X118) != s(bool,X120) | s(bool,'F_0') = s(bool,X119) | s(bool,X118) = s(bool,X122) | s(bool,'T_0') = s(bool,X122)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f662,f5003])).
fof(f11101,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,X114)) | s(bool,X113) = s(bool,X115) | s(bool,X115) != s(bool,X116) | s(bool,'F_0') = s(bool,X116) | s(bool,X113) = s(bool,X117) | s(bool,'T_0') = s(bool,X117)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9891,f436])).
fof(f9891,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,X114)) | s(bool,X113) = s(bool,X115) | s(bool,X115) != s(bool,X116) | s(bool,'F_0') = s(bool,X116) | s(bool,'F_0') = s(bool,X114) | s(bool,X113) = s(bool,X117) | s(bool,'T_0') = s(bool,X117)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f642,f5003])).
fof(f11081,plain,(
  ( ! [X484,X486,X488,X485,X487] : (~p(s(bool,X486)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X485),s('type.hreal.hreal',X487))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X484),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X485,X484)),s('type.hreal.hreal',X487))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X484),s('type.hreal.hreal',X485))) | s(bool,X486) = s(bool,X488)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9849,f273])).
fof(f9849,plain,(
  ( ! [X484,X486,X488,X485,X487] : (~p(s(bool,X486)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X485),s('type.hreal.hreal',X487))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X484),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X485,X484)),s('type.hreal.hreal',X487))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X484),s('type.hreal.hreal',X485))) | s(bool,X486) = s(bool,X488) | s(bool,'T_0') = s(bool,X488)) ) | $spl24),
  inference(superposition,[],[f5121,f5003])).
fof(f11068,plain,(
  ( ! [X350,X352,X349,X351,X348] : (~p(s(bool,X349)) | s(bool,X348) != s(bool,X350) | s(bool,X350) != s(bool,X351) | s(bool,'F_0') != s(bool,X351) | s(bool,'F_0') = s(bool,X348) | s(bool,X349) = s(bool,X352)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9809,f273])).
fof(f9809,plain,(
  ( ! [X350,X352,X349,X351,X348] : (~p(s(bool,X349)) | s(bool,X348) != s(bool,X350) | s(bool,X350) != s(bool,X351) | s(bool,'F_0') != s(bool,X351) | s(bool,'F_0') = s(bool,X348) | s(bool,X349) = s(bool,X352) | s(bool,'T_0') = s(bool,X352)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8131,f5003])).
fof(f10854,plain,(
  ( ! [X333,X335,X332,X334] : (s(bool,X333) = s(bool,X334) | s(bool,'F_0') = s(bool,X332) | s(bool,X334) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X333) = s(bool,X335) | s(bool,'T_0') = s(bool,X335)) ) | ($spl6 | $spl24)),
  inference(duplicate_literal_removal,[],[f9805])).
fof(f9805,plain,(
  ( ! [X333,X335,X332,X334] : (s(bool,X333) = s(bool,X334) | s(bool,'F_0') = s(bool,X332) | s(bool,X334) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X332) | s(bool,X333) = s(bool,X335) | s(bool,'T_0') = s(bool,X335)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f5704,f5003])).
fof(f10856,plain,(
  ( ! [X302,X304,X306,X303,X305] : (s(bool,X303) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X304),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X304),s('type.hreal.hreal',X305))))) | s(bool,'F_0') = s(bool,X302) | s(bool,X303) = s(bool,X306) | s(bool,'T_0') = s(bool,X306)) ) | ($spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f9798])).
fof(f9798,plain,(
  ( ! [X302,X304,X306,X303,X305] : (s(bool,X303) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X304),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X304),s('type.hreal.hreal',X305))))) | s(bool,'F_0') = s(bool,X302) | s(bool,'F_0') = s(bool,X302) | s(bool,X303) = s(bool,X306) | s(bool,'T_0') = s(bool,X306)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f4294,f5003])).
fof(f11055,plain,(
  ( ! [X261,X263,X260,X262,X264] : (~p(s(bool,X261)) | s(bool,X262) = s(bool,X263) | s(bool,X260) = s(bool,X263) | s(bool,'F_0') != s(bool,X262) | s(bool,'F_0') = s(bool,X260) | s(bool,X261) = s(bool,X264)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f9788,f273])).
fof(f9788,plain,(
  ( ! [X261,X263,X260,X262,X264] : (~p(s(bool,X261)) | s(bool,X262) = s(bool,X263) | s(bool,X260) = s(bool,X263) | s(bool,'F_0') != s(bool,X262) | s(bool,'F_0') = s(bool,X260) | s(bool,X261) = s(bool,X264) | s(bool,'T_0') = s(bool,X264)) ) | ($spl24 | $spl32)),
  inference(superposition,[],[f2853,f5003])).
fof(f11053,plain,(
  ( ! [X255,X253,X251,X254,X252] : (~p(s(bool,X252)) | s(bool,'F_0') = s(bool,X253) | s(bool,X253) = s(bool,X254) | s(bool,X251) != s(bool,X254) | s(bool,'F_0') = s(bool,X251) | s(bool,X252) = s(bool,X255)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9786,f273])).
fof(f9786,plain,(
  ( ! [X255,X253,X251,X254,X252] : (~p(s(bool,X252)) | s(bool,'F_0') = s(bool,X253) | s(bool,X253) = s(bool,X254) | s(bool,X251) != s(bool,X254) | s(bool,'F_0') = s(bool,X251) | s(bool,X252) = s(bool,X255) | s(bool,'T_0') = s(bool,X255)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f1996,f5003])).
fof(f11041,plain,(
  ( ! [X165,X163,X166,X164,X162] : (~p(s(bool,X163)) | s(bool,'T_0') != s(bool,X164) | s(bool,X162) = s(bool,X165) | s(bool,X164) != s(bool,X165) | s(bool,'F_0') = s(bool,X162) | s(bool,X163) = s(bool,X166)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9765,f273])).
fof(f9765,plain,(
  ( ! [X165,X163,X166,X164,X162] : (~p(s(bool,X163)) | s(bool,'T_0') != s(bool,X164) | s(bool,X162) = s(bool,X165) | s(bool,X164) != s(bool,X165) | s(bool,'F_0') = s(bool,X162) | s(bool,X163) = s(bool,X166) | s(bool,'T_0') = s(bool,X166)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f696,f5003])).
fof(f11034,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,X114)) | s(bool,X113) = s(bool,X115) | s(bool,X115) != s(bool,X116) | s(bool,'F_0') = s(bool,X116) | s(bool,'F_0') = s(bool,X113) | s(bool,X114) = s(bool,X117)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9754,f273])).
fof(f9754,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,X114)) | s(bool,X113) = s(bool,X115) | s(bool,X115) != s(bool,X116) | s(bool,'F_0') = s(bool,X116) | s(bool,'F_0') = s(bool,X113) | s(bool,X114) = s(bool,X117) | s(bool,'T_0') = s(bool,X117)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f642,f5003])).
fof(f10864,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,X56) | s(bool,X57) = s(bool,X58) | s(bool,'F_0') = s(bool,X57) | s(bool,X55) = s(bool,X58) | s(bool,X55) = s(bool,X59) | s(bool,'F_0') = s(bool,X59)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f9706])).
fof(f9706,plain,(
  ( ! [X59,X57,X58,X56,X55] : (s(bool,'T_0') = s(bool,X56) | s(bool,X57) = s(bool,X58) | s(bool,'F_0') = s(bool,X57) | s(bool,X55) = s(bool,X58) | s(bool,X55) = s(bool,X59) | s(bool,'F_0') = s(bool,X59) | s(bool,'T_0') = s(bool,X56)) ) | $spl24),
  inference(superposition,[],[f5003,f5003])).
fof(f10867,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X85) = s(bool,X86) | s(bool,X87) = s(bool,X88) | s(bool,'F_0') = s(bool,X87) | s(bool,X86) = s(bool,X88) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X89)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f9690])).
fof(f9690,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s(bool,X85) = s(bool,X86) | s(bool,X87) = s(bool,X88) | s(bool,'F_0') = s(bool,X87) | s(bool,X86) = s(bool,X88) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) ) | $spl24),
  inference(superposition,[],[f5003,f5003])).
fof(f10870,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s(bool,X73) = s(bool,X74) | s(bool,X74) = s(bool,X75) | s(bool,'F_0') = s(bool,X75) | s(bool,'T_0') = s(bool,X72) | s(bool,X73) = s(bool,X76) | s(bool,'F_0') = s(bool,X76)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f9648])).
fof(f9648,plain,(
  ( ! [X76,X74,X72,X75,X73] : (s(bool,X73) = s(bool,X74) | s(bool,X74) = s(bool,X75) | s(bool,'F_0') = s(bool,X75) | s(bool,'T_0') = s(bool,X72) | s(bool,X73) = s(bool,X76) | s(bool,'F_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X72)) ) | $spl24),
  inference(superposition,[],[f5003,f5003])).
fof(f10876,plain,(
  ( ! [X154,X153] : (s(bool,'T_0') = s(bool,X153) | s(bool,X154) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X154)) ) | ($spl4 | $spl24)),
  inference(duplicate_literal_removal,[],[f9627])).
fof(f9627,plain,(
  ( ! [X154,X153] : (s(bool,'T_0') = s(bool,X153) | s(bool,X154) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | s(bool,'F_0') = s(bool,X154) | s(bool,'T_0') = s(bool,X153)) ) | ($spl4 | $spl24)),
  inference(superposition,[],[f5003,f112])).
fof(f10880,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,X129) = s(bool,X130) | s(bool,'F_0') = s(bool,X131) | s(bool,'T_0') = s(bool,X130) | s(bool,X132) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X129) = s(bool,X132)) ) | ($spl6 | $spl24)),
  inference(duplicate_literal_removal,[],[f9620])).
fof(f9620,plain,(
  ( ! [X132,X130,X131,X129] : (s(bool,X129) = s(bool,X130) | s(bool,'F_0') = s(bool,X131) | s(bool,'F_0') = s(bool,X131) | s(bool,'T_0') = s(bool,X130) | s(bool,X132) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X129) = s(bool,X132)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f5003,f5704])).
fof(f10881,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,X124) = s(bool,X125) | s(bool,'F_0') = s(bool,X126) | s(bool,'T_0') = s(bool,X125) | s(bool,X124) = s(bool,X127) | s(bool,X127) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f9619])).
fof(f9619,plain,(
  ( ! [X127,X125,X128,X126,X124] : (s(bool,X124) = s(bool,X125) | s(bool,'F_0') = s(bool,X126) | s(bool,'F_0') = s(bool,X126) | s(bool,'T_0') = s(bool,X125) | s(bool,X124) = s(bool,X127) | s(bool,X127) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) ) | $spl24),
  inference(superposition,[],[f5003,f5003])).
fof(f10882,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s(bool,'T_0') = s(bool,X63) | s(bool,X62) = s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,X65) = s(bool,X66) | s(bool,'F_0') = s(bool,X65) | s(bool,X62) = s(bool,X66)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f9605])).
fof(f9605,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s(bool,'T_0') = s(bool,X63) | s(bool,X62) = s(bool,X64) | s(bool,'F_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X63) | s(bool,X65) = s(bool,X66) | s(bool,'F_0') = s(bool,X65) | s(bool,X62) = s(bool,X66)) ) | $spl24),
  inference(superposition,[],[f5003,f5003])).
fof(f10886,plain,(
  ( ! [X111,X114,X112,X115,X113] : (s(bool,'F_0') = s(bool,X113) | s(bool,X114) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X111),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X111),s('type.hreal.hreal',X112))))) | s(bool,X114) = s(bool,X115) | s(bool,'T_0') = s(bool,X115)) ) | ($spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f9587])).
fof(f9587,plain,(
  ( ! [X111,X114,X112,X115,X113] : (s(bool,'F_0') = s(bool,X113) | s(bool,X114) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X111),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X111),s('type.hreal.hreal',X112))))) | s(bool,X114) = s(bool,X115) | s(bool,'T_0') = s(bool,X115) | s(bool,'F_0') = s(bool,X113)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f5003,f4294])).
fof(f10887,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'F_0') = s(bool,X90) | s(bool,X89) = s(bool,X91) | s(bool,X91) = s(bool,X92) | s(bool,'T_0') = s(bool,X92) | s(bool,X89) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl24)),
  inference(duplicate_literal_removal,[],[f9583])).
fof(f9583,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'F_0') = s(bool,X90) | s(bool,X89) = s(bool,X91) | s(bool,X91) = s(bool,X92) | s(bool,'T_0') = s(bool,X92) | s(bool,'F_0') = s(bool,X90) | s(bool,X89) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f5003,f5704])).
fof(f10888,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,'F_0') = s(bool,X75) | s(bool,X74) = s(bool,X76) | s(bool,X76) = s(bool,X77) | s(bool,'T_0') = s(bool,X77) | s(bool,X74) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f9580])).
fof(f9580,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,'F_0') = s(bool,X75) | s(bool,X74) = s(bool,X76) | s(bool,X76) = s(bool,X77) | s(bool,'T_0') = s(bool,X77) | s(bool,'F_0') = s(bool,X75) | s(bool,X74) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) ) | $spl24),
  inference(superposition,[],[f5003,f5003])).
fof(f10891,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X86) = s(bool,X87) | s(bool,X87) = s(bool,X88) | s(bool,X88) = s(bool,X89) | s(bool,'T_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X86)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f9555])).
fof(f9555,plain,(
  ( ! [X90,X88,X87,X89,X86] : (s(bool,X86) = s(bool,X87) | s(bool,X87) = s(bool,X88) | s(bool,X88) = s(bool,X89) | s(bool,'T_0') = s(bool,X89) | s(bool,'F_0') = s(bool,X90) | s(bool,'F_0') = s(bool,X90) | s(bool,'T_0') = s(bool,X86)) ) | $spl24),
  inference(superposition,[],[f5003,f5003])).
fof(f10898,plain,(
  ( ! [X111,X109,X112,X110,X108] : (s(bool,X108) = s(bool,X109) | s(bool,'F_0') = s(bool,X109) | s(bool,'T_0') = s(bool,X110) | s(bool,X111) = s(bool,X112) | s(bool,'F_0') = s(bool,X111) | s(bool,X108) = s(bool,X112)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f9520])).
fof(f9520,plain,(
  ( ! [X111,X109,X112,X110,X108] : (s(bool,X108) = s(bool,X109) | s(bool,'F_0') = s(bool,X109) | s(bool,'T_0') = s(bool,X110) | s(bool,'T_0') = s(bool,X110) | s(bool,X111) = s(bool,X112) | s(bool,'F_0') = s(bool,X111) | s(bool,X108) = s(bool,X112)) ) | $spl24),
  inference(superposition,[],[f5003,f5003])).
fof(f10899,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'F_0') = s(bool,X97) | s(bool,X96) = s(bool,X98) | s(bool,'T_0') = s(bool,X98) | s(bool,X99) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X96) = s(bool,X99)) ) | ($spl6 | $spl24)),
  inference(duplicate_literal_removal,[],[f9517])).
fof(f9517,plain,(
  ( ! [X99,X97,X98,X96] : (s(bool,'F_0') = s(bool,X97) | s(bool,'F_0') = s(bool,X97) | s(bool,X96) = s(bool,X98) | s(bool,'T_0') = s(bool,X98) | s(bool,X99) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X96) = s(bool,X99)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f5003,f5704])).
fof(f10900,plain,(
  ( ! [X90,X88,X87,X91,X89] : (s(bool,'F_0') = s(bool,X88) | s(bool,X87) = s(bool,X89) | s(bool,'T_0') = s(bool,X89) | s(bool,X87) = s(bool,X90) | s(bool,X90) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f9515])).
fof(f9515,plain,(
  ( ! [X90,X88,X87,X91,X89] : (s(bool,'F_0') = s(bool,X88) | s(bool,'F_0') = s(bool,X88) | s(bool,X87) = s(bool,X89) | s(bool,'T_0') = s(bool,X89) | s(bool,X87) = s(bool,X90) | s(bool,X90) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) ) | $spl24),
  inference(superposition,[],[f5003,f5003])).
fof(f10904,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,X101) = s(bool,X102) | s(bool,'F_0') = s(bool,X100) | s(bool,X102) = s(bool,X103) | s(bool,'T_0') = s(bool,X103) | s(bool,X101) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl24)),
  inference(duplicate_literal_removal,[],[f9477])).
fof(f9477,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,X101) = s(bool,X102) | s(bool,'F_0') = s(bool,X100) | s(bool,X102) = s(bool,X103) | s(bool,'T_0') = s(bool,X103) | s(bool,'F_0') = s(bool,X100) | s(bool,X101) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f5003,f5704])).
fof(f10905,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,X78) = s(bool,X79) | s(bool,'F_0') = s(bool,X77) | s(bool,X79) = s(bool,X80) | s(bool,'T_0') = s(bool,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f9472])).
fof(f9472,plain,(
  ( ! [X80,X78,X81,X79,X77] : (s(bool,X78) = s(bool,X79) | s(bool,'F_0') = s(bool,X77) | s(bool,X79) = s(bool,X80) | s(bool,'T_0') = s(bool,X80) | s(bool,'F_0') = s(bool,X77) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) ) | $spl24),
  inference(superposition,[],[f5003,f5003])).
fof(f10906,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s(bool,X60) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X58),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X58),s('type.hreal.hreal',X59))))) | s(bool,'F_0') = s(bool,X57) | s(bool,X60) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) ) | ($spl16 | $spl24)),
  inference(duplicate_literal_removal,[],[f9468])).
fof(f9468,plain,(
  ( ! [X61,X59,X57,X60,X58] : (s(bool,X60) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X58),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X58),s('type.hreal.hreal',X59))))) | s(bool,'F_0') = s(bool,X57) | s(bool,X60) = s(bool,X61) | s(bool,'T_0') = s(bool,X61) | s(bool,'F_0') = s(bool,X57)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f5003,f4294])).
fof(f9375,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'F_0') = s(bool,X68) | s(bool,X66) != s(bool,X68) | s(bool,'T_0') = s(bool,X67) | s(bool,X69) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X66) = s(bool,X69)) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f9315,f68])).
fof(f9315,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'F_0') != s(bool,X67) | s(bool,'F_0') = s(bool,X68) | s(bool,X66) != s(bool,X68) | s(bool,'T_0') = s(bool,X67) | s(bool,X69) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X66) = s(bool,X69)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f753,f5704])).
fof(f9374,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'F_0') = s(bool,X64) | s(bool,X62) != s(bool,X64) | s(bool,'T_0') = s(bool,X63) | s(bool,'F_0') = s(bool,X65) | s(bool,X62) = s(bool,X65)) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f9314,f695])).
fof(f9314,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X63) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X64) | s(bool,X62) != s(bool,X64) | s(bool,'T_0') = s(bool,X63) | s(bool,'F_0') = s(bool,X65) | s(bool,X62) = s(bool,X65)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f753,f5704])).
fof(f9372,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'F_0') = s(bool,X47) | s(bool,X45) != s(bool,X47) | s(bool,'T_0') = s(bool,X46) | s(bool,X45) = s(bool,X48) | s(bool,'F_0') = s(bool,X48)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9310,f68])).
fof(f9310,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'F_0') != s(bool,X46) | s(bool,'F_0') = s(bool,X47) | s(bool,X45) != s(bool,X47) | s(bool,'T_0') = s(bool,X46) | s(bool,X45) = s(bool,X48) | s(bool,'F_0') = s(bool,X48)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f753,f782])).
fof(f9371,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') = s(bool,X43) | s(bool,X41) != s(bool,X43) | s(bool,'T_0') = s(bool,X42) | s(bool,X41) = s(bool,X44) | s(bool,'F_0') = s(bool,X44)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9309,f68])).
fof(f9309,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,X42) | s(bool,'F_0') = s(bool,X43) | s(bool,X41) != s(bool,X43) | s(bool,'T_0') = s(bool,X42) | s(bool,X41) = s(bool,X44) | s(bool,'F_0') = s(bool,X44)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f753,f782])).
fof(f9296,plain,(
  ( ! [X118,X116,X119,X117,X115] : (s(bool,'T_0') != s(bool,X117) | s(bool,X118) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal',X116))))) | s(bool,X118) != s(bool,X119) | s(bool,'T_0') = s(bool,X119)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9263,f705])).
fof(f9263,plain,(
  ( ! [X118,X116,X119,X117,X115] : (s(bool,'T_0') != s(bool,X117) | s(bool,X118) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X115),s('type.hreal.hreal',X116))))) | s(bool,X118) != s(bool,X119) | s(bool,'T_0') = s(bool,X119) | s(bool,'F_0') = s(bool,X117)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f701,f4294])).
fof(f9290,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'T_0') != s(bool,X75) | s(bool,X74) != s(bool,X76) | s(bool,X76) != s(bool,X77) | s(bool,'T_0') = s(bool,X77) | s(bool,X74) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9251,f705])).
fof(f9251,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'T_0') != s(bool,X75) | s(bool,X74) != s(bool,X76) | s(bool,X76) != s(bool,X77) | s(bool,'T_0') = s(bool,X77) | s(bool,'F_0') = s(bool,X75) | s(bool,X74) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f701,f5704])).
fof(f3560,plain,(
  ( ! [X30,X33,X32] : (s(bool,'F_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,X32) != s(bool,X33) | s(bool,X30) != s(bool,X33)) ) | $spl58),
  inference(cnf_transformation,[],[f3560_D])).
fof(f3560_D,plain,(
  ( ! [X30,X33,X32] : (s(bool,'F_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,X32) != s(bool,X33) | s(bool,X30) != s(bool,X33)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f9289,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') != s(bool,X71) | s(bool,X70) != s(bool,X72) | s(bool,X72) != s(bool,X73) | s(bool,'T_0') = s(bool,X73) | s(bool,'F_0') = s(bool,X70)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9250,f727])).
fof(f9250,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') != s(bool,X71) | s(bool,X70) != s(bool,X72) | s(bool,X72) != s(bool,X73) | s(bool,'T_0') = s(bool,X73) | s(bool,'F_0') = s(bool,X70) | s(bool,X71) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f701,f5704])).
fof(f9287,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,'T_0') != s(bool,X54) | s(bool,X53) != s(bool,X55) | s(bool,X55) != s(bool,X56) | s(bool,'T_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9246,f705])).
fof(f9246,plain,(
  ( ! [X54,X56,X55,X53] : (s(bool,'T_0') != s(bool,X54) | s(bool,X53) != s(bool,X55) | s(bool,X55) != s(bool,X56) | s(bool,'T_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f701,f782])).
fof(f9286,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'T_0') != s(bool,X50) | s(bool,X49) != s(bool,X51) | s(bool,X51) != s(bool,X52) | s(bool,'T_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X49)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9245,f705])).
fof(f9245,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'T_0') != s(bool,X50) | s(bool,X49) != s(bool,X51) | s(bool,X51) != s(bool,X52) | s(bool,'T_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f701,f782])).
fof(f9198,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'T_0') = s(bool,X68) | s(bool,X66) = s(bool,X68) | s(bool,'T_0') = s(bool,X67) | s(bool,X69) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X66) = s(bool,X69)) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f9134,f68])).
fof(f9134,plain,(
  ( ! [X68,X66,X69,X67] : (s(bool,'F_0') != s(bool,X67) | s(bool,'T_0') = s(bool,X68) | s(bool,X66) = s(bool,X68) | s(bool,'T_0') = s(bool,X67) | s(bool,X69) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X66) = s(bool,X69)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f671,f5704])).
fof(f9197,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'T_0') = s(bool,X64) | s(bool,X62) = s(bool,X64) | s(bool,'T_0') = s(bool,X63) | s(bool,'F_0') = s(bool,X65) | s(bool,X62) = s(bool,X65)) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f9133,f695])).
fof(f9133,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,X63) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X64) | s(bool,X62) = s(bool,X64) | s(bool,'T_0') = s(bool,X63) | s(bool,'F_0') = s(bool,X65) | s(bool,X62) = s(bool,X65)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f671,f5704])).
fof(f9194,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'T_0') = s(bool,X47) | s(bool,X45) = s(bool,X47) | s(bool,'T_0') = s(bool,X46) | s(bool,X45) = s(bool,X48) | s(bool,'F_0') = s(bool,X48)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9129,f68])).
fof(f9129,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'F_0') != s(bool,X46) | s(bool,'T_0') = s(bool,X47) | s(bool,X45) = s(bool,X47) | s(bool,'T_0') = s(bool,X46) | s(bool,X45) = s(bool,X48) | s(bool,'F_0') = s(bool,X48)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f671,f782])).
fof(f9193,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'T_0') = s(bool,X43) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X42) | s(bool,X41) = s(bool,X44) | s(bool,'F_0') = s(bool,X44)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f9128,f68])).
fof(f9128,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,X42) | s(bool,'T_0') = s(bool,X43) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X42) | s(bool,X41) = s(bool,X44) | s(bool,'F_0') = s(bool,X44)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f671,f782])).
fof(f8997,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X37)) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X34),s('type.hreal.hreal',X35))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X34),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X35),s('type.hreal.hreal',X36)))))))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f8977,f133])).
fof(f8977,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X37)) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X34),s('type.hreal.hreal',X35))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X34),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X35),s('type.hreal.hreal',X36)))))))) ) | $spl24),
  inference(superposition,[],[f8633,f672])).
fof(f8995,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X33)) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X30),s('type.hreal.hreal',X31))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X30),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X31),s('type.hreal.hreal',X32)))))))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f8976,f133])).
fof(f8976,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X33)) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X30),s('type.hreal.hreal',X31))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X30),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X31),s('type.hreal.hreal',X32))))))) | s(bool,'T_0') = s(bool,X33)) ) | $spl24),
  inference(superposition,[],[f8633,f672])).
fof(f8835,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal',X89))) | ~p(s(bool,X90)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X89),s('type.hreal.hreal',X91))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X89,X88)),s('type.hreal.hreal',X91)))))) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f8834,f113])).
fof(f8834,plain,(
  ( ! [X90,X88,X91,X89] : (~p(s(bool,X90)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X89),s('type.hreal.hreal',X91))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X89,X88)),s('type.hreal.hreal',X91))))) | s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X88),s('type.hreal.hreal',X89)))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f8797,f436])).
fof(f8833,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'F_0') = s(bool,X86) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X85),s('type.hreal.hreal',X87))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X85,X84)),s('type.hreal.hreal',X87))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal',X85)))) ) | $spl6),
  inference(subsumption_resolution,[],[f8832,f4678])).
fof(f8832,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'F_0') = s(bool,X86) | ~p(s(bool,X86)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X85),s('type.hreal.hreal',X87))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X85,X84)),s('type.hreal.hreal',X87))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal',X85)))) ) | $spl6),
  inference(forward_demodulation,[],[f8796,f113])).
fof(f8796,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,X86)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X85),s('type.hreal.hreal',X87))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X85,X84)),s('type.hreal.hreal',X87))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal',X85))) | s(bool,X86) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(superposition,[],[f5121,f5704])).
fof(f8828,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X69)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X68),s('type.hreal.hreal',X70))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X68,X67)),s('type.hreal.hreal',X70))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal',X68)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f8792,f436])).
fof(f8792,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X69)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X68),s('type.hreal.hreal',X70))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X68,X67)),s('type.hreal.hreal',X70))))) | s(bool,'F_0') = s(bool,X69) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X67),s('type.hreal.hreal',X68)))) ) | $spl16),
  inference(superposition,[],[f5121,f782])).
fof(f8826,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X64),s('type.hreal.hreal',X66))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X64,X63)),s('type.hreal.hreal',X66))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal',X64)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f8791,f436])).
fof(f8791,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X65)) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X64),s('type.hreal.hreal',X66))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK3(X64,X63)),s('type.hreal.hreal',X66))))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X63),s('type.hreal.hreal',X64))) | s(bool,'F_0') = s(bool,X65)) ) | $spl16),
  inference(superposition,[],[f5121,f782])).
fof(f8502,plain,(
  ( ! [X57,X56,X55] : (s(bool,'T_0') != s(bool,X56) | s(bool,X57) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X55) = s(bool,X57) | s(bool,X55) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl3 | $spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f8459,f705])).
fof(f8300,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'F_0') = s(bool,X74) | ~p(s(bool,X75)) | s(bool,X74) != s(bool,X76) | s(bool,X76) != s(bool,X77) | s(bool,'F_0') != s(bool,X77)) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f8299,f113])).
fof(f8299,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X75)) | s(bool,X74) != s(bool,X76) | s(bool,X76) != s(bool,X77) | s(bool,'F_0') != s(bool,X77) | s(bool,X74) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f8271,f436])).
fof(f8271,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X75)) | s(bool,X74) != s(bool,X76) | s(bool,X76) != s(bool,X77) | s(bool,'F_0') != s(bool,X77) | s(bool,'F_0') = s(bool,X75) | s(bool,X74) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8131,f5704])).
fof(f8298,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'F_0') = s(bool,X71) | s(bool,X70) != s(bool,X72) | s(bool,X72) != s(bool,X73) | s(bool,'F_0') != s(bool,X73) | s(bool,'F_0') = s(bool,X70)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f8297,f4678])).
fof(f8297,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'F_0') = s(bool,X71) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,X72) | s(bool,X72) != s(bool,X73) | s(bool,'F_0') != s(bool,X73) | s(bool,'F_0') = s(bool,X70)) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f8270,f113])).
fof(f8270,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X71)) | s(bool,X70) != s(bool,X72) | s(bool,X72) != s(bool,X73) | s(bool,'F_0') != s(bool,X73) | s(bool,'F_0') = s(bool,X70) | s(bool,X71) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8131,f5704])).
fof(f8293,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') != s(bool,X56) | s(bool,'F_0') = s(bool,X53)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f8266,f436])).
fof(f8266,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') != s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8131,f782])).
fof(f8291,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X49) != s(bool,X51) | s(bool,X51) != s(bool,X52) | s(bool,'F_0') != s(bool,X52) | s(bool,'F_0') = s(bool,X49)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f8265,f436])).
fof(f8265,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X49) != s(bool,X51) | s(bool,X51) != s(bool,X52) | s(bool,'F_0') != s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f8131,f782])).
fof(f8205,plain,(
  $false | ($spl4 | $spl84)),
  inference(trivial_inequality_removal,[],[f8200])).
fof(f8200,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl4 | $spl84)),
  inference(superposition,[],[f8147,f112])).
fof(f8204,plain,(
  $false | ($spl24 | $spl84)),
  inference(trivial_inequality_removal,[],[f8202])).
fof(f8202,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl24 | $spl84)),
  inference(superposition,[],[f8147,f4167])).
fof(f8203,plain,(
  $false | $spl84),
  inference(equality_resolution,[],[f8147])).
fof(f8152,plain,(
  ( ! [X87,X85,X84] : (s(bool,'F_0') != s(bool,X87) | s(bool,X87) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal',X85)))))) ) | $spl86),
  inference(cnf_transformation,[],[f8152_D])).
fof(f8152_D,plain,(
  ( ! [X87,X85,X84] : (s(bool,'F_0') != s(bool,X87) | s(bool,X87) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal',X85)))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f8151,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'F_0') != s(bool,X87) | s(bool,'T_0') != s(bool,X86) | s(bool,X87) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal',X85)))))) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f8150,f113])).
fof(f8150,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X86) | s(bool,X87) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X87) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal',X85)))))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f8114,f705])).
fof(f8114,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') != s(bool,X86) | s(bool,X87) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X87) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X84),s('type.hreal.hreal',X85))))) | s(bool,'F_0') = s(bool,X86)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3455,f4294])).
fof(f7573,plain,(
  ( ! [X81,X79] : (s(bool,'T_0') != s(bool,X79) | s(bool,X79) = s(bool,X81) | s(bool,X81) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl78),
  inference(cnf_transformation,[],[f7573_D])).
fof(f7573_D,plain,(
  ( ! [X81,X79] : (s(bool,'T_0') != s(bool,X79) | s(bool,X79) = s(bool,X81) | s(bool,X81) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f8148,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') != s(bool,X66) | s(bool,'F_0') != s(bool,X67) | s(bool,X68) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X66) = s(bool,X68)) ) | ($spl6 | $spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f8106,f2900])).
fof(f8106,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') != s(bool,X66) | s(bool,X67) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') != s(bool,X67) | s(bool,X68) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X66) = s(bool,X68)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3455,f5704])).
fof(f8146,plain,(
  ( ! [X57,X55] : (s(bool,X57) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X55) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X55) != s(bool,X57)) ) | $spl82),
  inference(cnf_transformation,[],[f8146_D])).
fof(f8146_D,plain,(
  ( ! [X57,X55] : (s(bool,X57) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X55) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X55) != s(bool,X57)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f8145,plain,(
  ( ! [X57,X56,X55] : (s(bool,'T_0') != s(bool,X56) | s(bool,X57) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X55) != s(bool,X57) | s(bool,X55) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f8102,f705])).
fof(f8102,plain,(
  ( ! [X57,X56,X55] : (s(bool,'T_0') != s(bool,X56) | s(bool,X57) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X55) != s(bool,X57) | s(bool,'F_0') = s(bool,X56) | s(bool,X55) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3455,f5704])).
fof(f7891,plain,(
  ( ! [X173,X171,X172,X170] : (s(bool,'F_0') = s(bool,X170) | ~p(s(bool,X171)) | s(bool,X172) = s(bool,X173) | s(bool,X170) = s(bool,X173) | s(bool,'F_0') != s(bool,X172)) ) | ($spl6 | $spl24 | $spl32)),
  inference(forward_demodulation,[],[f7890,f113])).
fof(f7890,plain,(
  ( ! [X173,X171,X172,X170] : (~p(s(bool,X171)) | s(bool,X172) = s(bool,X173) | s(bool,X170) = s(bool,X173) | s(bool,'F_0') != s(bool,X172) | s(bool,X170) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f7352,f436])).
fof(f7889,plain,(
  ( ! [X165,X163,X166,X164] : (s(bool,'F_0') = s(bool,X163) | ~p(s(bool,X164)) | s(bool,'F_0') = s(bool,X165) | s(bool,X165) = s(bool,X166) | s(bool,X163) != s(bool,X166)) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f7888,f113])).
fof(f7888,plain,(
  ( ! [X165,X163,X166,X164] : (~p(s(bool,X164)) | s(bool,'F_0') = s(bool,X165) | s(bool,X165) = s(bool,X166) | s(bool,X163) != s(bool,X166) | s(bool,X163) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f7350,f436])).
fof(f7350,plain,(
  ( ! [X165,X163,X166,X164] : (~p(s(bool,X164)) | s(bool,'F_0') = s(bool,X165) | s(bool,X165) = s(bool,X166) | s(bool,X163) != s(bool,X166) | s(bool,'F_0') = s(bool,X164) | s(bool,X163) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16)),
  inference(superposition,[],[f1996,f5704])).
fof(f7887,plain,(
  ( ! [X158,X156,X159,X157] : (s(bool,'F_0') = s(bool,X156) | ~p(s(bool,X157)) | s(bool,'T_0') = s(bool,X158) | s(bool,X158) != s(bool,X159) | s(bool,X156) != s(bool,X159)) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f7886,f113])).
fof(f7886,plain,(
  ( ! [X158,X156,X159,X157] : (~p(s(bool,X157)) | s(bool,'T_0') = s(bool,X158) | s(bool,X158) != s(bool,X159) | s(bool,X156) != s(bool,X159) | s(bool,X156) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f7348,f436])).
fof(f7348,plain,(
  ( ! [X158,X156,X159,X157] : (~p(s(bool,X157)) | s(bool,'T_0') = s(bool,X158) | s(bool,X158) != s(bool,X159) | s(bool,X156) != s(bool,X159) | s(bool,'F_0') = s(bool,X157) | s(bool,X156) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f1148,f5704])).
fof(f7880,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,'F_0') = s(bool,X113) | ~p(s(bool,X114)) | s(bool,'T_0') != s(bool,X115) | s(bool,X113) = s(bool,X116) | s(bool,X115) != s(bool,X116)) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f7879,f113])).
fof(f7879,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,X114)) | s(bool,'T_0') != s(bool,X115) | s(bool,X113) = s(bool,X116) | s(bool,X115) != s(bool,X116) | s(bool,X113) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f7334,f436])).
fof(f7334,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,X114)) | s(bool,'T_0') != s(bool,X115) | s(bool,X113) = s(bool,X116) | s(bool,X115) != s(bool,X116) | s(bool,'F_0') = s(bool,X114) | s(bool,X113) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f696,f5704])).
fof(f7875,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'F_0') = s(bool,X91) | ~p(s(bool,X92)) | s(bool,'T_0') = s(bool,X93) | s(bool,X91) = s(bool,X94) | s(bool,X93) = s(bool,X94)) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f7874,f113])).
fof(f7874,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X92)) | s(bool,'T_0') = s(bool,X93) | s(bool,X91) = s(bool,X94) | s(bool,X93) = s(bool,X94) | s(bool,X91) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f7327,f436])).
fof(f7327,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X92)) | s(bool,'T_0') = s(bool,X93) | s(bool,X91) = s(bool,X94) | s(bool,X93) = s(bool,X94) | s(bool,'F_0') = s(bool,X92) | s(bool,X91) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f666,f5704])).
fof(f7871,plain,(
  ( ! [X85,X83,X86,X84] : (s(bool,'F_0') = s(bool,X83) | ~p(s(bool,X84)) | s(bool,X83) = s(bool,X85) | s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X86)) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f7870,f113])).
fof(f7870,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | s(bool,X83) = s(bool,X85) | s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X86) | s(bool,X83) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f7325,f436])).
fof(f7325,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | s(bool,X83) = s(bool,X85) | s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X86) | s(bool,'F_0') = s(bool,X84) | s(bool,X83) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16)),
  inference(superposition,[],[f642,f5704])).
fof(f7815,plain,(
  ( ! [X173,X171,X172,X170] : (s(bool,'F_0') = s(bool,X171) | s(bool,X172) = s(bool,X173) | s(bool,X170) = s(bool,X173) | s(bool,'F_0') != s(bool,X172) | s(bool,'F_0') = s(bool,X170)) ) | ($spl6 | $spl32)),
  inference(subsumption_resolution,[],[f7814,f4678])).
fof(f7814,plain,(
  ( ! [X173,X171,X172,X170] : (s(bool,'F_0') = s(bool,X171) | ~p(s(bool,X171)) | s(bool,X172) = s(bool,X173) | s(bool,X170) = s(bool,X173) | s(bool,'F_0') != s(bool,X172) | s(bool,'F_0') = s(bool,X170)) ) | ($spl6 | $spl32)),
  inference(forward_demodulation,[],[f7245,f113])).
fof(f7245,plain,(
  ( ! [X173,X171,X172,X170] : (~p(s(bool,X171)) | s(bool,X172) = s(bool,X173) | s(bool,X170) = s(bool,X173) | s(bool,'F_0') != s(bool,X172) | s(bool,'F_0') = s(bool,X170) | s(bool,X171) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl32)),
  inference(superposition,[],[f2853,f5704])).
fof(f7811,plain,(
  ( ! [X165,X163,X166,X164] : (s(bool,'F_0') = s(bool,X164) | s(bool,'F_0') = s(bool,X165) | s(bool,X165) = s(bool,X166) | s(bool,X163) != s(bool,X166) | s(bool,'F_0') = s(bool,X163)) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f7810,f4678])).
fof(f7810,plain,(
  ( ! [X165,X163,X166,X164] : (s(bool,'F_0') = s(bool,X164) | ~p(s(bool,X164)) | s(bool,'F_0') = s(bool,X165) | s(bool,X165) = s(bool,X166) | s(bool,X163) != s(bool,X166) | s(bool,'F_0') = s(bool,X163)) ) | ($spl6 | $spl16)),
  inference(forward_demodulation,[],[f7243,f113])).
fof(f7243,plain,(
  ( ! [X165,X163,X166,X164] : (~p(s(bool,X164)) | s(bool,'F_0') = s(bool,X165) | s(bool,X165) = s(bool,X166) | s(bool,X163) != s(bool,X166) | s(bool,'F_0') = s(bool,X163) | s(bool,X164) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16)),
  inference(superposition,[],[f1996,f5704])).
fof(f7808,plain,(
  ( ! [X158,X156,X159,X157] : (s(bool,'F_0') = s(bool,X157) | s(bool,'T_0') = s(bool,X158) | s(bool,X158) != s(bool,X159) | s(bool,X156) != s(bool,X159) | s(bool,'F_0') = s(bool,X156)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f7807,f4678])).
fof(f7807,plain,(
  ( ! [X158,X156,X159,X157] : (s(bool,'F_0') = s(bool,X157) | ~p(s(bool,X157)) | s(bool,'T_0') = s(bool,X158) | s(bool,X158) != s(bool,X159) | s(bool,X156) != s(bool,X159) | s(bool,'F_0') = s(bool,X156)) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f7241,f113])).
fof(f7241,plain,(
  ( ! [X158,X156,X159,X157] : (~p(s(bool,X157)) | s(bool,'T_0') = s(bool,X158) | s(bool,X158) != s(bool,X159) | s(bool,X156) != s(bool,X159) | s(bool,'F_0') = s(bool,X156) | s(bool,X157) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f1148,f5704])).
fof(f7800,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,'F_0') = s(bool,X114) | s(bool,'T_0') != s(bool,X115) | s(bool,X113) = s(bool,X116) | s(bool,X115) != s(bool,X116) | s(bool,'F_0') = s(bool,X113)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f7799,f4678])).
fof(f7799,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,'F_0') = s(bool,X114) | ~p(s(bool,X114)) | s(bool,'T_0') != s(bool,X115) | s(bool,X113) = s(bool,X116) | s(bool,X115) != s(bool,X116) | s(bool,'F_0') = s(bool,X113)) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f7227,f113])).
fof(f7227,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,X114)) | s(bool,'T_0') != s(bool,X115) | s(bool,X113) = s(bool,X116) | s(bool,X115) != s(bool,X116) | s(bool,'F_0') = s(bool,X113) | s(bool,X114) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f696,f5704])).
fof(f7795,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'F_0') = s(bool,X92) | s(bool,'T_0') = s(bool,X93) | s(bool,X91) = s(bool,X94) | s(bool,X93) = s(bool,X94) | s(bool,'F_0') = s(bool,X91)) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f7794,f4678])).
fof(f7794,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'F_0') = s(bool,X92) | ~p(s(bool,X92)) | s(bool,'T_0') = s(bool,X93) | s(bool,X91) = s(bool,X94) | s(bool,X93) = s(bool,X94) | s(bool,'F_0') = s(bool,X91)) ) | ($spl6 | $spl24)),
  inference(forward_demodulation,[],[f7220,f113])).
fof(f7220,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X92)) | s(bool,'T_0') = s(bool,X93) | s(bool,X91) = s(bool,X94) | s(bool,X93) = s(bool,X94) | s(bool,'F_0') = s(bool,X91) | s(bool,X92) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f666,f5704])).
fof(f7793,plain,(
  ( ! [X85,X83,X86,X84] : (s(bool,'F_0') = s(bool,X84) | s(bool,X83) = s(bool,X85) | s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X86) | s(bool,'F_0') = s(bool,X83)) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f7792,f4678])).
fof(f7792,plain,(
  ( ! [X85,X83,X86,X84] : (s(bool,'F_0') = s(bool,X84) | ~p(s(bool,X84)) | s(bool,X83) = s(bool,X85) | s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X86) | s(bool,'F_0') = s(bool,X83)) ) | ($spl6 | $spl16)),
  inference(forward_demodulation,[],[f7218,f113])).
fof(f7218,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | s(bool,X83) = s(bool,X85) | s(bool,X85) != s(bool,X86) | s(bool,'F_0') = s(bool,X86) | s(bool,'F_0') = s(bool,X83) | s(bool,X84) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl16)),
  inference(superposition,[],[f642,f5704])).
fof(f7632,plain,(
  ( ! [X180,X178,X179] : (s(bool,X178) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X179) | s(bool,X180) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X178) = s(bool,X180)) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f7006,f68])).
fof(f7006,plain,(
  ( ! [X180,X178,X179] : (s(bool,'F_0') != s(bool,X179) | s(bool,X178) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X179) | s(bool,X180) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X178) = s(bool,X180)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f3149,f5704])).
fof(f7461,plain,(
  ( ! [X175,X173,X174] : (s(bool,'T_0') != s(bool,X173) | s(bool,'F_0') = s(bool,X174) | s(bool,X175) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X173) = s(bool,X175)) ) | ($spl6 | $spl16)),
  inference(duplicate_literal_removal,[],[f6929])).
fof(f6929,plain,(
  ( ! [X175,X173,X174] : (s(bool,'T_0') != s(bool,X173) | s(bool,'F_0') = s(bool,X174) | s(bool,'F_0') = s(bool,X174) | s(bool,X175) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X173) = s(bool,X175)) ) | ($spl6 | $spl16)),
  inference(superposition,[],[f781,f5704])).
fof(f7582,plain,(
  ( ! [X114,X115,X113] : (s(bool,'T_0') != s(bool,X113) | s(bool,'F_0') != s(bool,X114) | s(bool,X115) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X113) = s(bool,X115)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f6911,f705])).
fof(f6911,plain,(
  ( ! [X114,X115,X113] : (s(bool,'T_0') != s(bool,X113) | s(bool,'T_0') = s(bool,X114) | s(bool,'F_0') != s(bool,X114) | s(bool,X115) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X113) = s(bool,X115)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f586,f5704])).
fof(f7572,plain,(
  ( ! [X80,X81,X79] : (s(bool,'T_0') != s(bool,X79) | s(bool,'F_0') = s(bool,X80) | s(bool,X81) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X79) = s(bool,X81)) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f6900,f4678])).
fof(f6900,plain,(
  ( ! [X80,X81,X79] : (s(bool,'T_0') != s(bool,X79) | s(bool,'F_0') = s(bool,X80) | ~p(s(bool,X80)) | s(bool,X81) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X79) = s(bool,X81)) ) | ($spl6 | $spl16)),
  inference(superposition,[],[f273,f5704])).
fof(f7466,plain,(
  ( ! [X50,X51,X49] : (s(bool,'F_0') = s(bool,X50) | s(bool,X49) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X51) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X49) = s(bool,X51)) ) | $spl6),
  inference(duplicate_literal_removal,[],[f6857])).
fof(f6857,plain,(
  ( ! [X50,X51,X49] : (s(bool,'F_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X50) | s(bool,X49) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X51) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X49) = s(bool,X51)) ) | $spl6),
  inference(superposition,[],[f5704,f5704])).
fof(f7470,plain,(
  ( ! [X54,X52,X53] : (s(bool,X53) = s(bool,X54) | s(bool,'F_0') = s(bool,X52) | s(bool,X54) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X53) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f6825])).
fof(f6825,plain,(
  ( ! [X54,X52,X53] : (s(bool,X53) = s(bool,X54) | s(bool,'F_0') = s(bool,X52) | s(bool,X54) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X52) | s(bool,X53) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(superposition,[],[f5704,f5704])).
fof(f7482,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') = s(bool,X49) | s(bool,X50) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X48) = s(bool,X50) | s(bool,X48) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f6754])).
fof(f6754,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') = s(bool,X49) | s(bool,X50) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X48) = s(bool,X50) | s(bool,'F_0') = s(bool,X49) | s(bool,X48) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | $spl6),
  inference(superposition,[],[f5704,f5704])).
fof(f6705,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,X34) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X36),s('type.hreal.hreal',X37))) | s('type.hreal.hreal',X37) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X36),s('type.hreal.hreal',sK3(X37,X36)))) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f6682,f133])).
fof(f6682,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,X34) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X36),s('type.hreal.hreal',X37))) | s('type.hreal.hreal',X37) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X36),s('type.hreal.hreal',sK3(X37,X36)))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl24),
  inference(superposition,[],[f5061,f672])).
fof(f6703,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,X30) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X32),s('type.hreal.hreal',X33))) | s('type.hreal.hreal',X33) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X32),s('type.hreal.hreal',sK3(X33,X32)))) | s(bool,'T_0') = s(bool,X30)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f6681,f133])).
fof(f6681,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,X30) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X32),s('type.hreal.hreal',X33))) | s('type.hreal.hreal',X33) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X32),s('type.hreal.hreal',sK3(X33,X32)))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl24),
  inference(superposition,[],[f5061,f672])).
fof(f6648,plain,(
  ( ! [X30,X32] : (s(bool,X30) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X32) | s(bool,X30) = s(bool,X32)) ) | $spl72),
  inference(cnf_transformation,[],[f6648_D])).
fof(f6648_D,plain,(
  ( ! [X30,X32] : (s(bool,X30) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,X32) | s(bool,X30) = s(bool,X32)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f6649,plain,(
  ( ! [X35,X33,X34] : (s(bool,X33) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X34) | s(bool,X33) = s(bool,X35) | s(bool,'F_0') = s(bool,X35)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f6609,f68])).
fof(f6609,plain,(
  ( ! [X35,X33,X34] : (s(bool,'F_0') != s(bool,X34) | s(bool,X33) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X34) | s(bool,X33) = s(bool,X35) | s(bool,'F_0') = s(bool,X35)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3149,f782])).
fof(f6647,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X31) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f6608,f68])).
fof(f6608,plain,(
  ( ! [X30,X31,X32] : (s(bool,'F_0') != s(bool,X31) | s(bool,X30) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,X31) | s(bool,X30) = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f3149,f782])).
fof(f6188,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'F_0') = s(bool,X55) | s(bool,X55) = s(bool,X56) | s(bool,X53) != s(bool,X56) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f6170,f436])).
fof(f6170,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'F_0') = s(bool,X55) | s(bool,X55) = s(bool,X56) | s(bool,X53) != s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | $spl16),
  inference(superposition,[],[f1996,f782])).
fof(f6187,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'F_0') = s(bool,X51) | s(bool,X51) = s(bool,X52) | s(bool,X49) != s(bool,X52) | s(bool,'F_0') = s(bool,X49)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f6169,f436])).
fof(f6169,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'F_0') = s(bool,X51) | s(bool,X51) = s(bool,X52) | s(bool,X49) != s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | $spl16),
  inference(superposition,[],[f1996,f782])).
fof(f6071,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,X36) = s(bool,X37) | s(bool,X34) != s(bool,X37) | s(bool,'F_0') != s(bool,X36) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f6051,f133])).
fof(f6051,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,X36) = s(bool,X37) | s(bool,X34) != s(bool,X37) | s(bool,'F_0') != s(bool,X36) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | ($spl24 | $spl32)),
  inference(superposition,[],[f2861,f672])).
fof(f6070,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,X32) = s(bool,X33) | s(bool,X30) != s(bool,X33) | s(bool,'F_0') != s(bool,X32) | s(bool,'T_0') = s(bool,X30)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f6050,f133])).
fof(f6050,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,X32) = s(bool,X33) | s(bool,X30) != s(bool,X33) | s(bool,'F_0') != s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | ($spl24 | $spl32)),
  inference(superposition,[],[f2861,f672])).
fof(f6016,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X55) = s(bool,X56) | s(bool,X53) = s(bool,X56) | s(bool,'F_0') != s(bool,X55) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f5998,f436])).
fof(f5998,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X55) = s(bool,X56) | s(bool,X53) = s(bool,X56) | s(bool,'F_0') != s(bool,X55) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl32)),
  inference(superposition,[],[f2853,f782])).
fof(f6015,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X51) = s(bool,X52) | s(bool,X49) = s(bool,X52) | s(bool,'F_0') != s(bool,X51) | s(bool,'F_0') = s(bool,X49)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f5997,f436])).
fof(f5997,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X51) = s(bool,X52) | s(bool,X49) = s(bool,X52) | s(bool,'F_0') != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl16 | $spl32)),
  inference(superposition,[],[f2853,f782])).
fof(f5854,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'F_0') = s(bool,X36) | s(bool,X34) != s(bool,X37) | s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f5834,f133])).
fof(f5834,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'F_0') = s(bool,X36) | s(bool,X34) != s(bool,X37) | s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl24),
  inference(superposition,[],[f755,f672])).
fof(f5853,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'F_0') = s(bool,X32) | s(bool,X30) != s(bool,X33) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,X30)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f5833,f133])).
fof(f5833,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'F_0') = s(bool,X32) | s(bool,X30) != s(bool,X33) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl24),
  inference(superposition,[],[f755,f672])).
fof(f5800,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X53) = s(bool,X55) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f5782,f436])).
fof(f5782,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X53) = s(bool,X55) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | $spl16),
  inference(superposition,[],[f642,f782])).
fof(f5799,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X49) = s(bool,X51) | s(bool,X51) != s(bool,X52) | s(bool,'F_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X49)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f5781,f436])).
fof(f5781,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X49) = s(bool,X51) | s(bool,X51) != s(bool,X52) | s(bool,'F_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | $spl16),
  inference(superposition,[],[f642,f782])).
fof(f5694,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'F_0') = s(bool,X36) | s(bool,X34) = s(bool,X37) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f5674,f133])).
fof(f5674,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'F_0') = s(bool,X36) | s(bool,X34) = s(bool,X37) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl24),
  inference(superposition,[],[f4997,f672])).
fof(f5693,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'F_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X30)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f5673,f133])).
fof(f5673,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'F_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl24),
  inference(superposition,[],[f4997,f672])).
fof(f5176,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,X44)) | s('type.hreal.hreal',X43) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X42),s('type.hreal.hreal',sK3(X43,X42)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X42),s('type.hreal.hreal',X43)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f5116,f436])).
fof(f5116,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,X44)) | s('type.hreal.hreal',X43) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X42),s('type.hreal.hreal',sK3(X43,X42)))) | s(bool,'F_0') = s(bool,X44) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X42),s('type.hreal.hreal',X43)))) ) | $spl16),
  inference(superposition,[],[f60,f782])).
fof(f5175,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,X41)) | s('type.hreal.hreal',X40) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X39),s('type.hreal.hreal',sK3(X40,X39)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X39),s('type.hreal.hreal',X40)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f5115,f436])).
fof(f5115,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,X41)) | s('type.hreal.hreal',X40) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X39),s('type.hreal.hreal',sK3(X40,X39)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X39),s('type.hreal.hreal',X40))) | s(bool,'F_0') = s(bool,X41)) ) | $spl16),
  inference(superposition,[],[f60,f782])).
fof(f131,plain,(
  ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) | $spl15),
  inference(cnf_transformation,[],[f131_D])).
fof(f131_D,plain,(
  ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f4763,plain,(
  ( ! [X35] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl6 | $spl24)),
  inference(subsumption_resolution,[],[f4745,f68])).
fof(f4745,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,X35) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f436,f695])).
fof(f4677,plain,(
  ~p(s(bool,'F_0')) | $spl71),
  inference(cnf_transformation,[],[f4677_D])).
fof(f4677_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f70,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X11)))),
  file('thm.realax.HREAL_LT_LADD','thm.bool.IMP_CLAUSES|split|4')).
fof(f4675,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(subsumption_resolution,[],[f80,f79])).
fof(f80,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f53])).
fof(f4640,plain,(
  $false | ($spl6 | $spl16 | $spl24 | $spl32 | $spl68)),
  inference(subsumption_resolution,[],[f4639,f4391])).
fof(f4639,plain,(
  ( ! [X60] : (s(bool,'T_0') != s(bool,X60)) ) | ($spl6 | $spl16 | $spl24 | $spl32 | $spl68)),
  inference(subsumption_resolution,[],[f4418,f727])).
fof(f4418,plain,(
  ( ! [X60] : (s(bool,'T_0') != s(bool,X60) | s(bool,X60) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) ) | ($spl6 | $spl32 | $spl68)),
  inference(backward_demodulation,[],[f4391,f2900])).
fof(f4616,plain,(
  $false | $spl68),
  inference(subsumption_resolution,[],[f4396,f86])).
fof(f4396,plain,(
  ~p(s(bool,'T_0')) | $spl68),
  inference(backward_demodulation,[],[f4391,f84])).
fof(f4587,plain,(
  $false | ($spl3 | $spl68)),
  inference(trivial_inequality_removal,[],[f4575])).
fof(f4575,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl3 | $spl68)),
  inference(superposition,[],[f130,f4391])).
fof(f4585,plain,(
  $false | ($spl3 | $spl68)),
  inference(trivial_inequality_removal,[],[f4584])).
fof(f4584,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl3 | $spl68)),
  inference(superposition,[],[f110,f4391])).
fof(f4390,plain,(
  ( ! [X47,X45,X44] : (s(bool,X47) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X44),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X44),s('type.hreal.hreal',X45))))) | s(bool,'F_0') = s(bool,X47)) ) | $spl66),
  inference(cnf_transformation,[],[f4390_D])).
fof(f4390_D,plain,(
  ( ! [X47,X45,X44] : (s(bool,X47) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X44),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X44),s('type.hreal.hreal',X45))))) | s(bool,'F_0') = s(bool,X47)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4392,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'T_0') = s(bool,X50) | s(bool,X51) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X48),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X48),s('type.hreal.hreal',X49))))) | s(bool,'F_0') = s(bool,X51)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f4363,f68])).
fof(f4363,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,'F_0') != s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,X51) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X48),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X48),s('type.hreal.hreal',X49))))) | s(bool,'F_0') = s(bool,X51)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f4271,f782])).
fof(f4389,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'T_0') = s(bool,X46) | s(bool,X47) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X44),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X44),s('type.hreal.hreal',X45))))) | s(bool,'F_0') = s(bool,X47)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f4362,f68])).
fof(f4362,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,'F_0') != s(bool,X46) | s(bool,'T_0') = s(bool,X46) | s(bool,X47) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X44),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X44),s('type.hreal.hreal',X45))))) | s(bool,'F_0') = s(bool,X47)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f4271,f782])).
fof(f4044,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'F_0') = s(bool,X55) | s(bool,X55) = s(bool,X56) | s(bool,X53) != s(bool,X56) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f4027,f436])).
fof(f4027,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'F_0') = s(bool,X55) | s(bool,X55) = s(bool,X56) | s(bool,X53) != s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | $spl16),
  inference(superposition,[],[f1996,f782])).
fof(f4043,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'F_0') = s(bool,X51) | s(bool,X51) = s(bool,X52) | s(bool,X49) != s(bool,X52) | s(bool,'F_0') = s(bool,X49)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f4026,f436])).
fof(f4026,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'F_0') = s(bool,X51) | s(bool,X51) = s(bool,X52) | s(bool,X49) != s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | $spl16),
  inference(superposition,[],[f1996,f782])).
fof(f3879,plain,(
  ( ! [X72,X73] : (s(bool,'F_0') != s(bool,X73) | s(bool,'F_0') != s(bool,X72) | s(bool,X72) = s(bool,X73)) ) | $spl62),
  inference(cnf_transformation,[],[f3879_D])).
fof(f3879_D,plain,(
  ( ! [X72,X73] : (s(bool,'F_0') != s(bool,X73) | s(bool,'F_0') != s(bool,X72) | s(bool,X72) = s(bool,X73)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f3894,plain,(
  ( ! [X80,X81,X79] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X79)) | s(bool,X80) = s(bool,X81) | s(bool,'F_0') != s(bool,X80)) ) | ($spl6 | $spl16 | $spl24 | $spl32)),
  inference(forward_demodulation,[],[f3893,f113])).
fof(f3893,plain,(
  ( ! [X80,X81,X79] : (p(s(bool,X79)) | s(bool,X80) = s(bool,X81) | s(bool,X81) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') != s(bool,X80)) ) | ($spl6 | $spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f3872,f133])).
fof(f3872,plain,(
  ( ! [X80,X81,X79] : (p(s(bool,X79)) | s(bool,X80) = s(bool,X81) | s(bool,X81) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') != s(bool,X80) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl24 | $spl32)),
  inference(superposition,[],[f2861,f695])).
fof(f3892,plain,(
  ( ! [X78,X76,X77] : (p(s(bool,X76)) | s(bool,X77) = s(bool,X78) | s(bool,'F_0') != s(bool,X78) | s(bool,'F_0') != s(bool,X77)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f3871,f133])).
fof(f3871,plain,(
  ( ! [X78,X76,X77] : (p(s(bool,X76)) | s(bool,X77) = s(bool,X78) | s(bool,'F_0') != s(bool,X78) | s(bool,'F_0') != s(bool,X77) | s(bool,'T_0') = s(bool,X76)) ) | ($spl24 | $spl32)),
  inference(superposition,[],[f2861,f68])).
fof(f3891,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X72)) | s(bool,X73) = s(bool,X74) | s(bool,'F_0') != s(bool,X74) | s(bool,'F_0') != s(bool,X73) | s(bool,'F_0') = s(bool,X75)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f3870,f561])).
fof(f3870,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X72)) | s(bool,X73) = s(bool,X74) | s(bool,'F_0') != s(bool,X74) | s(bool,'F_0') != s(bool,X73) | s(bool,X72) = s(bool,X75) | s(bool,'F_0') = s(bool,X75)) ) | ($spl16 | $spl24 | $spl32)),
  inference(superposition,[],[f2861,f782])).
fof(f3890,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X68)) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') != s(bool,X70) | s(bool,'F_0') != s(bool,X69) | s(bool,'F_0') = s(bool,X71)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f3869,f561])).
fof(f3869,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X68)) | s(bool,X69) = s(bool,X70) | s(bool,'F_0') != s(bool,X70) | s(bool,'F_0') != s(bool,X69) | s(bool,X68) = s(bool,X71) | s(bool,'F_0') = s(bool,X71)) ) | ($spl16 | $spl24 | $spl32)),
  inference(superposition,[],[f2861,f782])).
fof(f3883,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,X36) = s(bool,X37) | s(bool,X34) != s(bool,X37) | s(bool,'F_0') != s(bool,X36) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f3860,f133])).
fof(f3860,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,X36) = s(bool,X37) | s(bool,X34) != s(bool,X37) | s(bool,'F_0') != s(bool,X36) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | ($spl24 | $spl32)),
  inference(superposition,[],[f2861,f672])).
fof(f3881,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,X32) = s(bool,X33) | s(bool,X30) != s(bool,X33) | s(bool,'F_0') != s(bool,X32) | s(bool,'T_0') = s(bool,X30)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f3859,f133])).
fof(f3859,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,X32) = s(bool,X33) | s(bool,X30) != s(bool,X33) | s(bool,'F_0') != s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | ($spl24 | $spl32)),
  inference(superposition,[],[f2861,f672])).
fof(f3849,plain,(
  ( ! [X76,X77,X75] : (s(bool,X75) = s(bool,X76) | s(bool,'F_0') != s(bool,X76) | s(bool,'F_0') != s(bool,X75) | ~p(s(bool,X77))) ) | ($spl24 | $spl32)),
  inference(resolution,[],[f2861,f70])).
fof(f3878,plain,(
  ( ! [X74,X72,X73] : (s(bool,X72) = s(bool,X73) | s(bool,'F_0') != s(bool,X73) | s(bool,'F_0') != s(bool,X72) | p(s(bool,X74))) ) | ($spl24 | $spl32)),
  inference(subsumption_resolution,[],[f3848,f85])).
fof(f3848,plain,(
  ( ! [X74,X72,X73] : (s(bool,X72) = s(bool,X73) | s(bool,'F_0') != s(bool,X73) | s(bool,'F_0') != s(bool,X72) | p(s(bool,X74)) | ~p(s(bool,'T_0'))) ) | ($spl24 | $spl32)),
  inference(resolution,[],[f2861,f80])).
fof(f3789,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X55) = s(bool,X56) | s(bool,X53) = s(bool,X56) | s(bool,'F_0') != s(bool,X55) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f3771,f436])).
fof(f3771,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X55) = s(bool,X56) | s(bool,X53) = s(bool,X56) | s(bool,'F_0') != s(bool,X55) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl32)),
  inference(superposition,[],[f2853,f782])).
fof(f3787,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X51) = s(bool,X52) | s(bool,X49) = s(bool,X52) | s(bool,'F_0') != s(bool,X51) | s(bool,'F_0') = s(bool,X49)) ) | ($spl16 | $spl24 | $spl32)),
  inference(subsumption_resolution,[],[f3770,f436])).
fof(f3770,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X51) = s(bool,X52) | s(bool,X49) = s(bool,X52) | s(bool,'F_0') != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl16 | $spl32)),
  inference(superposition,[],[f2853,f782])).
fof(f3686,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'F_0') = s(bool,X36) | s(bool,X34) = s(bool,X37) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3666,f133])).
fof(f3666,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'F_0') = s(bool,X36) | s(bool,X34) = s(bool,X37) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f775,f672])).
fof(f775,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1)) ) | $spl16),
  inference(resolution,[],[f317,f62])).
fof(f3685,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'F_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X30)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3665,f133])).
fof(f3665,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'F_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f775,f672])).
fof(f3632,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'T_0') = s(bool,X55) | s(bool,X55) != s(bool,X56) | s(bool,X53) != s(bool,X56) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3615,f436])).
fof(f3615,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'T_0') = s(bool,X55) | s(bool,X55) != s(bool,X56) | s(bool,X53) != s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f1148,f782])).
fof(f3631,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'T_0') = s(bool,X51) | s(bool,X51) != s(bool,X52) | s(bool,X49) != s(bool,X52) | s(bool,'F_0') = s(bool,X49)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3614,f436])).
fof(f3614,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'T_0') = s(bool,X51) | s(bool,X51) != s(bool,X52) | s(bool,X49) != s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f1148,f782])).
fof(f3561,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'F_0') = s(bool,X36) | s(bool,X34) != s(bool,X37) | s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3540,f133])).
fof(f3540,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'F_0') = s(bool,X36) | s(bool,X34) != s(bool,X37) | s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl24),
  inference(superposition,[],[f755,f672])).
fof(f3559,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'F_0') = s(bool,X32) | s(bool,X30) != s(bool,X33) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,X30)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3539,f133])).
fof(f3539,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'F_0') = s(bool,X32) | s(bool,X30) != s(bool,X33) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl24),
  inference(superposition,[],[f755,f672])).
fof(f3439,plain,(
  ( ! [X62,X63] : (s(bool,'F_0') != s(bool,X63) | s(bool,X62) != s(bool,X63) | s(bool,'T_0') != s(bool,X62)) ) | $spl54),
  inference(cnf_transformation,[],[f3439_D])).
fof(f3439_D,plain,(
  ( ! [X62,X63] : (s(bool,'F_0') != s(bool,X63) | s(bool,X62) != s(bool,X63) | s(bool,'T_0') != s(bool,X62)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f3454,plain,(
  ( ! [X80,X81,X79] : (s(bool,'F_0') != s(bool,X81) | p(s(bool,X79)) | s(bool,'T_0') != s(bool,X80) | s(bool,X80) != s(bool,X81)) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f3453,f113])).
fof(f3453,plain,(
  ( ! [X80,X81,X79] : (p(s(bool,X79)) | s(bool,'T_0') != s(bool,X80) | s(bool,X81) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X80) != s(bool,X81)) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3432,f133])).
fof(f3432,plain,(
  ( ! [X80,X81,X79] : (p(s(bool,X79)) | s(bool,'T_0') != s(bool,X80) | s(bool,X81) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,X80) != s(bool,X81) | s(bool,'T_0') = s(bool,X79)) ) | ($spl6 | $spl16 | $spl24)),
  inference(superposition,[],[f703,f695])).
fof(f3452,plain,(
  ( ! [X78,X76,X77] : (p(s(bool,X76)) | s(bool,'T_0') != s(bool,X77) | s(bool,'F_0') != s(bool,X78) | s(bool,X77) != s(bool,X78)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3431,f133])).
fof(f3431,plain,(
  ( ! [X78,X76,X77] : (p(s(bool,X76)) | s(bool,'T_0') != s(bool,X77) | s(bool,'F_0') != s(bool,X78) | s(bool,X77) != s(bool,X78) | s(bool,'T_0') = s(bool,X76)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f703,f68])).
fof(f3451,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X72)) | s(bool,'T_0') != s(bool,X73) | s(bool,'F_0') != s(bool,X74) | s(bool,X73) != s(bool,X74) | s(bool,'F_0') = s(bool,X75)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3430,f561])).
fof(f3430,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X72)) | s(bool,'T_0') != s(bool,X73) | s(bool,'F_0') != s(bool,X74) | s(bool,X73) != s(bool,X74) | s(bool,X72) = s(bool,X75) | s(bool,'F_0') = s(bool,X75)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f703,f782])).
fof(f3450,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X68)) | s(bool,'T_0') != s(bool,X69) | s(bool,'F_0') != s(bool,X70) | s(bool,X69) != s(bool,X70) | s(bool,'F_0') = s(bool,X71)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3429,f561])).
fof(f3429,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X68)) | s(bool,'T_0') != s(bool,X69) | s(bool,'F_0') != s(bool,X70) | s(bool,X69) != s(bool,X70) | s(bool,X68) = s(bool,X71) | s(bool,'F_0') = s(bool,X71)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f703,f782])).
fof(f3442,plain,(
  ( ! [X30,X33,X32] : (s(bool,'T_0') != s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,X32) != s(bool,X33) | s(bool,X30) != s(bool,X33)) ) | $spl56),
  inference(cnf_transformation,[],[f3442_D])).
fof(f3442_D,plain,(
  ( ! [X30,X33,X32] : (s(bool,'T_0') != s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,X32) != s(bool,X33) | s(bool,X30) != s(bool,X33)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f3443,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'T_0') != s(bool,X36) | s(bool,X34) != s(bool,X37) | s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3420,f133])).
fof(f3420,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'T_0') != s(bool,X36) | s(bool,X34) != s(bool,X37) | s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f703,f672])).
fof(f3441,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'T_0') != s(bool,X32) | s(bool,X30) != s(bool,X33) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,X30)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3419,f133])).
fof(f3419,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'T_0') != s(bool,X32) | s(bool,X30) != s(bool,X33) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f703,f672])).
fof(f3409,plain,(
  ( ! [X66,X67,X65] : (s(bool,'T_0') != s(bool,X65) | s(bool,'F_0') != s(bool,X66) | s(bool,X65) != s(bool,X66) | ~p(s(bool,X67))) ) | ($spl16 | $spl24)),
  inference(resolution,[],[f703,f70])).
fof(f3438,plain,(
  ( ! [X64,X62,X63] : (s(bool,'T_0') != s(bool,X62) | s(bool,'F_0') != s(bool,X63) | s(bool,X62) != s(bool,X63) | p(s(bool,X64))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3408,f85])).
fof(f3408,plain,(
  ( ! [X64,X62,X63] : (s(bool,'T_0') != s(bool,X62) | s(bool,'F_0') != s(bool,X63) | s(bool,X62) != s(bool,X63) | p(s(bool,X64)) | ~p(s(bool,'T_0'))) ) | ($spl16 | $spl24)),
  inference(resolution,[],[f703,f80])).
fof(f3386,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'T_0') != s(bool,X55) | s(bool,X53) = s(bool,X56) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3368,f436])).
fof(f3368,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'T_0') != s(bool,X55) | s(bool,X53) = s(bool,X56) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f696,f782])).
fof(f3384,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'T_0') != s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,X51) != s(bool,X52) | s(bool,'F_0') = s(bool,X49)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3367,f436])).
fof(f3367,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'T_0') != s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,X51) != s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f696,f782])).
fof(f3138,plain,(
  ( ! [X30,X33,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,X32) != s(bool,X33) | s(bool,X30) = s(bool,X33)) ) | $spl46),
  inference(cnf_transformation,[],[f3138_D])).
fof(f3138_D,plain,(
  ( ! [X30,X33,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30) | s(bool,X32) != s(bool,X33) | s(bool,X30) = s(bool,X33)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f3319,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'T_0') = s(bool,X36) | s(bool,X34) != s(bool,X37) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3299,f133])).
fof(f3299,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'T_0') = s(bool,X36) | s(bool,X34) != s(bool,X37) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl24),
  inference(superposition,[],[f673,f672])).
fof(f3318,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'T_0') = s(bool,X32) | s(bool,X30) != s(bool,X33) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X30)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3298,f133])).
fof(f3298,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'T_0') = s(bool,X32) | s(bool,X30) != s(bool,X33) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl24),
  inference(superposition,[],[f673,f672])).
fof(f3266,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'T_0') = s(bool,X55) | s(bool,X53) = s(bool,X56) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3248,f436])).
fof(f3248,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,'T_0') = s(bool,X55) | s(bool,X53) = s(bool,X56) | s(bool,X55) = s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f666,f782])).
fof(f3264,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'T_0') = s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X49)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3247,f436])).
fof(f3247,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,'T_0') = s(bool,X51) | s(bool,X49) = s(bool,X52) | s(bool,X51) = s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f666,f782])).
fof(f3205,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X53) = s(bool,X55) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3187,f436])).
fof(f3187,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X53) = s(bool,X55) | s(bool,X55) != s(bool,X56) | s(bool,'F_0') = s(bool,X56) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | $spl16),
  inference(superposition,[],[f642,f782])).
fof(f3203,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X49) = s(bool,X51) | s(bool,X51) != s(bool,X52) | s(bool,'F_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X49)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3186,f436])).
fof(f3186,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X49) = s(bool,X51) | s(bool,X51) != s(bool,X52) | s(bool,'F_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | $spl16),
  inference(superposition,[],[f642,f782])).
fof(f3139,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'T_0') = s(bool,X36) | s(bool,X34) = s(bool,X37) | s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,X34)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3118,f133])).
fof(f3118,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X35)) | s(bool,'T_0') = s(bool,X36) | s(bool,X34) = s(bool,X37) | s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl24),
  inference(superposition,[],[f580,f672])).
fof(f3137,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'T_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,X30)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3117,f133])).
fof(f3117,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'T_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl24),
  inference(superposition,[],[f580,f672])).
fof(f3083,plain,(
  ( ! [X52,X51,X49] : (s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | p(s(bool,X52)) | s(bool,X51) != s(bool,X52)) ) | $spl44),
  inference(cnf_transformation,[],[f3083_D])).
fof(f3083_D,plain,(
  ( ! [X52,X51,X49] : (s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | p(s(bool,X52)) | s(bool,X51) != s(bool,X52)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f3084,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X55) != s(bool,X56) | p(s(bool,X56)) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3062,f436])).
fof(f3062,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X55) != s(bool,X56) | p(s(bool,X56)) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f907,f782])).
fof(f3082,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X49) != s(bool,X51) | s(bool,X51) != s(bool,X52) | p(s(bool,X52)) | s(bool,'F_0') = s(bool,X49)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f3061,f436])).
fof(f3061,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X49) != s(bool,X51) | s(bool,X51) != s(bool,X52) | p(s(bool,X52)) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f907,f782])).
fof(f3019,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X55) = s(bool,X56) | ~p(s(bool,X56)) | s(bool,X53) != s(bool,X55) | s(bool,'F_0') = s(bool,X53)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f2997,f436])).
fof(f2997,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | s(bool,X55) = s(bool,X56) | ~p(s(bool,X56)) | s(bool,X53) != s(bool,X55) | s(bool,'F_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X53)) ) | $spl16),
  inference(superposition,[],[f662,f782])).
fof(f3017,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X51) = s(bool,X52) | ~p(s(bool,X52)) | s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X49)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f2996,f436])).
fof(f2996,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | s(bool,X51) = s(bool,X52) | ~p(s(bool,X52)) | s(bool,X49) != s(bool,X51) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50)) ) | $spl16),
  inference(superposition,[],[f662,f782])).
fof(f2847,plain,(
  ( ! [X57,X60,X58] : (s(bool,X57) != s(bool,X58) | s(bool,'T_0') = s(bool,X60) | s(bool,X57) = s(bool,X60) | p(s(bool,X58))) ) | $spl40),
  inference(cnf_transformation,[],[f2847_D])).
fof(f2847_D,plain,(
  ( ! [X57,X60,X58] : (s(bool,X57) != s(bool,X58) | s(bool,'T_0') = s(bool,X60) | s(bool,X57) = s(bool,X60) | p(s(bool,X58))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2848,plain,(
  ( ! [X61,X64,X62,X63] : (s(bool,X61) != s(bool,X62) | p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f2806,f133])).
fof(f2806,plain,(
  ( ! [X61,X64,X62,X63] : (s(bool,X61) != s(bool,X62) | p(s(bool,X62)) | s(bool,'T_0') = s(bool,X63) | p(s(bool,X63)) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) ) | $spl24),
  inference(superposition,[],[f532,f672])).
fof(f2846,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,X57) != s(bool,X58) | p(s(bool,X58)) | p(s(bool,X59)) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f2805,f133])).
fof(f2805,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,X57) != s(bool,X58) | p(s(bool,X58)) | s(bool,'T_0') = s(bool,X59) | p(s(bool,X59)) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) ) | $spl24),
  inference(superposition,[],[f532,f672])).
fof(f2844,plain,(
  ( ! [X43,X41,X44] : (s(bool,X41) = s(bool,X43) | s(bool,'F_0') = s(bool,X44) | s(bool,X41) = s(bool,X44) | p(s(bool,X43))) ) | $spl38),
  inference(cnf_transformation,[],[f2844_D])).
fof(f2844_D,plain,(
  ( ! [X43,X41,X44] : (s(bool,X41) = s(bool,X43) | s(bool,'F_0') = s(bool,X44) | s(bool,X41) = s(bool,X44) | p(s(bool,X43))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2845,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,X46)) | s(bool,X45) = s(bool,X47) | p(s(bool,X47)) | s(bool,X45) = s(bool,X48) | s(bool,'F_0') = s(bool,X48)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f2802,f289])).
fof(f2802,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'F_0') != s(bool,X46) | p(s(bool,X46)) | s(bool,X45) = s(bool,X47) | p(s(bool,X47)) | s(bool,X45) = s(bool,X48) | s(bool,'F_0') = s(bool,X48)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f532,f782])).
fof(f2843,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,X42)) | s(bool,X41) = s(bool,X43) | p(s(bool,X43)) | s(bool,X41) = s(bool,X44) | s(bool,'F_0') = s(bool,X44)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f2801,f289])).
fof(f2801,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,X42) | p(s(bool,X42)) | s(bool,X41) = s(bool,X43) | p(s(bool,X43)) | s(bool,X41) = s(bool,X44) | s(bool,'F_0') = s(bool,X44)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f532,f782])).
fof(f2829,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'F_0') != s(bool,X46) | p(s(bool,X45)) | s(bool,X46) = s(bool,X47) | p(s(bool,X47)) | s(bool,'F_0') = s(bool,X48)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f2778,f561])).
fof(f2778,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'F_0') != s(bool,X46) | p(s(bool,X45)) | s(bool,X46) = s(bool,X47) | p(s(bool,X47)) | s(bool,X45) = s(bool,X48) | s(bool,'F_0') = s(bool,X48)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f532,f782])).
fof(f2827,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,X42) | p(s(bool,X41)) | s(bool,X42) = s(bool,X43) | p(s(bool,X43)) | s(bool,'F_0') = s(bool,X44)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f2777,f561])).
fof(f2777,plain,(
  ( ! [X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,X42) | p(s(bool,X41)) | s(bool,X42) = s(bool,X43) | p(s(bool,X43)) | s(bool,X41) = s(bool,X44) | s(bool,'F_0') = s(bool,X44)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f532,f782])).
fof(f2826,plain,(
  ( ! [X39,X38,X40] : (s(bool,'F_0') != s(bool,X39) | p(s(bool,X38)) | s(bool,X39) = s(bool,X40) | p(s(bool,X40))) ) | ($spl6 | $spl16 | $spl24)),
  inference(forward_demodulation,[],[f2825,f113])).
fof(f2825,plain,(
  ( ! [X39,X38,X40] : (s(bool,X39) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | p(s(bool,X38)) | s(bool,X39) = s(bool,X40) | p(s(bool,X40))) ) | ($spl6 | $spl16 | $spl24)),
  inference(subsumption_resolution,[],[f2776,f133])).
fof(f2776,plain,(
  ( ! [X39,X38,X40] : (s(bool,X39) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | p(s(bool,X38)) | s(bool,X39) = s(bool,X40) | p(s(bool,X40)) | s(bool,'T_0') = s(bool,X38)) ) | ($spl6 | $spl24)),
  inference(superposition,[],[f532,f695])).
fof(f2823,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) != s(bool,X32) | s(bool,'T_0') = s(bool,X31) | p(s(bool,X33)) | s(bool,X32) = s(bool,X33)) ) | $spl34),
  inference(cnf_transformation,[],[f2823_D])).
fof(f2823_D,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) != s(bool,X32) | s(bool,'T_0') = s(bool,X31) | p(s(bool,X33)) | s(bool,X32) = s(bool,X33)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2824,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X35) != s(bool,X36) | p(s(bool,X34)) | s(bool,X36) = s(bool,X37) | p(s(bool,X37)) | s(bool,'T_0') = s(bool,X35)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f2775,f133])).
fof(f2775,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,X35) != s(bool,X36) | p(s(bool,X34)) | s(bool,X36) = s(bool,X37) | p(s(bool,X37)) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X34)) ) | $spl24),
  inference(superposition,[],[f532,f672])).
fof(f2822,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X31) != s(bool,X32) | p(s(bool,X30)) | s(bool,X32) = s(bool,X33) | p(s(bool,X33)) | s(bool,'T_0') = s(bool,X31)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f2774,f133])).
fof(f2774,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X31) != s(bool,X32) | p(s(bool,X30)) | s(bool,X32) = s(bool,X33) | p(s(bool,X33)) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X31)) ) | $spl24),
  inference(superposition,[],[f532,f672])).
fof(f2818,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | p(s(bool,X19)) | s(bool,X20) = s(bool,X21) | p(s(bool,X21))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f2771,f133])).
fof(f2771,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | p(s(bool,X19)) | s(bool,X20) = s(bool,X21) | p(s(bool,X21)) | s(bool,'T_0') = s(bool,X19)) ) | $spl24),
  inference(superposition,[],[f532,f68])).
fof(f2648,plain,(
  ( ! [X165,X166,X164] : (~p(s(bool,X166)) | s('type.hreal.hreal',X165) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X164),s('type.hreal.hreal',sK3(X165,X164)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X164),s('type.hreal.hreal',X165)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f2253,f436])).
fof(f2253,plain,(
  ( ! [X165,X166,X164] : (~p(s(bool,X166)) | s('type.hreal.hreal',X165) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X164),s('type.hreal.hreal',sK3(X165,X164)))) | s(bool,'F_0') = s(bool,X166) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X164),s('type.hreal.hreal',X165)))) ) | $spl16),
  inference(superposition,[],[f60,f782])).
fof(f2626,plain,(
  ( ! [X165,X166,X164] : (~p(s(bool,X166)) | s('type.hreal.hreal',X165) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X164),s('type.hreal.hreal',sK3(X165,X164)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X164),s('type.hreal.hreal',X165)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f2178,f436])).
fof(f2178,plain,(
  ( ! [X165,X166,X164] : (~p(s(bool,X166)) | s('type.hreal.hreal',X165) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X164),s('type.hreal.hreal',sK3(X165,X164)))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X164),s('type.hreal.hreal',X165))) | s(bool,'F_0') = s(bool,X166)) ) | $spl16),
  inference(superposition,[],[f60,f782])).
fof(f1904,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f1879,f434])).
fof(f1879,plain,(
  ( ! [X13] : (~p(s(bool,X13))) ) | $spl22),
  inference(resolution,[],[f434,f70])).
fof(f1877,plain,(
  $false | $spl22),
  inference(resolution,[],[f434,f84])).
fof(f1796,plain,(
  ( ! [X140,X138,X137] : (s('type.hreal.hreal',X138) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X137),s('type.hreal.hreal',X140))) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X137),s('type.hreal.hreal',X138)))) ) | $spl28),
  inference(cnf_transformation,[],[f1796_D])).
fof(f1796_D,plain,(
  ( ! [X140,X138,X137] : (s('type.hreal.hreal',X138) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X137),s('type.hreal.hreal',X140))) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X137),s('type.hreal.hreal',X138)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1816,plain,(
  ( ! [X140,X138,X139,X137] : (p(s(bool,X139)) | s('type.hreal.hreal',X138) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X137),s('type.hreal.hreal',X140))) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X137),s('type.hreal.hreal',X138)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f1408,f133])).
fof(f1408,plain,(
  ( ! [X140,X138,X139,X137] : (p(s(bool,X139)) | s('type.hreal.hreal',X138) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X137),s('type.hreal.hreal',X140))) | s(bool,'T_0') = s(bool,X139) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X137),s('type.hreal.hreal',X138)))) ) | $spl24),
  inference(superposition,[],[f61,f672])).
fof(f1795,plain,(
  ( ! [X140,X138,X139,X137] : (p(s(bool,X139)) | s('type.hreal.hreal',X138) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X137),s('type.hreal.hreal',X140))) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X137),s('type.hreal.hreal',X138)))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f1343,f133])).
fof(f1343,plain,(
  ( ! [X140,X138,X139,X137] : (p(s(bool,X139)) | s('type.hreal.hreal',X138) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X137),s('type.hreal.hreal',X140))) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X137),s('type.hreal.hreal',X138))) | s(bool,'T_0') = s(bool,X139)) ) | $spl24),
  inference(superposition,[],[f61,f672])).
fof(f728,plain,(
  $false | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f108,f364])).
fof(f364,plain,(
  s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | $spl6),
  inference(backward_demodulation,[],[f113,f56])).
fof(f726,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,X19)) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f716,f133])).
fof(f716,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,X19) | s(bool,'T_0') != s(bool,X19) | s(bool,'T_0') = s(bool,X18)) ) | ($spl16 | $spl24)),
  inference(superposition,[],[f538,f68])).
fof(f707,plain,(
  ( ! [X26,X25] : (s(bool,'F_0') != s(bool,X25) | s(bool,'T_0') != s(bool,X25) | ~p(s(bool,X26))) ) | ($spl16 | $spl24)),
  inference(resolution,[],[f538,f70])).
fof(f723,plain,(
  ( ! [X24,X23] : (s(bool,'F_0') != s(bool,X23) | s(bool,'T_0') != s(bool,X23) | p(s(bool,X24))) ) | ($spl16 | $spl24)),
  inference(subsumption_resolution,[],[f706,f85])).
fof(f706,plain,(
  ( ! [X24,X23] : (s(bool,'F_0') != s(bool,X23) | s(bool,'T_0') != s(bool,X23) | p(s(bool,X24)) | ~p(s(bool,'T_0'))) ) | ($spl16 | $spl24)),
  inference(resolution,[],[f538,f80])).
fof(f531,plain,(
  $false | ($spl6 | $spl10 | $spl16)),
  inference(subsumption_resolution,[],[f515,f505])).
fof(f505,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl6 | $spl10 | $spl16)),
  inference(subsumption_resolution,[],[f504,f289])).
fof(f504,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0))) ) | ($spl6 | $spl10)),
  inference(forward_demodulation,[],[f117,f113])).
fof(f515,plain,(
  ( ! [X7] : (~p(s(bool,X7))) ) | ($spl6 | $spl10 | $spl16)),
  inference(resolution,[],[f505,f70])).
fof(f513,plain,(
  $false | ($spl6 | $spl10 | $spl16)),
  inference(resolution,[],[f505,f84])).
fof(f431,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) | ~p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f364,f65])).
fof(f430,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | ~p(s(bool,X2)) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))))) ) | $spl6),
  inference(superposition,[],[f364,f65])).
fof(f435,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))))) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f429,f289])).
fof(f429,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))))) ) | $spl6),
  inference(superposition,[],[f364,f62])).
fof(f433,plain,(
  ( ! [X0] : (p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) | p(s(bool,X0))) ) | ($spl6 | $spl16)),
  inference(subsumption_resolution,[],[f428,f289])).
fof(f428,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) | p(s(bool,X0))) ) | $spl6),
  inference(superposition,[],[f364,f62])).
fof(f359,plain,(
  $false | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f358,f110])).
fof(f358,plain,(
  s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | $spl7),
  inference(trivial_inequality_removal,[],[f353])).
fof(f353,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | $spl7),
  inference(superposition,[],[f114,f68])).
fof(f313,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) | ~p(s(bool,X3))) )),
  inference(superposition,[],[f56,f65])).
fof(f312,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | ~p(s(bool,X2)) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))))) )),
  inference(superposition,[],[f56,f65])).
fof(f311,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | p(s(bool,X1)) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))))) )),
  inference(superposition,[],[f56,f62])).
fof(f310,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) | p(s(bool,X0))) )),
  inference(superposition,[],[f56,f62])).
fof(f309,plain,(
  s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))),
  inference(superposition,[],[f56,f68])).
fof(f308,plain,(
  s(bool,'T_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))),
  inference(superposition,[],[f56,f68])).
fof(f240,plain,(
  ( ! [X3] : (~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) | ~p(s(bool,X3))) ) | $spl2),
  inference(subsumption_resolution,[],[f236,f176])).
fof(f176,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,'T_0') = s(bool,X7)) ) | $spl2),
  inference(forward_demodulation,[],[f175,f109])).
fof(f175,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | ~p(s(bool,X7))) ) | $spl2),
  inference(subsumption_resolution,[],[f153,f79])).
fof(f153,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | s(bool,X7) = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | ~p(s(bool,X7))) ) | $spl2),
  inference(superposition,[],[f65,f109])).
fof(f236,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) | ~p(s(bool,X3))) ) | $spl2),
  inference(superposition,[],[f135,f65])).
fof(f135,plain,(
  s(bool,'T_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))) | $spl2),
  inference(backward_demodulation,[],[f109,f56])).
fof(f238,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f235,f176])).
fof(f235,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2)) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))))) ) | $spl2),
  inference(superposition,[],[f135,f65])).
fof(f234,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))))) ) | $spl2),
  inference(superposition,[],[f135,f62])).
fof(f233,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) | p(s(bool,X0))) ) | $spl2),
  inference(superposition,[],[f135,f62])).
fof(f134,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f134_D])).
fof(f134_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f128,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) | ~p(s(bool,X3))) ) | $spl3),
  inference(superposition,[],[f110,f65])).
fof(f127,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2)) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))))) ) | $spl3),
  inference(superposition,[],[f110,f65])).
fof(f132,plain,(
  p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) | $spl14),
  inference(cnf_transformation,[],[f132_D])).
fof(f132_D,plain,(
  p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f126,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))))) ) | $spl3),
  inference(superposition,[],[f110,f62])).
fof(f125,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2)))) | p(s(bool,X0))) ) | $spl3),
  inference(superposition,[],[f110,f62])).
fof(f106,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) | ~p(s(bool,X3))) )),
  inference(superposition,[],[f56,f65])).
fof(f105,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | ~p(s(bool,X2)) | ~p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))))) )),
  inference(superposition,[],[f56,f65])).
fof(f104,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | p(s(bool,X1)) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2))))))) )),
  inference(superposition,[],[f56,f62])).
fof(f103,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | p(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))) | p(s(bool,X0))) )),
  inference(superposition,[],[f56,f62])).
fof(f102,plain,(
  s(bool,'F_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'T_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))),
  inference(superposition,[],[f56,f68])).
fof(f101,plain,(
  s(bool,'T_0') != s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',sK1),s('type.hreal.hreal',sK2))) | s(bool,'F_0') = s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK1))),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',sK0),s('type.hreal.hreal',sK2)))))),
  inference(superposition,[],[f56,f68])).
% SZS output end Proof for thm.realax
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 265
% Passive clauses: 4755
% Generated clauses: 21620
% Final active clauses: 84
% Final passive clauses: 1957
% Input formulas: 16
% Initial clauses: 31
% 
% Selected by SInE selection: 15
% SInE iterations: 1
% 
% Duplicate literals: 861
% Trivial inequalities: 483
% Fw subsumption resolutions: 3042
% Fw demodulations: 2187
% Bw demodulations: 55
% 
% Simple tautologies: 628
% Equational tautologies: 109
% Forward subsumptions: 9374
% Fw demodulations to eq. taut.: 23
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 1127
% Forward superposition: 7679
% Backward superposition: 5184
% Self superposition: 206
% Equality factoring: 170
% Equality resolution: 130
% 
% Split clauses: 490
% Split components: 994
% SAT solver clauses: 533
% SAT solver unit clauses: 10
% SAT solver binary clauses: 50
% 
% Sat splits: 490
% Sat splitting refutations: 40
% 
% Memory used [KB]: 9594
% Time elapsed: 0.650 s
% ------------------------------
% ------------------------------
% Success in time 12.376 s
