Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.rat
% (84196)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 593
% Passive clauses: 8561
% Generated clauses: 18506
% Final active clauses: 592
% Final passive clauses: 7615
% Input formulas: 85
% Initial clauses: 253
% 
% Duplicate literals: 275
% Trivial inequalities: 138
% Fw subsumption resolutions: 1929
% Fw demodulations: 676
% Bw demodulations: 1
% 
% Simple tautologies: 244
% Equational tautologies: 59
% Forward subsumptions: 6372
% 
% Binary resolution: 8993
% Factoring: 948
% Forward superposition: 2145
% Backward superposition: 3013
% Self superposition: 33
% Equality factoring: 26
% Equality resolution: 22
% 
% Split clauses: 219
% Split components: 439
% SAT solver clauses: 443
% SAT solver unit clauses: 17
% SAT solver binary clauses: 94
% 
% Sat splits: 219
% 
% TWLsolver clauses: 443
% TWLsolver calls for satisfiability: 266
% 
% Memory used [KB]: 8827
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (84210)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 716
% Passive clauses: 2995
% Generated clauses: 19813
% Final active clauses: 544
% Final passive clauses: 1910
% Input formulas: 85
% Initial clauses: 253
% 
% Duplicate literals: 267
% Trivial inequalities: 72
% Fw subsumption resolutions: 1672
% Fw demodulations: 1796
% Bw demodulations: 356
% Condensations: 210
% Global subsumptions: 22
% 
% Simple tautologies: 320
% Equational tautologies: 20
% Forward subsumptions: 8775
% Fw demodulations to eq. taut.: 3
% Bw demodulations to eq. taut.: 35
% 
% Binary resolution: 4759
% Factoring: 254
% Forward superposition: 3100
% Backward superposition: 3232
% Self superposition: 37
% Equality factoring: 26
% Equality resolution: 14
% 
% Split clauses: 319
% Split components: 638
% SAT solver clauses: 3800
% SAT solver unit clauses: 274
% SAT solver binary clauses: 96
% SAT solver learnt clauses: 45
% SAT solver learnt literals: 14
% 
% Sat splits: 319
% Sat splitting refutations: 31
% 
% TWLsolver clauses: 3711
% TWLsolver calls for satisfiability: 1951
% 
% Memory used [KB]: 15735
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f22081,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f344,f611,f609,f343,f610,f608,f2397,f2403,f2399,f2683,f2687,f2674,f2685,f2401,f3458,f3486,f3488,f3480,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4841,f4884,f4793,f4886,f4797,f4888,f4800,f4890,f4803,f4892,f4806,f4894,f4809,f4896,f4812,f4898,f4815,f4900,f4818,f4902,f4821,f4904,f4824,f4906,f4827,f4908,f4830,f4910,f4833,f4912,f4836,f4914,f4839,f4916,f4794,f4842,f5076,f5077,f5079,f5081,f5083,f5085,f5087,f5089,f5091,f5093,f5095,f5097,f5099,f5101,f5103,f5105,f5107,f5109,f5111,f5113,f5115,f5117,f5119,f5121,f5123,f5125,f5127,f5129,f5131,f5133,f5135,f5137,f5139,f5141,f5143,f5145,f5147,f5149,f5151,f5153,f5155,f5157,f5159,f5161,f5163,f5165,f5167,f5169,f5171,f5173,f5175,f5177,f5179,f5181,f5183,f5185,f5187,f5189,f5191,f5193,f5195,f5197,f5199,f5201,f5203,f5205,f5207,f5209,f5211,f5213,f5215,f5217,f5219,f8440,f3478,f8472,f8474,f8465,f10677,f10638,f10641,f10955,f10956,f10957,f15517,f15518,f15525,f15522,f15519,f15520,f15528,f15521,f15530,f15523,f15527,f16063,f16064,f16076,f16074,f19549,f19653,f19651,f22080])).
fof(f22080,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f22079,f612])).
fof(f612,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f349,f345])).
fof(f345,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0',aHL_TRUTH)).
fof(f349,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f31])).
fof(f31,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.bool.NOT_CLAUSES|split|1')).
fof(f22079,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f22061,f10896])).
fof(f10896,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f645,f611,f879])).
fof(f879,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f875,f612])).
fof(f875,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f462,f371])).
fof(f371,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : (s(bool,'T_0') = s(bool,X7) | s(bool,'F_0') = s(bool,X7))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.bool.BOOL_CASES_AX')).
fof(f462,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f207])).
fof(f207,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f645,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f612,f382])).
fof(f382,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X7] : (s(bool,'T_0') = s(bool,X7) <=> p(s(bool,X7)))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f22061,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl0),
  inference(equality_resolution,[],[f11308])).
fof(f11308,plain,(
  ( ! [X1] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X1))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl0),
  inference(superposition,[],[f11019,f1638])).
fof(f1638,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X0)))) )),
  inference(superposition,[],[f454,f355])).
fof(f355,plain,(
  ( ! [X0] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X0)))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X0)))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X19] : s('type.rat.rat',X19) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X19)))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.rat.RAT_ADD_LID')).
fof(f454,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X1)))))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X1)))))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X6,X24] : s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X6),s('type.rat.rat',X24))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X24)))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.rat.RAT_SUB_ADDAINV')).
fof(f11019,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))) ) | $spl0),
  inference(backward_demodulation,[],[f11013,f469])).
fof(f469,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))) )),
  inference(cnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))) & (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))))),
  inference(nnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) <=> s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X6,X24] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat',X24)))) <=> s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X24),s('type.rat.rat',X6))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.rat.rat_les_def')).
fof(f11013,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl0),
  inference(forward_demodulation,[],[f10967,f8299])).
fof(f8299,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0)))))) )),
  inference(forward_demodulation,[],[f8265,f3771])).
fof(f3771,plain,(
  ( ! [X6] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X6))))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X6)))))))) )),
  inference(superposition,[],[f1277,f1264])).
fof(f1264,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X0)))))))) )),
  inference(superposition,[],[f358,f353])).
fof(f353,plain,(
  ( ! [X0] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X0)))))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : s('type.rat.rat',X0) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X0)))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X23] : s('type.rat.rat',X23) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X23)))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.rat.RAT')).
fof(f358,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0))))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0)))))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0))))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0)))))),
  inference(rectify,[],[f72])).
fof(f72,axiom,(
  ! [X18] : s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X18))))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X18)))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.rat.RAT_AINV_CALCULATE')).
fof(f1277,plain,(
  ( ! [X1] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X1))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1)))))) )),
  inference(superposition,[],[f359,f354])).
fof(f354,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X0)))))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X0)))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X23] : s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X23))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X23)))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.rat.rat_sgn_def')).
fof(f359,plain,(
  ( ! [X0] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0)))))))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0] : s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0)))))))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X18] : s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X18))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X18)))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.rat.RAT_SGN_CONG')).
fof(f8265,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X0)))))))) )),
  inference(unit_resulting_resolution,[],[f354,f1404])).
fof(f1404,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0)))))) )),
  inference(superposition,[],[f467,f360])).
fof(f360,plain,(
  ( ! [X0] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0)))))))) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0] : s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0)))))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X18] : s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X18))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X18)))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.frac.FRAC_SGN_AINV')).
fof(f467,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1] : ((s('type.integer.int',X1) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X12,X13] : (s('type.integer.int',X13) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X12))) <=> s('type.integer.int',X12) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X13))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.integer.INT_NEG_EQ')).
fof(f10967,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f608,f467])).
fof(f19651,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12)))))) ) | $spl72),
  inference(cnf_transformation,[],[f19651_D])).
fof(f19651_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12)))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f19653,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))))) | $spl75),
  inference(cnf_transformation,[],[f19653_D])).
fof(f19653_D,plain,(
  s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f19549,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))))) | s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12)))))) ) | $spl0),
  inference(equality_factoring,[],[f11742])).
fof(f11742,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(subsumption_resolution,[],[f11703,f645])).
fof(f11703,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f10988,f817])).
fof(f817,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f371,f371])).
fof(f10988,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))))) | $spl0),
  inference(superposition,[],[f364,f608])).
fof(f364,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X21] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21)))))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.integer.INT_LT_REDUCE|split|3')).
fof(f16074,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool))))) ) | $spl68),
  inference(cnf_transformation,[],[f16074_D])).
fof(f16074_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f16076,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl71),
  inference(cnf_transformation,[],[f16076_D])).
fof(f16076_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f16064,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK16(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2452,f817])).
fof(f2452,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK16(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK16(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f509,f371])).
fof(f509,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK16(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK16(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK16(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK16(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f215])).
fof(f215,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0',aHL_EXT)).
fof(f16063,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK16(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK16(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2452,f817])).
fof(f15527,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl62),
  inference(cnf_transformation,[],[f15527_D])).
fof(f15527_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f15523,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool))) ) | $spl58),
  inference(cnf_transformation,[],[f15523_D])).
fof(f15523_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f15530,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl67),
  inference(cnf_transformation,[],[f15530_D])).
fof(f15530_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f15521,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK16(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2451,f371])).
fof(f2451,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,X2,bool))) )),
  inference(superposition,[],[f509,f371])).
fof(f15528,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl64),
  inference(cnf_transformation,[],[f15528_D])).
fof(f15528_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f15520,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK16(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2451,f817])).
fof(f15519,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK16(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f2451,f817])).
fof(f15522,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl56),
  inference(cnf_transformation,[],[f15522_D])).
fof(f15522_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f15525,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl61),
  inference(cnf_transformation,[],[f15525_D])).
fof(f15525_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f15518,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK16(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2451,f817])).
fof(f15517,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK16(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2451,f817])).
fof(f10957,plain,(
  $false | ($spl3 | $spl50)),
  inference(subsumption_resolution,[],[f10868,f10936])).
fof(f10936,plain,(
  ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | ($spl3 | $spl50)),
  inference(unit_resulting_resolution,[],[f611,f4842])).
fof(f10868,plain,(
  sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f611,f345,f611,f590])).
fof(f590,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f341])).
fof(f341,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f216,f233,f232])).
fof(f232,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f233,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f216,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X23,X27,X28] : ((p(s(bool,X28)) <=> (p(s(bool,X27)) => p(s(bool,X23)))) <=> ((p(s(bool,X28)) | p(s(bool,X27))) & (p(s(bool,X28)) | ~p(s(bool,X23))) & (~p(s(bool,X27)) | p(s(bool,X23)) | ~p(s(bool,X28)))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.sat.sat_27_prove_ltb')).
fof(f10956,plain,(
  $false | ($spl3 | $spl50)),
  inference(subsumption_resolution,[],[f10870,f6515])).
fof(f6515,plain,(
  ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0') | $spl50),
  inference(unit_resulting_resolution,[],[f612,f4842])).
fof(f10870,plain,(
  sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f612,f345,f611,f590])).
fof(f10955,plain,(
  $false | ($spl3 | $spl50)),
  inference(subsumption_resolution,[],[f10871,f10936])).
fof(f10871,plain,(
  sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f611,f345,f611,f590])).
fof(f10641,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2630,f10524,f467])).
fof(f10524,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0)))))))) ) | $spl2),
  inference(forward_demodulation,[],[f10523,f354])).
fof(f10523,plain,(
  ( ! [X0] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X0))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0)))))))) ) | $spl2),
  inference(forward_demodulation,[],[f10496,f2621])).
fof(f2621,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12)))))))) ) | $spl2),
  inference(backward_demodulation,[],[f2620,f363])).
fof(f363,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X12] : s('type.integer.int',X12) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.integer.INT_MUL_RID')).
fof(f2620,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl2),
  inference(forward_demodulation,[],[f2619,f1638])).
fof(f2619,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK12))))) | $spl2),
  inference(subsumption_resolution,[],[f2607,f345])).
fof(f2607,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK12))))) | $spl2),
  inference(superposition,[],[f468,f648])).
fof(f648,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl2),
  inference(unit_resulting_resolution,[],[f610,f383])).
fof(f383,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f239])).
fof(f468,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))) )),
  inference(cnf_transformation,[],[f285])).
fof(f10496,plain,(
  ( ! [X0] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X0))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0))))),s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12)))))))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2621,f8300])).
fof(f8300,plain,(
  ( ! [X4,X5] : (s('type.integer.int',X5) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X4))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X4)))))) )),
  inference(backward_demodulation,[],[f8299,f1400])).
fof(f1400,plain,(
  ( ! [X4,X5] : (s('type.integer.int',X5) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X4))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X4)))))) )),
  inference(superposition,[],[f467,f354])).
fof(f2630,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))))) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f2620,f1390])).
fof(f1390,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f609,f467])).
fof(f10638,plain,(
  $false | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2629,f10524])).
fof(f2629,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))) | ($spl1 | $spl2)),
  inference(backward_demodulation,[],[f2620,f609])).
fof(f10677,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f10654])).
fof(f10654,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))) | ($spl1 | $spl2)),
  inference(superposition,[],[f2629,f10524])).
fof(f8465,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))) | $spl53),
  inference(cnf_transformation,[],[f8465_D])).
fof(f8465_D,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f8474,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl54),
  inference(cnf_transformation,[],[f8474_D])).
fof(f8474_D,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f8472,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))) | $spl2),
  inference(forward_demodulation,[],[f8471,f8304])).
fof(f8304,plain,(
  ( ! [X1] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))))) ) | $spl2),
  inference(backward_demodulation,[],[f8299,f2636])).
fof(f2636,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))))) ) | $spl2),
  inference(backward_demodulation,[],[f2620,f2582])).
fof(f2582,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))))) )),
  inference(subsumption_resolution,[],[f2569,f345])).
fof(f2569,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))))) )),
  inference(superposition,[],[f395,f367])).
fof(f367,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X21] : s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X21))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.integer.INT_LT_REDUCE|split|5')).
fof(f395,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0] : ((s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) & (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X12] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X12))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.intExtension.INT_SGN_CLAUSES|split|0')).
fof(f8471,plain,(
  ( ! [X19] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X19))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl2),
  inference(forward_demodulation,[],[f8453,f8320])).
fof(f8320,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))))) | $spl2),
  inference(backward_demodulation,[],[f8299,f3456])).
fof(f3456,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))))) | $spl2),
  inference(superposition,[],[f2636,f2620])).
fof(f8453,plain,(
  ( ! [X19] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X19))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl2),
  inference(superposition,[],[f2623,f8304])).
fof(f2623,plain,(
  ( ! [X0] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl2),
  inference(backward_demodulation,[],[f2620,f392])).
fof(f392,plain,(
  ( ! [X0] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0] : ((s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) & (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X12] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X12))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.intExtension.INT_SGN_CLAUSES|split|2')).
fof(f3478,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl10),
  inference(cnf_transformation,[],[f3478_D])).
fof(f3478_D,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f8440,plain,(
  ( ! [X1] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl2),
  inference(superposition,[],[f2623,f8304])).
fof(f5219,plain,(
  $false | ($spl2 | $spl18)),
  inference(global_subsumption,[],[f5218,f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188,f5190,f5192,f5194,f5196,f5198,f5200,f5202,f5204,f5206,f5208,f5210])).
fof(f5210,plain,(
  ( ! [X85,X83,X86,X84] : (~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X83))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))))),X85,X86) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X83))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))))),X85,X86)) ) | $spl18),
  inference(subsumption_resolution,[],[f5051,f345])).
fof(f5051,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X83))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))))),X85,X86) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X83))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))))),X85,X86)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3019])).
fof(f5208,plain,(
  ( ! [X80,X81,X79,X82] : (~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X79))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))))),X81,X82) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X79))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))))),X81,X82)) ) | $spl18),
  inference(subsumption_resolution,[],[f5050,f345])).
fof(f5050,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X79))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))))),X81,X82) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X79))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))))),X81,X82)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3018])).
fof(f5206,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X77)) | p(s(bool,X78)) | ~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X75))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X76)))))))),X78,X77)) ) | $spl18),
  inference(subsumption_resolution,[],[f5049,f345])).
fof(f5049,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X77)) | p(s(bool,X78)) | ~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X75))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X76)))))))),X78,X77)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3017])).
fof(f5204,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X73)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X71))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X72)))))))),X73,X74)) ) | $spl18),
  inference(subsumption_resolution,[],[f5048,f345])).
fof(f5048,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X73)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X71))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X72)))))))),X73,X74)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3016])).
fof(f5202,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X69)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X67))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X68)))))))),X70,X69)) ) | $spl18),
  inference(subsumption_resolution,[],[f5047,f345])).
fof(f5047,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X69)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X67))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X68)))))))),X70,X69)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3015])).
fof(f5200,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X65)) | ~p(s(bool,X66)) | ~sP10(X66,X65,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X63))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X64)))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5046,f345])).
fof(f5046,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X65)) | ~p(s(bool,X66)) | ~sP10(X66,X65,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X63))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X64)))))))))) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3014])).
fof(f5198,plain,(
  ( ! [X61,X59,X62,X60] : (p(s(bool,X61)) | p(s(bool,X62)) | ~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X59))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X60)))))))),X62,X61)) ) | $spl18),
  inference(subsumption_resolution,[],[f5045,f345])).
fof(f5045,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | p(s(bool,X61)) | p(s(bool,X62)) | ~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X59))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X60)))))))),X62,X61)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3013])).
fof(f5196,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X57)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X55))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56)))))))),X57,X58)) ) | $spl18),
  inference(subsumption_resolution,[],[f5044,f345])).
fof(f5044,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X55))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56)))))))),X57,X58)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3012])).
fof(f5194,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X53)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X52)))))))),X54,X53)) ) | $spl18),
  inference(subsumption_resolution,[],[f5043,f345])).
fof(f5043,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X53)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X52)))))))),X54,X53)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3011])).
fof(f5192,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X49)) | p(s(bool,X50)) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X47))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48)))))))),X49,X50)) ) | $spl18),
  inference(subsumption_resolution,[],[f5042,f345])).
fof(f5042,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X49)) | p(s(bool,X50)) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X47))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48)))))))),X49,X50)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3010])).
fof(f5190,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP7(X46,X45,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5041,f345])).
fof(f5041,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP7(X46,X45,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))))))))) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3009])).
fof(f5188,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP6(X41,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40)))))))),X42)) ) | $spl18),
  inference(subsumption_resolution,[],[f5040,f345])).
fof(f5040,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP6(X41,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40)))))))),X42)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3008])).
fof(f5186,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP3(X37,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36)))))))),X38)) ) | $spl18),
  inference(subsumption_resolution,[],[f5039,f345])).
fof(f5039,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP3(X37,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36)))))))),X38)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3007])).
fof(f5184,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP4(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32)))))))),X34)) ) | $spl18),
  inference(subsumption_resolution,[],[f5038,f345])).
fof(f5038,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~sP4(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32)))))))),X34)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3006])).
fof(f5182,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X30)) | sP1(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29)))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5037,f345])).
fof(f5037,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP1(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29)))))))))) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3005])).
fof(f5180,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X27)) | sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26)))))))),X27)) ) | $spl18),
  inference(subsumption_resolution,[],[f5036,f345])).
fof(f5036,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26)))))))),X27)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3004])).
fof(f5178,plain,(
  ( ! [X24,X23,X22] : (sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23)))))))),X24)) ) | $spl18),
  inference(subsumption_resolution,[],[f5035,f345])).
fof(f5035,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23)))))))),X24)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3003])).
fof(f5176,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20)))))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5034,f345])).
fof(f5034,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X21,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20)))))))))) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f3002])).
fof(f5166,plain,(
  ( ! [X85,X83,X86,X84] : (~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X83))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))),X85,X86) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X83))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))),X85,X86)) ) | $spl18),
  inference(subsumption_resolution,[],[f5029,f345])).
fof(f5029,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X83))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))),X85,X86) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X83))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))),X85,X86)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2354])).
fof(f5164,plain,(
  ( ! [X80,X81,X79,X82] : (~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X79))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))),X81,X82) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X79))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))),X81,X82)) ) | $spl18),
  inference(subsumption_resolution,[],[f5028,f345])).
fof(f5028,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X79))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))),X81,X82) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X79))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))),X81,X82)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2353])).
fof(f5162,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X77)) | p(s(bool,X78)) | ~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X75))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X76)))))),X78,X77)) ) | $spl18),
  inference(subsumption_resolution,[],[f5027,f345])).
fof(f5027,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X77)) | p(s(bool,X78)) | ~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X75))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X76)))))),X78,X77)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2352])).
fof(f5160,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X73)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X71))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X72)))))),X73,X74)) ) | $spl18),
  inference(subsumption_resolution,[],[f5026,f345])).
fof(f5026,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X73)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X71))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X72)))))),X73,X74)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2351])).
fof(f5158,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X69)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X67))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X68)))))),X70,X69)) ) | $spl18),
  inference(subsumption_resolution,[],[f5025,f345])).
fof(f5025,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X69)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X67))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X68)))))),X70,X69)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2350])).
fof(f5156,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X65)) | ~p(s(bool,X66)) | ~sP10(X66,X65,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X63))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X64)))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5024,f345])).
fof(f5024,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X65)) | ~p(s(bool,X66)) | ~sP10(X66,X65,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X63))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X64)))))))) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2349])).
fof(f5154,plain,(
  ( ! [X61,X59,X62,X60] : (p(s(bool,X61)) | p(s(bool,X62)) | ~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X59))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X60)))))),X62,X61)) ) | $spl18),
  inference(subsumption_resolution,[],[f5023,f345])).
fof(f5023,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | p(s(bool,X61)) | p(s(bool,X62)) | ~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X59))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X60)))))),X62,X61)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2348])).
fof(f5152,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X57)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X55))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56)))))),X57,X58)) ) | $spl18),
  inference(subsumption_resolution,[],[f5022,f345])).
fof(f5022,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X55))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56)))))),X57,X58)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2347])).
fof(f5150,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X53)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X52)))))),X54,X53)) ) | $spl18),
  inference(subsumption_resolution,[],[f5021,f345])).
fof(f5021,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X53)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X52)))))),X54,X53)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2346])).
fof(f5148,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X49)) | p(s(bool,X50)) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X47))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48)))))),X49,X50)) ) | $spl18),
  inference(subsumption_resolution,[],[f5020,f345])).
fof(f5020,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X49)) | p(s(bool,X50)) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X47))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48)))))),X49,X50)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2345])).
fof(f5146,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP7(X46,X45,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5019,f345])).
fof(f5019,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP7(X46,X45,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))))))) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2344])).
fof(f5144,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP6(X41,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40)))))),X42)) ) | $spl18),
  inference(subsumption_resolution,[],[f5018,f345])).
fof(f5018,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP6(X41,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40)))))),X42)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2343])).
fof(f5142,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP3(X37,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36)))))),X38)) ) | $spl18),
  inference(subsumption_resolution,[],[f5017,f345])).
fof(f5017,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP3(X37,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36)))))),X38)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2342])).
fof(f5140,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP4(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32)))))),X34)) ) | $spl18),
  inference(subsumption_resolution,[],[f5016,f345])).
fof(f5016,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~sP4(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32)))))),X34)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2341])).
fof(f5138,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X30)) | sP1(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29)))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5015,f345])).
fof(f5015,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP1(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29)))))))) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2340])).
fof(f5136,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X27)) | sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26)))))),X27)) ) | $spl18),
  inference(subsumption_resolution,[],[f5014,f345])).
fof(f5014,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26)))))),X27)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2339])).
fof(f5134,plain,(
  ( ! [X24,X23,X22] : (sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23)))))),X24)) ) | $spl18),
  inference(subsumption_resolution,[],[f5013,f345])).
fof(f5013,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23)))))),X24)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2338])).
fof(f5132,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20)))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f5012,f345])).
fof(f5012,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X21,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20)))))))) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f2337])).
fof(f5128,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f4995,f345])).
fof(f4995,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int',X0) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f460])).
fof(f5126,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f4994,f345])).
fof(f4994,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int',X1) = s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f461])).
fof(f5112,plain,(
  ( ! [X85,X83,X86,X84] : (~sP8('const.rat.rat_gre_2'(s('type.rat.rat',X83),s('type.rat.rat',X84)),X85,X86) | sP9('const.rat.rat_gre_2'(s('type.rat.rat',X83),s('type.rat.rat',X84)),X85,X86)) ) | $spl18),
  inference(subsumption_resolution,[],[f4957,f345])).
fof(f4957,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP8('const.rat.rat_gre_2'(s('type.rat.rat',X83),s('type.rat.rat',X84)),X85,X86) | sP9('const.rat.rat_gre_2'(s('type.rat.rat',X83),s('type.rat.rat',X84)),X85,X86)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1346])).
fof(f5110,plain,(
  ( ! [X80,X81,X79,X82] : (~sP10('const.rat.rat_gre_2'(s('type.rat.rat',X79),s('type.rat.rat',X80)),X81,X82) | sP11('const.rat.rat_gre_2'(s('type.rat.rat',X79),s('type.rat.rat',X80)),X81,X82)) ) | $spl18),
  inference(subsumption_resolution,[],[f4956,f345])).
fof(f4956,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~sP10('const.rat.rat_gre_2'(s('type.rat.rat',X79),s('type.rat.rat',X80)),X81,X82) | sP11('const.rat.rat_gre_2'(s('type.rat.rat',X79),s('type.rat.rat',X80)),X81,X82)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1345])).
fof(f5108,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X77)) | p(s(bool,X78)) | ~sP11('const.rat.rat_gre_2'(s('type.rat.rat',X75),s('type.rat.rat',X76)),X78,X77)) ) | $spl18),
  inference(subsumption_resolution,[],[f4955,f345])).
fof(f4955,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X77)) | p(s(bool,X78)) | ~sP11('const.rat.rat_gre_2'(s('type.rat.rat',X75),s('type.rat.rat',X76)),X78,X77)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1344])).
fof(f5106,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X73)) | sP11('const.rat.rat_gre_2'(s('type.rat.rat',X71),s('type.rat.rat',X72)),X73,X74)) ) | $spl18),
  inference(subsumption_resolution,[],[f4954,f345])).
fof(f4954,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X73)) | sP11('const.rat.rat_gre_2'(s('type.rat.rat',X71),s('type.rat.rat',X72)),X73,X74)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1343])).
fof(f5104,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X69)) | sP11('const.rat.rat_gre_2'(s('type.rat.rat',X67),s('type.rat.rat',X68)),X70,X69)) ) | $spl18),
  inference(subsumption_resolution,[],[f4953,f345])).
fof(f4953,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X69)) | sP11('const.rat.rat_gre_2'(s('type.rat.rat',X67),s('type.rat.rat',X68)),X70,X69)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1342])).
fof(f5102,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X65)) | ~p(s(bool,X66)) | ~sP10(X66,X65,'const.rat.rat_gre_2'(s('type.rat.rat',X63),s('type.rat.rat',X64)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4952,f345])).
fof(f4952,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X65)) | ~p(s(bool,X66)) | ~sP10(X66,X65,'const.rat.rat_gre_2'(s('type.rat.rat',X63),s('type.rat.rat',X64)))) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1341])).
fof(f5100,plain,(
  ( ! [X61,X59,X62,X60] : (p(s(bool,X61)) | p(s(bool,X62)) | ~sP9('const.rat.rat_gre_2'(s('type.rat.rat',X59),s('type.rat.rat',X60)),X62,X61)) ) | $spl18),
  inference(subsumption_resolution,[],[f4951,f345])).
fof(f4951,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | p(s(bool,X61)) | p(s(bool,X62)) | ~sP9('const.rat.rat_gre_2'(s('type.rat.rat',X59),s('type.rat.rat',X60)),X62,X61)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1340])).
fof(f5098,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X57)) | sP9('const.rat.rat_gre_2'(s('type.rat.rat',X55),s('type.rat.rat',X56)),X57,X58)) ) | $spl18),
  inference(subsumption_resolution,[],[f4950,f345])).
fof(f4950,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | sP9('const.rat.rat_gre_2'(s('type.rat.rat',X55),s('type.rat.rat',X56)),X57,X58)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1339])).
fof(f5096,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X53)) | sP9('const.rat.rat_gre_2'(s('type.rat.rat',X51),s('type.rat.rat',X52)),X54,X53)) ) | $spl18),
  inference(subsumption_resolution,[],[f4949,f345])).
fof(f4949,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X53)) | sP9('const.rat.rat_gre_2'(s('type.rat.rat',X51),s('type.rat.rat',X52)),X54,X53)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1338])).
fof(f5094,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X49)) | p(s(bool,X50)) | ~sP8('const.rat.rat_gre_2'(s('type.rat.rat',X47),s('type.rat.rat',X48)),X49,X50)) ) | $spl18),
  inference(subsumption_resolution,[],[f4948,f345])).
fof(f4948,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X49)) | p(s(bool,X50)) | ~sP8('const.rat.rat_gre_2'(s('type.rat.rat',X47),s('type.rat.rat',X48)),X49,X50)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1337])).
fof(f5092,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP7(X46,X45,'const.rat.rat_gre_2'(s('type.rat.rat',X43),s('type.rat.rat',X44)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4947,f345])).
fof(f4947,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP7(X46,X45,'const.rat.rat_gre_2'(s('type.rat.rat',X43),s('type.rat.rat',X44)))) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1336])).
fof(f5090,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP6(X41,'const.rat.rat_gre_2'(s('type.rat.rat',X39),s('type.rat.rat',X40)),X42)) ) | $spl18),
  inference(subsumption_resolution,[],[f4946,f345])).
fof(f4946,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP6(X41,'const.rat.rat_gre_2'(s('type.rat.rat',X39),s('type.rat.rat',X40)),X42)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1335])).
fof(f5088,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP3(X37,'const.rat.rat_gre_2'(s('type.rat.rat',X35),s('type.rat.rat',X36)),X38)) ) | $spl18),
  inference(subsumption_resolution,[],[f4945,f345])).
fof(f4945,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP3(X37,'const.rat.rat_gre_2'(s('type.rat.rat',X35),s('type.rat.rat',X36)),X38)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1334])).
fof(f5086,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X33)) | p(s(bool,X34)) | ~sP4(X33,'const.rat.rat_gre_2'(s('type.rat.rat',X31),s('type.rat.rat',X32)),X34)) ) | $spl18),
  inference(subsumption_resolution,[],[f4944,f345])).
fof(f4944,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~sP4(X33,'const.rat.rat_gre_2'(s('type.rat.rat',X31),s('type.rat.rat',X32)),X34)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1333])).
fof(f5084,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X30)) | sP1(X30,'const.rat.rat_gre_2'(s('type.rat.rat',X28),s('type.rat.rat',X29)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4943,f345])).
fof(f4943,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP1(X30,'const.rat.rat_gre_2'(s('type.rat.rat',X28),s('type.rat.rat',X29)))) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1332])).
fof(f5082,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X27)) | sP1('const.rat.rat_gre_2'(s('type.rat.rat',X25),s('type.rat.rat',X26)),X27)) ) | $spl18),
  inference(subsumption_resolution,[],[f4942,f345])).
fof(f4942,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1('const.rat.rat_gre_2'(s('type.rat.rat',X25),s('type.rat.rat',X26)),X27)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1331])).
fof(f5080,plain,(
  ( ! [X24,X23,X22] : (sP0('const.rat.rat_gre_2'(s('type.rat.rat',X22),s('type.rat.rat',X23)),X24)) ) | $spl18),
  inference(subsumption_resolution,[],[f4941,f345])).
fof(f4941,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0('const.rat.rat_gre_2'(s('type.rat.rat',X22),s('type.rat.rat',X23)),X24)) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1330])).
fof(f5078,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.rat.rat_gre_2'(s('type.rat.rat',X19),s('type.rat.rat',X20)))) ) | $spl18),
  inference(subsumption_resolution,[],[f4940,f345])).
fof(f4940,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X21,'const.rat.rat_gre_2'(s('type.rat.rat',X19),s('type.rat.rat',X20)))) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f1329])).
fof(f5075,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl18),
  inference(subsumption_resolution,[],[f4920,f4794])).
fof(f4920,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl18),
  inference(backward_demodulation,[],[f4794,f386])).
fof(f4512,plain,(
  ( ! [X206,X204,X205,X203] : (s(X205,f(s(bool,'T_0'))) != s(X205,g(s(bool,'T_0'))) | s(fun(bool,X205),X203) = s(fun(bool,X205),X204) | s(bool,X206) = s(bool,sK16(X203,X204,X205,bool)) | s(bool,'T_0') = s(bool,X206)) )),
  inference(superposition,[],[f509,f817])).
fof(f4511,plain,(
  ( ! [X202,X200,X199,X201] : (~p(s(bool,happ(s(fun(bool,bool),X199),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X200),s(bool,X201)))) | ~p(s(bool,happ(s(fun(bool,bool),X200),s(bool,sK14(X199,X200,bool))))) | s(bool,X202) = s(bool,sK15(X199,X200,bool)) | s(bool,'T_0') = s(bool,X202)) )),
  inference(superposition,[],[f451,f817])).
fof(f4510,plain,(
  ( ! [X198,X196,X197,X195] : (~p(s(bool,happ(s(fun(bool,bool),X195),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X195),s(bool,X197)))) | ~p(s(bool,happ(s(fun(bool,bool),X196),s(bool,sK14(X195,X196,bool))))) | s(bool,X198) = s(bool,sK15(X195,X196,bool)) | s(bool,'T_0') = s(bool,X198)) )),
  inference(superposition,[],[f452,f817])).
fof(f4509,plain,(
  ( ! [X191,X194,X192,X193] : (~p(s(bool,happ(s(fun(bool,bool),X191),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X192),s(bool,X193)))) | ~p(s(bool,happ(s(fun(bool,bool),X192),s(bool,sK13(X191,X192,bool))))) | s(bool,X194) = s(bool,sK13(X191,X192,bool)) | s(bool,'T_0') = s(bool,X194)) )),
  inference(superposition,[],[f449,f817])).
fof(f4508,plain,(
  ( ! [X189,X187,X190,X188] : (~p(s(bool,happ(s(fun(bool,bool),X187),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X187),s(bool,X189)))) | ~p(s(bool,happ(s(fun(bool,bool),X188),s(bool,sK13(X187,X188,bool))))) | s(bool,X190) = s(bool,sK13(X187,X188,bool)) | s(bool,'T_0') = s(bool,X190)) )),
  inference(superposition,[],[f450,f817])).
fof(f4507,plain,(
  ( ! [X185,X186,X184,X183] : (s(bool,'T_0') != s(bool,f(s(X183,sK16(X184,X185,bool,X183)))) | s(fun(X183,bool),X184) = s(fun(X183,bool),X185) | s(bool,X186) = s(bool,g(s(X183,sK16(X184,X185,bool,X183)))) | s(bool,'T_0') = s(bool,X186)) )),
  inference(superposition,[],[f509,f817])).
fof(f4877,plain,(
  ( ! [X152,X151,X150] : (s('type.integer.int',X151) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X151),s('type.integer.int',X150))))))) | s(bool,X152) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X150))) | s(bool,'T_0') = s(bool,X152)) )),
  inference(subsumption_resolution,[],[f4498,f345])).
fof(f4498,plain,(
  ( ! [X152,X151,X150] : (~p(s(bool,'T_0')) | s('type.integer.int',X151) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X151),s('type.integer.int',X150))))))) | s(bool,X152) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X150))) | s(bool,'T_0') = s(bool,X152)) )),
  inference(superposition,[],[f460,f817])).
fof(f4876,plain,(
  ( ! [X149,X147,X148] : (s('type.integer.int',X147) = s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X148),s('type.integer.int',X147))))))) | s(bool,X149) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X147))) | s(bool,'T_0') = s(bool,X149)) )),
  inference(subsumption_resolution,[],[f4497,f345])).
fof(f4497,plain,(
  ( ! [X149,X147,X148] : (~p(s(bool,'T_0')) | s('type.integer.int',X147) = s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X148),s('type.integer.int',X147))))))) | s(bool,X149) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X147))) | s(bool,'T_0') = s(bool,X149)) )),
  inference(superposition,[],[f461,f817])).
fof(f4874,plain,(
  ( ! [X142,X140,X143,X141,X139] : (p(s(bool,happ(s(fun(X139,bool),X141),s(X139,X142)))) | ~p(s(bool,happ(s(fun(X139,bool),X141),s(X139,sK14(X140,X141,X139))))) | s(bool,X143) = s(bool,happ(s(fun(X139,bool),X140),s(X139,sK15(X140,X141,X139)))) | s(bool,'T_0') = s(bool,X143)) )),
  inference(subsumption_resolution,[],[f4494,f345])).
fof(f4494,plain,(
  ( ! [X142,X140,X143,X141,X139] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X139,bool),X141),s(X139,X142)))) | ~p(s(bool,happ(s(fun(X139,bool),X141),s(X139,sK14(X140,X141,X139))))) | s(bool,X143) = s(bool,happ(s(fun(X139,bool),X140),s(X139,sK15(X140,X141,X139)))) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f451,f817])).
fof(f4873,plain,(
  ( ! [X134,X138,X136,X135,X137] : (p(s(bool,happ(s(fun(X134,bool),X135),s(X134,X137)))) | ~p(s(bool,happ(s(fun(X134,bool),X136),s(X134,sK14(X135,X136,X134))))) | s(bool,X138) = s(bool,happ(s(fun(X134,bool),X135),s(X134,sK15(X135,X136,X134)))) | s(bool,'T_0') = s(bool,X138)) )),
  inference(subsumption_resolution,[],[f4493,f345])).
fof(f4493,plain,(
  ( ! [X134,X138,X136,X135,X137] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X134,bool),X135),s(X134,X137)))) | ~p(s(bool,happ(s(fun(X134,bool),X136),s(X134,sK14(X135,X136,X134))))) | s(bool,X138) = s(bool,happ(s(fun(X134,bool),X135),s(X134,sK15(X135,X136,X134)))) | s(bool,'T_0') = s(bool,X138)) )),
  inference(superposition,[],[f452,f817])).
fof(f4872,plain,(
  ( ! [X132,X130,X133,X131,X129] : (p(s(bool,happ(s(fun(X129,bool),X131),s(X129,X132)))) | ~p(s(bool,happ(s(fun(X129,bool),X131),s(X129,sK13(X130,X131,X129))))) | s(bool,X133) = s(bool,happ(s(fun(X129,bool),X130),s(X129,sK13(X130,X131,X129)))) | s(bool,'T_0') = s(bool,X133)) )),
  inference(subsumption_resolution,[],[f4492,f345])).
fof(f4492,plain,(
  ( ! [X132,X130,X133,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X129,bool),X131),s(X129,X132)))) | ~p(s(bool,happ(s(fun(X129,bool),X131),s(X129,sK13(X130,X131,X129))))) | s(bool,X133) = s(bool,happ(s(fun(X129,bool),X130),s(X129,sK13(X130,X131,X129)))) | s(bool,'T_0') = s(bool,X133)) )),
  inference(superposition,[],[f449,f817])).
fof(f4871,plain,(
  ( ! [X127,X125,X128,X126,X124] : (p(s(bool,happ(s(fun(X124,bool),X125),s(X124,X127)))) | ~p(s(bool,happ(s(fun(X124,bool),X126),s(X124,sK13(X125,X126,X124))))) | s(bool,X128) = s(bool,happ(s(fun(X124,bool),X125),s(X124,sK13(X125,X126,X124)))) | s(bool,'T_0') = s(bool,X128)) )),
  inference(subsumption_resolution,[],[f4491,f345])).
fof(f4491,plain,(
  ( ! [X127,X125,X128,X126,X124] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X124,bool),X125),s(X124,X127)))) | ~p(s(bool,happ(s(fun(X124,bool),X126),s(X124,sK13(X125,X126,X124))))) | s(bool,X128) = s(bool,happ(s(fun(X124,bool),X125),s(X124,sK13(X125,X126,X124)))) | s(bool,'T_0') = s(bool,X128)) )),
  inference(superposition,[],[f450,f817])).
fof(f4868,plain,(
  ( ! [X116,X114,X115,X113] : (sP0(X115,'const.rat.rat_gre_2'(s('type.rat.rat',X114),s('type.rat.rat',X113))) | s(bool,X116) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X113),s('type.rat.rat',X114))) | s(bool,'T_0') = s(bool,X116)) )),
  inference(subsumption_resolution,[],[f4488,f345])).
fof(f4488,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,'T_0')) | sP0(X115,'const.rat.rat_gre_2'(s('type.rat.rat',X114),s('type.rat.rat',X113))) | s(bool,X116) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X113),s('type.rat.rat',X114))) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f1329,f817])).
fof(f4867,plain,(
  ( ! [X111,X109,X112,X110] : (sP0('const.rat.rat_gre_2'(s('type.rat.rat',X110),s('type.rat.rat',X109)),X111) | s(bool,X112) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X109),s('type.rat.rat',X110))) | s(bool,'T_0') = s(bool,X112)) )),
  inference(subsumption_resolution,[],[f4487,f345])).
fof(f4487,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,'T_0')) | sP0('const.rat.rat_gre_2'(s('type.rat.rat',X110),s('type.rat.rat',X109)),X111) | s(bool,X112) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X109),s('type.rat.rat',X110))) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f1330,f817])).
fof(f4862,plain,(
  ( ! [X85,X83,X86,X84] : (~sP8(X83,X84,X85) | sP9(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f4476,f345])).
fof(f4476,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP8(X83,X84,X85) | sP9(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f597,f817])).
fof(f4861,plain,(
  ( ! [X80,X81,X79,X82] : (~sP10(X79,X80,X81) | sP11(X79,X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f4475,f345])).
fof(f4475,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~sP10(X79,X80,X81) | sP11(X79,X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f596,f817])).
fof(f4860,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | p(s(bool,X77)) | ~sP11(X75,X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f4474,f345])).
fof(f4474,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | p(s(bool,X77)) | ~sP11(X75,X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f593,f817])).
fof(f4859,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X72)) | sP11(X71,X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f4473,f345])).
fof(f4473,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X72)) | sP11(X71,X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f592,f817])).
fof(f4858,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | sP11(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f4472,f345])).
fof(f4472,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | sP11(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f591,f817])).
fof(f4857,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | ~p(s(bool,X65)) | ~sP10(X65,X64,X63) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f4471,f345])).
fof(f4471,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | ~p(s(bool,X65)) | ~sP10(X65,X64,X63) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f586,f817])).
fof(f4856,plain,(
  ( ! [X61,X59,X62,X60] : (p(s(bool,X60)) | p(s(bool,X61)) | ~sP9(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f4470,f345])).
fof(f4470,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | p(s(bool,X60)) | p(s(bool,X61)) | ~sP9(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f576,f817])).
fof(f4855,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | sP9(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f4469,f345])).
fof(f4469,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP9(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f575,f817])).
fof(f4854,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | sP9(X51,X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f4468,f345])).
fof(f4468,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | sP9(X51,X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f574,f817])).
fof(f4853,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X48)) | p(s(bool,X49)) | ~sP8(X47,X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f4467,f345])).
fof(f4467,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X48)) | p(s(bool,X49)) | ~sP8(X47,X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f569,f817])).
fof(f4852,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP7(X45,X44,X43) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f4466,f345])).
fof(f4466,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP7(X45,X44,X43) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f561,f817])).
fof(f4851,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP6(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f4465,f345])).
fof(f4465,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP6(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f552,f817])).
fof(f4850,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP3(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f4464,f345])).
fof(f4464,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP3(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f533,f817])).
fof(f4849,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~sP4(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4463,f345])).
fof(f4463,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP4(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f529,f817])).
fof(f4848,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f4462,f345])).
fof(f4462,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f496,f817])).
fof(f4847,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f4461,f345])).
fof(f4461,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f493,f817])).
fof(f4846,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f4460,f345])).
fof(f4460,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f478,f817])).
fof(f4845,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f4459,f345])).
fof(f4459,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f477,f817])).
fof(f4844,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f4458,f345])).
fof(f4458,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f465,f817])).
fof(f4843,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4454,f345])).
fof(f4454,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f386,f817])).
fof(f4419,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f382,f817])).
fof(f4348,plain,(
  ( ! [X216,X215,X213,X214] : (s(X215,f(s(bool,X216))) != s(X215,g(s(bool,X216))) | s(fun(bool,X215),X213) = s(fun(bool,X215),X214) | s(bool,'T_0') = s(bool,X216) | s(bool,'T_0') = s(bool,sK16(X213,X214,X215,bool))) )),
  inference(superposition,[],[f509,f817])).
fof(f4347,plain,(
  ( ! [X212,X210,X211,X209] : (~p(s(bool,happ(s(fun(bool,bool),X209),s(bool,X211)))) | p(s(bool,happ(s(fun(bool,bool),X210),s(bool,X212)))) | ~p(s(bool,happ(s(fun(bool,bool),X210),s(bool,sK14(X209,X210,bool))))) | s(bool,'T_0') = s(bool,X211) | s(bool,'T_0') = s(bool,sK15(X209,X210,bool))) )),
  inference(superposition,[],[f451,f817])).
fof(f4346,plain,(
  ( ! [X208,X206,X207,X205] : (~p(s(bool,happ(s(fun(bool,bool),X205),s(bool,X207)))) | p(s(bool,happ(s(fun(bool,bool),X205),s(bool,X208)))) | ~p(s(bool,happ(s(fun(bool,bool),X206),s(bool,sK14(X205,X206,bool))))) | s(bool,'T_0') = s(bool,X207) | s(bool,'T_0') = s(bool,sK15(X205,X206,bool))) )),
  inference(superposition,[],[f452,f817])).
fof(f4345,plain,(
  ( ! [X204,X202,X203,X201] : (~p(s(bool,happ(s(fun(bool,bool),X201),s(bool,X203)))) | p(s(bool,happ(s(fun(bool,bool),X202),s(bool,X204)))) | ~p(s(bool,happ(s(fun(bool,bool),X202),s(bool,sK13(X201,X202,bool))))) | s(bool,'T_0') = s(bool,X203) | s(bool,'T_0') = s(bool,sK13(X201,X202,bool))) )),
  inference(superposition,[],[f449,f817])).
fof(f4344,plain,(
  ( ! [X198,X200,X199,X197] : (~p(s(bool,happ(s(fun(bool,bool),X197),s(bool,X199)))) | p(s(bool,happ(s(fun(bool,bool),X197),s(bool,X200)))) | ~p(s(bool,happ(s(fun(bool,bool),X198),s(bool,sK13(X197,X198,bool))))) | s(bool,'T_0') = s(bool,X199) | s(bool,'T_0') = s(bool,sK13(X197,X198,bool))) )),
  inference(superposition,[],[f450,f817])).
fof(f4343,plain,(
  ( ! [X196,X194,X195,X193] : (s(bool,X196) != s(bool,f(s(X193,sK16(X194,X195,bool,X193)))) | s(fun(X193,bool),X194) = s(fun(X193,bool),X195) | s(bool,'T_0') = s(bool,X196) | s(bool,'T_0') = s(bool,g(s(X193,sK16(X194,X195,bool,X193))))) )),
  inference(superposition,[],[f509,f817])).
fof(f4260,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X38),s('type.rat.rat',X37))) | s(bool,X39) = s(bool,'const.rat.rat_gre_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f817,f453])).
fof(f4258,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X34),s('type.num.num',X33))) | s(bool,X35) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X34))))))))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f817,f459])).
fof(f4255,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num',X20))) | s(bool,X21) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20))))))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f817,f458])).
fof(f4243,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f817,f817])).
fof(f4218,plain,(
  ( ! [X52,X53,X51] : (s(bool,X53) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X52),s('type.rat.rat',X51))) | s(bool,'T_0') = s(bool,X53) | s(bool,'T_0') = s(bool,'const.rat.rat_gre_2'(s('type.rat.rat',X51),s('type.rat.rat',X52)))) )),
  inference(superposition,[],[f817,f453])).
fof(f4216,plain,(
  ( ! [X47,X48,X49] : (s(bool,X49) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X48),s('type.num.num',X47))) | s(bool,'T_0') = s(bool,X49) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X47))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48)))))))))) )),
  inference(superposition,[],[f817,f459])).
fof(f4208,plain,(
  ( ! [X30,X28,X29] : (s(bool,X30) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29)))))))) )),
  inference(superposition,[],[f817,f458])).
fof(f4035,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f1638,f356])).
fof(f4009,plain,(
  ( ! [X0] : (sP4('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1547,f524])).
fof(f3995,plain,(
  ( ! [X0] : (~sP2('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f612,f625,f621,f1546,f527])).
fof(f3962,plain,(
  ( ! [X0] : (~sP2('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f612,f624,f620,f1544,f528])).
fof(f3773,plain,(
  ( ! [X8] : (s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X8))))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X8)))))))))) )),
  inference(superposition,[],[f358,f1264])).
fof(f3769,plain,(
  ( ! [X2,X3] : (s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X3))) != s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X2))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X3))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X2))))))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X2))))))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X3)))))) )),
  inference(superposition,[],[f470,f1264])).
fof(f3739,plain,(
  ( ! [X19,X17,X18] : (s(bool,X19) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X18))))))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f3725,f612])).
fof(f3725,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,'F_0')) | s(bool,X19) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X18))))))))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f879,f455])).
fof(f3738,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15))))))))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f3724,f612])).
fof(f3724,plain,(
  ( ! [X15,X16] : (p(s(bool,'F_0')) | s(bool,X16) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15))))))))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f879,f364])).
fof(f3736,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f3720,f612])).
fof(f3720,plain,(
  ( ! [X6,X7] : (p(s(bool,'F_0')) | s(bool,X7) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f879,f362])).
fof(f3735,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f3719,f612])).
fof(f3719,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f879,f347])).
fof(f3123,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',X0) != s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X0))))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X0))))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1)))))) )),
  inference(superposition,[],[f470,f353])).
fof(f3122,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))) )),
  inference(unit_resulting_resolution,[],[f353,f470])).
fof(f3109,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK14(X0,X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool))) )),
  inference(superposition,[],[f452,f371])).
fof(f3098,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK14(X0,X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool))) )),
  inference(superposition,[],[f451,f371])).
fof(f3086,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK13(X0,X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK13(X0,X1,bool))) )),
  inference(superposition,[],[f450,f371])).
fof(f3039,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,sK13(X0,X1,bool))))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK13(X0,X1,bool))) )),
  inference(superposition,[],[f449,f371])).
fof(f3019,plain,(
  ( ! [X85,X83,X86,X84] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X83))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))))),X85,X86) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X83))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))))),X85,X86) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X84),s('type.num.num',X83))))) )),
  inference(superposition,[],[f597,f459])).
fof(f3018,plain,(
  ( ! [X80,X81,X79,X82] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X79))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))))),X81,X82) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X79))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))))),X81,X82) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X80),s('type.num.num',X79))))) )),
  inference(superposition,[],[f596,f459])).
fof(f3017,plain,(
  ( ! [X78,X76,X77,X75] : (~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X75))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X76)))))))),X78,X77) | ~p(s(bool,X77)) | p(s(bool,X78)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X76),s('type.num.num',X75))))) )),
  inference(superposition,[],[f593,f459])).
fof(f3016,plain,(
  ( ! [X74,X72,X71,X73] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X71))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X72)))))))),X73,X74) | ~p(s(bool,X73)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X72),s('type.num.num',X71))))) )),
  inference(superposition,[],[f592,f459])).
fof(f3015,plain,(
  ( ! [X70,X68,X69,X67] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X67))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X68)))))))),X70,X69) | p(s(bool,X69)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X68),s('type.num.num',X67))))) )),
  inference(superposition,[],[f591,f459])).
fof(f3014,plain,(
  ( ! [X66,X64,X65,X63] : (~sP10(X66,X65,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X63))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X64))))))))) | p(s(bool,X65)) | ~p(s(bool,X66)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X64),s('type.num.num',X63))))) )),
  inference(superposition,[],[f586,f459])).
fof(f3013,plain,(
  ( ! [X61,X59,X62,X60] : (~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X59))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X60)))))))),X62,X61) | p(s(bool,X61)) | p(s(bool,X62)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X60),s('type.num.num',X59))))) )),
  inference(superposition,[],[f576,f459])).
fof(f3012,plain,(
  ( ! [X57,X58,X56,X55] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X55))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56)))))))),X57,X58) | ~p(s(bool,X57)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X56),s('type.num.num',X55))))) )),
  inference(superposition,[],[f575,f459])).
fof(f3011,plain,(
  ( ! [X54,X52,X53,X51] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X52)))))))),X54,X53) | ~p(s(bool,X53)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X52),s('type.num.num',X51))))) )),
  inference(superposition,[],[f574,f459])).
fof(f3010,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X48),s('type.num.num',X47)))) | p(s(bool,X49)) | p(s(bool,X50)) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X47))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48)))))))),X49,X50)) )),
  inference(superposition,[],[f569,f459])).
fof(f3009,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X44),s('type.num.num',X43)))) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP7(X46,X45,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))))))))) )),
  inference(superposition,[],[f561,f459])).
fof(f3008,plain,(
  ( ! [X39,X41,X42,X40] : (~sP6(X41,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40)))))))),X42) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X40),s('type.num.num',X39))))) )),
  inference(superposition,[],[f552,f459])).
fof(f3007,plain,(
  ( ! [X37,X35,X38,X36] : (~sP3(X37,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36)))))))),X38) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num',X35))))) )),
  inference(superposition,[],[f533,f459])).
fof(f3006,plain,(
  ( ! [X33,X31,X34,X32] : (~sP4(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32)))))))),X34) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num',X31))))) )),
  inference(superposition,[],[f529,f459])).
fof(f3005,plain,(
  ( ! [X30,X28,X29] : (sP1(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29))))))))) | p(s(bool,X30)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X29),s('type.num.num',X28))))) )),
  inference(superposition,[],[f496,f459])).
fof(f3004,plain,(
  ( ! [X26,X27,X25] : (sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26)))))))),X27) | p(s(bool,X27)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num',X25))))) )),
  inference(superposition,[],[f493,f459])).
fof(f3003,plain,(
  ( ! [X24,X23,X22] : (sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23)))))))),X24) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X23),s('type.num.num',X22))))) )),
  inference(superposition,[],[f478,f459])).
fof(f3002,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X20),s('type.num.num',X19))))) )),
  inference(superposition,[],[f477,f459])).
fof(f2993,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))))))))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) )),
  inference(superposition,[],[f459,f371])).
fof(f2987,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f612,f459,f544])).
fof(f2988,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f345,f459,f542])).
fof(f2989,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f612,f459,f544])).
fof(f2990,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f345,f459,f542])).
fof(f2583,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))))))))))) )),
  inference(subsumption_resolution,[],[f2570,f345])).
fof(f2570,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))))))))))) )),
  inference(superposition,[],[f395,f368])).
fof(f2539,plain,(
  ( ! [X0] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f394,f357])).
fof(f2550,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(subsumption_resolution,[],[f2549,f612])).
fof(f2549,plain,(
  p(s(bool,'F_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(forward_demodulation,[],[f2548,f347])).
fof(f2548,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(forward_demodulation,[],[f2547,f361])).
fof(f2547,plain,(
  ( ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f2537,f1725])).
fof(f2537,plain,(
  ( ! [X6] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f394,f361])).
fof(f2526,plain,(
  ( ! [X4,X3] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3))))))) = s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3)))))))))))))) )),
  inference(subsumption_resolution,[],[f2512,f345])).
fof(f2512,plain,(
  ( ! [X4,X3] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3))))))) = s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3)))))))))))))) )),
  inference(superposition,[],[f461,f366])).
fof(f2525,plain,(
  ( ! [X2,X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))))))) )),
  inference(subsumption_resolution,[],[f2511,f345])).
fof(f2511,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) = s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))))))) )),
  inference(superposition,[],[f461,f365])).
fof(f2488,plain,(
  ( ! [X4,X3] : (s('type.integer.int',X4) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3)))))))))))))) )),
  inference(subsumption_resolution,[],[f2474,f345])).
fof(f2474,plain,(
  ( ! [X4,X3] : (~p(s(bool,'T_0')) | s('type.integer.int',X4) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3)))))))))))))) )),
  inference(superposition,[],[f460,f366])).
fof(f2487,plain,(
  ( ! [X2,X1] : (s('type.integer.int',X2) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))))))) )),
  inference(subsumption_resolution,[],[f2473,f345])).
fof(f2473,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s('type.integer.int',X2) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))))))))))) )),
  inference(superposition,[],[f460,f365])).
fof(f2388,plain,(
  ( ! [X0] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f392,f357])).
fof(f2354,plain,(
  ( ! [X85,X83,X86,X84] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X83))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))),X85,X86) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X83))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))),X85,X86) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X83),s('type.num.num',X84))))) )),
  inference(superposition,[],[f597,f458])).
fof(f2353,plain,(
  ( ! [X80,X81,X79,X82] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X79))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))),X81,X82) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X79))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))),X81,X82) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X79),s('type.num.num',X80))))) )),
  inference(superposition,[],[f596,f458])).
fof(f2352,plain,(
  ( ! [X78,X76,X77,X75] : (~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X75))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X76)))))),X78,X77) | ~p(s(bool,X77)) | p(s(bool,X78)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X75),s('type.num.num',X76))))) )),
  inference(superposition,[],[f593,f458])).
fof(f2351,plain,(
  ( ! [X74,X72,X71,X73] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X71))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X72)))))),X73,X74) | ~p(s(bool,X73)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X71),s('type.num.num',X72))))) )),
  inference(superposition,[],[f592,f458])).
fof(f2350,plain,(
  ( ! [X70,X68,X69,X67] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X67))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X68)))))),X70,X69) | p(s(bool,X69)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X67),s('type.num.num',X68))))) )),
  inference(superposition,[],[f591,f458])).
fof(f2349,plain,(
  ( ! [X66,X64,X65,X63] : (~sP10(X66,X65,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X63))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X64))))))) | p(s(bool,X65)) | ~p(s(bool,X66)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X63),s('type.num.num',X64))))) )),
  inference(superposition,[],[f586,f458])).
fof(f2348,plain,(
  ( ! [X61,X59,X62,X60] : (~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X59))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X60)))))),X62,X61) | p(s(bool,X61)) | p(s(bool,X62)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X59),s('type.num.num',X60))))) )),
  inference(superposition,[],[f576,f458])).
fof(f2347,plain,(
  ( ! [X57,X58,X56,X55] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X55))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56)))))),X57,X58) | ~p(s(bool,X57)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X55),s('type.num.num',X56))))) )),
  inference(superposition,[],[f575,f458])).
fof(f2346,plain,(
  ( ! [X54,X52,X53,X51] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X52)))))),X54,X53) | ~p(s(bool,X53)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X51),s('type.num.num',X52))))) )),
  inference(superposition,[],[f574,f458])).
fof(f2345,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X47),s('type.num.num',X48)))) | p(s(bool,X49)) | p(s(bool,X50)) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X47))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48)))))),X49,X50)) )),
  inference(superposition,[],[f569,f458])).
fof(f2344,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X43),s('type.num.num',X44)))) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP7(X46,X45,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))))))) )),
  inference(superposition,[],[f561,f458])).
fof(f2343,plain,(
  ( ! [X39,X41,X42,X40] : (~sP6(X41,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40)))))),X42) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X39),s('type.num.num',X40))))) )),
  inference(superposition,[],[f552,f458])).
fof(f2342,plain,(
  ( ! [X37,X35,X38,X36] : (~sP3(X37,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36)))))),X38) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X35),s('type.num.num',X36))))) )),
  inference(superposition,[],[f533,f458])).
fof(f2341,plain,(
  ( ! [X33,X31,X34,X32] : (~sP4(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32)))))),X34) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X31),s('type.num.num',X32))))) )),
  inference(superposition,[],[f529,f458])).
fof(f2340,plain,(
  ( ! [X30,X28,X29] : (sP1(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29))))))) | p(s(bool,X30)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X29))))) )),
  inference(superposition,[],[f496,f458])).
fof(f2339,plain,(
  ( ! [X26,X27,X25] : (sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26)))))),X27) | p(s(bool,X27)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X25),s('type.num.num',X26))))) )),
  inference(superposition,[],[f493,f458])).
fof(f2338,plain,(
  ( ! [X24,X23,X22] : (sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23)))))),X24) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num',X23))))) )),
  inference(superposition,[],[f478,f458])).
fof(f2337,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num',X20))))) )),
  inference(superposition,[],[f477,f458])).
fof(f2328,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))))))) | s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3)))) )),
  inference(superposition,[],[f458,f371])).
fof(f2324,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f612,f458,f544])).
fof(f2325,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f345,f458,f542])).
fof(f2326,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f612,f458,f544])).
fof(f2327,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f345,f458,f542])).
fof(f2286,plain,(
  ( ! [X85,X83,X86,X84] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X83))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))),X85,X86) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X83))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))),X85,X86)) )),
  inference(subsumption_resolution,[],[f2257,f345])).
fof(f2257,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X83))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))),X85,X86) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X83))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))),X85,X86)) )),
  inference(superposition,[],[f597,f457])).
fof(f2285,plain,(
  ( ! [X80,X81,X79,X82] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X79))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))),X81,X82) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X79))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))),X81,X82)) )),
  inference(subsumption_resolution,[],[f2256,f345])).
fof(f2256,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X79))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))),X81,X82) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X79))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))),X81,X82)) )),
  inference(superposition,[],[f596,f457])).
fof(f2284,plain,(
  ( ! [X78,X76,X77,X75] : (~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X75))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X76)))))),X78,X77) | p(s(bool,X78)) | ~p(s(bool,X77))) )),
  inference(subsumption_resolution,[],[f2255,f345])).
fof(f2255,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X77)) | p(s(bool,X78)) | ~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X75))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X76)))))),X78,X77)) )),
  inference(superposition,[],[f593,f457])).
fof(f2283,plain,(
  ( ! [X74,X72,X71,X73] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X71))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X72)))))),X73,X74) | ~p(s(bool,X73))) )),
  inference(subsumption_resolution,[],[f2254,f345])).
fof(f2254,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X73)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X71))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X72)))))),X73,X74)) )),
  inference(superposition,[],[f592,f457])).
fof(f2282,plain,(
  ( ! [X70,X68,X69,X67] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X67))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X68)))))),X70,X69) | p(s(bool,X69))) )),
  inference(subsumption_resolution,[],[f2253,f345])).
fof(f2253,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X69)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X67))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X68)))))),X70,X69)) )),
  inference(superposition,[],[f591,f457])).
fof(f2281,plain,(
  ( ! [X66,X64,X65,X63] : (~sP10(X66,X65,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X63))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X64))))))) | ~p(s(bool,X66)) | p(s(bool,X65))) )),
  inference(subsumption_resolution,[],[f2252,f345])).
fof(f2252,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X65)) | ~p(s(bool,X66)) | ~sP10(X66,X65,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X63))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X64)))))))) )),
  inference(superposition,[],[f586,f457])).
fof(f2280,plain,(
  ( ! [X61,X59,X62,X60] : (~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X59))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X60)))))),X62,X61) | p(s(bool,X62)) | p(s(bool,X61))) )),
  inference(subsumption_resolution,[],[f2251,f345])).
fof(f2251,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | p(s(bool,X61)) | p(s(bool,X62)) | ~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X59))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X60)))))),X62,X61)) )),
  inference(superposition,[],[f576,f457])).
fof(f2279,plain,(
  ( ! [X57,X58,X56,X55] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X55))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56)))))),X57,X58) | ~p(s(bool,X57))) )),
  inference(subsumption_resolution,[],[f2250,f345])).
fof(f2250,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X55))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56)))))),X57,X58)) )),
  inference(superposition,[],[f575,f457])).
fof(f2278,plain,(
  ( ! [X54,X52,X53,X51] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X51))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X52)))))),X54,X53) | ~p(s(bool,X53))) )),
  inference(subsumption_resolution,[],[f2249,f345])).
fof(f2249,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X53)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X51))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X52)))))),X54,X53)) )),
  inference(superposition,[],[f574,f457])).
fof(f2277,plain,(
  ( ! [X47,X50,X48,X49] : (~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X47))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48)))))),X49,X50) | p(s(bool,X50)) | p(s(bool,X49))) )),
  inference(subsumption_resolution,[],[f2248,f345])).
fof(f2248,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X49)) | p(s(bool,X50)) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X47))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48)))))),X49,X50)) )),
  inference(superposition,[],[f569,f457])).
fof(f2276,plain,(
  ( ! [X45,X43,X46,X44] : (~sP7(X46,X45,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X43))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))))))) | p(s(bool,X46)) | ~p(s(bool,X45))) )),
  inference(subsumption_resolution,[],[f2247,f345])).
fof(f2247,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP7(X46,X45,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X43))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))))))) )),
  inference(superposition,[],[f561,f457])).
fof(f2275,plain,(
  ( ! [X39,X41,X42,X40] : (~sP6(X41,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X39))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40)))))),X42) | p(s(bool,X42)) | ~p(s(bool,X41))) )),
  inference(subsumption_resolution,[],[f2246,f345])).
fof(f2246,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP6(X41,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X39))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40)))))),X42)) )),
  inference(superposition,[],[f552,f457])).
fof(f2274,plain,(
  ( ! [X37,X35,X38,X36] : (~sP3(X37,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X35))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36)))))),X38) | p(s(bool,X38)) | ~p(s(bool,X37))) )),
  inference(subsumption_resolution,[],[f2245,f345])).
fof(f2245,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP3(X37,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X35))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36)))))),X38)) )),
  inference(superposition,[],[f533,f457])).
fof(f2273,plain,(
  ( ! [X33,X31,X34,X32] : (~sP4(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X31))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32)))))),X34) | p(s(bool,X34)) | ~p(s(bool,X33))) )),
  inference(subsumption_resolution,[],[f2244,f345])).
fof(f2244,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~sP4(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X31))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32)))))),X34)) )),
  inference(superposition,[],[f529,f457])).
fof(f2272,plain,(
  ( ! [X30,X28,X29] : (sP1(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X28))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29))))))) | p(s(bool,X30))) )),
  inference(subsumption_resolution,[],[f2243,f345])).
fof(f2243,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP1(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X28))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29)))))))) )),
  inference(superposition,[],[f496,f457])).
fof(f2271,plain,(
  ( ! [X26,X27,X25] : (sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X25))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26)))))),X27) | p(s(bool,X27))) )),
  inference(subsumption_resolution,[],[f2242,f345])).
fof(f2242,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X25))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26)))))),X27)) )),
  inference(superposition,[],[f493,f457])).
fof(f2270,plain,(
  ( ! [X24,X23,X22] : (sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X22))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23)))))),X24)) )),
  inference(subsumption_resolution,[],[f2241,f345])).
fof(f2241,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X22))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23)))))),X24)) )),
  inference(superposition,[],[f478,f457])).
fof(f2269,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X19))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20)))))))) )),
  inference(subsumption_resolution,[],[f2240,f345])).
fof(f2240,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X21,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X19))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20)))))))) )),
  inference(superposition,[],[f477,f457])).
fof(f2227,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f612,f457,f544])).
fof(f2228,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f345,f457,f542])).
fof(f2229,plain,(
  ( ! [X0,X1] : (~sP5('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f612,f457,f544])).
fof(f2230,plain,(
  ( ! [X0,X1] : (sP5('T_0','T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f345,f457,f542])).
fof(f2217,plain,(
  ( ! [X85,X83,X86,X84] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X83))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))),X85,X86) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X83))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))),X85,X86)) )),
  inference(subsumption_resolution,[],[f2188,f345])).
fof(f2188,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X83))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))),X85,X86) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X83))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X84)))))),X85,X86)) )),
  inference(superposition,[],[f597,f456])).
fof(f2216,plain,(
  ( ! [X80,X81,X79,X82] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X79))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))),X81,X82) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X79))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))),X81,X82)) )),
  inference(subsumption_resolution,[],[f2187,f345])).
fof(f2187,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X79))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))),X81,X82) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X79))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X80)))))),X81,X82)) )),
  inference(superposition,[],[f596,f456])).
fof(f2215,plain,(
  ( ! [X78,X76,X77,X75] : (~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X75))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X76)))))),X78,X77) | p(s(bool,X78)) | ~p(s(bool,X77))) )),
  inference(subsumption_resolution,[],[f2186,f345])).
fof(f2186,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X77)) | p(s(bool,X78)) | ~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X75))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X76)))))),X78,X77)) )),
  inference(superposition,[],[f593,f456])).
fof(f2214,plain,(
  ( ! [X74,X72,X71,X73] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X71))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X72)))))),X73,X74) | ~p(s(bool,X73))) )),
  inference(subsumption_resolution,[],[f2185,f345])).
fof(f2185,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X73)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X71))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X72)))))),X73,X74)) )),
  inference(superposition,[],[f592,f456])).
fof(f2213,plain,(
  ( ! [X70,X68,X69,X67] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X67))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X68)))))),X70,X69) | p(s(bool,X69))) )),
  inference(subsumption_resolution,[],[f2184,f345])).
fof(f2184,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X69)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X67))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X68)))))),X70,X69)) )),
  inference(superposition,[],[f591,f456])).
fof(f2212,plain,(
  ( ! [X66,X64,X65,X63] : (~sP10(X66,X65,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X63))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X64))))))) | ~p(s(bool,X66)) | p(s(bool,X65))) )),
  inference(subsumption_resolution,[],[f2183,f345])).
fof(f2183,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X65)) | ~p(s(bool,X66)) | ~sP10(X66,X65,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X63))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X64)))))))) )),
  inference(superposition,[],[f586,f456])).
fof(f2211,plain,(
  ( ! [X61,X59,X62,X60] : (~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X59))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X60)))))),X62,X61) | p(s(bool,X62)) | p(s(bool,X61))) )),
  inference(subsumption_resolution,[],[f2182,f345])).
fof(f2182,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | p(s(bool,X61)) | p(s(bool,X62)) | ~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X59))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X60)))))),X62,X61)) )),
  inference(superposition,[],[f576,f456])).
fof(f2210,plain,(
  ( ! [X57,X58,X56,X55] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X55))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56)))))),X57,X58) | ~p(s(bool,X57))) )),
  inference(subsumption_resolution,[],[f2181,f345])).
fof(f2181,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X57)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X55))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56)))))),X57,X58)) )),
  inference(superposition,[],[f575,f456])).
fof(f2209,plain,(
  ( ! [X54,X52,X53,X51] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X51))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X52)))))),X54,X53) | ~p(s(bool,X53))) )),
  inference(subsumption_resolution,[],[f2180,f345])).
fof(f2180,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X53)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X51))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X52)))))),X54,X53)) )),
  inference(superposition,[],[f574,f456])).
fof(f2208,plain,(
  ( ! [X47,X50,X48,X49] : (~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X47))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48)))))),X49,X50) | p(s(bool,X50)) | p(s(bool,X49))) )),
  inference(subsumption_resolution,[],[f2179,f345])).
fof(f2179,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X49)) | p(s(bool,X50)) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X47))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48)))))),X49,X50)) )),
  inference(superposition,[],[f569,f456])).
fof(f2207,plain,(
  ( ! [X45,X43,X46,X44] : (~sP7(X46,X45,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X43))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))))))) | p(s(bool,X46)) | ~p(s(bool,X45))) )),
  inference(subsumption_resolution,[],[f2178,f345])).
fof(f2178,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP7(X46,X45,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X43))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))))))) )),
  inference(superposition,[],[f561,f456])).
fof(f2206,plain,(
  ( ! [X39,X41,X42,X40] : (~sP6(X41,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X39))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40)))))),X42) | p(s(bool,X42)) | ~p(s(bool,X41))) )),
  inference(subsumption_resolution,[],[f2177,f345])).
fof(f2177,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP6(X41,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X39))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40)))))),X42)) )),
  inference(superposition,[],[f552,f456])).
fof(f2205,plain,(
  ( ! [X37,X35,X38,X36] : (~sP3(X37,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X35))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36)))))),X38) | p(s(bool,X38)) | ~p(s(bool,X37))) )),
  inference(subsumption_resolution,[],[f2176,f345])).
fof(f2176,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP3(X37,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X35))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36)))))),X38)) )),
  inference(superposition,[],[f533,f456])).
fof(f2204,plain,(
  ( ! [X33,X31,X34,X32] : (~sP4(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X31))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32)))))),X34) | p(s(bool,X34)) | ~p(s(bool,X33))) )),
  inference(subsumption_resolution,[],[f2175,f345])).
fof(f2175,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~sP4(X33,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X31))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32)))))),X34)) )),
  inference(superposition,[],[f529,f456])).
fof(f2203,plain,(
  ( ! [X30,X28,X29] : (sP1(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X28))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29))))))) | p(s(bool,X30))) )),
  inference(subsumption_resolution,[],[f2174,f345])).
fof(f2174,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP1(X30,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X28))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29)))))))) )),
  inference(superposition,[],[f496,f456])).
fof(f2202,plain,(
  ( ! [X26,X27,X25] : (sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X25))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26)))))),X27) | p(s(bool,X27))) )),
  inference(subsumption_resolution,[],[f2173,f345])).
fof(f2173,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X25))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26)))))),X27)) )),
  inference(superposition,[],[f493,f456])).
fof(f2201,plain,(
  ( ! [X24,X23,X22] : (sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X22))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23)))))),X24)) )),
  inference(subsumption_resolution,[],[f2172,f345])).
fof(f2172,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X22))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23)))))),X24)) )),
  inference(superposition,[],[f478,f456])).
fof(f2200,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X19))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20)))))))) )),
  inference(subsumption_resolution,[],[f2171,f345])).
fof(f2171,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X21,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X19))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20)))))))) )),
  inference(superposition,[],[f477,f456])).
fof(f2158,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f612,f456,f544])).
fof(f2159,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f345,f456,f542])).
fof(f2160,plain,(
  ( ! [X0,X1] : (~sP5('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f612,f456,f544])).
fof(f2161,plain,(
  ( ! [X0,X1] : (sP5('T_0','T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f345,f456,f542])).
fof(f2115,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f612,f455,f544])).
fof(f2116,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f345,f455,f542])).
fof(f2117,plain,(
  ( ! [X0,X1] : (~sP5('F_0','F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))) )),
  inference(unit_resulting_resolution,[],[f612,f455,f544])).
fof(f2118,plain,(
  ( ! [X0,X1] : (sP5('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))) )),
  inference(unit_resulting_resolution,[],[f345,f455,f542])).
fof(f2103,plain,(
  ( ! [X59,X60,X58] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X58))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X59,X60) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X58))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X59,X60)) )),
  inference(subsumption_resolution,[],[f2074,f345])).
fof(f2074,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'T_0')) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X58))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X59,X60) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X58))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X59,X60)) )),
  inference(superposition,[],[f597,f368])).
fof(f2102,plain,(
  ( ! [X57,X56,X55] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X55))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X56,X57) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X55))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X56,X57)) )),
  inference(subsumption_resolution,[],[f2073,f345])).
fof(f2073,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,'T_0')) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X55))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X56,X57) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X55))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X56,X57)) )),
  inference(superposition,[],[f596,f368])).
fof(f2101,plain,(
  ( ! [X54,X52,X53] : (~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X52))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X54,X53) | p(s(bool,X54)) | ~p(s(bool,X53))) )),
  inference(subsumption_resolution,[],[f2072,f345])).
fof(f2072,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X52))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X54,X53)) )),
  inference(superposition,[],[f593,f368])).
fof(f2100,plain,(
  ( ! [X50,X51,X49] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X49))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X50,X51) | ~p(s(bool,X50))) )),
  inference(subsumption_resolution,[],[f2071,f345])).
fof(f2071,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X49))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X50,X51)) )),
  inference(superposition,[],[f592,f368])).
fof(f2099,plain,(
  ( ! [X47,X48,X46] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X46))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X48,X47) | p(s(bool,X47))) )),
  inference(subsumption_resolution,[],[f2070,f345])).
fof(f2070,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'T_0')) | p(s(bool,X47)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X46))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X48,X47)) )),
  inference(superposition,[],[f591,f368])).
fof(f2098,plain,(
  ( ! [X45,X43,X44] : (~sP10(X45,X44,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X43))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X45)) | p(s(bool,X44))) )),
  inference(subsumption_resolution,[],[f2069,f345])).
fof(f2069,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'T_0')) | p(s(bool,X44)) | ~p(s(bool,X45)) | ~sP10(X45,X44,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X43))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f586,f368])).
fof(f2097,plain,(
  ( ! [X41,X42,X40] : (~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X40))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X42,X41) | p(s(bool,X42)) | p(s(bool,X41))) )),
  inference(subsumption_resolution,[],[f2068,f345])).
fof(f2068,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,X41)) | p(s(bool,X42)) | ~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X40))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X42,X41)) )),
  inference(superposition,[],[f576,f368])).
fof(f2096,plain,(
  ( ! [X39,X37,X38] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X37))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X38,X39) | ~p(s(bool,X38))) )),
  inference(subsumption_resolution,[],[f2067,f345])).
fof(f2067,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'T_0')) | ~p(s(bool,X38)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X37))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X38,X39)) )),
  inference(superposition,[],[f575,f368])).
fof(f2095,plain,(
  ( ! [X35,X36,X34] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X34))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X36,X35) | ~p(s(bool,X35))) )),
  inference(subsumption_resolution,[],[f2066,f345])).
fof(f2066,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X34))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X36,X35)) )),
  inference(superposition,[],[f574,f368])).
fof(f2094,plain,(
  ( ! [X33,X31,X32] : (~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X31))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X32,X33) | p(s(bool,X33)) | p(s(bool,X32))) )),
  inference(subsumption_resolution,[],[f2065,f345])).
fof(f2065,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X32)) | p(s(bool,X33)) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X31))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X32,X33)) )),
  inference(superposition,[],[f569,f368])).
fof(f2093,plain,(
  ( ! [X30,X28,X29] : (~sP7(X30,X29,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X28))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X30)) | ~p(s(bool,X29))) )),
  inference(subsumption_resolution,[],[f2064,f345])).
fof(f2064,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~sP7(X30,X29,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X28))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f561,f368])).
fof(f2092,plain,(
  ( ! [X26,X27,X25] : (~sP6(X26,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X25))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X27) | p(s(bool,X27)) | ~p(s(bool,X26))) )),
  inference(subsumption_resolution,[],[f2063,f345])).
fof(f2063,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | p(s(bool,X27)) | ~sP6(X26,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X25))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X27)) )),
  inference(superposition,[],[f552,f368])).
fof(f2091,plain,(
  ( ! [X24,X23,X22] : (~sP3(X23,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X22))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24) | p(s(bool,X24)) | ~p(s(bool,X23))) )),
  inference(subsumption_resolution,[],[f2062,f345])).
fof(f2062,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | p(s(bool,X24)) | ~sP3(X23,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X22))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24)) )),
  inference(superposition,[],[f533,f368])).
fof(f2090,plain,(
  ( ! [X21,X19,X20] : (~sP4(X20,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X19))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X21) | p(s(bool,X21)) | ~p(s(bool,X20))) )),
  inference(subsumption_resolution,[],[f2061,f345])).
fof(f2061,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | p(s(bool,X21)) | ~sP4(X20,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X19))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X21)) )),
  inference(superposition,[],[f529,f368])).
fof(f2089,plain,(
  ( ! [X17,X18] : (sP1(X18,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X18))) )),
  inference(subsumption_resolution,[],[f2060,f345])).
fof(f2060,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,X18)) | sP1(X18,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f496,f368])).
fof(f2088,plain,(
  ( ! [X15,X16] : (sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X15))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X16) | p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f2059,f345])).
fof(f2059,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | p(s(bool,X16)) | sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X15))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X16)) )),
  inference(superposition,[],[f493,f368])).
fof(f2087,plain,(
  ( ! [X14,X13] : (sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X13))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X14)) )),
  inference(subsumption_resolution,[],[f2058,f345])).
fof(f2058,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X13))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X14)) )),
  inference(superposition,[],[f478,f368])).
fof(f2086,plain,(
  ( ! [X12,X11] : (sP0(X12,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X11))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f2057,f345])).
fof(f2057,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | sP0(X12,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X11))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f477,f368])).
fof(f2044,plain,(
  ( ! [X0] : (~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f612,f368,f544])).
fof(f2045,plain,(
  ( ! [X0] : (sP5('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f345,f368,f542])).
fof(f2046,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f612,f368,f544])).
fof(f2047,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f345,f368,f542])).
fof(f2026,plain,(
  ( ! [X59,X60,X58] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X58))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X59,X60) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X58))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X59,X60)) )),
  inference(subsumption_resolution,[],[f1997,f345])).
fof(f1997,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'T_0')) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X58))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X59,X60) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X58))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X59,X60)) )),
  inference(superposition,[],[f597,f367])).
fof(f2025,plain,(
  ( ! [X57,X56,X55] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X55))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X56,X57) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X55))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X56,X57)) )),
  inference(subsumption_resolution,[],[f1996,f345])).
fof(f1996,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,'T_0')) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X55))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X56,X57) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X55))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X56,X57)) )),
  inference(superposition,[],[f596,f367])).
fof(f2024,plain,(
  ( ! [X54,X52,X53] : (~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X52))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X54,X53) | p(s(bool,X54)) | ~p(s(bool,X53))) )),
  inference(subsumption_resolution,[],[f1995,f345])).
fof(f1995,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X52))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X54,X53)) )),
  inference(superposition,[],[f593,f367])).
fof(f2023,plain,(
  ( ! [X50,X51,X49] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X49))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X50,X51) | ~p(s(bool,X50))) )),
  inference(subsumption_resolution,[],[f1994,f345])).
fof(f1994,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X49))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X50,X51)) )),
  inference(superposition,[],[f592,f367])).
fof(f2022,plain,(
  ( ! [X47,X48,X46] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X46))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X48,X47) | p(s(bool,X47))) )),
  inference(subsumption_resolution,[],[f1993,f345])).
fof(f1993,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'T_0')) | p(s(bool,X47)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X46))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X48,X47)) )),
  inference(superposition,[],[f591,f367])).
fof(f2021,plain,(
  ( ! [X45,X43,X44] : (~sP10(X45,X44,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X43))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,X45)) | p(s(bool,X44))) )),
  inference(subsumption_resolution,[],[f1992,f345])).
fof(f1992,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'T_0')) | p(s(bool,X44)) | ~p(s(bool,X45)) | ~sP10(X45,X44,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X43))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f586,f367])).
fof(f2020,plain,(
  ( ! [X41,X42,X40] : (~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X40))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X42,X41) | p(s(bool,X42)) | p(s(bool,X41))) )),
  inference(subsumption_resolution,[],[f1991,f345])).
fof(f1991,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,X41)) | p(s(bool,X42)) | ~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X40))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X42,X41)) )),
  inference(superposition,[],[f576,f367])).
fof(f2019,plain,(
  ( ! [X39,X37,X38] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X37))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X38,X39) | ~p(s(bool,X38))) )),
  inference(subsumption_resolution,[],[f1990,f345])).
fof(f1990,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'T_0')) | ~p(s(bool,X38)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X37))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X38,X39)) )),
  inference(superposition,[],[f575,f367])).
fof(f2018,plain,(
  ( ! [X35,X36,X34] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X34))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X36,X35) | ~p(s(bool,X35))) )),
  inference(subsumption_resolution,[],[f1989,f345])).
fof(f1989,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X34))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X36,X35)) )),
  inference(superposition,[],[f574,f367])).
fof(f2017,plain,(
  ( ! [X33,X31,X32] : (~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X31))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X32,X33) | p(s(bool,X33)) | p(s(bool,X32))) )),
  inference(subsumption_resolution,[],[f1988,f345])).
fof(f1988,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X32)) | p(s(bool,X33)) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X31))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X32,X33)) )),
  inference(superposition,[],[f569,f367])).
fof(f2016,plain,(
  ( ! [X30,X28,X29] : (~sP7(X30,X29,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X28))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X30)) | ~p(s(bool,X29))) )),
  inference(subsumption_resolution,[],[f1987,f345])).
fof(f1987,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~sP7(X30,X29,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X28))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f561,f367])).
fof(f2015,plain,(
  ( ! [X26,X27,X25] : (~sP6(X26,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X25))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X27) | p(s(bool,X27)) | ~p(s(bool,X26))) )),
  inference(subsumption_resolution,[],[f1986,f345])).
fof(f1986,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | p(s(bool,X27)) | ~sP6(X26,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X25))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X27)) )),
  inference(superposition,[],[f552,f367])).
fof(f2014,plain,(
  ( ! [X24,X23,X22] : (~sP3(X23,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X22))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24) | p(s(bool,X24)) | ~p(s(bool,X23))) )),
  inference(subsumption_resolution,[],[f1985,f345])).
fof(f1985,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | p(s(bool,X24)) | ~sP3(X23,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X22))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X24)) )),
  inference(superposition,[],[f533,f367])).
fof(f2013,plain,(
  ( ! [X21,X19,X20] : (~sP4(X20,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X19))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X21) | p(s(bool,X21)) | ~p(s(bool,X20))) )),
  inference(subsumption_resolution,[],[f1984,f345])).
fof(f1984,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | p(s(bool,X21)) | ~sP4(X20,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X19))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X21)) )),
  inference(superposition,[],[f529,f367])).
fof(f2012,plain,(
  ( ! [X17,X18] : (sP1(X18,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X17))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X18))) )),
  inference(subsumption_resolution,[],[f1983,f345])).
fof(f1983,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,X18)) | sP1(X18,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X17))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f496,f367])).
fof(f2011,plain,(
  ( ! [X15,X16] : (sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X15))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X16) | p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f1982,f345])).
fof(f1982,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | p(s(bool,X16)) | sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X15))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X16)) )),
  inference(superposition,[],[f493,f367])).
fof(f2010,plain,(
  ( ! [X14,X13] : (sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X14)) )),
  inference(subsumption_resolution,[],[f1981,f345])).
fof(f1981,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),X14)) )),
  inference(superposition,[],[f478,f367])).
fof(f2009,plain,(
  ( ! [X12,X11] : (sP0(X12,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X11))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f1980,f345])).
fof(f1980,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | sP0(X12,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X11))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f477,f367])).
fof(f1967,plain,(
  ( ! [X0] : (~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f612,f367,f544])).
fof(f1968,plain,(
  ( ! [X0] : (sP5('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f345,f367,f542])).
fof(f1969,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f612,f367,f544])).
fof(f1970,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f345,f367,f542])).
fof(f1949,plain,(
  ( ! [X59,X60,X58] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X58)))))))),X59,X60) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X58)))))))),X59,X60)) )),
  inference(subsumption_resolution,[],[f1920,f345])).
fof(f1920,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'T_0')) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X58)))))))),X59,X60) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X58)))))))),X59,X60)) )),
  inference(superposition,[],[f597,f366])).
fof(f1948,plain,(
  ( ! [X57,X56,X55] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X55)))))))),X56,X57) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X55)))))))),X56,X57)) )),
  inference(subsumption_resolution,[],[f1919,f345])).
fof(f1919,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,'T_0')) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X55)))))))),X56,X57) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X55)))))))),X56,X57)) )),
  inference(superposition,[],[f596,f366])).
fof(f1947,plain,(
  ( ! [X54,X52,X53] : (~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X52)))))))),X54,X53) | p(s(bool,X54)) | ~p(s(bool,X53))) )),
  inference(subsumption_resolution,[],[f1918,f345])).
fof(f1918,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X52)))))))),X54,X53)) )),
  inference(superposition,[],[f593,f366])).
fof(f1946,plain,(
  ( ! [X50,X51,X49] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X49)))))))),X50,X51) | ~p(s(bool,X50))) )),
  inference(subsumption_resolution,[],[f1917,f345])).
fof(f1917,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X49)))))))),X50,X51)) )),
  inference(superposition,[],[f592,f366])).
fof(f1945,plain,(
  ( ! [X47,X48,X46] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X46)))))))),X48,X47) | p(s(bool,X47))) )),
  inference(subsumption_resolution,[],[f1916,f345])).
fof(f1916,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'T_0')) | p(s(bool,X47)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X46)))))))),X48,X47)) )),
  inference(superposition,[],[f591,f366])).
fof(f1944,plain,(
  ( ! [X45,X43,X44] : (~sP10(X45,X44,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X43))))))))) | ~p(s(bool,X45)) | p(s(bool,X44))) )),
  inference(subsumption_resolution,[],[f1915,f345])).
fof(f1915,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'T_0')) | p(s(bool,X44)) | ~p(s(bool,X45)) | ~sP10(X45,X44,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X43)))))))))) )),
  inference(superposition,[],[f586,f366])).
fof(f1943,plain,(
  ( ! [X41,X42,X40] : (~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X40)))))))),X42,X41) | p(s(bool,X42)) | p(s(bool,X41))) )),
  inference(subsumption_resolution,[],[f1914,f345])).
fof(f1914,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,X41)) | p(s(bool,X42)) | ~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X40)))))))),X42,X41)) )),
  inference(superposition,[],[f576,f366])).
fof(f1942,plain,(
  ( ! [X39,X37,X38] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X37)))))))),X38,X39) | ~p(s(bool,X38))) )),
  inference(subsumption_resolution,[],[f1913,f345])).
fof(f1913,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'T_0')) | ~p(s(bool,X38)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X37)))))))),X38,X39)) )),
  inference(superposition,[],[f575,f366])).
fof(f1941,plain,(
  ( ! [X35,X36,X34] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X34)))))))),X36,X35) | ~p(s(bool,X35))) )),
  inference(subsumption_resolution,[],[f1912,f345])).
fof(f1912,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X34)))))))),X36,X35)) )),
  inference(superposition,[],[f574,f366])).
fof(f1940,plain,(
  ( ! [X33,X31,X32] : (~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X31)))))))),X32,X33) | p(s(bool,X33)) | p(s(bool,X32))) )),
  inference(subsumption_resolution,[],[f1911,f345])).
fof(f1911,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X32)) | p(s(bool,X33)) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X31)))))))),X32,X33)) )),
  inference(superposition,[],[f569,f366])).
fof(f1939,plain,(
  ( ! [X30,X28,X29] : (~sP7(X30,X29,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X28))))))))) | p(s(bool,X30)) | ~p(s(bool,X29))) )),
  inference(subsumption_resolution,[],[f1910,f345])).
fof(f1910,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~sP7(X30,X29,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X28)))))))))) )),
  inference(superposition,[],[f561,f366])).
fof(f1938,plain,(
  ( ! [X26,X27,X25] : (~sP6(X26,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X25)))))))),X27) | p(s(bool,X27)) | ~p(s(bool,X26))) )),
  inference(subsumption_resolution,[],[f1909,f345])).
fof(f1909,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | p(s(bool,X27)) | ~sP6(X26,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X25)))))))),X27)) )),
  inference(superposition,[],[f552,f366])).
fof(f1937,plain,(
  ( ! [X24,X23,X22] : (~sP3(X23,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X22)))))))),X24) | p(s(bool,X24)) | ~p(s(bool,X23))) )),
  inference(subsumption_resolution,[],[f1908,f345])).
fof(f1908,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | p(s(bool,X24)) | ~sP3(X23,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X22)))))))),X24)) )),
  inference(superposition,[],[f533,f366])).
fof(f1936,plain,(
  ( ! [X21,X19,X20] : (~sP4(X20,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X19)))))))),X21) | p(s(bool,X21)) | ~p(s(bool,X20))) )),
  inference(subsumption_resolution,[],[f1907,f345])).
fof(f1907,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | p(s(bool,X21)) | ~sP4(X20,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X19)))))))),X21)) )),
  inference(superposition,[],[f529,f366])).
fof(f1935,plain,(
  ( ! [X17,X18] : (sP1(X18,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17))))))))) | p(s(bool,X18))) )),
  inference(subsumption_resolution,[],[f1906,f345])).
fof(f1906,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,X18)) | sP1(X18,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17)))))))))) )),
  inference(superposition,[],[f496,f366])).
fof(f1934,plain,(
  ( ! [X15,X16] : (sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X15)))))))),X16) | p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f1905,f345])).
fof(f1905,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | p(s(bool,X16)) | sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X15)))))))),X16)) )),
  inference(superposition,[],[f493,f366])).
fof(f1933,plain,(
  ( ! [X14,X13] : (sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X13)))))))),X14)) )),
  inference(subsumption_resolution,[],[f1904,f345])).
fof(f1904,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X13)))))))),X14)) )),
  inference(superposition,[],[f478,f366])).
fof(f1932,plain,(
  ( ! [X12,X11] : (sP0(X12,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X11)))))))))) )),
  inference(subsumption_resolution,[],[f1903,f345])).
fof(f1903,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | sP0(X12,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X11)))))))))) )),
  inference(superposition,[],[f477,f366])).
fof(f1890,plain,(
  ( ! [X0] : (~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f612,f366,f544])).
fof(f1891,plain,(
  ( ! [X0] : (sP5('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f345,f366,f542])).
fof(f1892,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f612,f366,f544])).
fof(f1893,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f345,f366,f542])).
fof(f1874,plain,(
  ( ! [X59,X60,X58] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X58)))))))),X59,X60) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X58)))))))),X59,X60)) )),
  inference(subsumption_resolution,[],[f1845,f345])).
fof(f1845,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'T_0')) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X58)))))))),X59,X60) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X58)))))))),X59,X60)) )),
  inference(superposition,[],[f597,f365])).
fof(f1873,plain,(
  ( ! [X57,X56,X55] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X55)))))))),X56,X57) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X55)))))))),X56,X57)) )),
  inference(subsumption_resolution,[],[f1844,f345])).
fof(f1844,plain,(
  ( ! [X57,X56,X55] : (~p(s(bool,'T_0')) | ~sP10('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X55)))))))),X56,X57) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X55)))))))),X56,X57)) )),
  inference(superposition,[],[f596,f365])).
fof(f1872,plain,(
  ( ! [X54,X52,X53] : (~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X52)))))))),X54,X53) | p(s(bool,X54)) | ~p(s(bool,X53))) )),
  inference(subsumption_resolution,[],[f1843,f345])).
fof(f1843,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X52)))))))),X54,X53)) )),
  inference(superposition,[],[f593,f365])).
fof(f1871,plain,(
  ( ! [X50,X51,X49] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X49)))))))),X50,X51) | ~p(s(bool,X50))) )),
  inference(subsumption_resolution,[],[f1842,f345])).
fof(f1842,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X49)))))))),X50,X51)) )),
  inference(superposition,[],[f592,f365])).
fof(f1870,plain,(
  ( ! [X47,X48,X46] : (sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X46)))))))),X48,X47) | p(s(bool,X47))) )),
  inference(subsumption_resolution,[],[f1841,f345])).
fof(f1841,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'T_0')) | p(s(bool,X47)) | sP11('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X46)))))))),X48,X47)) )),
  inference(superposition,[],[f591,f365])).
fof(f1869,plain,(
  ( ! [X45,X43,X44] : (~sP10(X45,X44,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X43))))))))) | ~p(s(bool,X45)) | p(s(bool,X44))) )),
  inference(subsumption_resolution,[],[f1840,f345])).
fof(f1840,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'T_0')) | p(s(bool,X44)) | ~p(s(bool,X45)) | ~sP10(X45,X44,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X43)))))))))) )),
  inference(superposition,[],[f586,f365])).
fof(f1868,plain,(
  ( ! [X41,X42,X40] : (~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X40)))))))),X42,X41) | p(s(bool,X42)) | p(s(bool,X41))) )),
  inference(subsumption_resolution,[],[f1839,f345])).
fof(f1839,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,X41)) | p(s(bool,X42)) | ~sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X40)))))))),X42,X41)) )),
  inference(superposition,[],[f576,f365])).
fof(f1867,plain,(
  ( ! [X39,X37,X38] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X37)))))))),X38,X39) | ~p(s(bool,X38))) )),
  inference(subsumption_resolution,[],[f1838,f345])).
fof(f1838,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'T_0')) | ~p(s(bool,X38)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X37)))))))),X38,X39)) )),
  inference(superposition,[],[f575,f365])).
fof(f1866,plain,(
  ( ! [X35,X36,X34] : (sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X34)))))))),X36,X35) | ~p(s(bool,X35))) )),
  inference(subsumption_resolution,[],[f1837,f345])).
fof(f1837,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | sP9('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X34)))))))),X36,X35)) )),
  inference(superposition,[],[f574,f365])).
fof(f1865,plain,(
  ( ! [X33,X31,X32] : (~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X31)))))))),X32,X33) | p(s(bool,X33)) | p(s(bool,X32))) )),
  inference(subsumption_resolution,[],[f1836,f345])).
fof(f1836,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X32)) | p(s(bool,X33)) | ~sP8('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X31)))))))),X32,X33)) )),
  inference(superposition,[],[f569,f365])).
fof(f1864,plain,(
  ( ! [X30,X28,X29] : (~sP7(X30,X29,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X28))))))))) | p(s(bool,X30)) | ~p(s(bool,X29))) )),
  inference(subsumption_resolution,[],[f1835,f345])).
fof(f1835,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | p(s(bool,X30)) | ~sP7(X30,X29,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X28)))))))))) )),
  inference(superposition,[],[f561,f365])).
fof(f1863,plain,(
  ( ! [X26,X27,X25] : (~sP6(X26,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X25)))))))),X27) | p(s(bool,X27)) | ~p(s(bool,X26))) )),
  inference(subsumption_resolution,[],[f1834,f345])).
fof(f1834,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | p(s(bool,X27)) | ~sP6(X26,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X25)))))))),X27)) )),
  inference(superposition,[],[f552,f365])).
fof(f1862,plain,(
  ( ! [X24,X23,X22] : (~sP3(X23,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X22)))))))),X24) | p(s(bool,X24)) | ~p(s(bool,X23))) )),
  inference(subsumption_resolution,[],[f1833,f345])).
fof(f1833,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | p(s(bool,X24)) | ~sP3(X23,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X22)))))))),X24)) )),
  inference(superposition,[],[f533,f365])).
fof(f1861,plain,(
  ( ! [X21,X19,X20] : (~sP4(X20,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X19)))))))),X21) | p(s(bool,X21)) | ~p(s(bool,X20))) )),
  inference(subsumption_resolution,[],[f1832,f345])).
fof(f1832,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,X20)) | p(s(bool,X21)) | ~sP4(X20,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X19)))))))),X21)) )),
  inference(superposition,[],[f529,f365])).
fof(f1860,plain,(
  ( ! [X17,X18] : (sP1(X18,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X17))))))))) | p(s(bool,X18))) )),
  inference(subsumption_resolution,[],[f1831,f345])).
fof(f1831,plain,(
  ( ! [X17,X18] : (~p(s(bool,'T_0')) | p(s(bool,X18)) | sP1(X18,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X17)))))))))) )),
  inference(superposition,[],[f496,f365])).
fof(f1859,plain,(
  ( ! [X15,X16] : (sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X15)))))))),X16) | p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f1830,f345])).
fof(f1830,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | p(s(bool,X16)) | sP1('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X15)))))))),X16)) )),
  inference(superposition,[],[f493,f365])).
fof(f1858,plain,(
  ( ! [X14,X13] : (sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13)))))))),X14)) )),
  inference(subsumption_resolution,[],[f1829,f345])).
fof(f1829,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | sP0('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13)))))))),X14)) )),
  inference(superposition,[],[f478,f365])).
fof(f1857,plain,(
  ( ! [X12,X11] : (sP0(X12,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X11)))))))))) )),
  inference(subsumption_resolution,[],[f1828,f345])).
fof(f1828,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | sP0(X12,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X11)))))))))) )),
  inference(superposition,[],[f477,f365])).
fof(f1815,plain,(
  ( ! [X0] : (~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f612,f365,f544])).
fof(f1816,plain,(
  ( ! [X0] : (sP5('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f345,f365,f542])).
fof(f1817,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f612,f365,f544])).
fof(f1818,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f345,f365,f542])).
fof(f1758,plain,(
  ( ! [X0] : (~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f612,f364,f544])).
fof(f1759,plain,(
  ( ! [X0] : (sP5('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f345,f364,f542])).
fof(f1760,plain,(
  ( ! [X0] : (~sP5('F_0','F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f612,f364,f544])).
fof(f1761,plain,(
  ( ! [X0] : (sP5('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))))) )),
  inference(unit_resulting_resolution,[],[f345,f364,f542])).
fof(f1725,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f1714,f361])).
fof(f1714,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))),
  inference(unit_resulting_resolution,[],[f363,f391])).
fof(f1696,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.rat.rat_gre_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)),'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f1311,f524])).
fof(f1691,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))) )),
  inference(superposition,[],[f390,f357])).
fof(f1690,plain,(
  ( ! [X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X7))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f390,f363])).
fof(f1688,plain,(
  ( ! [X5] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X5))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X5)))))) )),
  inference(superposition,[],[f390,f357])).
fof(f1687,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X4))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X4)))))))) )),
  inference(superposition,[],[f390,f359])).
fof(f1684,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X1))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(superposition,[],[f390,f1405])).
fof(f1637,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0)))))))) )),
  inference(superposition,[],[f454,f358])).
fof(f1616,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)),'const.rat.rat_gre_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f1309,f524])).
fof(f1602,plain,(
  ( ! [X2,X1] : (s('type.integer.int',X2) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f467,f363])).
fof(f1544,plain,(
  ( ! [X0] : (~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f612,f362,f544])).
fof(f1545,plain,(
  ( ! [X0] : (sP5('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f345,f362,f542])).
fof(f1546,plain,(
  ( ! [X0] : (~sP5('F_0','F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f612,f362,f544])).
fof(f1547,plain,(
  ( ! [X0] : (sP5('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(unit_resulting_resolution,[],[f345,f362,f542])).
fof(f1532,plain,(
  ~sP2('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f624,f1215,f620,f612,f528])).
fof(f1496,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f626,f622,f620,f612,f527])).
fof(f1502,plain,(
  ~sP2('const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f625,f1217,f621,f612,f527])).
fof(f1471,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f630,f622,f618,f345,f526])).
fof(f1427,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))) )),
  inference(superposition,[],[f467,f1405])).
fof(f1405,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(equality_resolution,[],[f467])).
fof(f1402,plain,(
  ( ! [X8,X9] : (s('type.integer.int',X9) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X8))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X8)))))) )),
  inference(superposition,[],[f467,f357])).
fof(f1401,plain,(
  ( ! [X6,X7] : (s('type.integer.int',X7) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X6))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X6)))))))) )),
  inference(superposition,[],[f467,f359])).
fof(f1416,plain,(
  ( ! [X2,X3] : (s('type.integer.int',X3) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X2))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X2)))))))) )),
  inference(forward_demodulation,[],[f1399,f1391])).
fof(f1399,plain,(
  ( ! [X2,X3] : (s('type.integer.int',X3) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X2))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X2)))))))) )),
  inference(superposition,[],[f467,f360])).
fof(f1406,plain,(
  ( ! [X0] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0)))))))) )),
  inference(backward_demodulation,[],[f1391,f360])).
fof(f1391,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))))) = s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0)))))) )),
  inference(unit_resulting_resolution,[],[f360,f467])).
fof(f1346,plain,(
  ( ! [X85,X83,X86,X84] : (sP9('const.rat.rat_gre_2'(s('type.rat.rat',X83),s('type.rat.rat',X84)),X85,X86) | ~sP8('const.rat.rat_gre_2'(s('type.rat.rat',X83),s('type.rat.rat',X84)),X85,X86) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X84),s('type.rat.rat',X83))))) )),
  inference(superposition,[],[f597,f453])).
fof(f1345,plain,(
  ( ! [X80,X81,X79,X82] : (sP11('const.rat.rat_gre_2'(s('type.rat.rat',X79),s('type.rat.rat',X80)),X81,X82) | ~sP10('const.rat.rat_gre_2'(s('type.rat.rat',X79),s('type.rat.rat',X80)),X81,X82) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X80),s('type.rat.rat',X79))))) )),
  inference(superposition,[],[f596,f453])).
fof(f1344,plain,(
  ( ! [X78,X76,X77,X75] : (~sP11('const.rat.rat_gre_2'(s('type.rat.rat',X75),s('type.rat.rat',X76)),X78,X77) | ~p(s(bool,X77)) | p(s(bool,X78)) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X76),s('type.rat.rat',X75))))) )),
  inference(superposition,[],[f593,f453])).
fof(f1343,plain,(
  ( ! [X74,X72,X71,X73] : (sP11('const.rat.rat_gre_2'(s('type.rat.rat',X71),s('type.rat.rat',X72)),X73,X74) | ~p(s(bool,X73)) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X72),s('type.rat.rat',X71))))) )),
  inference(superposition,[],[f592,f453])).
fof(f1342,plain,(
  ( ! [X70,X68,X69,X67] : (sP11('const.rat.rat_gre_2'(s('type.rat.rat',X67),s('type.rat.rat',X68)),X70,X69) | p(s(bool,X69)) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X68),s('type.rat.rat',X67))))) )),
  inference(superposition,[],[f591,f453])).
fof(f1341,plain,(
  ( ! [X66,X64,X65,X63] : (~sP10(X66,X65,'const.rat.rat_gre_2'(s('type.rat.rat',X63),s('type.rat.rat',X64))) | p(s(bool,X65)) | ~p(s(bool,X66)) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X64),s('type.rat.rat',X63))))) )),
  inference(superposition,[],[f586,f453])).
fof(f1340,plain,(
  ( ! [X61,X59,X62,X60] : (~sP9('const.rat.rat_gre_2'(s('type.rat.rat',X59),s('type.rat.rat',X60)),X62,X61) | p(s(bool,X61)) | p(s(bool,X62)) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X60),s('type.rat.rat',X59))))) )),
  inference(superposition,[],[f576,f453])).
fof(f1339,plain,(
  ( ! [X57,X58,X56,X55] : (sP9('const.rat.rat_gre_2'(s('type.rat.rat',X55),s('type.rat.rat',X56)),X57,X58) | ~p(s(bool,X57)) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X56),s('type.rat.rat',X55))))) )),
  inference(superposition,[],[f575,f453])).
fof(f1338,plain,(
  ( ! [X54,X52,X53,X51] : (sP9('const.rat.rat_gre_2'(s('type.rat.rat',X51),s('type.rat.rat',X52)),X54,X53) | ~p(s(bool,X53)) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X52),s('type.rat.rat',X51))))) )),
  inference(superposition,[],[f574,f453])).
fof(f1337,plain,(
  ( ! [X47,X50,X48,X49] : (~sP8('const.rat.rat_gre_2'(s('type.rat.rat',X47),s('type.rat.rat',X48)),X49,X50) | p(s(bool,X49)) | p(s(bool,X50)) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X47))))) )),
  inference(superposition,[],[f569,f453])).
fof(f1336,plain,(
  ( ! [X45,X43,X46,X44] : (~sP7(X46,X45,'const.rat.rat_gre_2'(s('type.rat.rat',X43),s('type.rat.rat',X44))) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X44),s('type.rat.rat',X43))))) )),
  inference(superposition,[],[f561,f453])).
fof(f1335,plain,(
  ( ! [X39,X41,X42,X40] : (~sP6(X41,'const.rat.rat_gre_2'(s('type.rat.rat',X39),s('type.rat.rat',X40)),X42) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X39))))) )),
  inference(superposition,[],[f552,f453])).
fof(f1334,plain,(
  ( ! [X37,X35,X38,X36] : (~sP3(X37,'const.rat.rat_gre_2'(s('type.rat.rat',X35),s('type.rat.rat',X36)),X38) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X36),s('type.rat.rat',X35))))) )),
  inference(superposition,[],[f533,f453])).
fof(f1333,plain,(
  ( ! [X33,X31,X34,X32] : (~sP4(X33,'const.rat.rat_gre_2'(s('type.rat.rat',X31),s('type.rat.rat',X32)),X34) | ~p(s(bool,X33)) | p(s(bool,X34)) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X32),s('type.rat.rat',X31))))) )),
  inference(superposition,[],[f529,f453])).
fof(f1332,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X29),s('type.rat.rat',X28)))) | p(s(bool,X30)) | sP1(X30,'const.rat.rat_gre_2'(s('type.rat.rat',X28),s('type.rat.rat',X29)))) )),
  inference(superposition,[],[f496,f453])).
fof(f1331,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X26),s('type.rat.rat',X25)))) | p(s(bool,X27)) | sP1('const.rat.rat_gre_2'(s('type.rat.rat',X25),s('type.rat.rat',X26)),X27)) )),
  inference(superposition,[],[f493,f453])).
fof(f1330,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X22)))) | sP0('const.rat.rat_gre_2'(s('type.rat.rat',X22),s('type.rat.rat',X23)),X24)) )),
  inference(superposition,[],[f478,f453])).
fof(f1329,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X20),s('type.rat.rat',X19)))) | sP0(X21,'const.rat.rat_gre_2'(s('type.rat.rat',X19),s('type.rat.rat',X20)))) )),
  inference(superposition,[],[f477,f453])).
fof(f1320,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,'const.rat.rat_gre_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))) | s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X3),s('type.rat.rat',X2)))) )),
  inference(superposition,[],[f453,f371])).
fof(f1308,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.rat.rat_gre_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)),'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f612,f453,f544])).
fof(f1309,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.rat.rat_gre_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)),'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f345,f453,f542])).
fof(f1310,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)),'const.rat.rat_gre_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f612,f453,f544])).
fof(f1311,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)),'const.rat.rat_gre_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f345,f453,f542])).
fof(f1303,plain,(
  sP4('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f1218,f524])).
fof(f1215,plain,(
  ~sP5('F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f612,f347,f544])).
fof(f1216,plain,(
  sP5('T_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f345,f347,f542])).
fof(f1217,plain,(
  ~sP5('F_0','F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f612,f347,f544])).
fof(f1218,plain,(
  sP5('T_0','F_0','const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f345,f347,f542])).
fof(f1169,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1160,f581])).
fof(f1160,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f345,f612,f345,f586])).
fof(f1152,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f644,f345,f612,f585])).
fof(f1100,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1087,f564])).
fof(f1087,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f612,f612,f345,f569])).
fof(f1078,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f640,f612,f612,f568])).
fof(f1018,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1009,f545])).
fof(f1009,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f612,f345,f345,f552])).
fof(f999,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f632,f345,f345,f548])).
fof(f995,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f986,f522])).
fof(f986,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f612,f612,f612,f537])).
fof(f980,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f971,f523])).
fof(f971,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f612,f345,f345,f533])).
fof(f952,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f612,f345,f345,f529])).
fof(f912,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f345,f645,f543])).
fof(f914,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f345,f645,f543])).
fof(f897,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f612,f645,f541])).
fof(f898,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f612,f645,f541])).
fof(f752,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f345,f345,f592])).
fof(f741,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f612,f345,f591])).
fof(f733,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f345,f612,f580])).
fof(f729,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f612,f612,f579])).
fof(f720,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f345,f345,f575])).
fof(f715,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f345,f345,f574])).
fof(f712,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f345,f612,f563])).
fof(f708,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f345,f612,f562])).
fof(f707,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f612,f612,f557])).
fof(f701,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f612,f612,f556])).
fof(f696,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f345,f612,f547])).
fof(f694,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f345,f612,f546])).
fof(f687,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f612,f612,f613])).
fof(f677,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f612,f345,f496])).
fof(f669,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f612,f345,f493])).
fof(f667,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f345,f345,f492])).
fof(f666,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f612,f612,f491])).
fof(f644,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f612,f589])).
fof(f642,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f345,f588])).
fof(f641,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f612,f587])).
fof(f640,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f612,f572])).
fof(f638,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f345,f571])).
fof(f636,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f345,f570])).
fof(f635,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f612,f555])).
fof(f634,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f612,f554])).
fof(f632,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f345,f553])).
fof(f630,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f345,f540])).
fof(f628,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f345,f539])).
fof(f626,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f345,f538])).
fof(f625,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f612,f536])).
fof(f624,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f612,f535])).
fof(f622,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f345,f534])).
fof(f621,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f612,f532])).
fof(f620,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f612,f531])).
fof(f618,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f345,f530])).
fof(f616,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f345,f478])).
fof(f614,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f345,f477])).
fof(f591,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f342])).
fof(f592,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f342])).
fof(f593,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f342])).
fof(f586,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f339])).
fof(f339,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f338])).
fof(f338,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f232])).
fof(f587,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f340])).
fof(f588,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f340])).
fof(f589,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f340])).
fof(f579,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f337])).
fof(f337,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f336])).
fof(f336,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f335])).
fof(f335,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f233])).
fof(f580,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f337])).
fof(f581,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f337])).
fof(f596,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f582])).
fof(f582,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f337])).
fof(f583,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f337])).
fof(f584,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f337])).
fof(f585,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f337])).
fof(f573,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f334])).
fof(f334,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f333])).
fof(f333,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f188,f230,f229])).
fof(f229,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f230,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f188,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f84])).
fof(f84,axiom,(
  ! [X23,X27,X28] : ((p(s(bool,X28)) <=> (p(s(bool,X27)) | p(s(bool,X23)))) <=> ((p(s(bool,X28)) | ~p(s(bool,X27))) & (p(s(bool,X28)) | ~p(s(bool,X23))) & (p(s(bool,X27)) | p(s(bool,X23)) | ~p(s(bool,X28)))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.sat.sat_25_prove_ltb')).
fof(f574,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f334])).
fof(f575,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f334])).
fof(f576,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f334])).
fof(f569,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f332])).
fof(f332,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f331])).
fof(f331,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f330])).
fof(f330,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f229])).
fof(f570,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f332])).
fof(f571,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f332])).
fof(f572,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f332])).
fof(f562,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f328])).
fof(f328,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f327])).
fof(f327,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f230])).
fof(f563,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f329])).
fof(f564,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f329])).
fof(f597,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f565])).
fof(f565,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f329])).
fof(f566,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f329])).
fof(f567,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f329])).
fof(f568,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f329])).
fof(f556,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f186,f227,f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f227,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f186,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X23,X27,X28] : ((p(s(bool,X28)) <=> (p(s(bool,X27)) & p(s(bool,X23)))) <=> ((p(s(bool,X28)) | ~p(s(bool,X27)) | ~p(s(bool,X23))) & (p(s(bool,X27)) | ~p(s(bool,X28))) & (p(s(bool,X23)) | ~p(s(bool,X28)))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.sat.sat_23_prove_ltb')).
fof(f557,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f326])).
fof(f558,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f326])).
fof(f561,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f326])).
fof(f552,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f323])).
fof(f323,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f226])).
fof(f553,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f324])).
fof(f554,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f324])).
fof(f555,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f324])).
fof(f545,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f322])).
fof(f322,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f321])).
fof(f321,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f227])).
fof(f546,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f322])).
fof(f547,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f322])).
fof(f548,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f322])).
fof(f598,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f551])).
fof(f551,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f322])).
fof(f541,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f184,f224,f223,f222,f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f222,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f223,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f224,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f184,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X23,X27,X28] : ((p(s(bool,X28)) <=> s(bool,X23) = s(bool,X27)) <=> ((p(s(bool,X28)) | p(s(bool,X27)) | p(s(bool,X23))) & (p(s(bool,X28)) | ~p(s(bool,X23)) | ~p(s(bool,X27))) & (p(s(bool,X27)) | ~p(s(bool,X23)) | ~p(s(bool,X28))) & (p(s(bool,X23)) | ~p(s(bool,X27)) | ~p(s(bool,X28)))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.sat.sat_2|49|_prove_ltb')).
fof(f542,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f319])).
fof(f543,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f319])).
fof(f544,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f319])).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f318])).
fof(f318,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f317])).
fof(f317,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f221])).
fof(f538,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f318])).
fof(f539,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f318])).
fof(f540,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f318])).
fof(f533,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f315])).
fof(f315,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f314])).
fof(f314,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f222])).
fof(f534,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f316])).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f316])).
fof(f536,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f316])).
fof(f529,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f312])).
fof(f312,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f311])).
fof(f311,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f223])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f313])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f313])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f313])).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f310])).
fof(f310,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f309])).
fof(f309,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f308])).
fof(f308,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f224])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f310])).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f310])).
fof(f525,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f310])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | ~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f310])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f310])).
fof(f528,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f310])).
fof(f491,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f296])).
fof(f296,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f492,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f297])).
fof(f493,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f297])).
fof(f496,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f297])).
fof(f613,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f476,f612])).
fof(f476,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f289])).
fof(f289,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f477,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f290])).
fof(f478,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f290])).
fof(f470,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))) != s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1] : ((s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))) != s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))) & (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))) | s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1)))))),
  inference(nnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1] : (s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1))) <=> s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X18,X25] : (s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X18))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X25))) <=> s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X18))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X25))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X25))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X18))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.rat.RAT_EQ')).
fof(f471,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))) | s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X1)))) )),
  inference(cnf_transformation,[],[f286])).
fof(f465,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f208])).
fof(f461,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | s('type.integer.int',X1) = s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))) )),
  inference(cnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | s('type.integer.int',X1) = s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))))),
  inference(ennf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) => s('type.integer.int',X1) = s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X19,X20] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X20)))) => s('type.integer.int',X20) = s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X19),s('type.integer.int',X20))))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.frac.DNM')).
fof(f460,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))))) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))))),
  inference(ennf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) => s('type.integer.int',X0) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X19,X20] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X20)))) => s('type.integer.int',X19) = s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X19),s('type.integer.int',X20))))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.frac.NMR')).
fof(f459,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X21,X22] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num',X21))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22)))))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.integer.INT_LT_REDUCE|split|11')).
fof(f458,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X21,X22] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X21),s('type.num.num',X22))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22)))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.integer.INT_LT_REDUCE|split|7')).
fof(f457,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X21,X22] : s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X21))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22)))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.integer.INT_LT_REDUCE|split|8')).
fof(f456,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X21,X22] : s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X21))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22)))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.integer.INT_LT_REDUCE|split|9')).
fof(f455,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X21,X22] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22)))))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.integer.INT_LT_REDUCE|split|10')).
fof(f453,plain,(
  ( ! [X0,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))) = s(bool,'const.rat.rat_gre_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))) = s(bool,'const.rat.rat_gre_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X6,X24] : s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X24),s('type.rat.rat',X6))) = s(bool,'const.rat.rat_gre_2'(s('type.rat.rat',X6),s('type.rat.rat',X24)))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.rat.rat_gre_def')).
fof(f449,plain,(
  ( ! [X2,X0,X7,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK13(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK13(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK13(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK13(X2,X1,X0)))))) | (! [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK15(X2,X1,X0))))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15])],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2] : ((? [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | (! [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))))) & (? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X11] : ! [X14,X15] : (! [X12] : (p(s(bool,happ(s(fun(X11,bool),X14),s(X11,X12)))) & p(s(bool,happ(s(fun(X11,bool),X15),s(X11,X12))))) <=> (! [X12] : p(s(bool,happ(s(fun(X11,bool),X14),s(X11,X12)))) & ! [X12] : p(s(bool,happ(s(fun(X11,bool),X15),s(X11,X12))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.bool.FORALL_AND_THM')).
fof(f450,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK13(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK13(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f283])).
fof(f451,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK15(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f283])).
fof(f452,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK15(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f283])).
fof(f394,plain,(
  ( ! [X0] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(cnf_transformation,[],[f245])).
fof(f393,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(cnf_transformation,[],[f244])).
fof(f390,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0] : ((s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0)))))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X12] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X12))) <=> s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.intExtension.INT_SGN_CLAUSES|split|1')).
fof(f391,plain,(
  ( ! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f243])).
fof(f386,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X7] : (s(bool,'F_0') = s(bool,X7) <=> ~p(s(bool,X7)))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f387,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f241])).
fof(f369,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',X0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',X0))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X0))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X26] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num',X26))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X26))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.rat.RAT_OF_NUM_CALCULATE')).
fof(f368,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X21] : s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X21))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.integer.INT_LT_REDUCE|split|6')).
fof(f366,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X21] : s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X21)))))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.integer.INT_LT_REDUCE|split|1')).
fof(f365,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X21] : s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X21)))))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.integer.INT_LT_REDUCE|split|0')).
fof(f362,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X21] : s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.integer.INT_LT_REDUCE|split|4')).
fof(f361,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X12] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X12)))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.integer.INT_MUL_LZERO')).
fof(f357,plain,(
  ( ! [X0] : (s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X0))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X18] : s('type.integer.int','const.frac.frac_sgn_1'(s('type.frac.frac',X18))) = s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X18)))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.frac.frac_sgn_def')).
fof(f356,plain,(
  ( ! [X0] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X6] : s('type.rat.rat',X6) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.rat.RAT_SUB_RID')).
fof(f347,plain,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f56])).
fof(f56,axiom,(
  s(bool,'F_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.integer.INT_LT_REDUCE|split|2')).
fof(f5218,plain,(
  ( ! [X97,X98,X96] : (p(s(bool,X98)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X98,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X96))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X97)))))))))) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f5055,f345])).
fof(f5055,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,X98)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X98,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X96))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X97)))))))))) ) | ($spl2 | $spl18)),
  inference(backward_demodulation,[],[f4794,f3023])).
fof(f3023,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X97),s('type.num.num',X96)))) | p(s(bool,X98)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X98,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X96))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X97)))))))))) ) | $spl2),
  inference(superposition,[],[f1199,f459])).
fof(f1199,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,X0)) ) | $spl2),
  inference(subsumption_resolution,[],[f1195,f345])).
fof(f1195,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,X0)) ) | $spl2),
  inference(superposition,[],[f593,f648])).
fof(f5217,plain,(
  $false | ($spl2 | $spl18)),
  inference(global_subsumption,[],[f5216,f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188,f5190,f5192,f5194,f5196,f5198,f5200,f5202,f5204,f5206,f5208,f5210])).
fof(f5216,plain,(
  ( ! [X94,X95,X93] : (sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X93))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X94)))))))),X95)) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f5054,f345])).
fof(f5054,plain,(
  ( ! [X94,X95,X93] : (~p(s(bool,'T_0')) | sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X93))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X94)))))))),X95)) ) | ($spl2 | $spl18)),
  inference(backward_demodulation,[],[f4794,f3022])).
fof(f3022,plain,(
  ( ! [X94,X95,X93] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X94),s('type.num.num',X93)))) | sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X93))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X94)))))))),X95)) ) | $spl2),
  inference(superposition,[],[f754,f459])).
fof(f754,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,X1)) ) | $spl2),
  inference(subsumption_resolution,[],[f753,f345])).
fof(f753,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,X1)) ) | $spl2),
  inference(superposition,[],[f592,f648])).
fof(f5215,plain,(
  $false | ($spl2 | $spl18)),
  inference(global_subsumption,[],[f5214,f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188,f5190,f5192,f5194,f5196,f5198,f5200,f5202,f5204,f5206,f5208,f5210])).
fof(f5214,plain,(
  ( ! [X92,X90,X91] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X90))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X91)))))))),X92)) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f5053,f345])).
fof(f5053,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,'T_0')) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X90))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X91)))))))),X92)) ) | ($spl2 | $spl18)),
  inference(backward_demodulation,[],[f4794,f3021])).
fof(f3021,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X91),s('type.num.num',X90)))) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X90))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X91)))))))),X92)) ) | $spl2),
  inference(superposition,[],[f722,f459])).
fof(f722,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,X1)) ) | $spl2),
  inference(subsumption_resolution,[],[f721,f345])).
fof(f721,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X0,X1)) ) | $spl2),
  inference(superposition,[],[f575,f648])).
fof(f5213,plain,(
  $false | ($spl2 | $spl18)),
  inference(global_subsumption,[],[f5212,f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188,f5190,f5192,f5194,f5196,f5198,f5200,f5202,f5204,f5206,f5208,f5210])).
fof(f5212,plain,(
  ( ! [X88,X87,X89] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X89,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X87))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X88)))))))))) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f5052,f345])).
fof(f5052,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,'T_0')) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X89,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X87))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X88)))))))))) ) | ($spl2 | $spl18)),
  inference(backward_demodulation,[],[f4794,f3020])).
fof(f3020,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X88),s('type.num.num',X87)))) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X89,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X87))))))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X88)))))))))) ) | $spl2),
  inference(superposition,[],[f717,f459])).
fof(f717,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,X0)) ) | $spl2),
  inference(subsumption_resolution,[],[f716,f345])).
fof(f716,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,X0)) ) | $spl2),
  inference(superposition,[],[f574,f648])).
fof(f5211,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188,f5190,f5192,f5194,f5196,f5198,f5200,f5202,f5204,f5206,f5208,f5210])).
fof(f5209,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188,f5190,f5192,f5194,f5196,f5198,f5200,f5202,f5204,f5206,f5208])).
fof(f5207,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188,f5190,f5192,f5194,f5196,f5198,f5200,f5202,f5204,f5206])).
fof(f5205,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188,f5190,f5192,f5194,f5196,f5198,f5200,f5202,f5204])).
fof(f5203,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188,f5190,f5192,f5194,f5196,f5198,f5200,f5202])).
fof(f5201,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188,f5190,f5192,f5194,f5196,f5198,f5200])).
fof(f5199,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188,f5190,f5192,f5194,f5196,f5198])).
fof(f5197,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188,f5190,f5192,f5194,f5196])).
fof(f5195,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188,f5190,f5192,f5194])).
fof(f5193,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188,f5190,f5192])).
fof(f5191,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188,f5190])).
fof(f5189,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186,f5188])).
fof(f5187,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184,f5186])).
fof(f5185,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182,f5184])).
fof(f5183,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180,f5182])).
fof(f5181,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178,f5180])).
fof(f5179,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176,f5178])).
fof(f5177,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166,f5176])).
fof(f5175,plain,(
  $false | ($spl2 | $spl18)),
  inference(global_subsumption,[],[f5174,f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166])).
fof(f5174,plain,(
  ( ! [X97,X98,X96] : (p(s(bool,X98)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X98,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X96))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X97)))))))) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f5033,f345])).
fof(f5033,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,X98)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X98,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X96))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X97)))))))) ) | ($spl2 | $spl18)),
  inference(backward_demodulation,[],[f4794,f2358])).
fof(f2358,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X96),s('type.num.num',X97)))) | p(s(bool,X98)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X98,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X96))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X97)))))))) ) | $spl2),
  inference(superposition,[],[f1199,f458])).
fof(f5173,plain,(
  $false | ($spl2 | $spl18)),
  inference(global_subsumption,[],[f5172,f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166])).
fof(f5172,plain,(
  ( ! [X94,X95,X93] : (sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X93))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X94)))))),X95)) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f5032,f345])).
fof(f5032,plain,(
  ( ! [X94,X95,X93] : (~p(s(bool,'T_0')) | sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X93))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X94)))))),X95)) ) | ($spl2 | $spl18)),
  inference(backward_demodulation,[],[f4794,f2357])).
fof(f2357,plain,(
  ( ! [X94,X95,X93] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X93),s('type.num.num',X94)))) | sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X93))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X94)))))),X95)) ) | $spl2),
  inference(superposition,[],[f754,f458])).
fof(f5171,plain,(
  $false | ($spl2 | $spl18)),
  inference(global_subsumption,[],[f5170,f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166])).
fof(f5170,plain,(
  ( ! [X92,X90,X91] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X90))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X91)))))),X92)) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f5031,f345])).
fof(f5031,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,'T_0')) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X90))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X91)))))),X92)) ) | ($spl2 | $spl18)),
  inference(backward_demodulation,[],[f4794,f2356])).
fof(f2356,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X90),s('type.num.num',X91)))) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X90))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X91)))))),X92)) ) | $spl2),
  inference(superposition,[],[f722,f458])).
fof(f5169,plain,(
  $false | ($spl2 | $spl18)),
  inference(global_subsumption,[],[f5168,f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166])).
fof(f5168,plain,(
  ( ! [X88,X87,X89] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X89,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X87))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X88)))))))) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f5030,f345])).
fof(f5030,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,'T_0')) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X89,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X87))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X88)))))))) ) | ($spl2 | $spl18)),
  inference(backward_demodulation,[],[f4794,f2355])).
fof(f2355,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X87),s('type.num.num',X88)))) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X89,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X87))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X88)))))))) ) | $spl2),
  inference(superposition,[],[f717,f458])).
fof(f5167,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164,f5166])).
fof(f5165,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162,f5164])).
fof(f5163,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160,f5162])).
fof(f5161,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158,f5160])).
fof(f5159,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156,f5158])).
fof(f5157,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154,f5156])).
fof(f5155,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152,f5154])).
fof(f5153,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150,f5152])).
fof(f5151,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148,f5150])).
fof(f5149,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146,f5148])).
fof(f5147,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144,f5146])).
fof(f5145,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142,f5144])).
fof(f5143,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140,f5142])).
fof(f5141,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138,f5140])).
fof(f5139,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136,f5138])).
fof(f5137,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134,f5136])).
fof(f5135,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132,f5134])).
fof(f5133,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128,f5132])).
fof(f5131,plain,(
  $false | ($spl2 | $spl18)),
  inference(global_subsumption,[],[f5130,f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128])).
fof(f5130,plain,(
  ( ! [X0] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12)))))) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f5005,f345])).
fof(f5005,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12)))))) ) | ($spl2 | $spl18)),
  inference(backward_demodulation,[],[f4794,f2624])).
fof(f2624,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_gt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12)))))) ) | $spl2),
  inference(backward_demodulation,[],[f2620,f393])).
fof(f5129,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126,f5128])).
fof(f5127,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112,f5126])).
fof(f5125,plain,(
  $false | ($spl2 | $spl18)),
  inference(global_subsumption,[],[f5124,f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112])).
fof(f5124,plain,(
  ( ! [X0] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12)))))))) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f4989,f345])).
fof(f4989,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12)))))))) ) | ($spl2 | $spl18)),
  inference(backward_demodulation,[],[f4794,f2626])).
fof(f2626,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12)))))))) ) | $spl2),
  inference(backward_demodulation,[],[f2620,f395])).
fof(f5123,plain,(
  $false | ($spl2 | $spl18)),
  inference(global_subsumption,[],[f5122,f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112])).
fof(f5122,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f4962,f345])).
fof(f4962,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))) ) | ($spl2 | $spl18)),
  inference(backward_demodulation,[],[f4794,f2627])).
fof(f2627,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))) ) | $spl2),
  inference(backward_demodulation,[],[f2620,f468])).
fof(f5121,plain,(
  $false | ($spl2 | $spl18)),
  inference(global_subsumption,[],[f5120,f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112])).
fof(f5120,plain,(
  ( ! [X97,X98,X96] : (p(s(bool,X98)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X98,'const.rat.rat_gre_2'(s('type.rat.rat',X96),s('type.rat.rat',X97)))) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f4961,f345])).
fof(f4961,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,X98)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X98,'const.rat.rat_gre_2'(s('type.rat.rat',X96),s('type.rat.rat',X97)))) ) | ($spl2 | $spl18)),
  inference(backward_demodulation,[],[f4794,f1350])).
fof(f1350,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X97),s('type.rat.rat',X96)))) | p(s(bool,X98)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X98,'const.rat.rat_gre_2'(s('type.rat.rat',X96),s('type.rat.rat',X97)))) ) | $spl2),
  inference(superposition,[],[f1199,f453])).
fof(f5119,plain,(
  $false | ($spl2 | $spl18)),
  inference(global_subsumption,[],[f5118,f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112])).
fof(f5118,plain,(
  ( ! [X94,X95,X93] : (sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.rat.rat_gre_2'(s('type.rat.rat',X93),s('type.rat.rat',X94)),X95)) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f4960,f345])).
fof(f4960,plain,(
  ( ! [X94,X95,X93] : (~p(s(bool,'T_0')) | sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.rat.rat_gre_2'(s('type.rat.rat',X93),s('type.rat.rat',X94)),X95)) ) | ($spl2 | $spl18)),
  inference(backward_demodulation,[],[f4794,f1349])).
fof(f1349,plain,(
  ( ! [X94,X95,X93] : (sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.rat.rat_gre_2'(s('type.rat.rat',X93),s('type.rat.rat',X94)),X95) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X94),s('type.rat.rat',X93))))) ) | $spl2),
  inference(superposition,[],[f754,f453])).
fof(f5117,plain,(
  $false | ($spl2 | $spl18)),
  inference(global_subsumption,[],[f5116,f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112])).
fof(f5116,plain,(
  ( ! [X92,X90,X91] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.rat.rat_gre_2'(s('type.rat.rat',X90),s('type.rat.rat',X91)),X92)) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f4959,f345])).
fof(f4959,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,'T_0')) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.rat.rat_gre_2'(s('type.rat.rat',X90),s('type.rat.rat',X91)),X92)) ) | ($spl2 | $spl18)),
  inference(backward_demodulation,[],[f4794,f1348])).
fof(f1348,plain,(
  ( ! [X92,X90,X91] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.rat.rat_gre_2'(s('type.rat.rat',X90),s('type.rat.rat',X91)),X92) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X91),s('type.rat.rat',X90))))) ) | $spl2),
  inference(superposition,[],[f722,f453])).
fof(f5115,plain,(
  $false | ($spl2 | $spl18)),
  inference(global_subsumption,[],[f5114,f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112])).
fof(f5114,plain,(
  ( ! [X88,X87,X89] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X89,'const.rat.rat_gre_2'(s('type.rat.rat',X87),s('type.rat.rat',X88)))) ) | ($spl2 | $spl18)),
  inference(subsumption_resolution,[],[f4958,f345])).
fof(f4958,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,'T_0')) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X89,'const.rat.rat_gre_2'(s('type.rat.rat',X87),s('type.rat.rat',X88)))) ) | ($spl2 | $spl18)),
  inference(backward_demodulation,[],[f4794,f1347])).
fof(f1347,plain,(
  ( ! [X88,X87,X89] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X89,'const.rat.rat_gre_2'(s('type.rat.rat',X87),s('type.rat.rat',X88))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X88),s('type.rat.rat',X87))))) ) | $spl2),
  inference(superposition,[],[f717,f453])).
fof(f5113,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110,f5112])).
fof(f5111,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108,f5110])).
fof(f5109,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106,f5108])).
fof(f5107,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104,f5106])).
fof(f5105,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102,f5104])).
fof(f5103,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100,f5102])).
fof(f5101,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098,f5100])).
fof(f5099,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096,f5098])).
fof(f5097,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094,f5096])).
fof(f5095,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092,f5094])).
fof(f5093,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090,f5092])).
fof(f5091,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088,f5090])).
fof(f5089,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086,f5088])).
fof(f5087,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084,f5086])).
fof(f5085,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082,f5084])).
fof(f5083,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080,f5082])).
fof(f5081,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078,f5080])).
fof(f5079,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075,f5078])).
fof(f5077,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f4923,f345])).
fof(f4923,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(backward_demodulation,[],[f4794,f612])).
fof(f5076,plain,(
  $false | $spl18),
  inference(global_subsumption,[],[f344,f343,f345,f347,f612,f353,f354,f355,f356,f357,f358,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f371,f383,f382,f387,f386,f391,f390,f393,f392,f395,f394,f452,f451,f450,f449,f453,f454,f455,f456,f457,f458,f459,f460,f461,f465,f462,f467,f469,f468,f471,f470,f478,f477,f613,f496,f493,f492,f491,f509,f528,f527,f526,f525,f524,f523,f522,f532,f531,f530,f529,f536,f535,f534,f533,f540,f539,f538,f537,f544,f543,f542,f541,f598,f548,f547,f546,f545,f555,f554,f553,f552,f561,f558,f557,f556,f568,f567,f566,f597,f564,f563,f562,f572,f571,f570,f569,f576,f575,f574,f573,f585,f584,f583,f596,f581,f580,f579,f589,f588,f587,f586,f593,f592,f591,f590,f614,f616,f618,f620,f621,f622,f624,f625,f626,f628,f630,f632,f634,f635,f636,f638,f640,f641,f642,f644,f645,f666,f667,f669,f677,f687,f694,f696,f701,f707,f708,f712,f715,f720,f729,f733,f741,f752,f817,f879,f898,f897,f914,f912,f952,f971,f980,f986,f995,f999,f1009,f1018,f1078,f1087,f1100,f1152,f1160,f1169,f1218,f1217,f1216,f1215,f1264,f1277,f1303,f1311,f1310,f1309,f1308,f1320,f1329,f1330,f1331,f1332,f1333,f1334,f1335,f1336,f1337,f1338,f1339,f1340,f1341,f1342,f1343,f1344,f1345,f1346,f1391,f1406,f1416,f1400,f1401,f1402,f1404,f1405,f1427,f1471,f1502,f1496,f1532,f1547,f1546,f1545,f1544,f1602,f1616,f1637,f1638,f1684,f1687,f1688,f1690,f1691,f1696,f1725,f1761,f1760,f1759,f1758,f1818,f1817,f1816,f1815,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1867,f1868,f1869,f1870,f1871,f1872,f1873,f1874,f1893,f1892,f1891,f1890,f1932,f1933,f1934,f1935,f1936,f1937,f1938,f1939,f1940,f1941,f1942,f1943,f1944,f1945,f1946,f1947,f1948,f1949,f1970,f1969,f1968,f1967,f2009,f2010,f2011,f2012,f2013,f2014,f2015,f2016,f2017,f2018,f2019,f2020,f2021,f2022,f2023,f2024,f2025,f2026,f2047,f2046,f2045,f2044,f2086,f2087,f2088,f2089,f2090,f2091,f2092,f2093,f2094,f2095,f2096,f2097,f2098,f2099,f2100,f2101,f2102,f2103,f2118,f2117,f2116,f2115,f2161,f2160,f2159,f2158,f2200,f2201,f2202,f2203,f2204,f2205,f2206,f2207,f2208,f2209,f2210,f2211,f2212,f2213,f2214,f2215,f2216,f2217,f2230,f2229,f2228,f2227,f2269,f2270,f2271,f2272,f2273,f2274,f2275,f2276,f2277,f2278,f2279,f2280,f2281,f2282,f2283,f2284,f2285,f2286,f2327,f2326,f2325,f2324,f2328,f2337,f2338,f2339,f2340,f2341,f2342,f2343,f2344,f2345,f2346,f2347,f2348,f2349,f2350,f2351,f2352,f2353,f2354,f2397,f2403,f2388,f2451,f2452,f2487,f2488,f2525,f2526,f2550,f2539,f2582,f2583,f2990,f2989,f2988,f2987,f2993,f3002,f3003,f3004,f3005,f3006,f3007,f3008,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3016,f3017,f3018,f3019,f3039,f3086,f3098,f3109,f3122,f3123,f3735,f3736,f3738,f3739,f3769,f3771,f3773,f3962,f3995,f4009,f4035,f4208,f4216,f4218,f4243,f4255,f4258,f4260,f4343,f4344,f4345,f4346,f4347,f4348,f4419,f4792,f4796,f4799,f4802,f4805,f4808,f4811,f4814,f4817,f4820,f4823,f4826,f4829,f4832,f4835,f4838,f4843,f4844,f4845,f4846,f4847,f4848,f4849,f4850,f4851,f4852,f4853,f4854,f4855,f4856,f4857,f4858,f4859,f4860,f4861,f4862,f4867,f4868,f4871,f4872,f4873,f4874,f4876,f4877,f4507,f4508,f4509,f4510,f4511,f4512,f4884,f4886,f4888,f4890,f4892,f4894,f4896,f4898,f4900,f4902,f4904,f4906,f4908,f4910,f4912,f4914,f4794,f5075])).
fof(f4842,plain,(
  ( ! [X1] : (~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,'T_0') | p(s(bool,X1))) ) | $spl50),
  inference(cnf_transformation,[],[f4842_D])).
fof(f4842_D,plain,(
  ( ! [X1] : (~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,'T_0') | p(s(bool,X1))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4794,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(cnf_transformation,[],[f4794_D])).
fof(f4794_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4916,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,'T_0')) ) | $spl2),
  inference(condensation,[],[f4915])).
fof(f4915,plain,(
  ( ! [X103,X105,X104] : (p(s(bool,X104)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X104,'T_0') | s(bool,X103) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) ) | $spl2),
  inference(subsumption_resolution,[],[f4544,f879])).
fof(f4544,plain,(
  ( ! [X103,X105,X104] : (~p(s(bool,X103)) | p(s(bool,X104)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X104,'T_0') | s(bool,X103) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) ) | $spl2),
  inference(superposition,[],[f1199,f817])).
fof(f4839,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl48),
  inference(cnf_transformation,[],[f4839_D])).
fof(f4839_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f4914,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4913])).
fof(f4913,plain,(
  ( ! [X90,X88,X87,X89] : (~sP8('T_0',X88,X89) | sP9('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f4539,f879])).
fof(f4539,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~sP8('T_0',X88,X89) | sP9('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f597,f817])).
fof(f4836,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl46),
  inference(cnf_transformation,[],[f4836_D])).
fof(f4836_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4912,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f4911])).
fof(f4911,plain,(
  ( ! [X85,X83,X86,X84] : (~sP10('T_0',X84,X85) | sP11('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f4538,f879])).
fof(f4538,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~sP10('T_0',X84,X85) | sP11('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f596,f817])).
fof(f4833,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl44),
  inference(cnf_transformation,[],[f4833_D])).
fof(f4833_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4910,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4909])).
fof(f4909,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f4537,f879])).
fof(f4537,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f593,f817])).
fof(f4830,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl42),
  inference(cnf_transformation,[],[f4830_D])).
fof(f4830_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4908,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f4907])).
fof(f4907,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | sP11('T_0',X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f4536,f879])).
fof(f4536,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | ~p(s(bool,X76)) | sP11('T_0',X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f592,f817])).
fof(f4827,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl40),
  inference(cnf_transformation,[],[f4827_D])).
fof(f4827_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4906,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4905])).
fof(f4905,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | sP11('T_0',X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f4535,f879])).
fof(f4535,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | sP11('T_0',X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f591,f817])).
fof(f4824,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl38),
  inference(cnf_transformation,[],[f4824_D])).
fof(f4824_D,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4904,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f4903])).
fof(f4903,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,'T_0') | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f4534,f879])).
fof(f4534,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,'T_0') | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f586,f817])).
fof(f4821,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl36),
  inference(cnf_transformation,[],[f4821_D])).
fof(f4821_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4902,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4901])).
fof(f4901,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | p(s(bool,X65)) | ~sP9('T_0',X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f4533,f879])).
fof(f4533,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP9('T_0',X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f576,f817])).
fof(f4818,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f4818_D])).
fof(f4818_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4900,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4899])).
fof(f4899,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9('T_0',X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f4532,f879])).
fof(f4532,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | sP9('T_0',X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f575,f817])).
fof(f4815,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f4815_D])).
fof(f4815_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4898,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4897])).
fof(f4897,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | sP9('T_0',X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f4531,f879])).
fof(f4531,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | ~p(s(bool,X56)) | sP9('T_0',X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f574,f817])).
fof(f4812,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl30),
  inference(cnf_transformation,[],[f4812_D])).
fof(f4812_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4896,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f4895])).
fof(f4895,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X52)) | p(s(bool,X53)) | ~sP8('T_0',X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f4530,f879])).
fof(f4530,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | p(s(bool,X53)) | ~sP8('T_0',X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f569,f817])).
fof(f4809,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f4809_D])).
fof(f4809_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f4894,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f4893])).
fof(f4893,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,'T_0') | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f4529,f879])).
fof(f4529,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,'T_0') | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f561,f817])).
fof(f4806,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl26),
  inference(cnf_transformation,[],[f4806_D])).
fof(f4806_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4892,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f4891])).
fof(f4891,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f4528,f879])).
fof(f4528,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f552,f817])).
fof(f4803,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl24),
  inference(cnf_transformation,[],[f4803_D])).
fof(f4803_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4890,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f4889])).
fof(f4889,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f4527,f879])).
fof(f4527,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f533,f817])).
fof(f4800,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl22),
  inference(cnf_transformation,[],[f4800_D])).
fof(f4800_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f4888,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f4887])).
fof(f4887,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f4526,f879])).
fof(f4526,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f529,f817])).
fof(f4797,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl20),
  inference(cnf_transformation,[],[f4797_D])).
fof(f4797_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f4886,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4885])).
fof(f4885,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4525,f879])).
fof(f4525,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f496,f817])).
fof(f4793,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl16),
  inference(cnf_transformation,[],[f4793_D])).
fof(f4793_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f4884,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4883])).
fof(f4883,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f4524,f879])).
fof(f4524,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f493,f817])).
fof(f4841,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X1,'T_0')) ) | $spl2),
  inference(condensation,[],[f4840])).
fof(f4840,plain,(
  ( ! [X103,X105,X104] : (p(s(bool,X104)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X104,'T_0') | s(bool,X103) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) ) | $spl2),
  inference(subsumption_resolution,[],[f4450,f879])).
fof(f4450,plain,(
  ( ! [X103,X105,X104] : (~p(s(bool,X103)) | p(s(bool,X104)) | ~sP11('const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X104,'T_0') | s(bool,X103) = s(bool,X105) | s(bool,'T_0') = s(bool,X105)) ) | $spl2),
  inference(superposition,[],[f1199,f817])).
fof(f4838,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4837])).
fof(f4837,plain,(
  ( ! [X90,X88,X87,X89] : (~sP8('T_0',X88,X89) | sP9('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f4445,f879])).
fof(f4445,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | ~sP8('T_0',X88,X89) | sP9('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f597,f817])).
fof(f4835,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f4834])).
fof(f4834,plain,(
  ( ! [X85,X83,X86,X84] : (~sP10('T_0',X84,X85) | sP11('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f4444,f879])).
fof(f4444,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~sP10('T_0',X84,X85) | sP11('T_0',X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f596,f817])).
fof(f4832,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4831])).
fof(f4831,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f4443,f879])).
fof(f4443,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP11('T_0',X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f593,f817])).
fof(f4829,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f4828])).
fof(f4828,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | sP11('T_0',X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f4442,f879])).
fof(f4442,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | ~p(s(bool,X76)) | sP11('T_0',X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f592,f817])).
fof(f4826,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4825])).
fof(f4825,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | sP11('T_0',X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f4441,f879])).
fof(f4441,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | p(s(bool,X72)) | sP11('T_0',X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f591,f817])).
fof(f4823,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f4822])).
fof(f4822,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,'T_0') | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f4440,f879])).
fof(f4440,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X67)) | p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP10(X69,X68,'T_0') | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f586,f817])).
fof(f4820,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4819])).
fof(f4819,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | p(s(bool,X65)) | ~sP9('T_0',X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f4439,f879])).
fof(f4439,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X63)) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP9('T_0',X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f576,f817])).
fof(f4817,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4816])).
fof(f4816,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9('T_0',X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f4438,f879])).
fof(f4438,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X59)) | ~p(s(bool,X60)) | sP9('T_0',X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f575,f817])).
fof(f4814,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4813])).
fof(f4813,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | sP9('T_0',X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f4437,f879])).
fof(f4437,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X55)) | ~p(s(bool,X56)) | sP9('T_0',X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f574,f817])).
fof(f4811,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f4810])).
fof(f4810,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X52)) | p(s(bool,X53)) | ~sP8('T_0',X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f4436,f879])).
fof(f4436,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | p(s(bool,X53)) | ~sP8('T_0',X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f569,f817])).
fof(f4808,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f4807])).
fof(f4807,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,'T_0') | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f4435,f879])).
fof(f4435,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,'T_0') | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f561,f817])).
fof(f4805,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f4804])).
fof(f4804,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f4434,f879])).
fof(f4434,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,'T_0',X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f552,f817])).
fof(f4802,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f4801])).
fof(f4801,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f4433,f879])).
fof(f4433,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,'T_0',X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f533,f817])).
fof(f4799,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f4798])).
fof(f4798,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f4432,f879])).
fof(f4432,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,'T_0',X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f529,f817])).
fof(f4796,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4795])).
fof(f4795,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f4431,f879])).
fof(f4431,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1(X33,'T_0') | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f496,f817])).
fof(f4792,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4791])).
fof(f4791,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f4430,f879])).
fof(f4430,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1('T_0',X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f493,f817])).
fof(f3480,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))) | $spl13),
  inference(cnf_transformation,[],[f3480_D])).
fof(f3480_D,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f3488,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl14),
  inference(cnf_transformation,[],[f3488_D])).
fof(f3488_D,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f3486,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))) | $spl2),
  inference(forward_demodulation,[],[f3485,f2636])).
fof(f3485,plain,(
  ( ! [X16] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X16))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl2),
  inference(forward_demodulation,[],[f3469,f3456])).
fof(f3469,plain,(
  ( ! [X16] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X16))))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl2),
  inference(superposition,[],[f2623,f2636])).
fof(f3458,plain,(
  ( ! [X1] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl2),
  inference(superposition,[],[f2623,f2636])).
fof(f2401,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl6),
  inference(cnf_transformation,[],[f2401_D])).
fof(f2401_D,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f2685,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) | $spl9),
  inference(cnf_transformation,[],[f2685_D])).
fof(f2685_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f2674,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl2),
  inference(superposition,[],[f2623,f1725])).
fof(f2687,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) | $spl2),
  inference(forward_demodulation,[],[f2686,f361])).
fof(f2686,plain,(
  ( ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl2),
  inference(forward_demodulation,[],[f2671,f1725])).
fof(f2671,plain,(
  ( ! [X6] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl2),
  inference(superposition,[],[f2623,f361])).
fof(f2683,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) | $spl2),
  inference(forward_demodulation,[],[f2682,f1725])).
fof(f2682,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl2),
  inference(forward_demodulation,[],[f2670,f1725])).
fof(f2670,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK12))))) != s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl2),
  inference(superposition,[],[f2623,f1725])).
fof(f2399,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl5),
  inference(cnf_transformation,[],[f2399_D])).
fof(f2399_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f2403,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(forward_demodulation,[],[f2402,f361])).
fof(f2402,plain,(
  ( ! [X6] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f2386,f1725])).
fof(f2386,plain,(
  ( ! [X6] : (s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f392,f361])).
fof(f2397,plain,(
  p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(forward_demodulation,[],[f2396,f1725])).
fof(f2396,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(forward_demodulation,[],[f2385,f1725])).
fof(f2385,plain,(
  s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.integer.int_gt_2'(s('type.integer.int','const.intExtension.SGN_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(superposition,[],[f392,f1725])).
fof(f608,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl0),
  inference(cnf_transformation,[],[f608_D])).
fof(f608_D,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f610,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl2),
  inference(cnf_transformation,[],[f610_D])).
fof(f610_D,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f343,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) & (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f235])).
fof(f235,plain,(
  ? [X0] : ((s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) & (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ? [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <~> p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ~! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X6))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X6))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  file('thm.rat.RAT_SGN_CLAUSES|split|0','thm.rat.RAT_SGN_CLAUSES|split|0')).
fof(f609,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl1),
  inference(cnf_transformation,[],[f609_D])).
fof(f609_D,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f611,plain,(
  ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl3),
  inference(cnf_transformation,[],[f611_D])).
fof(f611_D,plain,(
  ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f344,plain,(
  ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK12),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK12))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(cnf_transformation,[],[f236])).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1781
% Passive clauses: 3391
% Generated clauses: 21799
% Final active clauses: 1322
% Final passive clauses: 963
% Input formulas: 85
% Initial clauses: 253
% 
% Duplicate literals: 282
% Trivial inequalities: 70
% Fw subsumption resolutions: 818
% Fw demodulations: 4861
% Bw demodulations: 244
% Condensations: 242
% Global subsumptions: 74
% 
% Simple tautologies: 112
% Equational tautologies: 80
% Forward subsumptions: 11724
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 8299
% Factoring: 1
% Forward superposition: 2779
% Backward superposition: 3581
% Self superposition: 48
% Equality factoring: 32
% Equality resolution: 19
% 
% Split clauses: 53
% Split components: 111
% SAT solver clauses: 3653
% SAT solver unit clauses: 363
% SAT solver binary clauses: 1974
% 
% Sat splits: 53
% Sat splitting refutations: 80
% 
% Memory used [KB]: 12792
% Time elapsed: 0.881 s
% ------------------------------
% ------------------------------
% Success in time 2.26 s
