Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.list
% (116037)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 660
% Passive clauses: 9573
% Generated clauses: 19024
% Final active clauses: 573
% Final passive clauses: 8389
% Input formulas: 69
% Initial clauses: 243
% 
% Duplicate literals: 274
% Trivial inequalities: 130
% Fw subsumption resolutions: 1584
% Fw demodulations: 380
% Bw demodulations: 86
% 
% Simple tautologies: 240
% Equational tautologies: 67
% Forward subsumptions: 6538
% 
% Binary resolution: 10830
% Factoring: 1002
% Forward superposition: 2015
% Backward superposition: 2204
% Self superposition: 27
% Equality factoring: 26
% Equality resolution: 39
% 
% Split clauses: 289
% Split components: 579
% SAT solver clauses: 590
% SAT solver unit clauses: 19
% SAT solver binary clauses: 143
% 
% Sat splits: 289
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 588
% TWLsolver calls for satisfiability: 335
% 
% Memory used [KB]: 8955
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (116053)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 474
% Passive clauses: 3324
% Generated clauses: 13027
% Final active clauses: 460
% Final passive clauses: 2804
% Input formulas: 69
% Initial clauses: 243
% 
% Duplicate literals: 186
% Trivial inequalities: 39
% Fw subsumption resolutions: 971
% Fw demodulations: 706
% Bw demodulations: 57
% Condensations: 231
% Global subsumptions: 22
% 
% Simple tautologies: 284
% Equational tautologies: 18
% Forward subsumptions: 3539
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 4182
% Factoring: 132
% Forward superposition: 1533
% Backward superposition: 1317
% Self superposition: 26
% Equality factoring: 26
% Equality resolution: 9
% 
% Split clauses: 40
% Split components: 80
% SAT solver clauses: 3958
% SAT solver unit clauses: 242
% SAT solver binary clauses: 94
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 40
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 3870
% TWLsolver calls for satisfiability: 1950
% 
% Memory used [KB]: 14839
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f11493,plain,(
  $false),
  inference(subsumption_resolution,[],[f11492,f5000])).
fof(f5000,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f356,f355,f357,f364,f366,f376,f375,f430,f441,f473,f470,f484,f483,f612,f502,f499,f498,f497,f516,f529,f528,f527,f526,f525,f524,f523,f533,f532,f531,f530,f537,f536,f535,f534,f541,f540,f539,f538,f545,f544,f543,f542,f599,f549,f548,f547,f546,f556,f555,f554,f553,f562,f559,f558,f557,f569,f568,f567,f598,f565,f564,f563,f573,f572,f571,f570,f577,f576,f575,f574,f586,f585,f584,f597,f582,f581,f580,f590,f589,f588,f587,f594,f593,f592,f591,f614,f613,f616,f615,f618,f617,f620,f619,f622,f621,f624,f623,f626,f625,f628,f627,f630,f629,f632,f631,f634,f633,f636,f635,f638,f637,f640,f639,f642,f641,f644,f643,f646,f645,f648,f647,f650,f649,f652,f651,f654,f653,f657,f687,f688,f691,f694,f697,f700,f703,f712,f717,f718,f721,f724,f727,f734,f735,f737,f740,f747,f750,f754,f755,f761,f762,f765,f768,f779,f782,f786,f797,f795,f800,f815,f813,f818,f828,f826,f836,f882,f891,f889,f906,f933,f931,f940,f939,f938,f937,f950,f949,f948,f947,f963,f962,f959,f958,f976,f975,f974,f973,f991,f990,f989,f988,f993,f1034,f1033,f1032,f1031,f1063,f1064,f1065,f1066,f1067,f1068,f1069,f1070,f1071,f1072,f1089,f1090,f1099,f1115,f1124,f1130,f1139,f1144,f1143,f1154,f1163,f1173,f1191,f1192,f1195,f1202,f1229,f1228,f1238,f1247,f1297,f1298,f1299,f1309,f1322,f1344,f1343,f1353,f1362,f1401,f1402,f1403,f1416,f1417,f1421,f1432,f1464,f1462,f1461,f1458,f1457,f1455,f1452,f1451,f1517,f1515,f1511,f1510,f1503,f1501,f1498,f1581,f1579,f1570,f1562,f1559,f1548,f1644,f1645,f1646,f1643,f1713,f1801,f1846,f1857,f1882,f1887,f1915,f1922,f1937,f1951,f2022,f2073,f2100,f2148,f2149,f2641,f2640,f2639,f2638,f2637,f2636,f2635,f2634,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f2626,f2625,f2624,f2623,f2619,f2618,f2614,f2611,f2610,f2605,f2604,f2598,f2597,f2576,f2575,f2574,f2570,f2569,f2565,f2564,f2563,f2558,f2548,f2542,f2539,f2538,f2534,f2533,f2532,f2531,f2530,f2527,f2525,f2520,f2519,f2514,f2513,f2507,f2505,f2504,f2499,f2496,f2495,f2494,f2493,f2492,f2491,f2490,f2489,f2488,f2487,f2486,f2485,f2484,f2483,f2482,f2481,f2480,f2479,f2474,f2473,f2467,f2466,f2461,f2460,f2454,f2453,f2451,f2449,f2447,f2445,f2444,f2441,f2440,f2439,f2438,f2436,f2435,f2434,f2433,f2432,f2726,f2761,f2778,f2811,f2829,f2893,f2914,f2929,f2943,f2958,f2973,f2987,f3003,f3019,f3057,f3194,f3305,f3938,f3942,f3949,f3952,f3955,f3958,f3961,f3964,f3967,f3970,f3973,f3976,f3979,f3982,f3985,f3987,f3990,f3991,f3992,f3993,f3994,f3995,f3996,f3997,f3998,f3999,f4000,f4001,f4002,f4003,f4004,f4005,f4011,f3481,f4030,f4032,f4038,f4040,f4042,f4044,f4046,f4048,f4050,f4052,f4054,f4056,f4058,f4060,f4062,f4025,f460,f2792,f2008,f2927,f1743,f2925,f601,f4024,f433,f4023,f3339,f3480,f3478,f3477,f3441,f3193,f3191,f3190,f3154,f2002,f2408,f2391,f2004,f2003,f2001,f2000,f1677,f2407,f2390,f1799,f1765,f1608,f1593,f2988,f469,f2744,f2743,f2425,f2424,f2135,f2134,f2374,f2373,f446,f4022,f4027,f4026,f4021,f4020,f4019,f4018,f4017,f4016,f4015,f4014,f4013,f4012,f3989,f3988,f2793,f1839,f1838,f1820,f1819,f461,f459,f458,f457,f456,f455,f454,f452,f450,f448,f437,f436,f354,f4362,f2111,f4010,f1714,f3343,f3479,f3348,f1715,f3344,f3447,f3446,f3445,f3341,f3159,f3158,f3448,f3350,f3192,f3349,f3449,f3452,f3451,f3450,f3347,f3346,f3165,f3164,f3163,f3162,f3161,f3444,f3443,f3345,f3342,f3160,f3157,f3156,f2959,f468,f445,f1856,f2427,f463,f1712,f466,f444,f467,f465,f4007,f4008,f4028,f464,f462,f4006,f4009,f3189,f3476,f2868,f2853,f2843,f2833,f2077,f2076,f2034,f2033,f2026,f2025,f1918,f1917,f1893,f1892,f1869,f1868,f1850,f1706,f1705,f1699,f1698,f1675,f1637,f1636,f1630,f1629,f1606,f1591,f674,f610,f609,f380,f379,f4999])).
fof(f4999,plain,(
  ( ! [X9] : (~p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f4772,f366])).
fof(f4772,plain,(
  ( ! [X9] : (~p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f379,f366])).
fof(f379,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.list.EVERY_MEM_MONO','thm.bool.EQ_CLAUSES|split|3')).
fof(f380,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f236])).
fof(f609,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f360,f357])).
fof(f360,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f64])).
fof(f64,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.list.EVERY_MEM_MONO','thm.bool.NOT_CLAUSES|split|1')).
fof(f610,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(forward_demodulation,[],[f431,f364])).
fof(f431,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : ! [X1] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X10] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.NIL_0')))))),
  file('thm.list.EVERY_MEM_MONO','thm.list.MEM|split|0')).
fof(f674,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f356,f380])).
fof(f1591,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK18(X0,bool))) )),
  inference(superposition,[],[f436,f366])).
fof(f1606,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK19(X0,bool))) )),
  inference(superposition,[],[f437,f366])).
fof(f1629,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X1),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f463,f366])).
fof(f1630,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X4),s(fun(bool,bool),X5)))))) | s(bool,X3) != s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f463,f366])).
fof(f1636,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f463,f366])).
fof(f1637,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))))) | s(bool,X3) != s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f463,f366])).
fof(f1675,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X1),s('type.list.list'(bool),X0)))) | s(bool,'T_0') = s(bool,sK20(X0,X1,bool))) )),
  inference(superposition,[],[f446,f366])).
fof(f1698,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X1),s('type.list.list'(bool),X2)))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f466,f366])).
fof(f1699,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X4),s('type.list.list'(bool),X5)))))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f466,f366])).
fof(f1705,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X0),s('type.list.list'(bool),X2)))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f466,f366])).
fof(f1706,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X5)))))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f466,f366])).
fof(f1850,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X0)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X1),s('type.list.list'(bool),X0)))) | s(bool,'T_0') = s(bool,sK20(X0,X1,bool))) )),
  inference(superposition,[],[f445,f366])).
fof(f1868,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X1),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X2)))))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f459,f366])).
fof(f1869,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X4),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X3),s('type.list.list'(bool),X5)))))) | p(s(bool,happ(s(fun(bool,bool),X4),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f459,f366])).
fof(f1892,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X2),s(fun(bool,bool),X1)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f464,f366])).
fof(f1893,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X5),s(fun(bool,bool),X4)))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f464,f366])).
fof(f1917,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK29(X0,X1,X2,bool))) )),
  inference(superposition,[],[f516,f366])).
fof(f1918,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK29(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK29(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f516,f366])).
fof(f2025,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X1),s(fun(bool,bool),X2)))))) | s(bool,X0) = s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f462,f366])).
fof(f2026,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X4),s(fun(bool,bool),X5)))))) | s(bool,'F_0') = s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f462,f366])).
fof(f2033,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) | s(bool,X0) = s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f462,f366])).
fof(f2034,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X3),s(fun(bool,bool),X5)))))) | s(bool,'F_0') = s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f462,f366])).
fof(f2076,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X1)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X2),s('type.list.list'(bool),X1)))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f467,f366])).
fof(f2077,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X5),s('type.list.list'(bool),X4)))))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f467,f366])).
fof(f2833,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X7),s('type.list.list'(bool),X8)))))))) | s(bool,'F_0') = s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X8)))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f465,f366])).
fof(f2843,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X6),s('type.list.list'(bool),X8)))))))) | s(bool,'F_0') = s(bool,X7) | p(s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X8)))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f465,f366])).
fof(f2853,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X1),s('type.list.list'(bool),X2)))))))) | s(bool,X0) = s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X2)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f2831,f674])).
fof(f2831,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X1),s('type.list.list'(bool),X2)))))))) | s(bool,X0) = s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X2)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f465,f2149])).
fof(f2868,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X2)))))))) | s(bool,X0) = s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X2)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f2841,f674])).
fof(f2841,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15))),s('type.list.list'(bool),X2)))))))) | s(bool,X0) = s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X2)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f465,f2149])).
fof(f3476,plain,(
  ( ! [X352,X354,X351,X353] : (s(bool,'T_0') != s(bool,f(s(X351,sK29(X352,X353,bool,X351)))) | s(fun(X351,bool),X352) = s(fun(X351,bool),X353) | s(bool,X354) = s(bool,g(s(X351,sK29(X352,X353,bool,X351)))) | s(bool,'T_0') = s(bool,X354)) )),
  inference(superposition,[],[f516,f836])).
fof(f3189,plain,(
  ( ! [X356,X358,X357,X359] : (s(bool,X359) != s(bool,f(s(X356,sK29(X357,X358,bool,X356)))) | s(fun(X356,bool),X357) = s(fun(X356,bool),X358) | s(bool,'T_0') = s(bool,X359) | s(bool,'T_0') = s(bool,g(s(X356,sK29(X357,X358,bool,X356))))) )),
  inference(superposition,[],[f516,f836])).
fof(f4009,plain,(
  ( ! [X118,X116,X114,X117,X115] : (s(X114,X115) = s(X114,X116) | p(s(bool,'const.bool.IN_2'(s(X114,X115),s(fun(X114,bool),X117)))) | s(bool,X118) = s(bool,'const.bool.IN_2'(s(X114,X115),s(fun(X114,bool),'const.pred_set.INSERT_2'(s(X114,X116),s(fun(X114,bool),X117))))) | s(bool,'T_0') = s(bool,X118)) )),
  inference(subsumption_resolution,[],[f3410,f357])).
fof(f3410,plain,(
  ( ! [X118,X116,X114,X117,X115] : (~p(s(bool,'T_0')) | s(X114,X115) = s(X114,X116) | p(s(bool,'const.bool.IN_2'(s(X114,X115),s(fun(X114,bool),X117)))) | s(bool,X118) = s(bool,'const.bool.IN_2'(s(X114,X115),s(fun(X114,bool),'const.pred_set.INSERT_2'(s(X114,X116),s(fun(X114,bool),X117))))) | s(bool,'T_0') = s(bool,X118)) )),
  inference(superposition,[],[f462,f836])).
fof(f4006,plain,(
  ( ! [X99,X97,X100,X98,X96] : (p(s(bool,'const.bool.IN_2'(s(X96,X97),s(fun(X96,bool),'const.pred_set.INSERT_2'(s(X96,X99),s(fun(X96,bool),X98)))))) | s(bool,X100) = s(bool,'const.bool.IN_2'(s(X96,X97),s(fun(X96,bool),X98))) | s(bool,'T_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f3406,f357])).
fof(f3406,plain,(
  ( ! [X99,X97,X100,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X96,X97),s(fun(X96,bool),'const.pred_set.INSERT_2'(s(X96,X99),s(fun(X96,bool),X98)))))) | s(bool,X100) = s(bool,'const.bool.IN_2'(s(X96,X97),s(fun(X96,bool),X98))) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f464,f836])).
fof(f462,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(flattening,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ! [X10,X18,X19] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X18),s(fun(X6,bool),X19)))))) <=> (s(X6,X10) = s(X6,X18) | p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X19))))))),
  file('thm.list.EVERY_MEM_MONO','thm.pred_set.IN_INSERT')).
fof(f464,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))) )),
  inference(cnf_transformation,[],[f296])).
fof(f4028,plain,(
  ( ! [X208,X206,X209,X207,X205] : (~p(s(bool,'const.bool.IN_2'(s(X205,X208),s(fun(X205,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X205),X207)))))) | p(s(bool,happ(s(fun(X205,bool),X206),s(X205,X208)))) | s(bool,X209) = s(bool,'const.list.EVERY_2'(s(fun(X205,bool),X206),s('type.list.list'(X205),X207))) | s(bool,'T_0') = s(bool,X209)) )),
  inference(subsumption_resolution,[],[f3433,f357])).
fof(f3433,plain,(
  ( ! [X208,X206,X209,X207,X205] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X205,X208),s(fun(X205,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X205),X207)))))) | p(s(bool,happ(s(fun(X205,bool),X206),s(X205,X208)))) | s(bool,X209) = s(bool,'const.list.EVERY_2'(s(fun(X205,bool),X206),s('type.list.list'(X205),X207))) | s(bool,'T_0') = s(bool,X209)) )),
  inference(superposition,[],[f444,f836])).
fof(f4008,plain,(
  ( ! [X109,X107,X110,X108,X106] : (s(X106,X107) = s(X106,X108) | p(s(bool,'const.bool.IN_2'(s(X106,X107),s(fun(X106,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X106),X109)))))) | s(bool,X110) = s(bool,'const.bool.IN_2'(s(X106,X107),s(fun(X106,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X106),'const.list.CONS_2'(s(X106,X108),s('type.list.list'(X106),X109))))))) | s(bool,'T_0') = s(bool,X110)) )),
  inference(subsumption_resolution,[],[f3408,f357])).
fof(f3408,plain,(
  ( ! [X109,X107,X110,X108,X106] : (~p(s(bool,'T_0')) | s(X106,X107) = s(X106,X108) | p(s(bool,'const.bool.IN_2'(s(X106,X107),s(fun(X106,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X106),X109)))))) | s(bool,X110) = s(bool,'const.bool.IN_2'(s(X106,X107),s(fun(X106,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X106),'const.list.CONS_2'(s(X106,X108),s('type.list.list'(X106),X109))))))) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f465,f836])).
fof(f4007,plain,(
  ( ! [X103,X101,X105,X102,X104] : (p(s(bool,'const.bool.IN_2'(s(X101,X102),s(fun(X101,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X101),'const.list.CONS_2'(s(X101,X104),s('type.list.list'(X101),X103)))))))) | s(bool,X105) = s(bool,'const.bool.IN_2'(s(X101,X102),s(fun(X101,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X101),X103))))) | s(bool,'T_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f3407,f357])).
fof(f3407,plain,(
  ( ! [X103,X101,X105,X102,X104] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X101,X102),s(fun(X101,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X101),'const.list.CONS_2'(s(X101,X104),s('type.list.list'(X101),X103)))))))) | s(bool,X105) = s(bool,'const.bool.IN_2'(s(X101,X102),s(fun(X101,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X101),X103))))) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f467,f836])).
fof(f465,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3)))))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X10,X17,X15] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X17),s('type.list.list'(X6),X15)))))))) <=> (s(X6,X10) = s(X6,X17) | p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X15))))))))),
  file('thm.list.EVERY_MEM_MONO','thm.list.MEM|split|1')).
fof(f467,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))) )),
  inference(cnf_transformation,[],[f298])).
fof(f444,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK20(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X2,X1,X0)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK20])],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(rectify,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(ennf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X7,X9] : (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X9)))) <=> ! [X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X9)))))) => p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X11))))))),
  file('thm.list.EVERY_MEM_MONO','thm.list.EVERY_MEM')).
fof(f466,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))) )),
  inference(cnf_transformation,[],[f298])).
fof(f1712,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))))))) )),
  inference(equality_resolution,[],[f466])).
fof(f463,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))) )),
  inference(cnf_transformation,[],[f296])).
fof(f2427,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X3))))))))) )),
  inference(unit_resulting_resolution,[],[f1643,f464])).
fof(f1856,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK20(X1,X2,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) )),
  inference(subsumption_resolution,[],[f1852,f609])).
fof(f1852,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK20(X1,X2,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) )),
  inference(superposition,[],[f445,f366])).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK20(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f280])).
fof(f468,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | p(s(bool,'const.bool.IN_2'(s(X0,sK28(X4,X3,X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | (p(s(bool,'const.bool.IN_2'(s(X0,sK28(X4,X3,X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & s(bool,happ(s(fun(X0,bool),X3),s(X0,sK28(X4,X3,X2,X1,X0)))) != s(bool,happ(s(fun(X0,bool),X4),s(X0,sK28(X4,X3,X2,X1,X0))))) | s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK28])],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4] : (s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) != s(bool,happ(s(fun(X0,bool),X4),s(X0,X5)))) | s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) != s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) | s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))))),
  inference(ennf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2,X3,X4] : ((s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) => s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : ((s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) => s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X12,X13,X7,X14] : ((s('type.list.list'(X6),X12) = s('type.list.list'(X6),X13) & ! [X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X13)))))) => s(bool,happ(s(fun(X6,bool),X7),s(X6,X10))) = s(bool,happ(s(fun(X6,bool),X14),s(X6,X10))))) => s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X12))) = s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X14),s('type.list.list'(X6),X13))))),
  file('thm.list.EVERY_MEM_MONO','thm.list.EVERY_CONG')).
fof(f2959,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK28(X1,X2,X3,X3,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3)))))) | s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X3)))) )),
  inference(equality_resolution,[],[f468])).
fof(f3156,plain,(
  ( ! [X251,X249,X252,X250] : (~p(s(bool,'const.bool.IN_2'(s(bool,X250),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X251),s(fun(bool,bool),X252)))))) | s(bool,X249) = s(bool,X251) | p(s(bool,'const.bool.IN_2'(s(bool,X249),s(fun(bool,bool),X252)))) | s(bool,'T_0') = s(bool,X250) | s(bool,'T_0') = s(bool,X249)) )),
  inference(superposition,[],[f462,f836])).
fof(f3157,plain,(
  ( ! [X255,X253,X254,X256] : (~p(s(bool,'const.bool.IN_2'(s(bool,X255),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X254),s(fun(bool,bool),X256)))))) | s(bool,X253) = s(bool,X255) | p(s(bool,'const.bool.IN_2'(s(bool,X255),s(fun(bool,bool),X256)))) | s(bool,'T_0') = s(bool,X254) | s(bool,'T_0') = s(bool,X253)) )),
  inference(superposition,[],[f462,f836])).
fof(f3160,plain,(
  ( ! [X265,X267,X268,X266] : (~p(s(bool,'const.bool.IN_2'(s(bool,X266),s(fun(bool,bool),X267)))) | p(s(bool,'const.bool.IN_2'(s(bool,X265),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X268),s(fun(bool,bool),X267)))))) | s(bool,'T_0') = s(bool,X266) | s(bool,'T_0') = s(bool,X265)) )),
  inference(superposition,[],[f464,f836])).
fof(f3342,plain,(
  ( ! [X134,X132,X135,X133] : (~p(s(bool,'const.bool.IN_2'(s(bool,X133),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X132),s(fun(bool,bool),X134)))))) | s(bool,'T_0') = s(bool,X133) | p(s(bool,'const.bool.IN_2'(s(bool,X133),s(fun(bool,bool),X134)))) | s(bool,X132) = s(bool,X135) | s(bool,'T_0') = s(bool,X135)) )),
  inference(superposition,[],[f462,f836])).
fof(f3345,plain,(
  ( ! [X146,X144,X147,X145] : (~p(s(bool,'const.bool.IN_2'(s(bool,X144),s(fun(bool,bool),X145)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X146),s(fun(bool,bool),X145)))))) | s(bool,X144) = s(bool,X147) | s(bool,'T_0') = s(bool,X147)) )),
  inference(superposition,[],[f464,f836])).
fof(f3443,plain,(
  ( ! [X246,X248,X247,X245] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X246),s(fun(bool,bool),X247)))))) | s(bool,X245) = s(bool,X246) | p(s(bool,'const.bool.IN_2'(s(bool,X245),s(fun(bool,bool),X247)))) | s(bool,X245) = s(bool,X248) | s(bool,'T_0') = s(bool,X248)) )),
  inference(superposition,[],[f462,f836])).
fof(f3444,plain,(
  ( ! [X251,X249,X252,X250] : (~p(s(bool,'const.bool.IN_2'(s(bool,X250),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X251)))))) | s(bool,X249) = s(bool,X250) | p(s(bool,'const.bool.IN_2'(s(bool,X250),s(fun(bool,bool),X251)))) | s(bool,X249) = s(bool,X252) | s(bool,'T_0') = s(bool,X252)) )),
  inference(superposition,[],[f462,f836])).
fof(f3161,plain,(
  ( ! [X269,X271,X270,X272] : (~p(s(bool,'const.bool.IN_2'(s(bool,X270),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X271),s('type.list.list'(bool),X272)))))))) | s(bool,X269) = s(bool,X271) | p(s(bool,'const.bool.IN_2'(s(bool,X269),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X272)))))) | s(bool,'T_0') = s(bool,X270) | s(bool,'T_0') = s(bool,X269)) )),
  inference(superposition,[],[f465,f836])).
fof(f3162,plain,(
  ( ! [X273,X275,X276,X274] : (~p(s(bool,'const.bool.IN_2'(s(bool,X275),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X274),s('type.list.list'(bool),X276)))))))) | s(bool,X273) = s(bool,X275) | p(s(bool,'const.bool.IN_2'(s(bool,X275),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X276)))))) | s(bool,'T_0') = s(bool,X274) | s(bool,'T_0') = s(bool,X273)) )),
  inference(superposition,[],[f465,f836])).
fof(f3163,plain,(
  ( ! [X280,X277,X279,X278] : (s(bool,X278) != s(bool,X279) | p(s(bool,'const.bool.IN_2'(s(bool,X277),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X279),s('type.list.list'(bool),X280)))))))) | s(bool,'T_0') = s(bool,X278) | s(bool,'T_0') = s(bool,X277)) )),
  inference(superposition,[],[f466,f836])).
fof(f3164,plain,(
  ( ! [X282,X281,X283,X284] : (s(bool,X282) != s(bool,X283) | p(s(bool,'const.bool.IN_2'(s(bool,X283),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X281),s('type.list.list'(bool),X284)))))))) | s(bool,'T_0') = s(bool,X282) | s(bool,'T_0') = s(bool,X281)) )),
  inference(superposition,[],[f466,f836])).
fof(f3165,plain,(
  ( ! [X288,X285,X287,X286] : (~p(s(bool,'const.bool.IN_2'(s(bool,X286),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X287)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X285),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X288),s('type.list.list'(bool),X287)))))))) | s(bool,'T_0') = s(bool,X286) | s(bool,'T_0') = s(bool,X285)) )),
  inference(superposition,[],[f467,f836])).
fof(f3346,plain,(
  ( ! [X151,X149,X150,X148] : (~p(s(bool,'const.bool.IN_2'(s(bool,X148),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X149),s('type.list.list'(bool),X150)))))))) | s(bool,'T_0') = s(bool,X149) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X150)))))) | s(bool,X148) = s(bool,X151) | s(bool,'T_0') = s(bool,X151)) )),
  inference(superposition,[],[f465,f836])).
fof(f3347,plain,(
  ( ! [X154,X152,X155,X153] : (~p(s(bool,'const.bool.IN_2'(s(bool,X153),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X152),s('type.list.list'(bool),X154)))))))) | s(bool,'T_0') = s(bool,X153) | p(s(bool,'const.bool.IN_2'(s(bool,X153),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X154)))))) | s(bool,X152) = s(bool,X155) | s(bool,'T_0') = s(bool,X155)) )),
  inference(superposition,[],[f465,f836])).
fof(f3450,plain,(
  ( ! [X273,X275,X276,X274] : (s(bool,'T_0') != s(bool,X274) | p(s(bool,'const.bool.IN_2'(s(bool,X273),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X274),s('type.list.list'(bool),X275)))))))) | s(bool,X273) = s(bool,X276) | s(bool,'T_0') = s(bool,X276)) )),
  inference(superposition,[],[f466,f836])).
fof(f3451,plain,(
  ( ! [X280,X277,X279,X278] : (s(bool,'T_0') != s(bool,X278) | p(s(bool,'const.bool.IN_2'(s(bool,X278),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X277),s('type.list.list'(bool),X279)))))))) | s(bool,X277) = s(bool,X280) | s(bool,'T_0') = s(bool,X280)) )),
  inference(superposition,[],[f466,f836])).
fof(f3452,plain,(
  ( ! [X282,X281,X283,X284] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X282)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X281),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X283),s('type.list.list'(bool),X282)))))))) | s(bool,X281) = s(bool,X284) | s(bool,'T_0') = s(bool,X284)) )),
  inference(superposition,[],[f467,f836])).
fof(f3449,plain,(
  ( ! [X269,X271,X270,X272] : (~p(s(bool,'const.bool.IN_2'(s(bool,X270),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X271)))))))) | s(bool,X269) = s(bool,X270) | p(s(bool,'const.bool.IN_2'(s(bool,X270),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X271)))))) | s(bool,X269) = s(bool,X272) | s(bool,'T_0') = s(bool,X272)) )),
  inference(superposition,[],[f465,f836])).
fof(f3349,plain,(
  ( ! [X163,X161,X162,X160] : (s(bool,X160) != s(bool,X161) | p(s(bool,'const.bool.IN_2'(s(bool,X161),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X162)))))))) | s(bool,X160) = s(bool,X163) | s(bool,'T_0') = s(bool,X163)) )),
  inference(superposition,[],[f466,f836])).
fof(f3192,plain,(
  ( ! [X366,X368,X367] : (p(s(bool,'const.bool.IN_2'(s(bool,X368),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X366)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X367),s('type.list.list'(bool),X366)))) | s(bool,'T_0') = s(bool,X368) | s(bool,'T_0') = s(bool,sK20(X366,X367,bool))) )),
  inference(superposition,[],[f445,f836])).
fof(f3350,plain,(
  ( ! [X167,X165,X166,X164] : (~p(s(bool,'const.bool.IN_2'(s(bool,X164),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X165)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X166),s('type.list.list'(bool),X165)))))))) | s(bool,X164) = s(bool,X167) | s(bool,'T_0') = s(bool,X167)) )),
  inference(superposition,[],[f467,f836])).
fof(f3448,plain,(
  ( ! [X265,X267,X268,X266] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X266),s('type.list.list'(bool),X267)))))))) | s(bool,X265) = s(bool,X266) | p(s(bool,'const.bool.IN_2'(s(bool,X265),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X267)))))) | s(bool,X265) = s(bool,X268) | s(bool,'T_0') = s(bool,X268)) )),
  inference(superposition,[],[f465,f836])).
fof(f3158,plain,(
  ( ! [X257,X259,X260,X258] : (s(bool,X258) != s(bool,X259) | p(s(bool,'const.bool.IN_2'(s(bool,X257),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X259),s(fun(bool,bool),X260)))))) | s(bool,'T_0') = s(bool,X258) | s(bool,'T_0') = s(bool,X257)) )),
  inference(superposition,[],[f463,f836])).
fof(f3159,plain,(
  ( ! [X261,X263,X262,X264] : (s(bool,X262) != s(bool,X263) | p(s(bool,'const.bool.IN_2'(s(bool,X263),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X261),s(fun(bool,bool),X264)))))) | s(bool,'T_0') = s(bool,X262) | s(bool,'T_0') = s(bool,X261)) )),
  inference(superposition,[],[f463,f836])).
fof(f3341,plain,(
  ( ! [X130,X128,X131,X129] : (~p(s(bool,'const.bool.IN_2'(s(bool,X128),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X129),s(fun(bool,bool),X130)))))) | s(bool,'T_0') = s(bool,X129) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X130)))) | s(bool,X128) = s(bool,X131) | s(bool,'T_0') = s(bool,X131)) )),
  inference(superposition,[],[f462,f836])).
fof(f3445,plain,(
  ( ! [X255,X253,X254,X256] : (s(bool,'T_0') != s(bool,X254) | p(s(bool,'const.bool.IN_2'(s(bool,X253),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X254),s(fun(bool,bool),X255)))))) | s(bool,X253) = s(bool,X256) | s(bool,'T_0') = s(bool,X256)) )),
  inference(superposition,[],[f463,f836])).
fof(f3446,plain,(
  ( ! [X257,X259,X260,X258] : (s(bool,'T_0') != s(bool,X258) | p(s(bool,'const.bool.IN_2'(s(bool,X258),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X257),s(fun(bool,bool),X259)))))) | s(bool,X257) = s(bool,X260) | s(bool,'T_0') = s(bool,X260)) )),
  inference(superposition,[],[f463,f836])).
fof(f3447,plain,(
  ( ! [X261,X263,X262,X264] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X262)))) | p(s(bool,'const.bool.IN_2'(s(bool,X261),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X263),s(fun(bool,bool),X262)))))) | s(bool,X261) = s(bool,X264) | s(bool,'T_0') = s(bool,X264)) )),
  inference(superposition,[],[f464,f836])).
fof(f3344,plain,(
  ( ! [X142,X140,X143,X141] : (s(bool,X140) != s(bool,X141) | p(s(bool,'const.bool.IN_2'(s(bool,X141),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X142)))))) | s(bool,X140) = s(bool,X143) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f463,f836])).
fof(f1715,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X2))))))))) )),
  inference(forward_demodulation,[],[f1695,f430])).
fof(f1695,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X2))))))))) )),
  inference(unit_resulting_resolution,[],[f430,f466])).
fof(f3348,plain,(
  ( ! [X158,X156,X159,X157] : (s(bool,X156) != s(bool,X157) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X157),s('type.list.list'(bool),X158)))))))) | s(bool,X156) = s(bool,X159) | s(bool,'T_0') = s(bool,X159)) )),
  inference(superposition,[],[f466,f836])).
fof(f3479,plain,(
  ( ! [X362,X361,X363] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X361)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X362),s('type.list.list'(bool),X361)))) | s(bool,X363) = s(bool,sK20(X361,X362,bool)) | s(bool,'T_0') = s(bool,X363)) )),
  inference(superposition,[],[f445,f836])).
fof(f3343,plain,(
  ( ! [X138,X136,X139,X137] : (s(bool,X136) != s(bool,X137) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X137),s(fun(bool,bool),X138)))))) | s(bool,X136) = s(bool,X139) | s(bool,'T_0') = s(bool,X139)) )),
  inference(superposition,[],[f463,f836])).
fof(f1714,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(fun(X0,bool)),'const.list.CONS_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s('type.list.list'(fun(X0,bool)),X1))))))))) )),
  inference(forward_demodulation,[],[f1696,f364])).
fof(f1696,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0'))),s(fun(fun(X0,bool),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(fun(X0,bool)),'const.list.CONS_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s('type.list.list'(fun(X0,bool)),X1))))))))) )),
  inference(unit_resulting_resolution,[],[f364,f466])).
fof(f4010,plain,(
  ( ! [X123,X124] : (p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,X123)))) | s(bool,X124) = s(bool,'const.bool.IN_2'(s(sK12,X123),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK15))))) | s(bool,'T_0') = s(bool,X124)) )),
  inference(subsumption_resolution,[],[f3412,f357])).
fof(f3412,plain,(
  ( ! [X123,X124] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,X123)))) | s(bool,X124) = s(bool,'const.bool.IN_2'(s(sK12,X123),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK15))))) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f2111,f836])).
fof(f2111,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK15)))))) | p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,X0))))) )),
  inference(subsumption_resolution,[],[f2105,f357])).
fof(f2105,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK15)))))) | p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,X0))))) )),
  inference(superposition,[],[f444,f657])).
fof(f4362,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK15)))))) | p(s(bool,happ(s(fun(sK12,bool),sK14),s(sK12,X4))))) )),
  inference(global_subsumption,[],[f356,f355,f357,f609,f364,f366,f376,f375,f380,f379,f430,f610,f441,f445,f444,f464,f463,f462,f467,f466,f465,f468,f473,f470,f484,f483,f612,f502,f499,f498,f497,f516,f529,f528,f527,f526,f525,f524,f523,f533,f532,f531,f530,f537,f536,f535,f534,f541,f540,f539,f538,f545,f544,f543,f542,f599,f549,f548,f547,f546,f556,f555,f554,f553,f562,f559,f558,f557,f569,f568,f567,f598,f565,f564,f563,f573,f572,f571,f570,f577,f576,f575,f574,f586,f585,f584,f597,f582,f581,f580,f590,f589,f588,f587,f594,f593,f592,f591,f614,f613,f616,f615,f618,f617,f620,f619,f622,f621,f624,f623,f626,f625,f628,f627,f630,f629,f632,f631,f634,f633,f636,f635,f638,f637,f640,f639,f642,f641,f644,f643,f646,f645,f648,f647,f650,f649,f652,f651,f654,f653,f657,f674,f687,f688,f691,f694,f697,f700,f703,f712,f717,f718,f721,f724,f727,f734,f735,f737,f740,f747,f750,f754,f755,f761,f762,f765,f768,f779,f782,f786,f797,f795,f800,f815,f813,f818,f828,f826,f836,f882,f891,f889,f906,f933,f931,f940,f939,f938,f937,f950,f949,f948,f947,f963,f962,f959,f958,f976,f975,f974,f973,f991,f990,f989,f988,f993,f1034,f1033,f1032,f1031,f1063,f1064,f1065,f1066,f1067,f1068,f1069,f1070,f1071,f1072,f1089,f1090,f1099,f1115,f1124,f1130,f1139,f1144,f1143,f1154,f1163,f1173,f1191,f1192,f1195,f1202,f1229,f1228,f1238,f1247,f1297,f1298,f1299,f1309,f1322,f1344,f1343,f1353,f1362,f1401,f1402,f1403,f1416,f1417,f1421,f1432,f1464,f1462,f1461,f1458,f1457,f1455,f1452,f1451,f1517,f1515,f1511,f1510,f1503,f1501,f1498,f1581,f1579,f1570,f1562,f1559,f1548,f1591,f1606,f1644,f1645,f1646,f1629,f1630,f1636,f1637,f1643,f1675,f1713,f1714,f1715,f1698,f1699,f1705,f1706,f1712,f1801,f1846,f1850,f1856,f1857,f1868,f1869,f1882,f1887,f1892,f1893,f1915,f1917,f1918,f1922,f1937,f1951,f2022,f2025,f2026,f2033,f2034,f2073,f2076,f2077,f2100,f2111,f2148,f2149,f2641,f2640,f2639,f2638,f2637,f2636,f2635,f2634,f2633,f2632,f2631,f2630,f2629,f2628,f2627,f2626,f2625,f2624,f2623,f2619,f2618,f2614,f2611,f2610,f2605,f2604,f2598,f2597,f2576,f2575,f2574,f2570,f2569,f2565,f2564,f2563,f2558,f2548,f2542,f2539,f2538,f2534,f2533,f2532,f2531,f2530,f2527,f2525,f2520,f2519,f2514,f2513,f2507,f2505,f2504,f2499,f2496,f2495,f2494,f2493,f2492,f2491,f2490,f2489,f2488,f2487,f2486,f2485,f2484,f2483,f2482,f2481,f2480,f2479,f2474,f2473,f2467,f2466,f2461,f2460,f2454,f2453,f2451,f2449,f2447,f2445,f2444,f2441,f2440,f2439,f2438,f2436,f2435,f2434,f2433,f2432,f2427,f2726,f2761,f2778,f2811,f2829,f2853,f2833,f2868,f2843,f2893,f2914,f2929,f2943,f2958,f2959,f2973,f2987,f3003,f3019,f3057,f3156,f3157,f3158,f3159,f3160,f3161,f3162,f3163,f3164,f3165,f3189,f3192,f3194,f3305,f3938,f3942,f3949,f3952,f3955,f3958,f3961,f3964,f3967,f3970,f3973,f3976,f3979,f3982,f3985,f3341,f3342,f3343,f3344,f3345,f3346,f3347,f3348,f3349,f3350,f3987,f3990,f3991,f3992,f3993,f3994,f3995,f3996,f3997,f3998,f3999,f4000,f4001,f4002,f4003,f4004,f4005,f4006,f4007,f4008,f4009,f4010,f4011,f4028,f3443,f3444,f3445,f3446,f3447,f3448,f3449,f3450,f3451,f3452,f3476,f3479,f3481,f4030,f4032,f4038,f4040,f4042,f4044,f4046,f4048,f4050,f4052,f4054,f4056,f4058,f4060,f4062,f4025,f460,f2792,f2008,f2927,f1743,f2925,f601,f4024,f433,f4023,f3339,f3480,f3478,f3477,f3441,f3193,f3191,f3190,f3154,f2002,f2408,f2391,f2004,f2003,f2001,f2000,f1677,f2407,f2390,f1799,f1765,f1608,f1593,f2988,f469,f2744,f2743,f2425,f2424,f2135,f2134,f2374,f2373,f446,f4022,f4027,f4026,f4021,f4020,f4019,f4018,f4017,f4016,f4015,f4014,f4013,f4012,f3989,f3988,f2793,f1839,f1838,f1820,f1819,f461,f459,f458,f457,f456,f455,f454,f452,f450,f448,f437,f436,f354])).
fof(f354,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,X4)))) | p(s(bool,happ(s(fun(sK12,bool),sK14),s(sK12,X4)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK15))))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK15)))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,X4)))) | p(s(bool,happ(s(fun(sK12,bool),sK14),s(sK12,X4))))) & p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15])],[f177])).
fof(f177,plain,(
  ? [X0,X1,X2,X3] : (! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) & ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X3)))))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ? [X0,X1,X2,X3] : ((! [X4] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))) & ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X3)))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : ((! [X4] : ((p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))) => p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X3)))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3] : ((! [X4] : ((p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))) => p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X3)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ((! [X10] : ((p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X9)))))) & p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X10))))) => p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))) & p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X9))))) => p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X8),s('type.list.list'(X6),X9)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ((! [X10] : ((p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X9)))))) & p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X10))))) => p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))) & p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X9))))) => p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X8),s('type.list.list'(X6),X9)))))),
  file('thm.list.EVERY_MEM_MONO','thm.list.EVERY_MEM_MONO')).
fof(f436,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19])],[f273])).
fof(f273,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f272])).
fof(f272,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X6] : ! [X7] : (~! [X10] : p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X10)))) <=> ? [X10] : ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X10)))))),
  file('thm.list.EVERY_MEM_MONO','thm.bool.NOT_FORALL_THM')).
fof(f437,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f274])).
fof(f448,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21,sK22])],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,X7)) | ? [X10] : p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))) <=> ? [X10] : (p(s(bool,X7)) | p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))))),
  file('thm.list.EVERY_MEM_MONO','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f450,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f284])).
fof(f452,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f288])).
fof(f288,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23,sK24])],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & ? [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X6] : ! [X7,X8] : (? [X10] : (p(s(bool,X7)) & p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))) <=> (p(s(bool,X7)) & ? [X10] : p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))))),
  file('thm.list.EVERY_MEM_MONO','thm.bool.RIGHT_EXISTS_AND_THM')).
fof(f454,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f288])).
fof(f455,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK25(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK26(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK25(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK26(X2,X1,X0)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK27(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK27(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK25,sK26,sK27])],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ? [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X6] : ! [X7,X8] : (? [X10] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X10)))) | p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))) <=> (? [X10] : p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X10)))) | ? [X10] : p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10))))))),
  file('thm.list.EVERY_MEM_MONO','thm.bool.EXISTS_OR_THM')).
fof(f456,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK25(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK26(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f292])).
fof(f457,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK27(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK27(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f292])).
fof(f458,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK27(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK27(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f292])).
fof(f459,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))),
  inference(flattening,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))))) & ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))),
  inference(nnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) <=> (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) <=> (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X7,X17,X15] : (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X17),s('type.list.list'(X6),X15)))))) <=> (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X17)))) & p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X15))))))),
  file('thm.list.EVERY_MEM_MONO','thm.list.EVERY_DEF|split|1')).
fof(f461,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f294])).
fof(f1819,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,X15)))) | p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK23(X14,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X13)))))) )),
  inference(subsumption_resolution,[],[f1816,f357])).
fof(f1816,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,X15)))) | p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK23(X14,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X13)))))) )),
  inference(superposition,[],[f452,f657])).
fof(f1820,plain,(
  ( ! [X23,X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(X21,bool),X22),s(X21,X23)))) | p(s(bool,happ(s(fun(X21,bool),X22),s(X21,sK23(X22,'const.list.EVERY_2'(s(fun(X19,bool),X20),s('type.list.list'(X19),'const.list.NIL_0')),X21)))))) )),
  inference(subsumption_resolution,[],[f1818,f357])).
fof(f1818,plain,(
  ( ! [X23,X21,X19,X22,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X21,bool),X22),s(X21,X23)))) | p(s(bool,happ(s(fun(X21,bool),X22),s(X21,sK23(X22,'const.list.EVERY_2'(s(fun(X19,bool),X20),s('type.list.list'(X19),'const.list.NIL_0')),X21)))))) )),
  inference(superposition,[],[f452,f430])).
fof(f1838,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,X15)))) | p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK24(X14,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X13)))))) )),
  inference(subsumption_resolution,[],[f1835,f357])).
fof(f1835,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,X15)))) | p(s(bool,happ(s(fun(X13,bool),X14),s(X13,sK24(X14,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X13)))))) )),
  inference(superposition,[],[f454,f657])).
fof(f1839,plain,(
  ( ! [X23,X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(X21,bool),X22),s(X21,X23)))) | p(s(bool,happ(s(fun(X21,bool),X22),s(X21,sK24(X22,'const.list.EVERY_2'(s(fun(X19,bool),X20),s('type.list.list'(X19),'const.list.NIL_0')),X21)))))) )),
  inference(subsumption_resolution,[],[f1837,f357])).
fof(f1837,plain,(
  ( ! [X23,X21,X19,X22,X20] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X21,bool),X22),s(X21,X23)))) | p(s(bool,happ(s(fun(X21,bool),X22),s(X21,sK24(X22,'const.list.EVERY_2'(s(fun(X19,bool),X20),s('type.list.list'(X19),'const.list.NIL_0')),X21)))))) )),
  inference(superposition,[],[f454,f430])).
fof(f2793,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6)))) | p(s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.CONS_2'(s(X4,X6),s('type.list.list'(X4),'const.list.NIL_0'))))))) )),
  inference(subsumption_resolution,[],[f2785,f357])).
fof(f2785,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(X4,bool),X5),s('type.list.list'(X4),'const.list.CONS_2'(s(X4,X6),s('type.list.list'(X4),'const.list.NIL_0')))))) | ~p(s(bool,happ(s(fun(X4,bool),X5),s(X4,X6))))) )),
  inference(superposition,[],[f461,f430])).
fof(f3988,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,happ(s(fun(X11,bool),X12),s(X11,X13)))) | p(s(bool,happ(s(fun(X11,bool),X12),s(X11,sK23(X12,X10,X11))))) | s(bool,X10) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f3379,f357])).
fof(f3379,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X11,bool),X12),s(X11,X13)))) | p(s(bool,happ(s(fun(X11,bool),X12),s(X11,sK23(X12,X10,X11))))) | s(bool,X10) = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f452,f836])).
fof(f3989,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK24(X17,X15,X16))))) | s(bool,X15) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f3380,f357])).
fof(f3380,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK24(X17,X15,X16))))) | s(bool,X15) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f454,f836])).
fof(f4012,plain,(
  ( ! [X127,X130,X128,X131,X129] : (p(s(bool,X130)) | p(s(bool,happ(s(fun(X127,bool),X128),s(X127,sK21(X128,X130,X127))))) | s(bool,X131) = s(bool,happ(s(fun(X127,bool),X128),s(X127,X129))) | s(bool,'T_0') = s(bool,X131)) )),
  inference(subsumption_resolution,[],[f3414,f357])).
fof(f3414,plain,(
  ( ! [X127,X130,X128,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,X130)) | p(s(bool,happ(s(fun(X127,bool),X128),s(X127,sK21(X128,X130,X127))))) | s(bool,X131) = s(bool,happ(s(fun(X127,bool),X128),s(X127,X129))) | s(bool,'T_0') = s(bool,X131)) )),
  inference(superposition,[],[f448,f836])).
fof(f4013,plain,(
  ( ! [X134,X132,X136,X135,X133] : (p(s(bool,X135)) | p(s(bool,happ(s(fun(X132,bool),X133),s(X132,sK22(X133,X135,X132))))) | s(bool,X136) = s(bool,happ(s(fun(X132,bool),X133),s(X132,X134))) | s(bool,'T_0') = s(bool,X136)) )),
  inference(subsumption_resolution,[],[f3415,f357])).
fof(f3415,plain,(
  ( ! [X134,X132,X136,X135,X133] : (~p(s(bool,'T_0')) | p(s(bool,X135)) | p(s(bool,happ(s(fun(X132,bool),X133),s(X132,sK22(X133,X135,X132))))) | s(bool,X136) = s(bool,happ(s(fun(X132,bool),X133),s(X132,X134))) | s(bool,'T_0') = s(bool,X136)) )),
  inference(superposition,[],[f450,f836])).
fof(f4014,plain,(
  ( ! [X140,X138,X141,X139,X137] : (p(s(bool,happ(s(fun(X137,bool),X138),s(X137,sK25(X140,X138,X137))))) | p(s(bool,happ(s(fun(X137,bool),X140),s(X137,sK26(X140,X138,X137))))) | s(bool,X141) = s(bool,happ(s(fun(X137,bool),X138),s(X137,X139))) | s(bool,'T_0') = s(bool,X141)) )),
  inference(subsumption_resolution,[],[f3416,f357])).
fof(f3416,plain,(
  ( ! [X140,X138,X141,X139,X137] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X137,bool),X138),s(X137,sK25(X140,X138,X137))))) | p(s(bool,happ(s(fun(X137,bool),X140),s(X137,sK26(X140,X138,X137))))) | s(bool,X141) = s(bool,happ(s(fun(X137,bool),X138),s(X137,X139))) | s(bool,'T_0') = s(bool,X141)) )),
  inference(superposition,[],[f455,f836])).
fof(f4015,plain,(
  ( ! [X146,X144,X142,X145,X143] : (p(s(bool,happ(s(fun(X142,bool),X145),s(X142,sK25(X143,X145,X142))))) | p(s(bool,happ(s(fun(X142,bool),X143),s(X142,sK26(X143,X145,X142))))) | s(bool,X146) = s(bool,happ(s(fun(X142,bool),X143),s(X142,X144))) | s(bool,'T_0') = s(bool,X146)) )),
  inference(subsumption_resolution,[],[f3417,f357])).
fof(f3417,plain,(
  ( ! [X146,X144,X142,X145,X143] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X142,bool),X145),s(X142,sK25(X143,X145,X142))))) | p(s(bool,happ(s(fun(X142,bool),X143),s(X142,sK26(X143,X145,X142))))) | s(bool,X146) = s(bool,happ(s(fun(X142,bool),X143),s(X142,X144))) | s(bool,'T_0') = s(bool,X146)) )),
  inference(superposition,[],[f456,f836])).
fof(f4016,plain,(
  ( ! [X151,X149,X147,X150,X148] : (p(s(bool,happ(s(fun(X147,bool),X148),s(X147,sK27(X150,X148,X147))))) | p(s(bool,happ(s(fun(X147,bool),X150),s(X147,sK27(X150,X148,X147))))) | s(bool,X151) = s(bool,happ(s(fun(X147,bool),X148),s(X147,X149))) | s(bool,'T_0') = s(bool,X151)) )),
  inference(subsumption_resolution,[],[f3418,f357])).
fof(f3418,plain,(
  ( ! [X151,X149,X147,X150,X148] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X147,bool),X148),s(X147,sK27(X150,X148,X147))))) | p(s(bool,happ(s(fun(X147,bool),X150),s(X147,sK27(X150,X148,X147))))) | s(bool,X151) = s(bool,happ(s(fun(X147,bool),X148),s(X147,X149))) | s(bool,'T_0') = s(bool,X151)) )),
  inference(superposition,[],[f457,f836])).
fof(f4017,plain,(
  ( ! [X156,X154,X152,X155,X153] : (p(s(bool,happ(s(fun(X152,bool),X155),s(X152,sK27(X153,X155,X152))))) | p(s(bool,happ(s(fun(X152,bool),X153),s(X152,sK27(X153,X155,X152))))) | s(bool,X156) = s(bool,happ(s(fun(X152,bool),X153),s(X152,X154))) | s(bool,'T_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f3419,f357])).
fof(f3419,plain,(
  ( ! [X156,X154,X152,X155,X153] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X152,bool),X155),s(X152,sK27(X153,X155,X152))))) | p(s(bool,happ(s(fun(X152,bool),X153),s(X152,sK27(X153,X155,X152))))) | s(bool,X156) = s(bool,happ(s(fun(X152,bool),X153),s(X152,X154))) | s(bool,'T_0') = s(bool,X156)) )),
  inference(superposition,[],[f458,f836])).
fof(f4018,plain,(
  ( ! [X158,X159,X157,X160] : (p(s(bool,happ(s(fun(X157,bool),X158),s(X157,sK23(X158,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X157))))) | s(bool,X160) = s(bool,happ(s(fun(X157,bool),X158),s(X157,X159))) | s(bool,'T_0') = s(bool,X160)) )),
  inference(subsumption_resolution,[],[f3420,f357])).
fof(f3420,plain,(
  ( ! [X158,X159,X157,X160] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X157,bool),X158),s(X157,sK23(X158,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X157))))) | s(bool,X160) = s(bool,happ(s(fun(X157,bool),X158),s(X157,X159))) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f1819,f836])).
fof(f4019,plain,(
  ( ! [X163,X161,X164,X162] : (p(s(bool,happ(s(fun(X161,bool),X162),s(X161,sK24(X162,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X161))))) | s(bool,X164) = s(bool,happ(s(fun(X161,bool),X162),s(X161,X163))) | s(bool,'T_0') = s(bool,X164)) )),
  inference(subsumption_resolution,[],[f3421,f357])).
fof(f3421,plain,(
  ( ! [X163,X161,X164,X162] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X161,bool),X162),s(X161,sK24(X162,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X161))))) | s(bool,X164) = s(bool,happ(s(fun(X161,bool),X162),s(X161,X163))) | s(bool,'T_0') = s(bool,X164)) )),
  inference(superposition,[],[f1838,f836])).
fof(f4020,plain,(
  ( ! [X167,X165,X166,X168] : (p(s(bool,happ(s(fun(X165,bool),X166),s(X165,X167)))) | s(bool,X168) = s(bool,happ(s(fun(X165,bool),X166),s(X165,sK18(X166,X165)))) | s(bool,'T_0') = s(bool,X168)) )),
  inference(subsumption_resolution,[],[f3422,f357])).
fof(f3422,plain,(
  ( ! [X167,X165,X166,X168] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X165,bool),X166),s(X165,X167)))) | s(bool,X168) = s(bool,happ(s(fun(X165,bool),X166),s(X165,sK18(X166,X165)))) | s(bool,'T_0') = s(bool,X168)) )),
  inference(superposition,[],[f436,f836])).
fof(f4021,plain,(
  ( ! [X171,X169,X172,X170] : (p(s(bool,happ(s(fun(X169,bool),X170),s(X169,X171)))) | s(bool,X172) = s(bool,happ(s(fun(X169,bool),X170),s(X169,sK19(X170,X169)))) | s(bool,'T_0') = s(bool,X172)) )),
  inference(subsumption_resolution,[],[f3423,f357])).
fof(f3423,plain,(
  ( ! [X171,X169,X172,X170] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X169,bool),X170),s(X169,X171)))) | s(bool,X172) = s(bool,happ(s(fun(X169,bool),X170),s(X169,sK19(X170,X169)))) | s(bool,'T_0') = s(bool,X172)) )),
  inference(superposition,[],[f437,f836])).
fof(f4026,plain,(
  ( ! [X198,X196,X199,X197,X195] : (p(s(bool,happ(s(fun(X195,bool),X196),s(X195,X197)))) | s(bool,X199) = s(bool,'const.list.EVERY_2'(s(fun(X195,bool),X196),s('type.list.list'(X195),'const.list.CONS_2'(s(X195,X197),s('type.list.list'(X195),X198))))) | s(bool,'T_0') = s(bool,X199)) )),
  inference(subsumption_resolution,[],[f3431,f357])).
fof(f3431,plain,(
  ( ! [X198,X196,X199,X197,X195] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X195,bool),X196),s(X195,X197)))) | s(bool,X199) = s(bool,'const.list.EVERY_2'(s(fun(X195,bool),X196),s('type.list.list'(X195),'const.list.CONS_2'(s(X195,X197),s('type.list.list'(X195),X198))))) | s(bool,'T_0') = s(bool,X199)) )),
  inference(superposition,[],[f459,f836])).
fof(f4027,plain,(
  ( ! [X204,X202,X200,X203,X201] : (p(s(bool,'const.list.EVERY_2'(s(fun(X200,bool),X201),s('type.list.list'(X200),'const.list.CONS_2'(s(X200,X203),s('type.list.list'(X200),X202)))))) | ~p(s(bool,happ(s(fun(X200,bool),X201),s(X200,X203)))) | s(bool,X204) = s(bool,'const.list.EVERY_2'(s(fun(X200,bool),X201),s('type.list.list'(X200),X202))) | s(bool,'T_0') = s(bool,X204)) )),
  inference(subsumption_resolution,[],[f3432,f357])).
fof(f3432,plain,(
  ( ! [X204,X202,X200,X203,X201] : (~p(s(bool,'T_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(X200,bool),X201),s('type.list.list'(X200),'const.list.CONS_2'(s(X200,X203),s('type.list.list'(X200),X202)))))) | ~p(s(bool,happ(s(fun(X200,bool),X201),s(X200,X203)))) | s(bool,X204) = s(bool,'const.list.EVERY_2'(s(fun(X200,bool),X201),s('type.list.list'(X200),X202))) | s(bool,'T_0') = s(bool,X204)) )),
  inference(superposition,[],[f461,f836])).
fof(f4022,plain,(
  ( ! [X175,X173,X176,X174] : (p(s(bool,'const.list.EVERY_2'(s(fun(X173,bool),X174),s('type.list.list'(X173),X175)))) | s(bool,X176) = s(bool,happ(s(fun(X173,bool),X174),s(X173,sK20(X175,X174,X173)))) | s(bool,'T_0') = s(bool,X176)) )),
  inference(subsumption_resolution,[],[f3424,f357])).
fof(f3424,plain,(
  ( ! [X175,X173,X176,X174] : (~p(s(bool,'T_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(X173,bool),X174),s('type.list.list'(X173),X175)))) | s(bool,X176) = s(bool,happ(s(fun(X173,bool),X174),s(X173,sK20(X175,X174,X173)))) | s(bool,'T_0') = s(bool,X176)) )),
  inference(superposition,[],[f446,f836])).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X2,X1,X0))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f280])).
fof(f2373,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,sK26('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X2,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK25('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X2,X0)))))) )),
  inference(forward_demodulation,[],[f2359,f364])).
fof(f2359,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK25('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X2,X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0'))),s(X0,sK26('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X2,X0)))))) )),
  inference(superposition,[],[f456,f364])).
fof(f2374,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,happ(s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s(X3,sK26('const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X7,X3))))) | ~p(s(bool,happ(s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s(X3,X6)))) | p(s(bool,happ(s(fun(X3,bool),X7),s(X3,sK25('const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X7,X3)))))) )),
  inference(forward_demodulation,[],[f2360,f441])).
fof(f2360,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s(X3,X6)))) | p(s(bool,happ(s(fun(X3,bool),X7),s(X3,sK25('const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X7,X3))))) | p(s(bool,happ(s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5))))),s(X3,sK26('const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X7,X3)))))) )),
  inference(superposition,[],[f456,f441])).
fof(f2134,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,sK25(X2,'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK26(X2,'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X0)))))) )),
  inference(forward_demodulation,[],[f2123,f364])).
fof(f2123,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) | p(s(bool,happ(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0'))),s(X0,sK25(X2,'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK26(X2,'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X0)))))) )),
  inference(superposition,[],[f455,f364])).
fof(f2135,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,happ(s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s(X3,sK25(X7,'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X3))))) | ~p(s(bool,happ(s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s(X3,X6)))) | p(s(bool,happ(s(fun(X3,bool),X7),s(X3,sK26(X7,'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X3)))))) )),
  inference(forward_demodulation,[],[f2124,f441])).
fof(f2124,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s(X3,X6)))) | p(s(bool,happ(s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5))))),s(X3,sK25(X7,'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X3))))) | p(s(bool,happ(s(fun(X3,bool),X7),s(X3,sK26(X7,'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X3)))))) )),
  inference(superposition,[],[f455,f441])).
fof(f2424,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,sK27(X2,'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK27(X2,'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X0)))))) )),
  inference(forward_demodulation,[],[f2410,f364])).
fof(f2410,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) | p(s(bool,happ(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0'))),s(X0,sK27(X2,'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK27(X2,'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X0)))))) )),
  inference(superposition,[],[f457,f364])).
fof(f2425,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,happ(s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s(X3,sK27(X7,'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X3))))) | ~p(s(bool,happ(s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s(X3,X6)))) | p(s(bool,happ(s(fun(X3,bool),X7),s(X3,sK27(X7,'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X3)))))) )),
  inference(forward_demodulation,[],[f2411,f441])).
fof(f2411,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s(X3,X6)))) | p(s(bool,happ(s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5))))),s(X3,sK27(X7,'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X3))))) | p(s(bool,happ(s(fun(X3,bool),X7),s(X3,sK27(X7,'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X3)))))) )),
  inference(superposition,[],[f457,f441])).
fof(f2743,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,sK27('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X2,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK27('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X2,X0)))))) )),
  inference(forward_demodulation,[],[f2729,f364])).
fof(f2729,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK27('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X2,X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0'))),s(X0,sK27('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X2,X0)))))) )),
  inference(superposition,[],[f458,f364])).
fof(f2744,plain,(
  ( ! [X6,X4,X7,X5,X3] : (p(s(bool,happ(s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s(X3,sK27('const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X7,X3))))) | ~p(s(bool,happ(s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s(X3,X6)))) | p(s(bool,happ(s(fun(X3,bool),X7),s(X3,sK27('const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X7,X3)))))) )),
  inference(forward_demodulation,[],[f2730,f441])).
fof(f2730,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s(X3,X6)))) | p(s(bool,happ(s(fun(X3,bool),X7),s(X3,sK27('const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X7,X3))))) | p(s(bool,happ(s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5))))),s(X3,sK27('const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X7,X3)))))) )),
  inference(superposition,[],[f458,f441])).
fof(f469,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s('type.list.list'(X0),X1) != s('type.list.list'(X0),X2) | s(bool,happ(s(fun(X0,bool),X3),s(X0,sK28(X4,X3,X2,X1,X0)))) != s(bool,happ(s(fun(X0,bool),X4),s(X0,sK28(X4,X3,X2,X1,X0)))) | s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f299])).
fof(f2988,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X2,X1,X3,X3,X0)))) != s(bool,happ(s(fun(X0,bool),X2),s(X0,sK28(X2,X1,X3,X3,X0)))) | s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))) = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X3)))) )),
  inference(equality_resolution,[],[f469])).
fof(f1593,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,sK18('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f1590,f364])).
fof(f1590,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,sK18('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0'))),s(X0,X1))))) )),
  inference(superposition,[],[f436,f364])).
fof(f1608,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,sK19('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f1605,f364])).
fof(f1605,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,sK19('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0'))),s(X0,X1))))) )),
  inference(superposition,[],[f437,f364])).
fof(f1765,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) | p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,sK21('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X2,X0))))) | p(s(bool,X2))) )),
  inference(forward_demodulation,[],[f1756,f364])).
fof(f1756,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) | p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0'))),s(X0,sK21('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X2,X0)))))) )),
  inference(superposition,[],[f448,f364])).
fof(f1799,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) | p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,sK22('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X2,X0))))) | p(s(bool,X2))) )),
  inference(forward_demodulation,[],[f1790,f364])).
fof(f1790,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) | p(s(bool,X2)) | p(s(bool,happ(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0'))),s(X0,sK22('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X2,X0)))))) )),
  inference(superposition,[],[f450,f364])).
fof(f2390,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,sK23('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f2376,f364])).
fof(f2376,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) | p(s(bool,happ(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0'))),s(X0,sK23('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0)))))) )),
  inference(superposition,[],[f1819,f364])).
fof(f2407,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,sK24('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1))))) )),
  inference(forward_demodulation,[],[f2393,f364])).
fof(f2393,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,X1)))) | p(s(bool,happ(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0'))),s(X0,sK24('const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0)))))) )),
  inference(superposition,[],[f1838,f364])).
fof(f1677,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,sK20(X1,'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X0))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s('type.list.list'(X0),X1))))) )),
  inference(forward_demodulation,[],[f1674,f364])).
fof(f1674,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X0,sK20(X1,'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')),X0))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0'))),s('type.list.list'(X0),X1))))) )),
  inference(superposition,[],[f446,f364])).
fof(f2000,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s(X3,X6)))) | ~p(s(bool,happ(s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s(X3,sK18('const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X3)))))) )),
  inference(forward_demodulation,[],[f1953,f441])).
fof(f1953,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),'const.list.CONS_2'(s(X3,X4),s('type.list.list'(X3),X5))))),s(X3,sK18('const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5)))),X3))))) | p(s(bool,happ(s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X3),X5))))),s(X3,X6))))) )),
  inference(superposition,[],[f436,f441])).
fof(f2001,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,happ(s(fun(X7,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X8),s('type.list.list'(X7),X9))))),s(X7,X10)))) | ~p(s(bool,happ(s(fun(X7,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X8),s('type.list.list'(X7),X9))))),s(X7,sK19('const.pred_set.INSERT_2'(s(X7,X8),s(fun(X7,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X7),X9)))),X7)))))) )),
  inference(forward_demodulation,[],[f1954,f441])).
fof(f1954,plain,(
  ( ! [X10,X8,X7,X9] : (~p(s(bool,happ(s(fun(X7,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X8),s('type.list.list'(X7),X9))))),s(X7,sK19('const.pred_set.INSERT_2'(s(X7,X8),s(fun(X7,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X7),X9)))),X7))))) | p(s(bool,happ(s(fun(X7,bool),'const.pred_set.INSERT_2'(s(X7,X8),s(fun(X7,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X7),X9))))),s(X7,X10))))) )),
  inference(superposition,[],[f437,f441])).
fof(f2003,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,happ(s(fun(X15,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X15),'const.list.CONS_2'(s(X15,X16),s('type.list.list'(X15),X17))))),s(X15,sK21('const.pred_set.INSERT_2'(s(X15,X16),s(fun(X15,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X15),X17)))),X19,X15))))) | ~p(s(bool,happ(s(fun(X15,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X15),'const.list.CONS_2'(s(X15,X16),s('type.list.list'(X15),X17))))),s(X15,X18)))) | p(s(bool,X19))) )),
  inference(forward_demodulation,[],[f1956,f441])).
fof(f1956,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,happ(s(fun(X15,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X15),'const.list.CONS_2'(s(X15,X16),s('type.list.list'(X15),X17))))),s(X15,X18)))) | p(s(bool,X19)) | p(s(bool,happ(s(fun(X15,bool),'const.pred_set.INSERT_2'(s(X15,X16),s(fun(X15,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X15),X17))))),s(X15,sK21('const.pred_set.INSERT_2'(s(X15,X16),s(fun(X15,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X15),X17)))),X19,X15)))))) )),
  inference(superposition,[],[f448,f441])).
fof(f2004,plain,(
  ( ! [X24,X23,X21,X22,X20] : (p(s(bool,happ(s(fun(X20,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X20),'const.list.CONS_2'(s(X20,X21),s('type.list.list'(X20),X22))))),s(X20,sK22('const.pred_set.INSERT_2'(s(X20,X21),s(fun(X20,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X20),X22)))),X24,X20))))) | ~p(s(bool,happ(s(fun(X20,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X20),'const.list.CONS_2'(s(X20,X21),s('type.list.list'(X20),X22))))),s(X20,X23)))) | p(s(bool,X24))) )),
  inference(forward_demodulation,[],[f1957,f441])).
fof(f1957,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,happ(s(fun(X20,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X20),'const.list.CONS_2'(s(X20,X21),s('type.list.list'(X20),X22))))),s(X20,X23)))) | p(s(bool,X24)) | p(s(bool,happ(s(fun(X20,bool),'const.pred_set.INSERT_2'(s(X20,X21),s(fun(X20,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X20),X22))))),s(X20,sK22('const.pred_set.INSERT_2'(s(X20,X21),s(fun(X20,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X20),X22)))),X24,X20)))))) )),
  inference(superposition,[],[f450,f441])).
fof(f2391,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun(X2,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X2),'const.list.CONS_2'(s(X2,X3),s('type.list.list'(X2),X4))))),s(X2,sK23('const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X2),X4)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X2))))) | ~p(s(bool,happ(s(fun(X2,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X2),'const.list.CONS_2'(s(X2,X3),s('type.list.list'(X2),X4))))),s(X2,X5))))) )),
  inference(forward_demodulation,[],[f2377,f441])).
fof(f2377,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun(X2,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X2),'const.list.CONS_2'(s(X2,X3),s('type.list.list'(X2),X4))))),s(X2,X5)))) | p(s(bool,happ(s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X2),X4))))),s(X2,sK23('const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X2),X4)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X2)))))) )),
  inference(superposition,[],[f1819,f441])).
fof(f2408,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun(X2,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X2),'const.list.CONS_2'(s(X2,X3),s('type.list.list'(X2),X4))))),s(X2,sK24('const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X2),X4)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X2))))) | ~p(s(bool,happ(s(fun(X2,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X2),'const.list.CONS_2'(s(X2,X3),s('type.list.list'(X2),X4))))),s(X2,X5))))) )),
  inference(forward_demodulation,[],[f2394,f441])).
fof(f2394,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun(X2,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X2),'const.list.CONS_2'(s(X2,X3),s('type.list.list'(X2),X4))))),s(X2,X5)))) | p(s(bool,happ(s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X2),X4))))),s(X2,sK24('const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X2),X4)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X2)))))) )),
  inference(superposition,[],[f1838,f441])).
fof(f2002,plain,(
  ( ! [X14,X12,X13,X11] : (p(s(bool,'const.list.EVERY_2'(s(fun(X11,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X11),'const.list.CONS_2'(s(X11,X12),s('type.list.list'(X11),X13))))),s('type.list.list'(X11),X14)))) | ~p(s(bool,happ(s(fun(X11,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X11),'const.list.CONS_2'(s(X11,X12),s('type.list.list'(X11),X13))))),s(X11,sK20(X14,'const.pred_set.INSERT_2'(s(X11,X12),s(fun(X11,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X11),X13)))),X11)))))) )),
  inference(forward_demodulation,[],[f1955,f441])).
fof(f1955,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,happ(s(fun(X11,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X11),'const.list.CONS_2'(s(X11,X12),s('type.list.list'(X11),X13))))),s(X11,sK20(X14,'const.pred_set.INSERT_2'(s(X11,X12),s(fun(X11,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X11),X13)))),X11))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X11,bool),'const.pred_set.INSERT_2'(s(X11,X12),s(fun(X11,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X11),X13))))),s('type.list.list'(X11),X14))))) )),
  inference(superposition,[],[f446,f441])).
fof(f3154,plain,(
  ( ! [X243,X241,X244,X242] : (~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X243),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X242),s('type.list.list'(bool),X244)))))) | p(s(bool,happ(s(fun(bool,bool),X243),s(bool,X241)))) | s(bool,'T_0') = s(bool,X242) | s(bool,'T_0') = s(bool,X241)) )),
  inference(superposition,[],[f459,f836])).
fof(f3190,plain,(
  ( ! [X360,X362,X361] : (~p(s(bool,happ(s(fun(bool,bool),X360),s(bool,X361)))) | p(s(bool,happ(s(fun(bool,bool),X360),s(bool,X362)))) | s(bool,'T_0') = s(bool,X361) | s(bool,'T_0') = s(bool,sK18(X360,bool))) )),
  inference(superposition,[],[f436,f836])).
fof(f3191,plain,(
  ( ! [X364,X365,X363] : (~p(s(bool,happ(s(fun(bool,bool),X363),s(bool,X364)))) | p(s(bool,happ(s(fun(bool,bool),X363),s(bool,X365)))) | s(bool,'T_0') = s(bool,X364) | s(bool,'T_0') = s(bool,sK19(X363,bool))) )),
  inference(superposition,[],[f437,f836])).
fof(f3193,plain,(
  ( ! [X370,X369,X371] : (~p(s(bool,happ(s(fun(bool,bool),X370),s(bool,X371)))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X370),s('type.list.list'(bool),X369)))) | s(bool,'T_0') = s(bool,X371) | s(bool,'T_0') = s(bool,sK20(X369,X370,bool))) )),
  inference(superposition,[],[f446,f836])).
fof(f3441,plain,(
  ( ! [X239,X237,X240,X238] : (~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X238),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'T_0'),s('type.list.list'(bool),X239)))))) | p(s(bool,happ(s(fun(bool,bool),X238),s(bool,X237)))) | s(bool,X237) = s(bool,X240) | s(bool,'T_0') = s(bool,X240)) )),
  inference(superposition,[],[f459,f836])).
fof(f3477,plain,(
  ( ! [X356,X357,X355] : (~p(s(bool,happ(s(fun(bool,bool),X355),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X355),s(bool,X356)))) | s(bool,X357) = s(bool,sK18(X355,bool)) | s(bool,'T_0') = s(bool,X357)) )),
  inference(superposition,[],[f436,f836])).
fof(f3478,plain,(
  ( ! [X358,X360,X359] : (~p(s(bool,happ(s(fun(bool,bool),X358),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X358),s(bool,X359)))) | s(bool,X360) = s(bool,sK19(X358,bool)) | s(bool,'T_0') = s(bool,X360)) )),
  inference(superposition,[],[f437,f836])).
fof(f3480,plain,(
  ( ! [X364,X366,X365] : (~p(s(bool,happ(s(fun(bool,bool),X365),s(bool,'T_0')))) | p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X365),s('type.list.list'(bool),X364)))) | s(bool,X366) = s(bool,sK20(X364,X365,bool)) | s(bool,'T_0') = s(bool,X366)) )),
  inference(superposition,[],[f446,f836])).
fof(f3339,plain,(
  ( ! [X123,X121,X122,X120] : (~p(s(bool,'const.list.EVERY_2'(s(fun(bool,bool),X121),s('type.list.list'(bool),'const.list.CONS_2'(s(bool,X120),s('type.list.list'(bool),X122)))))) | p(s(bool,happ(s(fun(bool,bool),X121),s(bool,'T_0')))) | s(bool,X120) = s(bool,X123) | s(bool,'T_0') = s(bool,X123)) )),
  inference(superposition,[],[f459,f836])).
fof(f4023,plain,(
  ( ! [X177,X180,X178,X179] : (~p(s(bool,happ(s(fun('type.list.list'(X177),bool),X178),s('type.list.list'(X177),'const.list.CONS_2'(s(X177,sK17(X178,X177)),s('type.list.list'(X177),sK16(X178,X177))))))) | p(s(bool,happ(s(fun('type.list.list'(X177),bool),X178),s('type.list.list'(X177),X179)))) | s(bool,X180) = s(bool,happ(s(fun('type.list.list'(X177),bool),X178),s('type.list.list'(X177),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X180)) )),
  inference(subsumption_resolution,[],[f3425,f357])).
fof(f3425,plain,(
  ( ! [X177,X180,X178,X179] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun('type.list.list'(X177),bool),X178),s('type.list.list'(X177),'const.list.CONS_2'(s(X177,sK17(X178,X177)),s('type.list.list'(X177),sK16(X178,X177))))))) | p(s(bool,happ(s(fun('type.list.list'(X177),bool),X178),s('type.list.list'(X177),X179)))) | s(bool,X180) = s(bool,happ(s(fun('type.list.list'(X177),bool),X178),s('type.list.list'(X177),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X180)) )),
  inference(superposition,[],[f433,f836])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK17(X1,X0)),s('type.list.list'(X0),sK16(X1,X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK16(X1,X0))))) & ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK17(X1,X0)),s('type.list.list'(X0),sK16(X1,X0)))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16,sK17])],[f269])).
fof(f269,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X3] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X3)))) & ? [X4] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X4),s('type.list.list'(X0),X3))))))) | ! [X2] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | ? [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) & ? [X3] : ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) | ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(ennf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0] : ! [X1] : ((p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) & ! [X2] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) => ! [X3] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2)))))))) => ! [X4] : p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X4)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X7] : ((p(s(bool,happ(s(fun('type.list.list'(X6),bool),X7),s('type.list.list'(X6),'const.list.NIL_0')))) & ! [X15] : (p(s(bool,happ(s(fun('type.list.list'(X6),bool),X7),s('type.list.list'(X6),X15)))) => ! [X17] : p(s(bool,happ(s(fun('type.list.list'(X6),bool),X7),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X17),s('type.list.list'(X6),X15)))))))) => ! [X9] : p(s(bool,happ(s(fun('type.list.list'(X6),bool),X7),s('type.list.list'(X6),X9)))))),
  file('thm.list.EVERY_MEM_MONO','thm.list.list_INDUCT')).
fof(f4024,plain,(
  ( ! [X182,X183,X181] : (p(s(bool,happ(s(fun('type.list.list'(X181),bool),X182),s('type.list.list'(X181),sK16(X182,X181))))) | s(bool,X183) = s(bool,happ(s(fun('type.list.list'(X181),bool),X182),s('type.list.list'(X181),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X183)) )),
  inference(subsumption_resolution,[],[f3426,f357])).
fof(f3426,plain,(
  ( ! [X182,X183,X181] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.list.list'(X181),bool),X182),s('type.list.list'(X181),sK16(X182,X181))))) | s(bool,X183) = s(bool,happ(s(fun('type.list.list'(X181),bool),X182),s('type.list.list'(X181),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X183)) )),
  inference(superposition,[],[f601,f836])).
fof(f601,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK16(X1,X0)))))) )),
  inference(condensation,[],[f432])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),sK16(X1,X0))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f270])).
fof(f2925,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.pred_set.EMPTY_0'),s('type.list.list'(X0),X1)))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.pred_set.EMPTY_0'),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK17('const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X0)),'const.list.NIL_0')),X0)),s('type.list.list'(X0),sK16('const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X0)),'const.list.NIL_0')),X0))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.pred_set.EMPTY_0'),s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f2924,f364])).
fof(f2924,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.pred_set.EMPTY_0'),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK17('const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X0)),'const.list.NIL_0')),X0)),s('type.list.list'(X0),sK16('const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X0)),'const.list.NIL_0')),X0))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.pred_set.EMPTY_0'),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X0)),'const.list.NIL_0'))),s('type.list.list'(X0),X1))))) )),
  inference(forward_demodulation,[],[f2920,f364])).
fof(f2920,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.pred_set.EMPTY_0'),s('type.list.list'(X0),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X0)),'const.list.NIL_0'))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK17('const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X0)),'const.list.NIL_0')),X0)),s('type.list.list'(X0),sK16('const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X0)),'const.list.NIL_0')),X0))))))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X0)),'const.list.NIL_0'))),s('type.list.list'(X0),X1))))) )),
  inference(superposition,[],[f433,f364])).
fof(f1743,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.pred_set.EMPTY_0'),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.pred_set.EMPTY_0'),s('type.list.list'(X0),sK16('const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X0)),'const.list.NIL_0')),X0)))))) )),
  inference(forward_demodulation,[],[f1741,f364])).
fof(f1741,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.pred_set.EMPTY_0'),s('type.list.list'(X0),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X0),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X0)),'const.list.NIL_0'))),s('type.list.list'(X0),sK16('const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X0)),'const.list.NIL_0')),X0)))))) )),
  inference(superposition,[],[f601,f364])).
fof(f2927,plain,(
  ( ! [X4,X2,X5,X3] : (p(s(bool,happ(s(fun('type.list.list'(X2),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X2)),'const.list.CONS_2'(s('type.list.list'(X2),X3),s('type.list.list'('type.list.list'(X2)),X4))))),s('type.list.list'(X2),X5)))) | ~p(s(bool,happ(s(fun('type.list.list'(X2),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X2)),'const.list.CONS_2'(s('type.list.list'(X2),X3),s('type.list.list'('type.list.list'(X2)),X4))))),s('type.list.list'(X2),'const.list.CONS_2'(s(X2,sK17('const.pred_set.INSERT_2'(s('type.list.list'(X2),X3),s(fun('type.list.list'(X2),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X2)),X4)))),X2)),s('type.list.list'(X2),sK16('const.pred_set.INSERT_2'(s('type.list.list'(X2),X3),s(fun('type.list.list'(X2),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X2)),X4)))),X2))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X2),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X2)),'const.list.CONS_2'(s('type.list.list'(X2),X3),s('type.list.list'('type.list.list'(X2)),X4))))),s('type.list.list'(X2),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f2926,f441])).
fof(f2926,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X2),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X2)),'const.list.CONS_2'(s('type.list.list'(X2),X3),s('type.list.list'('type.list.list'(X2)),X4))))),s('type.list.list'(X2),'const.list.CONS_2'(s(X2,sK17('const.pred_set.INSERT_2'(s('type.list.list'(X2),X3),s(fun('type.list.list'(X2),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X2)),X4)))),X2)),s('type.list.list'(X2),sK16('const.pred_set.INSERT_2'(s('type.list.list'(X2),X3),s(fun('type.list.list'(X2),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X2)),X4)))),X2))))))) | ~p(s(bool,happ(s(fun('type.list.list'(X2),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X2)),'const.list.CONS_2'(s('type.list.list'(X2),X3),s('type.list.list'('type.list.list'(X2)),X4))))),s('type.list.list'(X2),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X2),bool),'const.pred_set.INSERT_2'(s('type.list.list'(X2),X3),s(fun('type.list.list'(X2),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X2)),X4))))),s('type.list.list'(X2),X5))))) )),
  inference(forward_demodulation,[],[f2921,f441])).
fof(f2921,plain,(
  ( ! [X4,X2,X5,X3] : (~p(s(bool,happ(s(fun('type.list.list'(X2),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X2)),'const.list.CONS_2'(s('type.list.list'(X2),X3),s('type.list.list'('type.list.list'(X2)),X4))))),s('type.list.list'(X2),'const.list.NIL_0')))) | ~p(s(bool,happ(s(fun('type.list.list'(X2),bool),'const.pred_set.INSERT_2'(s('type.list.list'(X2),X3),s(fun('type.list.list'(X2),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X2)),X4))))),s('type.list.list'(X2),'const.list.CONS_2'(s(X2,sK17('const.pred_set.INSERT_2'(s('type.list.list'(X2),X3),s(fun('type.list.list'(X2),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X2)),X4)))),X2)),s('type.list.list'(X2),sK16('const.pred_set.INSERT_2'(s('type.list.list'(X2),X3),s(fun('type.list.list'(X2),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X2)),X4)))),X2))))))) | p(s(bool,happ(s(fun('type.list.list'(X2),bool),'const.pred_set.INSERT_2'(s('type.list.list'(X2),X3),s(fun('type.list.list'(X2),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X2)),X4))))),s('type.list.list'(X2),X5))))) )),
  inference(superposition,[],[f433,f441])).
fof(f2008,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,happ(s(fun('type.list.list'(X98),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X98)),'const.list.CONS_2'(s('type.list.list'(X98),X99),s('type.list.list'('type.list.list'(X98)),X100))))),s('type.list.list'(X98),sK16('const.pred_set.INSERT_2'(s('type.list.list'(X98),X99),s(fun('type.list.list'(X98),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X98)),X100)))),X98))))) | ~p(s(bool,happ(s(fun('type.list.list'(X98),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X98)),'const.list.CONS_2'(s('type.list.list'(X98),X99),s('type.list.list'('type.list.list'(X98)),X100))))),s('type.list.list'(X98),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f1978,f441])).
fof(f1978,plain,(
  ( ! [X99,X100,X98] : (~p(s(bool,happ(s(fun('type.list.list'(X98),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X98)),'const.list.CONS_2'(s('type.list.list'(X98),X99),s('type.list.list'('type.list.list'(X98)),X100))))),s('type.list.list'(X98),'const.list.NIL_0')))) | p(s(bool,happ(s(fun('type.list.list'(X98),bool),'const.pred_set.INSERT_2'(s('type.list.list'(X98),X99),s(fun('type.list.list'(X98),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X98)),X100))))),s('type.list.list'(X98),sK16('const.pred_set.INSERT_2'(s('type.list.list'(X98),X99),s(fun('type.list.list'(X98),bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.list.list'(X98)),X100)))),X98)))))) )),
  inference(superposition,[],[f601,f441])).
fof(f2792,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,X0)))) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK15))))))) )),
  inference(subsumption_resolution,[],[f2782,f357])).
fof(f2782,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X0),s('type.list.list'(sK12),sK15)))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,X0))))) )),
  inference(superposition,[],[f461,f657])).
fof(f460,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))) )),
  inference(cnf_transformation,[],[f294])).
fof(f4025,plain,(
  ( ! [X191,X194,X192,X190,X193] : (p(s(bool,'const.list.EVERY_2'(s(fun(X190,bool),X191),s('type.list.list'(X190),X193)))) | s(bool,X194) = s(bool,'const.list.EVERY_2'(s(fun(X190,bool),X191),s('type.list.list'(X190),'const.list.CONS_2'(s(X190,X192),s('type.list.list'(X190),X193))))) | s(bool,'T_0') = s(bool,X194)) )),
  inference(subsumption_resolution,[],[f3430,f357])).
fof(f3430,plain,(
  ( ! [X191,X194,X192,X190,X193] : (~p(s(bool,'T_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(X190,bool),X191),s('type.list.list'(X190),X193)))) | s(bool,X194) = s(bool,'const.list.EVERY_2'(s(fun(X190,bool),X191),s('type.list.list'(X190),'const.list.CONS_2'(s(X190,X192),s('type.list.list'(X190),X193))))) | s(bool,'T_0') = s(bool,X194)) )),
  inference(superposition,[],[f460,f836])).
fof(f4062,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X1,'T_0')) )),
  inference(condensation,[],[f4061])).
fof(f4061,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X90)) | ~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X90,'T_0') | s(bool,X89) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f3508,f906])).
fof(f3508,plain,(
  ( ! [X90,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | ~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X90,'T_0') | s(bool,X89) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f1401,f836])).
fof(f4060,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4059])).
fof(f4059,plain,(
  ( ! [X85,X83,X84,X82] : (~sP8('T_0',X83,X84) | sP9('T_0',X83,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f3506,f906])).
fof(f3506,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,X82)) | ~sP8('T_0',X83,X84) | sP9('T_0',X83,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f598,f836])).
fof(f4058,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f4057])).
fof(f4057,plain,(
  ( ! [X80,X78,X81,X79] : (~sP10('T_0',X79,X80) | sP11('T_0',X79,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f3505,f906])).
fof(f3505,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X78)) | ~sP10('T_0',X79,X80) | sP11('T_0',X79,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f597,f836])).
fof(f4056,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4055])).
fof(f4055,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | ~sP11('T_0',X76,X75) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f3504,f906])).
fof(f3504,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X74)) | ~p(s(bool,X75)) | p(s(bool,X76)) | ~sP11('T_0',X76,X75) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f594,f836])).
fof(f4054,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f4053])).
fof(f4053,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X71)) | sP11('T_0',X71,X72) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f3503,f906])).
fof(f3503,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X70)) | ~p(s(bool,X71)) | sP11('T_0',X71,X72) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f593,f836])).
fof(f4052,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f4051])).
fof(f4051,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,X67)) | sP11('T_0',X68,X67) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f3502,f906])).
fof(f3502,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X66)) | p(s(bool,X67)) | sP11('T_0',X68,X67) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f592,f836])).
fof(f4050,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4049])).
fof(f4049,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X63)) | p(s(bool,X64)) | ~sP9('T_0',X64,X63) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f3501,f906])).
fof(f3501,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X62)) | p(s(bool,X63)) | p(s(bool,X64)) | ~sP9('T_0',X64,X63) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f577,f836])).
fof(f4048,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4047])).
fof(f4047,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X59)) | sP9('T_0',X59,X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f3500,f906])).
fof(f3500,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X58)) | ~p(s(bool,X59)) | sP9('T_0',X59,X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f576,f836])).
fof(f4046,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4045])).
fof(f4045,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X55)) | sP9('T_0',X56,X55) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f3499,f906])).
fof(f3499,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X54)) | ~p(s(bool,X55)) | sP9('T_0',X56,X55) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f575,f836])).
fof(f4044,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f4043])).
fof(f4043,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP7(X52,X51,'T_0') | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f3498,f906])).
fof(f3498,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP7(X52,X51,'T_0') | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f562,f836])).
fof(f4042,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f4041])).
fof(f4041,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | p(s(bool,X48)) | ~sP5(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f3497,f906])).
fof(f3497,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP5(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f526,f836])).
fof(f4040,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4039])).
fof(f4039,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X44)) | sP1(X44,'T_0') | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f3496,f906])).
fof(f3496,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,X43)) | p(s(bool,X44)) | sP1(X44,'T_0') | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f502,f836])).
fof(f4038,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4037])).
fof(f4037,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X41)) | sP1('T_0',X41) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f3495,f906])).
fof(f3495,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | sP1('T_0',X41) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f499,f836])).
fof(f4032,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK24(X2,'T_0',X1)))))) )),
  inference(condensation,[],[f4031])).
fof(f4031,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,happ(s(fun(X21,bool),X22),s(X21,X23)))) | p(s(bool,happ(s(fun(X21,bool),X22),s(X21,sK24(X22,'T_0',X21))))) | s(bool,X20) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f3489,f906])).
fof(f3489,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,X20)) | ~p(s(bool,happ(s(fun(X21,bool),X22),s(X21,X23)))) | p(s(bool,happ(s(fun(X21,bool),X22),s(X21,sK24(X22,'T_0',X21))))) | s(bool,X20) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f454,f836])).
fof(f4030,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK23(X2,'T_0',X1)))))) )),
  inference(condensation,[],[f4029])).
fof(f4029,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK23(X17,'T_0',X16))))) | s(bool,X15) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f3488,f906])).
fof(f3488,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,X15)) | ~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK23(X17,'T_0',X16))))) | s(bool,X15) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f452,f836])).
fof(f3481,plain,(
  ( ! [X368,X370,X367,X369] : (s(X369,f(s(bool,'T_0'))) != s(X369,g(s(bool,'T_0'))) | s(fun(bool,X369),X367) = s(fun(bool,X369),X368) | s(bool,X370) = s(bool,sK29(X367,X368,X369,bool)) | s(bool,'T_0') = s(bool,X370)) )),
  inference(superposition,[],[f516,f836])).
fof(f4011,plain,(
  ( ! [X125,X126] : (p(s(bool,happ(s(fun(sK12,bool),sK14),s(sK12,X125)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X125),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK15)))))) | s(bool,X126) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,X125))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(subsumption_resolution,[],[f3413,f357])).
fof(f3413,plain,(
  ( ! [X125,X126] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK14),s(sK12,X125)))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X125),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK15)))))) | s(bool,X126) = s(bool,happ(s(fun(sK12,bool),sK13),s(sK12,X125))) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f354,f836])).
fof(f4005,plain,(
  ( ! [X85,X86,X84] : (~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X85,X84) | p(s(bool,X85)) | s(bool,X84) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f3399,f357])).
fof(f3399,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,X85)) | ~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X85,X84) | s(bool,X84) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f1401,f836])).
fof(f4004,plain,(
  ( ! [X80,X78,X79,X77] : (~sP8(X77,X78,X79) | sP9(X77,X78,X79) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(subsumption_resolution,[],[f3397,f357])).
fof(f3397,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,'T_0')) | ~sP8(X77,X78,X79) | sP9(X77,X78,X79) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f598,f836])).
fof(f4003,plain,(
  ( ! [X76,X74,X75,X73] : (~sP10(X73,X74,X75) | sP11(X73,X74,X75) | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(subsumption_resolution,[],[f3396,f357])).
fof(f3396,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,'T_0')) | ~sP10(X73,X74,X75) | sP11(X73,X74,X75) | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f597,f836])).
fof(f4002,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,X70)) | p(s(bool,X71)) | ~sP11(X69,X71,X70) | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f3395,f357])).
fof(f3395,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,'T_0')) | ~p(s(bool,X70)) | p(s(bool,X71)) | ~sP11(X69,X71,X70) | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f594,f836])).
fof(f4001,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,X66)) | sP11(X65,X66,X67) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f3394,f357])).
fof(f3394,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'T_0')) | ~p(s(bool,X66)) | sP11(X65,X66,X67) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f593,f836])).
fof(f4000,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,X62)) | sP11(X61,X63,X62) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f3393,f357])).
fof(f3393,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'T_0')) | p(s(bool,X62)) | sP11(X61,X63,X62) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f592,f836])).
fof(f3999,plain,(
  ( ! [X59,X57,X60,X58] : (p(s(bool,X58)) | p(s(bool,X59)) | ~sP9(X57,X59,X58) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f3392,f357])).
fof(f3392,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'T_0')) | p(s(bool,X58)) | p(s(bool,X59)) | ~sP9(X57,X59,X58) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f577,f836])).
fof(f3998,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | sP9(X53,X54,X55) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f3391,f357])).
fof(f3391,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X54)) | sP9(X53,X54,X55) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f576,f836])).
fof(f3997,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | sP9(X49,X51,X50) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f3390,f357])).
fof(f3390,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | sP9(X49,X51,X50) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f575,f836])).
fof(f3996,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X46)) | p(s(bool,X47)) | ~sP7(X47,X46,X45) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f3389,f357])).
fof(f3389,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'T_0')) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP7(X47,X46,X45) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f562,f836])).
fof(f3995,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X42)) | p(s(bool,X43)) | ~sP5(X42,X41,X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f3388,f357])).
fof(f3388,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'T_0')) | ~p(s(bool,X42)) | p(s(bool,X43)) | ~sP5(X42,X41,X43) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f526,f836])).
fof(f3994,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X39)) | sP1(X39,X38) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f3387,f357])).
fof(f3387,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,'T_0')) | p(s(bool,X39)) | sP1(X39,X38) | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f502,f836])).
fof(f3993,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X36)) | sP1(X35,X36) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f3386,f357])).
fof(f3386,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,'T_0')) | p(s(bool,X36)) | sP1(X35,X36) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f499,f836])).
fof(f3992,plain,(
  ( ! [X33,X34,X32] : (sP0(X32,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f3385,f357])).
fof(f3385,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | sP0(X32,X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f484,f836])).
fof(f3991,plain,(
  ( ! [X30,X31,X29] : (sP0(X30,X29) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f3384,f357])).
fof(f3384,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | sP0(X30,X29) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f483,f836])).
fof(f3990,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,X27)) | s(bool,X26) = s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f3383,f357])).
fof(f3383,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,X27)) | s(bool,X26) = s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f473,f836])).
fof(f3987,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f3377,f357])).
fof(f3377,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f379,f836])).
fof(f3985,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X1,'T_0')) )),
  inference(condensation,[],[f3984])).
fof(f3984,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X90)) | ~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X90,'T_0') | s(bool,X89) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(subsumption_resolution,[],[f3331,f906])).
fof(f3331,plain,(
  ( ! [X90,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X90)) | ~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X90,'T_0') | s(bool,X89) = s(bool,X91) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f1401,f836])).
fof(f3982,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f3981])).
fof(f3981,plain,(
  ( ! [X85,X83,X84,X82] : (~sP8('T_0',X83,X84) | sP9('T_0',X83,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f3329,f906])).
fof(f3329,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,X82)) | ~sP8('T_0',X83,X84) | sP9('T_0',X83,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f598,f836])).
fof(f3979,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f3978])).
fof(f3978,plain,(
  ( ! [X80,X78,X81,X79] : (~sP10('T_0',X79,X80) | sP11('T_0',X79,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f3328,f906])).
fof(f3328,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X78)) | ~sP10('T_0',X79,X80) | sP11('T_0',X79,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f597,f836])).
fof(f3976,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f3975])).
fof(f3975,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X75)) | p(s(bool,X76)) | ~sP11('T_0',X76,X75) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f3327,f906])).
fof(f3327,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X74)) | ~p(s(bool,X75)) | p(s(bool,X76)) | ~sP11('T_0',X76,X75) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f594,f836])).
fof(f3973,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f3972])).
fof(f3972,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X71)) | sP11('T_0',X71,X72) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f3326,f906])).
fof(f3326,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X70)) | ~p(s(bool,X71)) | sP11('T_0',X71,X72) | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f593,f836])).
fof(f3970,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f3969])).
fof(f3969,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,X67)) | sP11('T_0',X68,X67) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f3325,f906])).
fof(f3325,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X66)) | p(s(bool,X67)) | sP11('T_0',X68,X67) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f592,f836])).
fof(f3967,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f3966])).
fof(f3966,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X63)) | p(s(bool,X64)) | ~sP9('T_0',X64,X63) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f3324,f906])).
fof(f3324,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X62)) | p(s(bool,X63)) | p(s(bool,X64)) | ~sP9('T_0',X64,X63) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f577,f836])).
fof(f3964,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f3963])).
fof(f3963,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X59)) | sP9('T_0',X59,X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f3323,f906])).
fof(f3323,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X58)) | ~p(s(bool,X59)) | sP9('T_0',X59,X60) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f576,f836])).
fof(f3961,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f3960])).
fof(f3960,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X55)) | sP9('T_0',X56,X55) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f3322,f906])).
fof(f3322,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X54)) | ~p(s(bool,X55)) | sP9('T_0',X56,X55) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f575,f836])).
fof(f3958,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f3957])).
fof(f3957,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP7(X52,X51,'T_0') | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f3321,f906])).
fof(f3321,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP7(X52,X51,'T_0') | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f562,f836])).
fof(f3955,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f3954])).
fof(f3954,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | p(s(bool,X48)) | ~sP5(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f3320,f906])).
fof(f3320,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP5(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f526,f836])).
fof(f3952,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f3951])).
fof(f3951,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X44)) | sP1(X44,'T_0') | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f3319,f906])).
fof(f3319,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,X43)) | p(s(bool,X44)) | sP1(X44,'T_0') | s(bool,X43) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f502,f836])).
fof(f3949,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f3948])).
fof(f3948,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X41)) | sP1('T_0',X41) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f3318,f906])).
fof(f3318,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | sP1('T_0',X41) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f499,f836])).
fof(f3942,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK24(X2,'T_0',X1)))))) )),
  inference(condensation,[],[f3941])).
fof(f3941,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,happ(s(fun(X21,bool),X22),s(X21,X23)))) | p(s(bool,happ(s(fun(X21,bool),X22),s(X21,sK24(X22,'T_0',X21))))) | s(bool,X20) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f3312,f906])).
fof(f3312,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,X20)) | ~p(s(bool,happ(s(fun(X21,bool),X22),s(X21,X23)))) | p(s(bool,happ(s(fun(X21,bool),X22),s(X21,sK24(X22,'T_0',X21))))) | s(bool,X20) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f454,f836])).
fof(f3938,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,happ(s(fun(X1,bool),X2),s(X1,X3)))) | p(s(bool,happ(s(fun(X1,bool),X2),s(X1,sK23(X2,'T_0',X1)))))) )),
  inference(condensation,[],[f3937])).
fof(f3937,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK23(X17,'T_0',X16))))) | s(bool,X15) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(subsumption_resolution,[],[f3311,f906])).
fof(f3311,plain,(
  ( ! [X19,X17,X15,X18,X16] : (~p(s(bool,X15)) | ~p(s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) | p(s(bool,happ(s(fun(X16,bool),X17),s(X16,sK23(X17,'T_0',X16))))) | s(bool,X15) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f452,f836])).
fof(f3305,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f375,f836])).
fof(f3194,plain,(
  ( ! [X372,X374,X373,X375] : (s(X374,f(s(bool,X375))) != s(X374,g(s(bool,X375))) | s(fun(bool,X374),X372) = s(fun(bool,X374),X373) | s(bool,'T_0') = s(bool,X375) | s(bool,'T_0') = s(bool,sK29(X372,X373,X374,bool))) )),
  inference(superposition,[],[f516,f836])).
fof(f3057,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f836,f836])).
fof(f3019,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X3)) )),
  inference(superposition,[],[f2628,f441])).
fof(f3003,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),X3,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(superposition,[],[f2626,f441])).
fof(f2987,plain,(
  ( ! [X2,X0,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))))) )),
  inference(superposition,[],[f2597,f441])).
fof(f2973,plain,(
  ( ! [X2,X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))))) )),
  inference(superposition,[],[f2504,f441])).
fof(f2958,plain,(
  ( ! [X2,X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),'T_0')) )),
  inference(superposition,[],[f2499,f441])).
fof(f2943,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'T_0')) )),
  inference(superposition,[],[f2494,f441])).
fof(f2929,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f357,f633,f628,f2494,f527])).
fof(f2914,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),'T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(superposition,[],[f2492,f441])).
fof(f2893,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(superposition,[],[f2489,f441])).
fof(f2829,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),'T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) )),
  inference(superposition,[],[f2488,f441])).
fof(f2811,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'F_0')) )),
  inference(superposition,[],[f2485,f441])).
fof(f2778,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'T_0')) )),
  inference(superposition,[],[f2484,f441])).
fof(f2761,plain,(
  ( ! [X2,X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))))) )),
  inference(superposition,[],[f2451,f441])).
fof(f2726,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))))) )),
  inference(superposition,[],[f2449,f441])).
fof(f2432,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f357,f1643,f473])).
fof(f2433,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))))) = s(bool,'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f1643,f473])).
fof(f2434,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f483])).
fof(f2435,plain,(
  ( ! [X2,X0,X3,X1] : (sP0('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f1643,f484])).
fof(f2436,plain,(
  ( ! [X2,X0,X1] : (~sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f1643,f498])).
fof(f2438,plain,(
  ( ! [X2,X0,X1] : (~sP1('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f357,f1643,f498])).
fof(f2439,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f1643,f498])).
fof(f2440,plain,(
  ( ! [X2,X0,X1] : (sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f609,f1643,f499])).
fof(f2441,plain,(
  ( ! [X2,X0,X1] : (sP1('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f609,f1643,f502])).
fof(f2444,plain,(
  ( ! [X2,X0,X1] : (~sP5('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f357,f609,f1643,f526])).
fof(f2445,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1643,f609,f1643,f526])).
fof(f2447,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f629,f617,f623,f1643,f527])).
fof(f2449,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f629,f618,f623,f1643,f527])).
fof(f2451,plain,(
  ( ! [X2,X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f630,f617,f624,f1643,f527])).
fof(f2453,plain,(
  ( ! [X2,X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f630,f618,f624,f1643,f527])).
fof(f2454,plain,(
  ( ! [X2,X0,X1] : (~sP4('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f357,f609,f1643,f530])).
fof(f2460,plain,(
  ( ! [X2,X0,X1] : (~sP4('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f357,f609,f1643,f530])).
fof(f2461,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP4('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1643,f609,f1643,f530])).
fof(f2466,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f531])).
fof(f2467,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f357,f609,f1643,f534])).
fof(f2473,plain,(
  ( ! [X2,X0,X1] : (~sP3('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f357,f609,f1643,f534])).
fof(f2474,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP3('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1643,f609,f1643,f534])).
fof(f2479,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f535])).
fof(f2480,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f539])).
fof(f2481,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f1643,f540])).
fof(f2482,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,X4)) )),
  inference(unit_resulting_resolution,[],[f1643,f541])).
fof(f2483,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f610,f1643,f543])).
fof(f2484,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f657,f1643,f543])).
fof(f2485,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f674,f1643,f543])).
fof(f2486,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(X3,bool),X4),s('type.list.list'(X3),'const.list.NIL_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f430,f1643,f543])).
fof(f2487,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f610,f1643,f543])).
fof(f2488,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f657,f1643,f543])).
fof(f2489,plain,(
  ( ! [X2,X0,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f674,f1643,f543])).
fof(f2490,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','const.list.EVERY_2'(s(fun(X3,bool),X4),s('type.list.list'(X3),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f430,f1643,f543])).
fof(f2491,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f653,f1643,f544])).
fof(f2492,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f654,f1643,f544])).
fof(f2493,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f653,f1643,f544])).
fof(f2494,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f654,f1643,f544])).
fof(f2495,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'F_0')) )),
  inference(unit_resulting_resolution,[],[f609,f1643,f547])).
fof(f2496,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0',X3)) )),
  inference(unit_resulting_resolution,[],[f609,f1643,f548])).
fof(f2499,plain,(
  ( ! [X2,X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f636,f1643,f549])).
fof(f2504,plain,(
  ( ! [X2,X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f357,f636,f1643,f549])).
fof(f2505,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f636,f1643,f549])).
fof(f2507,plain,(
  ( ! [X2,X0,X1] : (~sP6('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f357,f609,f1643,f553])).
fof(f2513,plain,(
  ( ! [X2,X0,X1] : (~sP6('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f357,f609,f1643,f553])).
fof(f2514,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP6('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1643,f609,f1643,f553])).
fof(f2519,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f554])).
fof(f2520,plain,(
  ( ! [X2,X0,X1] : (sP7('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f357,f1643,f559])).
fof(f2525,plain,(
  ( ! [X2,X0,X1] : (sP7('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f357,f357,f1643,f559])).
fof(f2527,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f357,f1643,f1643,f559])).
fof(f2530,plain,(
  ( ! [X2,X0,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f357,f1643,f559])).
fof(f2531,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1643,f357,f1643,f559])).
fof(f2532,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f357,f1643,f1643,f559])).
fof(f2533,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),X8)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f1643,f1643,f559])).
fof(f2534,plain,(
  ( ! [X2,X0,X1] : (~sP7('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f609,f1643,f562])).
fof(f2538,plain,(
  ( ! [X2,X0,X1] : (~sP7('F_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f357,f609,f1643,f562])).
fof(f2539,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP7('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f609,f1643,f562])).
fof(f2542,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('F_0',X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f609,f1643,f563])).
fof(f2548,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f609,f1643,f564])).
fof(f2558,plain,(
  ( ! [X2,X0,X1] : (~sP8('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f609,f609,f1643,f570])).
fof(f2563,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f571])).
fof(f2564,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP8(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f1643,f572])).
fof(f2565,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('T_0',X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f357,f1643,f575])).
fof(f2569,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.INSERT_2'(s(X4,X5),s(fun(X4,bool),X6)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f1643,f575])).
fof(f2570,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f357,f1643,f576])).
fof(f2574,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),X6)) )),
  inference(unit_resulting_resolution,[],[f1643,f1643,f576])).
fof(f2575,plain,(
  ( ! [X2,X0,X1] : (~sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f609,f609,f1643,f577])).
fof(f2576,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f609,f1643,f581])).
fof(f2597,plain,(
  ( ! [X2,X0,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f609,f652,f1643,f586])).
fof(f2598,plain,(
  ( ! [X2,X0,X1] : (~sP10('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f609,f357,f1643,f587])).
fof(f2604,plain,(
  ( ! [X2,X0,X1] : (~sP10('T_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f357,f609,f1643,f587])).
fof(f2605,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP10('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f609,f1643,f587])).
fof(f2610,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f1643,f589])).
fof(f2611,plain,(
  ( ! [X2,X0,X1] : (sP11('F_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f609,f609,f1643,f591])).
fof(f2614,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f357,f1643,f593])).
fof(f2618,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))),X6)) )),
  inference(unit_resulting_resolution,[],[f1643,f1643,f593])).
fof(f2619,plain,(
  ( ! [X2,X0,X1] : (~sP11('T_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f357,f609,f1643,f594])).
fof(f2623,plain,(
  ( ! [X2,X0,X1] : (~sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f609,f1643,f594])).
fof(f2624,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'F_0','const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f609,f1643,f594])).
fof(f2625,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'F_0')) )),
  inference(unit_resulting_resolution,[],[f647,f1643,f597])).
fof(f2626,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f648,f1643,f597])).
fof(f2627,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f649,f1643,f597])).
fof(f2628,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X3)) )),
  inference(unit_resulting_resolution,[],[f650,f1643,f597])).
fof(f2629,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'T_0')) )),
  inference(unit_resulting_resolution,[],[f641,f1643,f598])).
fof(f2630,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f642,f1643,f598])).
fof(f2631,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'T_0',X3)) )),
  inference(unit_resulting_resolution,[],[f643,f1643,f598])).
fof(f2632,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X3)) )),
  inference(unit_resulting_resolution,[],[f644,f1643,f598])).
fof(f2633,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.INSERT_2'(s(X1,X2),s(fun(X1,bool),X3)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f740])).
fof(f2634,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f1643,f750])).
fof(f2635,plain,(
  ( ! [X2,X0,X3,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),X3)) )),
  inference(unit_resulting_resolution,[],[f1643,f782])).
fof(f2636,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),X2,'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.INSERT_2'(s(X3,X4),s(fun(X3,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f1643,f1067])).
fof(f2637,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f1643,f1068])).
fof(f2638,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP11('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.INSERT_2'(s(X2,X3),s(fun(X2,bool),X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f1643,f1070])).
fof(f2639,plain,(
  ( ! [X2,X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f609,f1643,f1089])).
fof(f2640,plain,(
  ( ! [X2,X0,X1] : (~sP7('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f609,f1643,f1191])).
fof(f2641,plain,(
  ( ! [X2,X0,X1] : (~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(unit_resulting_resolution,[],[f609,f1643,f1401])).
fof(f2149,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f2145,f609])).
fof(f2145,plain,(
  ( ! [X8] : (p(s(bool,'F_0')) | s(bool,X8) = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f906,f674])).
fof(f2148,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f2143,f609])).
fof(f2143,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'F_0')) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f906,f610])).
fof(f2100,plain,(
  ( ! [X2,X0,X1] : (sP11('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'T_0',X2)) )),
  inference(unit_resulting_resolution,[],[f357,f1070])).
fof(f2073,plain,(
  ( ! [X2,X0,X1] : (sP11('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),X2,'F_0')) )),
  inference(unit_resulting_resolution,[],[f609,f1069])).
fof(f2022,plain,(
  ( ! [X2,X0,X1] : (sP9('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'T_0',X2)) )),
  inference(unit_resulting_resolution,[],[f357,f1068])).
fof(f1951,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))) = s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),'const.list.NIL_0')))))) )),
  inference(superposition,[],[f441,f364])).
fof(f1937,plain,(
  ( ! [X2,X0,X1] : (sP9('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),X2,'T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f1067])).
fof(f1922,plain,(
  ( ! [X0,X1] : (sP1('F_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f609,f1066])).
fof(f1915,plain,(
  ( ! [X0,X1] : (sP1('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f609,f1065])).
fof(f1887,plain,(
  ( ! [X0,X1] : (~sP10('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f1882,f800])).
fof(f1882,plain,(
  ( ! [X0,X1] : (~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f609,f1417])).
fof(f1857,plain,(
  ~sP10('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f1846,f800])).
fof(f1846,plain,(
  ~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f609,f1416])).
fof(f1801,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f609,f631,f617,f1033,f529])).
fof(f1713,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X2))))))))) )),
  inference(forward_demodulation,[],[f1697,f610])).
fof(f1697,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),'const.list.CONS_2'(s(bool,'F_0'),s('type.list.list'(bool),X2))))))))) )),
  inference(unit_resulting_resolution,[],[f610,f466])).
fof(f1643,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))))))) )),
  inference(equality_resolution,[],[f463])).
fof(f1646,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2))))))) )),
  inference(forward_demodulation,[],[f1626,f430])).
fof(f1626,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0'))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X2))))))) )),
  inference(unit_resulting_resolution,[],[f430,f463])).
fof(f1645,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),X1))))))) )),
  inference(forward_demodulation,[],[f1627,f364])).
fof(f1627,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0'))),s(fun(fun(X0,bool),bool),'const.pred_set.INSERT_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),X1))))))) )),
  inference(unit_resulting_resolution,[],[f364,f463])).
fof(f1644,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2))))))) )),
  inference(forward_demodulation,[],[f1628,f610])).
fof(f1628,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2))))))) )),
  inference(unit_resulting_resolution,[],[f610,f463])).
fof(f1548,plain,(
  ~sP3('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f634,f975,f617,f609,f529])).
fof(f1559,plain,(
  ( ! [X0,X1] : (~sP2('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f988,f619,f625,f609,f529])).
fof(f1562,plain,(
  sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f632,f618,f625,f609,f529])).
fof(f1570,plain,(
  sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f632,f618,f626,f609,f529])).
fof(f1579,plain,(
  sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f634,f621,f628,f609,f529])).
fof(f1581,plain,(
  ~sP2('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f974,f619,f628,f609,f529])).
fof(f1498,plain,(
  sP5('F_0','F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f634,f621,f627,f609,f528])).
fof(f1501,plain,(
  ( ! [X0,X1] : (~sP2('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f990,f621,f627,f609,f528])).
fof(f1503,plain,(
  sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f634,f622,f627,f609,f528])).
fof(f1510,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f629,f619,f625,f609,f528])).
fof(f1511,plain,(
  sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f630,f619,f625,f609,f528])).
fof(f1515,plain,(
  sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f630,f620,f626,f609,f528])).
fof(f1517,plain,(
  ~sP2('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f976,f621,f626,f609,f528])).
fof(f1451,plain,(
  ~sP3('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f633,f973,f621,f357,f527])).
fof(f1452,plain,(
  ( ! [X0,X1] : (~sP3('const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f633,f1031,f621,f357,f527])).
fof(f1455,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f633,f617,f623,f357,f527])).
fof(f1457,plain,(
  sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f633,f617,f624,f357,f527])).
fof(f1458,plain,(
  sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f633,f618,f624,f357,f527])).
fof(f1461,plain,(
  sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f633,f622,f627,f357,f527])).
fof(f1462,plain,(
  ~sP4('T_0','T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f633,f959,f628,f357,f527])).
fof(f1464,plain,(
  sP5('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f633,f622,f628,f357,f527])).
fof(f1432,plain,(
  ( ! [X0,X1] : (sP4('T_0','T_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f1032,f525])).
fof(f1421,plain,(
  ~sP10('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f357,f609,f1403,f586])).
fof(f1417,plain,(
  ( ! [X10,X11,X9] : (~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X11,'const.list.EVERY_2'(s(fun(X9,bool),X10),s('type.list.list'(X9),'const.list.NIL_0'))) | p(s(bool,X11))) )),
  inference(subsumption_resolution,[],[f1413,f357])).
fof(f1413,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'T_0')) | p(s(bool,X11)) | ~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X11,'const.list.EVERY_2'(s(fun(X9,bool),X10),s('type.list.list'(X9),'const.list.NIL_0')))) )),
  inference(superposition,[],[f1401,f430])).
fof(f1416,plain,(
  ( ! [X7] : (~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X7,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15))) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f1411,f357])).
fof(f1411,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X7,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) )),
  inference(superposition,[],[f1401,f657])).
fof(f1403,plain,(
  ~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f357,f609,f1401])).
fof(f1402,plain,(
  ( ! [X14,X17,X15,X16] : (~sP11('const.list.EVERY_2'(s(fun(X14,bool),X15),s('type.list.list'(X14),'const.list.NIL_0')),X17,X16) | p(s(bool,X17)) | ~p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f1397,f357])).
fof(f1397,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | p(s(bool,X17)) | ~sP11('const.list.EVERY_2'(s(fun(X14,bool),X15),s('type.list.list'(X14),'const.list.NIL_0')),X17,X16)) )),
  inference(superposition,[],[f594,f430])).
fof(f1401,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | ~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X11,X10)) )),
  inference(subsumption_resolution,[],[f1395,f357])).
fof(f1395,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | p(s(bool,X11)) | ~sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X11,X10)) )),
  inference(superposition,[],[f594,f657])).
fof(f1362,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1353,f582])).
fof(f1353,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f357,f357,f609,f587])).
fof(f1343,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f651,f357,f609,f586])).
fof(f1344,plain,(
  sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f652,f357,f609,f586])).
fof(f1322,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f991,f525])).
fof(f1309,plain,(
  ~sP8('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f609,f609,f1299,f569])).
fof(f1299,plain,(
  ~sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f609,f609,f1297])).
fof(f1298,plain,(
  ( ! [X14,X17,X15,X16] : (~sP9('const.list.EVERY_2'(s(fun(X14,bool),X15),s('type.list.list'(X14),'const.list.NIL_0')),X17,X16) | p(s(bool,X17)) | p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f1293,f357])).
fof(f1293,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'T_0')) | p(s(bool,X16)) | p(s(bool,X17)) | ~sP9('const.list.EVERY_2'(s(fun(X14,bool),X15),s('type.list.list'(X14),'const.list.NIL_0')),X17,X16)) )),
  inference(superposition,[],[f577,f430])).
fof(f1297,plain,(
  ( ! [X10,X11] : (~sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X11,X10) | p(s(bool,X11)) | p(s(bool,X10))) )),
  inference(subsumption_resolution,[],[f1291,f357])).
fof(f1291,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X10)) | p(s(bool,X11)) | ~sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X11,X10)) )),
  inference(superposition,[],[f577,f657])).
fof(f1247,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1238,f565])).
fof(f1238,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f357,f609,f609,f570])).
fof(f1228,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f645,f609,f609,f569])).
fof(f1229,plain,(
  sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f646,f609,f609,f569])).
fof(f1202,plain,(
  ~sP6('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f609,f1195,f599])).
fof(f1195,plain,(
  ~sP7('F_0','T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f357,f609,f1191])).
fof(f1192,plain,(
  ( ! [X14,X17,X15,X16] : (~sP7(X17,X16,'const.list.EVERY_2'(s(fun(X14,bool),X15),s('type.list.list'(X14),'const.list.NIL_0'))) | p(s(bool,X17)) | ~p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f1187,f357])).
fof(f1187,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | p(s(bool,X17)) | ~sP7(X17,X16,'const.list.EVERY_2'(s(fun(X14,bool),X15),s('type.list.list'(X14),'const.list.NIL_0')))) )),
  inference(superposition,[],[f562,f430])).
fof(f1191,plain,(
  ( ! [X10,X11] : (~sP7(X11,X10,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15))) | p(s(bool,X11)) | ~p(s(bool,X10))) )),
  inference(subsumption_resolution,[],[f1185,f357])).
fof(f1185,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | p(s(bool,X11)) | ~sP7(X11,X10,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) )),
  inference(superposition,[],[f562,f657])).
fof(f1173,plain,(
  ~sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f357,f609,f1089])).
fof(f1163,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1154,f546])).
fof(f1154,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f357,f609,f357,f553])).
fof(f1143,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f635,f357,f357,f549])).
fof(f1144,plain,(
  sP7('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f636,f357,f357,f549])).
fof(f1139,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1130,f523])).
fof(f1130,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f609,f609,f609,f538])).
fof(f1124,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1115,f524])).
fof(f1115,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f357,f609,f357,f534])).
fof(f1099,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f357,f609,f357,f530])).
fof(f1090,plain,(
  ( ! [X14,X17,X15,X16] : (~sP5(X16,'const.list.EVERY_2'(s(fun(X14,bool),X15),s('type.list.list'(X14),'const.list.NIL_0')),X17) | p(s(bool,X17)) | ~p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f1086,f357])).
fof(f1086,plain,(
  ( ! [X14,X17,X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | p(s(bool,X17)) | ~sP5(X16,'const.list.EVERY_2'(s(fun(X14,bool),X15),s('type.list.list'(X14),'const.list.NIL_0')),X17)) )),
  inference(superposition,[],[f526,f430])).
fof(f1089,plain,(
  ( ! [X10,X11] : (~sP5(X10,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X11) | p(s(bool,X11)) | ~p(s(bool,X10))) )),
  inference(subsumption_resolution,[],[f1084,f357])).
fof(f1084,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | p(s(bool,X11)) | ~sP5(X10,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X11)) )),
  inference(superposition,[],[f526,f657])).
fof(f1072,plain,(
  ( ! [X54,X52,X53,X51] : (sP9('const.list.EVERY_2'(s(fun(X51,bool),X52),s('type.list.list'(X51),'const.list.NIL_0')),X53,X54) | ~sP8('const.list.EVERY_2'(s(fun(X51,bool),X52),s('type.list.list'(X51),'const.list.NIL_0')),X53,X54)) )),
  inference(subsumption_resolution,[],[f1054,f357])).
fof(f1054,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~sP8('const.list.EVERY_2'(s(fun(X51,bool),X52),s('type.list.list'(X51),'const.list.NIL_0')),X53,X54) | sP9('const.list.EVERY_2'(s(fun(X51,bool),X52),s('type.list.list'(X51),'const.list.NIL_0')),X53,X54)) )),
  inference(superposition,[],[f598,f430])).
fof(f1071,plain,(
  ( ! [X47,X50,X48,X49] : (sP11('const.list.EVERY_2'(s(fun(X47,bool),X48),s('type.list.list'(X47),'const.list.NIL_0')),X49,X50) | ~sP10('const.list.EVERY_2'(s(fun(X47,bool),X48),s('type.list.list'(X47),'const.list.NIL_0')),X49,X50)) )),
  inference(subsumption_resolution,[],[f1053,f357])).
fof(f1053,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~sP10('const.list.EVERY_2'(s(fun(X47,bool),X48),s('type.list.list'(X47),'const.list.NIL_0')),X49,X50) | sP11('const.list.EVERY_2'(s(fun(X47,bool),X48),s('type.list.list'(X47),'const.list.NIL_0')),X49,X50)) )),
  inference(superposition,[],[f597,f430])).
fof(f1070,plain,(
  ( ! [X45,X43,X46,X44] : (sP11('const.list.EVERY_2'(s(fun(X43,bool),X44),s('type.list.list'(X43),'const.list.NIL_0')),X45,X46) | ~p(s(bool,X45))) )),
  inference(subsumption_resolution,[],[f1052,f357])).
fof(f1052,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | sP11('const.list.EVERY_2'(s(fun(X43,bool),X44),s('type.list.list'(X43),'const.list.NIL_0')),X45,X46)) )),
  inference(superposition,[],[f593,f430])).
fof(f1069,plain,(
  ( ! [X39,X41,X42,X40] : (sP11('const.list.EVERY_2'(s(fun(X39,bool),X40),s('type.list.list'(X39),'const.list.NIL_0')),X42,X41) | p(s(bool,X41))) )),
  inference(subsumption_resolution,[],[f1051,f357])).
fof(f1051,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,X41)) | sP11('const.list.EVERY_2'(s(fun(X39,bool),X40),s('type.list.list'(X39),'const.list.NIL_0')),X42,X41)) )),
  inference(superposition,[],[f592,f430])).
fof(f1068,plain,(
  ( ! [X37,X35,X38,X36] : (sP9('const.list.EVERY_2'(s(fun(X35,bool),X36),s('type.list.list'(X35),'const.list.NIL_0')),X37,X38) | ~p(s(bool,X37))) )),
  inference(subsumption_resolution,[],[f1050,f357])).
fof(f1050,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X37)) | sP9('const.list.EVERY_2'(s(fun(X35,bool),X36),s('type.list.list'(X35),'const.list.NIL_0')),X37,X38)) )),
  inference(superposition,[],[f576,f430])).
fof(f1067,plain,(
  ( ! [X33,X31,X34,X32] : (sP9('const.list.EVERY_2'(s(fun(X31,bool),X32),s('type.list.list'(X31),'const.list.NIL_0')),X34,X33) | ~p(s(bool,X33))) )),
  inference(subsumption_resolution,[],[f1049,f357])).
fof(f1049,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X33)) | sP9('const.list.EVERY_2'(s(fun(X31,bool),X32),s('type.list.list'(X31),'const.list.NIL_0')),X34,X33)) )),
  inference(superposition,[],[f575,f430])).
fof(f1066,plain,(
  ( ! [X30,X28,X29] : (sP1(X30,'const.list.EVERY_2'(s(fun(X28,bool),X29),s('type.list.list'(X28),'const.list.NIL_0'))) | p(s(bool,X30))) )),
  inference(subsumption_resolution,[],[f1048,f357])).
fof(f1048,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP1(X30,'const.list.EVERY_2'(s(fun(X28,bool),X29),s('type.list.list'(X28),'const.list.NIL_0')))) )),
  inference(superposition,[],[f502,f430])).
fof(f1065,plain,(
  ( ! [X26,X27,X25] : (sP1('const.list.EVERY_2'(s(fun(X25,bool),X26),s('type.list.list'(X25),'const.list.NIL_0')),X27) | p(s(bool,X27))) )),
  inference(subsumption_resolution,[],[f1047,f357])).
fof(f1047,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1('const.list.EVERY_2'(s(fun(X25,bool),X26),s('type.list.list'(X25),'const.list.NIL_0')),X27)) )),
  inference(superposition,[],[f499,f430])).
fof(f1064,plain,(
  ( ! [X24,X23,X22] : (sP0('const.list.EVERY_2'(s(fun(X22,bool),X23),s('type.list.list'(X22),'const.list.NIL_0')),X24)) )),
  inference(subsumption_resolution,[],[f1046,f357])).
fof(f1046,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0('const.list.EVERY_2'(s(fun(X22,bool),X23),s('type.list.list'(X22),'const.list.NIL_0')),X24)) )),
  inference(superposition,[],[f484,f430])).
fof(f1063,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.list.EVERY_2'(s(fun(X19,bool),X20),s('type.list.list'(X19),'const.list.NIL_0')))) )),
  inference(subsumption_resolution,[],[f1045,f357])).
fof(f1045,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X21,'const.list.EVERY_2'(s(fun(X19,bool),X20),s('type.list.list'(X19),'const.list.NIL_0')))) )),
  inference(superposition,[],[f483,f430])).
fof(f1031,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f609,f430,f545])).
fof(f1032,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f430,f543])).
fof(f1033,plain,(
  ( ! [X0,X1] : (~sP5('F_0','T_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f609,f430,f545])).
fof(f1034,plain,(
  ( ! [X0,X1] : (sP5('T_0','T_0','const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(unit_resulting_resolution,[],[f357,f430,f543])).
fof(f993,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f610,f366])).
fof(f988,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f609,f610,f545])).
fof(f989,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f357,f610,f543])).
fof(f990,plain,(
  ( ! [X0,X1] : (~sP5('F_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f609,f610,f545])).
fof(f991,plain,(
  ( ! [X0,X1] : (sP5('T_0','F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f357,f610,f543])).
fof(f973,plain,(
  ~sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f609,f657,f545])).
fof(f974,plain,(
  ~sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f609,f674,f545])).
fof(f975,plain,(
  ~sP5('F_0','T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f609,f657,f545])).
fof(f976,plain,(
  ~sP5('F_0','F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f609,f674,f545])).
fof(f958,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f357,f653,f544])).
fof(f959,plain,(
  ~sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f357,f654,f544])).
fof(f962,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f357,f653,f544])).
fof(f963,plain,(
  ~sP5('T_0','T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f357,f654,f544])).
fof(f947,plain,(
  sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f357,f657,f543])).
fof(f948,plain,(
  sP5('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f357,f674,f543])).
fof(f949,plain,(
  sP5('T_0','T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f357,f657,f543])).
fof(f950,plain,(
  sP5('T_0','F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f357,f674,f543])).
fof(f937,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f609,f653,f542])).
fof(f938,plain,(
  sP5('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),'T_0')),
  inference(unit_resulting_resolution,[],[f609,f654,f542])).
fof(f939,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f609,f653,f542])).
fof(f940,plain,(
  sP5('F_0','T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f609,f654,f542])).
fof(f931,plain,(
  ( ! [X0] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f642,f818])).
fof(f933,plain,(
  ( ! [X0] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f644,f818])).
fof(f906,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f899,f609])).
fof(f899,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f470,f366])).
fof(f889,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f648,f800])).
fof(f891,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f650,f800])).
fof(f882,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f357,f782])).
fof(f836,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f366,f366])).
fof(f826,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f638,f609,f599])).
fof(f828,plain,(
  ( ! [X0] : (sP7('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f640,f609,f599])).
fof(f818,plain,(
  ( ! [X0,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,X1) | ~sP8('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,X1)) )),
  inference(subsumption_resolution,[],[f816,f357])).
fof(f816,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP8('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,X1) | sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,X1)) )),
  inference(superposition,[],[f598,f657])).
fof(f813,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f642,f357,f598])).
fof(f815,plain,(
  ( ! [X0] : (sP9('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f644,f357,f598])).
fof(f800,plain,(
  ( ! [X0,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,X1) | ~sP10('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,X1)) )),
  inference(subsumption_resolution,[],[f798,f357])).
fof(f798,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP10('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,X1) | sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,X1)) )),
  inference(superposition,[],[f597,f657])).
fof(f795,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f648,f357,f597])).
fof(f797,plain,(
  ( ! [X0] : (sP11('T_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f650,f357,f597])).
fof(f786,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f609,f768])).
fof(f782,plain,(
  ( ! [X0,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,X1) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f780,f357])).
fof(f780,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,X1)) )),
  inference(superposition,[],[f593,f657])).
fof(f779,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f357,f357,f593])).
fof(f768,plain,(
  ( ! [X0,X1] : (sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X1,X0) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f766,f357])).
fof(f766,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP11('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X1,X0)) )),
  inference(superposition,[],[f592,f657])).
fof(f765,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f609,f357,f592])).
fof(f762,plain,(
  ( ! [X0] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f357,f750])).
fof(f761,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f357,f609,f581])).
fof(f755,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f609,f609,f580])).
fof(f754,plain,(
  ( ! [X0] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f740])).
fof(f750,plain,(
  ( ! [X0,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,X1) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f748,f357])).
fof(f748,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,X1)) )),
  inference(superposition,[],[f576,f657])).
fof(f747,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f357,f357,f576])).
fof(f740,plain,(
  ( ! [X0,X1] : (sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X1,X0) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f738,f357])).
fof(f738,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X1,X0)) )),
  inference(superposition,[],[f575,f657])).
fof(f737,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f357,f575])).
fof(f735,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f357,f609,f564])).
fof(f734,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f609,f563])).
fof(f727,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f609,f609,f558])).
fof(f724,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f609,f609,f557])).
fof(f721,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f357,f609,f548])).
fof(f718,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f357,f609,f547])).
fof(f717,plain,(
  sP1('F_0','const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f609,f703])).
fof(f712,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f609,f609,f612])).
fof(f703,plain,(
  ( ! [X0] : (sP1(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f701,f357])).
fof(f701,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) )),
  inference(superposition,[],[f502,f657])).
fof(f700,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f609,f357,f502])).
fof(f697,plain,(
  sP1('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),'F_0')),
  inference(unit_resulting_resolution,[],[f609,f694])).
fof(f694,plain,(
  ( ! [X0] : (sP1('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f692,f357])).
fof(f692,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0)) )),
  inference(superposition,[],[f499,f657])).
fof(f691,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f609,f357,f499])).
fof(f688,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f357,f357,f498])).
fof(f687,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f609,f609,f497])).
fof(f657,plain,(
  s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f355,f376])).
fof(f653,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f609,f375])).
fof(f654,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))),
  inference(unit_resulting_resolution,[],[f356,f375])).
fof(f651,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f609,f590])).
fof(f652,plain,(
  ( ! [X0,X1] : (sP10('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f356,f590])).
fof(f649,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f357,f589])).
fof(f650,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f355,f589])).
fof(f647,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f609,f588])).
fof(f648,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f356,f588])).
fof(f645,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f609,f573])).
fof(f646,plain,(
  ( ! [X0,X1] : (sP8('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f356,f573])).
fof(f643,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f357,f572])).
fof(f644,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f355,f572])).
fof(f641,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f571])).
fof(f642,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f355,f571])).
fof(f639,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f609,f556])).
fof(f640,plain,(
  ( ! [X0,X1] : (sP6('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f356,f556])).
fof(f637,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f609,f555])).
fof(f638,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f356,f555])).
fof(f635,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f554])).
fof(f636,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f355,f554])).
fof(f633,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f357,f541])).
fof(f634,plain,(
  ( ! [X0,X1] : (sP2('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f355,f541])).
fof(f631,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f357,f540])).
fof(f632,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f355,f540])).
fof(f629,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f539])).
fof(f630,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f355,f539])).
fof(f627,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f609,f537])).
fof(f628,plain,(
  ( ! [X0,X1] : (sP3('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f356,f537])).
fof(f625,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f609,f536])).
fof(f626,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f356,f536])).
fof(f623,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f535])).
fof(f624,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f355,f535])).
fof(f621,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f609,f533])).
fof(f622,plain,(
  ( ! [X0,X1] : (sP4('const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f356,f533])).
fof(f619,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f609,f532])).
fof(f620,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15)),X1)) )),
  inference(unit_resulting_resolution,[],[f356,f532])).
fof(f617,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f531])).
fof(f618,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f355,f531])).
fof(f615,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f357,f484])).
fof(f616,plain,(
  ( ! [X0] : (sP0('const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)),X0)) )),
  inference(unit_resulting_resolution,[],[f355,f484])).
fof(f613,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f357,f483])).
fof(f614,plain,(
  ( ! [X0] : (sP0(X0,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15)))) )),
  inference(unit_resulting_resolution,[],[f355,f483])).
fof(f591,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f353])).
fof(f353,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f352])).
fof(f352,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f212,f229,f228])).
fof(f228,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f229,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f212,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X24,X21,X20] : ((p(s(bool,X20)) <=> (p(s(bool,X21)) => p(s(bool,X24)))) <=> ((p(s(bool,X20)) | p(s(bool,X21))) & (p(s(bool,X20)) | ~p(s(bool,X24))) & (~p(s(bool,X21)) | p(s(bool,X24)) | ~p(s(bool,X20)))))),
  file('thm.list.EVERY_MEM_MONO','thm.sat.sat_27_prove_ltb')).
fof(f592,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f353])).
fof(f593,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f353])).
fof(f594,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f353])).
fof(f587,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f351])).
fof(f351,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f350])).
fof(f350,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f349])).
fof(f349,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f228])).
fof(f588,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f351])).
fof(f589,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f351])).
fof(f590,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f351])).
fof(f580,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f347])).
fof(f347,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f346])).
fof(f346,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f229])).
fof(f581,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f348])).
fof(f582,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f348])).
fof(f597,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f583])).
fof(f583,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f348])).
fof(f584,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f348])).
fof(f585,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f348])).
fof(f586,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f348])).
fof(f574,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f344])).
fof(f344,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f173,f226,f225])).
fof(f225,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f226,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f173,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X24,X21,X20] : ((p(s(bool,X20)) <=> (p(s(bool,X21)) | p(s(bool,X24)))) <=> ((p(s(bool,X20)) | ~p(s(bool,X21))) & (p(s(bool,X20)) | ~p(s(bool,X24))) & (p(s(bool,X21)) | p(s(bool,X24)) | ~p(s(bool,X20)))))),
  file('thm.list.EVERY_MEM_MONO','thm.sat.sat_25_prove_ltb')).
fof(f575,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f345])).
fof(f576,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f345])).
fof(f577,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f345])).
fof(f570,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f342])).
fof(f342,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f341])).
fof(f341,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f225])).
fof(f571,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f343])).
fof(f572,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f343])).
fof(f573,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f343])).
fof(f563,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f339])).
fof(f339,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f338])).
fof(f338,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f226])).
fof(f564,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f340])).
fof(f565,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f340])).
fof(f598,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f566])).
fof(f566,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f340])).
fof(f567,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f340])).
fof(f568,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f340])).
fof(f569,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f340])).
fof(f557,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f337])).
fof(f337,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f336])).
fof(f336,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f171,f223,f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f223,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f171,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X24,X21,X20] : ((p(s(bool,X20)) <=> (p(s(bool,X21)) & p(s(bool,X24)))) <=> ((p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X24))) & (p(s(bool,X21)) | ~p(s(bool,X20))) & (p(s(bool,X24)) | ~p(s(bool,X20)))))),
  file('thm.list.EVERY_MEM_MONO','thm.sat.sat_23_prove_ltb')).
fof(f558,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f337])).
fof(f559,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f337])).
fof(f562,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f337])).
fof(f553,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f335])).
fof(f335,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f334])).
fof(f334,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f222])).
fof(f554,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f335])).
fof(f555,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f335])).
fof(f556,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f335])).
fof(f546,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f332])).
fof(f332,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f331])).
fof(f331,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f223])).
fof(f547,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f333])).
fof(f548,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f333])).
fof(f549,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f333])).
fof(f599,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f552])).
fof(f552,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f333])).
fof(f542,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f330])).
fof(f330,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f169,f220,f219,f218,f217])).
fof(f217,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f218,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f219,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f220,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f169,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X24,X21,X20] : ((p(s(bool,X20)) <=> s(bool,X21) = s(bool,X24)) <=> ((p(s(bool,X20)) | p(s(bool,X21)) | p(s(bool,X24))) & (p(s(bool,X20)) | ~p(s(bool,X24)) | ~p(s(bool,X21))) & (p(s(bool,X21)) | ~p(s(bool,X24)) | ~p(s(bool,X20))) & (p(s(bool,X24)) | ~p(s(bool,X21)) | ~p(s(bool,X20)))))),
  file('thm.list.EVERY_MEM_MONO','thm.sat.sat_2|49|_prove_ltb')).
fof(f543,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f330])).
fof(f544,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f330])).
fof(f545,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f330])).
fof(f538,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f328])).
fof(f328,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f217])).
fof(f539,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f329])).
fof(f540,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f329])).
fof(f541,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f329])).
fof(f534,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f326])).
fof(f326,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f218])).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f327])).
fof(f536,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f327])).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f327])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f323])).
fof(f323,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f322])).
fof(f322,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f219])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f324])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f324])).
fof(f533,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f324])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f320])).
fof(f320,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f319])).
fof(f319,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f220])).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f321])).
fof(f525,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f321])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f321])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f321])).
fof(f528,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f321])).
fof(f529,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f321])).
fof(f516,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK29(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK29(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK29(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK29(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK29])],[f211])).
fof(f211,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.list.EVERY_MEM_MONO',aHL_EXT)).
fof(f497,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f310])).
fof(f310,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f309])).
fof(f309,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f308])).
fof(f308,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f498,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f310])).
fof(f499,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f310])).
fof(f502,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f310])).
fof(f612,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f482,f609])).
fof(f482,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f302])).
fof(f302,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f483,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f303])).
fof(f484,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f303])).
fof(f470,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f195])).
fof(f195,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X22,X23] : ((p(s(bool,X22)) => p(s(bool,X23))) => ((p(s(bool,X23)) => p(s(bool,X22))) => s(bool,X22) = s(bool,X23)))),
  file('thm.list.EVERY_MEM_MONO','thm.bool.IMP_ANTISYM_AX')).
fof(f473,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f196])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1))))) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2] : s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1))))) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0] : ! [X1,X2] : s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1))))) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X16] : ! [X15,X17] : s(fun(X16,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X16),'const.list.CONS_2'(s(X16,X17),s('type.list.list'(X16),X15))))) = s(fun(X16,bool),'const.pred_set.INSERT_2'(s(X16,X17),s(fun(X16,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X16),X15)))))),
  file('thm.list.EVERY_MEM_MONO','thm.list.LIST_TO_SET|split|1')).
fof(f430,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X7] : s(bool,'T_0') = s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.list.EVERY_MEM_MONO','thm.list.EVERY_DEF|split|0')).
fof(f375,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.list.EVERY_MEM_MONO','thm.bool.EQ_CLAUSES|split|1')).
fof(f376,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f234])).
fof(f366,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.list.EVERY_MEM_MONO',aHL_BOOL_CASES)).
fof(f364,plain,(
  ( ! [X0] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.list.EVERY_MEM_MONO','thm.list.LIST_TO_SET|split|0')).
fof(f357,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.EVERY_MEM_MONO',aHL_TRUTH)).
fof(f355,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK13),s('type.list.list'(sK12),sK15))))),
  inference(cnf_transformation,[],[f231])).
fof(f356,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15))))),
  inference(cnf_transformation,[],[f231])).
fof(f11492,plain,(
  p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f11171,f674])).
fof(f11171,plain,(
  p(s(bool,'const.list.EVERY_2'(s(fun(sK12,bool),sK14),s('type.list.list'(sK12),sK15))))),
  inference(unit_resulting_resolution,[],[f10643,f446])).
fof(f10643,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK14),s(sK12,sK20(sK15,sK14,sK12)))))),
  inference(subsumption_resolution,[],[f10461,f357])).
fof(f10461,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK14),s(sK12,sK20(sK15,sK14,sK12)))))),
  inference(superposition,[],[f4362,f10439])).
fof(f10439,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK20(sK15,sK14,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK15)))))),
  inference(subsumption_resolution,[],[f10430,f5000])).
fof(f10430,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK20(sK15,sK14,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK15)))))),
  inference(superposition,[],[f1856,f674])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 492
% Passive clauses: 1528
% Generated clauses: 11443
% Final active clauses: 448
% Final passive clauses: 722
% Input formulas: 69
% Initial clauses: 243
% 
% Duplicate literals: 116
% Trivial inequalities: 30
% Fw subsumption resolutions: 326
% Fw demodulations: 1434
% Bw demodulations: 353
% Condensations: 702
% Global subsumptions: 5
% 
% Simple tautologies: 118
% Equational tautologies: 60
% Forward subsumptions: 7063
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 3180
% Factoring: 1
% Forward superposition: 3136
% Backward superposition: 1481
% Self superposition: 30
% Equality factoring: 28
% Equality resolution: 12
% 
% Split clauses: 50
% Split components: 110
% SAT solver clauses: 1639
% SAT solver unit clauses: 673
% SAT solver binary clauses: 373
% 
% Sat splits: 50
% Sat splitting refutations: 4
% 
% Memory used [KB]: 7164
% Time elapsed: 0.365 s
% ------------------------------
% ------------------------------
% Success in time 1.7 s
