Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.extreal
% (43324)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1332
% Passive clauses: 7035
% Generated clauses: 37026
% Final active clauses: 1235
% Final passive clauses: 5369
% Input formulas: 261
% Initial clauses: 508
% 
% Duplicate literals: 211
% Trivial inequalities: 109
% Fw subsumption resolutions: 2139
% Fw demodulations: 3279
% Bw demodulations: 40
% 
% Simple tautologies: 305
% Equational tautologies: 29
% Forward subsumptions: 23607
% 
% Binary resolution: 6756
% Factoring: 599
% Forward superposition: 11526
% Backward superposition: 11601
% Self superposition: 46
% Equality factoring: 2
% Equality resolution: 25
% 
% Split clauses: 383
% Split components: 768
% SAT solver clauses: 553
% SAT solver unit clauses: 27
% SAT solver binary clauses: 90
% 
% Sat splits: 383
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 552
% TWLsolver calls for satisfiability: 450
% 
% Memory used [KB]: 18421
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (43339)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 707
% Passive clauses: 3611
% Generated clauses: 24833
% Final active clauses: 680
% Final passive clauses: 2825
% Input formulas: 261
% Initial clauses: 508
% 
% Duplicate literals: 100
% Trivial inequalities: 45
% Fw subsumption resolutions: 708
% Fw demodulations: 3125
% Bw demodulations: 105
% Condensations: 60
% Global subsumptions: 34
% 
% Simple tautologies: 284
% Equational tautologies: 18
% Forward subsumptions: 12634
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 2568
% Factoring: 110
% Forward superposition: 6882
% Backward superposition: 6869
% Self superposition: 68
% Equality factoring: 2
% Equality resolution: 24
% 
% Split clauses: 16
% Split components: 32
% SAT solver clauses: 4665
% SAT solver unit clauses: 668
% SAT solver binary clauses: 155
% SAT solver learnt clauses: 68
% SAT solver learnt literals: 10
% 
% Sat splits: 16
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 4529
% TWLsolver calls for satisfiability: 5193
% 
% Memory used [KB]: 23922
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f71526,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5399,f5401,f5405,f5403,f5406,f15940,f15955,f15954,f17172,f17175,f17177,f17179,f17181,f17185,f17187,f17191,f17193,f17215,f17217,f17239,f17241,f17264,f17287,f17310,f17450,f17452,f17454,f17489,f17501,f17568,f17570,f17572,f17584,f17596,f17919,f18029,f18828,f18830,f18832,f18837,f18849,f18851,f18883,f18894,f19006,f19020,f19170,f19184,f19920,f19922,f19924,f19959,f19971,f20038,f20040,f20042,f20054,f20066,f20389,f20499,f21298,f21300,f21302,f21307,f21319,f21321,f21353,f21364,f21476,f21490,f21640,f21654,f22289,f22291,f22293,f22321,f22323,f22325,f22390,f22392,f22421,f22423,f22488,f22490,f22519,f22521,f22586,f22588,f22617,f22619,f22641,f22643,f24560,f24561,f22666,f22670,f22694,f22699,f22702,f22706,f22730,f22735,f22760,f22785,f22813,f22815,f22817,f22836,f22838,f22857,f22859,f22878,f22880,f22899,f22901,f22933,f22969,f23005,f23267,f23269,f23271,f23273,f23275,f23277,f23306,f23308,f23310,f23312,f23314,f23316,f23318,f23320,f23322,f23324,f23326,f23328,f23421,f23425,f23430,f23432,f23434,f23436,f23438,f23459,f23461,f23463,f23465,f23467,f23469,f23471,f23473,f23475,f23477,f23479,f23481,f23483,f23485,f23487,f23489,f23491,f23493,f23495,f23497,f23529,f23552,f23584,f23607,f23609,f23647,f23649,f23669,f23671,f23709,f23711,f23731,f23733,f23773,f23775,f23795,f23797,f23832,f23868,f23937,f23939,f23968,f23970,f23999,f24001,f24030,f24115,f24117,f24119,f24121,f24123,f24125,f24152,f24154,f24156,f24191,f24193,f24213,f24231,f24249,f24281,f24284,f24286,f24288,f24290,f24292,f24310,f24312,f24314,f24316,f24318,f24320,f24322,f24324,f24326,f24328,f24330,f24381,f24383,f24385,f24387,f24389,f24391,f24398,f24400,f24402,f24404,f24406,f24408,f24444,f24446,f24482,f24484,f24486,f24488,f24490,f24495,f24497,f24499,f24501,f24503,f24505,f24507,f24509,f24511,f24513,f24515,f24517,f24519,f24521,f24523,f24525,f24562,f24563,f24564,f24565,f24566,f24567,f24581,f24582,f24583,f24584,f24585,f24586,f24587,f24588,f24589,f24590,f24591,f24592,f24593,f24594,f24595,f24596,f24597,f24598,f24599,f24600,f24601,f24602,f24603,f24604,f24605,f24606,f24607,f24608,f24609,f24610,f24611,f24612,f24613,f24614,f24615,f24616,f24617,f24618,f24619,f24620,f24621,f24622,f24623,f24624,f24625,f24626,f24627,f24628,f24629,f24630,f24631,f24632,f24633,f24634,f24635,f24636,f24637,f24638,f24639,f24640,f24641,f24642,f24643,f24644,f24645,f24646,f24647,f24648,f24649,f24650,f24651,f24652,f24653,f24654,f24655,f24656,f24657,f24658,f24659,f24660,f24661,f24662,f24663,f24664,f24665,f24666,f24667,f24668,f24669,f24670,f24671,f24672,f24673,f24674,f24675,f24676,f24677,f24678,f24679,f24680,f24681,f24682,f24683,f24684,f24685,f24686,f24687,f24688,f24689,f24690,f24691,f24692,f24693,f24694,f24695,f24696,f24697,f24698,f24699,f24700,f24701,f24702,f24703,f24704,f24705,f24706,f24707,f24708,f24709,f24710,f24711,f24712,f24713,f24714,f24715,f24716,f24717,f24718,f24719,f24720,f24721,f24722,f24723,f24724,f24725,f24726,f24727,f24728,f24729,f24730,f24731,f24732,f24733,f24734,f24735,f24736,f24737,f24738,f24739,f24740,f24741,f24742,f24743,f24744,f24745,f24746,f24747,f24748,f24749,f24750,f24751,f24752,f24753,f24754,f24755,f24756,f24757,f24758,f24759,f24760,f24761,f24762,f24763,f24764,f24765,f24766,f24767,f24768,f24769,f24770,f24771,f24772,f24773,f24774,f24775,f24776,f24777,f24778,f24779,f24780,f24781,f24782,f24783,f24784,f24785,f24786,f24787,f24788,f24789,f24790,f24791,f24792,f24793,f24794,f24795,f24796,f24797,f24798,f24799,f24800,f24801,f24802,f24803,f24804,f24805,f24806,f24807,f24808,f24809,f24810,f24811,f24812,f24813,f24814,f24815,f24816,f24817,f24818,f24819,f24820,f24821,f24822,f24823,f24824,f24825,f24826,f24827,f24828,f24829,f24830,f24831,f24832,f24833,f24834,f24835,f24836,f24837,f24838,f24839,f24840,f24841,f24842,f24843,f24844,f24845,f24846,f24847,f24848,f24849,f24850,f24851,f24852,f24853,f24854,f24855,f24856,f24857,f24858,f24859,f24860,f24861,f24862,f24863,f24864,f24865,f24866,f24867,f24868,f24869,f24870,f24871,f24872,f24873,f24874,f24875,f24876,f24877,f24878,f24879,f24880,f24881,f24882,f24883,f24884,f24885,f24886,f24887,f24888,f24889,f24890,f24891,f24892,f24893,f24894,f24895,f24896,f24897,f24898,f24899,f24900,f24901,f24902,f24903,f24904,f24905,f24906,f24907,f24908,f24909,f24910,f24911,f24912,f24913,f24914,f24915,f24916,f24917,f24918,f24919,f24920,f24921,f24922,f24923,f24924,f24925,f24926,f24927,f24928,f24929,f24930,f24931,f24932,f24933,f24934,f24935,f24936,f24937,f24938,f24939,f24940,f24941,f24942,f24943,f24944,f24945,f24946,f24947,f24948,f24949,f24950,f24951,f24952,f24953,f24954,f24955,f24956,f24957,f24958,f24959,f24960,f24961,f24962,f24963,f24964,f24965,f24966,f24967,f24968,f24969,f24970,f24971,f24972,f24973,f24974,f24975,f24976,f24977,f24978,f24979,f24980,f24981,f24982,f24983,f24984,f24985,f24986,f24987,f24988,f24989,f24990,f24991,f24992,f24993,f24994,f24995,f24996,f24997,f24998,f24999,f25000,f25001,f25002,f25003,f25004,f25005,f25006,f25007,f25008,f25009,f25010,f25011,f25012,f25013,f25014,f25015,f25016,f25017,f25018,f25019,f25020,f25021,f25022,f25023,f25024,f25025,f25026,f25027,f25028,f25029,f25030,f25031,f25032,f25033,f25034,f25035,f25036,f25037,f25038,f25039,f25040,f25041,f25042,f25043,f25044,f25045,f25046,f25047,f25048,f25049,f25050,f25051,f25052,f25053,f25054,f25055,f25056,f25057,f25058,f25059,f25060,f25061,f25062,f25063,f25064,f25065,f25066,f25067,f25068,f25069,f25070,f25071,f25072,f25073,f25074,f25075,f25076,f25077,f25078,f25079,f25080,f25081,f25082,f25083,f25084,f25085,f25086,f25087,f25088,f25089,f25090,f25091,f25092,f25093,f25094,f25095,f25096,f25097,f25098,f25099,f25100,f25101,f25102,f25103,f25104,f25105,f25106,f25107,f25108,f25109,f25110,f25111,f25112,f25113,f25114,f25115,f25116,f25117,f25118,f25119,f25120,f25121,f25122,f25123,f25124,f25125,f25126,f25127,f25128,f25129,f25130,f25131,f25132,f25133,f25134,f25135,f25136,f25137,f25138,f25139,f25140,f25141,f25142,f25143,f25144,f25145,f25146,f25147,f25148,f25149,f25150,f25151,f25152,f25153,f25154,f25155,f25156,f25157,f25158,f25159,f25160,f25161,f25162,f25163,f25164,f25165,f25166,f25167,f25168,f25169,f25170,f25171,f25172,f25173,f25174,f25175,f25176,f25177,f25178,f25179,f25180,f25181,f25182,f25183,f25184,f25185,f25186,f25187,f25188,f25189,f25190,f25191,f25192,f25193,f25194,f25195,f25196,f25197,f25198,f25199,f25200,f25201,f25202,f25203,f25204,f25205,f25206,f25207,f25208,f25209,f25210,f25211,f25212,f25213,f25214,f25215,f25216,f25217,f25218,f25219,f25220,f25221,f25222,f25223,f25224,f25225,f25226,f25227,f25228,f25229,f25230,f25231,f25232,f25233,f25234,f25235,f25236,f25237,f25238,f25239,f25240,f25241,f25242,f25243,f25244,f25245,f25246,f25247,f25248,f25249,f25250,f25251,f25252,f25253,f25254,f25255,f25256,f25257,f25258,f25259,f25260,f25261,f25262,f25263,f25264,f25265,f25266,f25267,f25268,f25269,f25270,f25271,f25272,f25273,f25274,f25275,f25276,f25277,f25278,f25279,f25280,f25281,f25282,f25283,f25284,f25285,f25286,f25287,f25288,f25289,f25290,f25291,f25292,f25293,f25294,f25295,f25296,f25297,f25298,f25299,f25300,f25301,f25302,f25303,f25304,f25305,f25306,f25307,f25308,f25309,f25310,f25311,f25312,f25313,f25314,f25315,f25316,f25317,f25318,f25319,f25320,f25321,f25322,f25323,f25324,f25325,f25326,f25327,f25328,f25329,f25330,f25331,f25332,f25333,f25334,f25335,f25336,f25337,f25338,f25339,f25340,f25341,f25342,f25343,f25344,f25345,f25346,f25347,f25348,f25349,f25350,f25351,f25352,f25353,f25354,f25355,f25356,f25357,f25358,f25359,f25360,f25361,f25362,f25363,f25364,f25365,f25366,f25367,f25368,f25369,f25370,f25371,f25372,f25373,f25374,f25375,f25376,f25377,f25378,f25379,f25380,f25381,f25382,f25383,f25384,f25385,f25386,f25387,f25388,f25389,f25390,f25391,f25392,f25393,f25394,f25395,f25396,f25397,f25398,f25399,f25400,f25401,f25402,f25403,f25404,f25405,f25406,f25407,f25408,f25409,f25410,f25411,f25412,f25413,f25414,f25415,f25416,f25417,f25418,f25419,f25420,f25421,f25422,f25423,f25424,f25425,f25426,f25427,f25428,f25429,f25430,f25431,f25432,f25433,f25434,f25435,f25436,f25437,f25438,f25439,f25440,f25441,f25442,f25443,f25444,f25445,f25446,f25447,f25448,f25449,f25450,f25451,f25452,f25453,f25454,f25455,f25456,f25457,f25458,f25459,f25460,f25461,f25462,f25463,f25464,f25465,f25466,f25467,f25468,f25469,f25470,f25471,f25472,f25473,f25474,f25475,f25476,f25477,f25478,f25479,f25480,f25481,f25482,f25483,f25484,f25485,f25486,f25487,f25488,f25489,f25490,f25491,f25492,f25493,f25494,f25495,f25496,f25497,f25498,f25499,f25500,f25501,f25502,f25503,f25504,f25505,f25506,f25507,f25508,f25509,f25510,f25511,f25512,f25513,f25514,f25515,f25516,f25517,f25518,f25519,f25520,f25521,f25522,f25523,f25524,f25525,f25526,f25527,f25778,f25780,f27501,f27503,f29660,f29661,f29796,f29797,f29932,f29933,f30071,f30072,f35326,f35329,f35332,f35334,f35327,f38196,f38197,f38362,f38363,f38528,f38529,f38695,f38696,f38917,f38921,f38919,f39483,f40667,f40671,f40669,f41270,f41273,f67101,f67103,f67102,f67191,f70116,f70117,f17173,f71492])).
fof(f71492,plain,(
  $false | $spl41),
  inference(subsumption_resolution,[],[f71491,f1710])).
fof(f1710,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f1451,f1038])).
fof(f1038,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f750])).
fof(f750,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X21] : (s(bool,'T_0') = s(bool,X21) <=> p(s(bool,X21)))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f1451,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f941,f919])).
fof(f919,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f23])).
fof(f23,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.TRUTH')).
fof(f941,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f727])).
fof(f727,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f265])).
fof(f265,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f47])).
fof(f47,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.NOT_CLAUSES|split|2')).
fof(f71491,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl41),
  inference(backward_demodulation,[],[f71490,f6033])).
fof(f6033,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(unit_resulting_resolution,[],[f1451,f5976,f1242])).
fof(f1242,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f658])).
fof(f658,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f657])).
fof(f657,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f500])).
fof(f500,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X19,X20] : ((p(s(bool,X19)) => p(s(bool,X20))) => ((p(s(bool,X20)) => p(s(bool,X19))) => s(bool,X19) = s(bool,X20)))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f5976,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(unit_resulting_resolution,[],[f916,f1058])).
fof(f1058,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f760])).
fof(f760,plain,(
  ! [X0] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) & (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f352])).
fof(f352,plain,(
  ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  ! [X6] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.le_infty|split|2')).
fof(f916,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK13)),
  inference(cnf_transformation,[],[f726])).
fof(f726,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK13) & s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK14) & s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14])],[f589])).
fof(f589,plain,(
  ? [X0,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(flattening,[],[f588])).
fof(f588,plain,(
  ? [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1)) & s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(ennf_transformation,[],[f264])).
fof(f264,plain,(
  ~! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1)) => s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1)) => ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X7)) => ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X7)) => ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.add_not_infty|split|0')).
fof(f71490,plain,(
  ( ! [X23] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',X23)))) ) | $spl41),
  inference(forward_demodulation,[],[f71489,f957])).
fof(f957,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f136])).
fof(f136,axiom,(
  ! [X12] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X12)))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_le_def|split|1')).
fof(f71489,plain,(
  ( ! [X23] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X23),s('type.extreal.extreal',sK14)))))) ) | $spl41),
  inference(subsumption_resolution,[],[f71488,f917])).
fof(f917,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',sK14)),
  inference(cnf_transformation,[],[f726])).
fof(f71488,plain,(
  ( ! [X23] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X23),s('type.extreal.extreal',sK14))))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK14)) ) | $spl41),
  inference(subsumption_resolution,[],[f71288,f38918])).
fof(f38918,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK14) | $spl41),
  inference(cnf_transformation,[],[f38918_D])).
fof(f38918_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',sK14) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f71288,plain,(
  ( ! [X23] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',X23))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X23),s('type.extreal.extreal',sK14))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK14) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK14)) )),
  inference(superposition,[],[f1319,f918])).
fof(f918,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))),
  inference(cnf_transformation,[],[f726])).
fof(f1319,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f682])).
fof(f682,plain,(
  ! [X0,X1,X2] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))))),
  inference(flattening,[],[f681])).
fof(f681,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))))),
  inference(ennf_transformation,[],[f553])).
fof(f553,plain,(
  ! [X0,X1,X2] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))))),
  inference(flattening,[],[f552])).
fof(f552,plain,(
  ! [X0,X1,X2] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X2))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X0))))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X6,X7,X38] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X38))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X38),s('type.extreal.extreal',X6))))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.le_radd')).
fof(f17173,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl10),
  inference(cnf_transformation,[],[f17173_D])).
fof(f17173_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f70117,plain,(
  ( ! [X8] : (p(s(bool,sK60(X8,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) ) | $spl52),
  inference(cnf_transformation,[],[f70117_D])).
fof(f70117_D,plain,(
  ( ! [X8] : (p(s(bool,sK60(X8,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f70116,plain,(
  ( ! [X8,X7] : (p(s(bool,sK60(X8,'const.combin.I_0',bool))) | ~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) )),
  inference(forward_demodulation,[],[f70074,f1462])).
fof(f1462,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f1128,f1127])).
fof(f1127,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f399])).
fof(f399,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f398])).
fof(f398,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8] : ! [X6] : s(X8,X6) = s(X8,'const.combin.I_1'(s(X8,X6)))),
  file('thm.extreal.add_not_infty|split|0','thm.combin.I_THM')).
fof(f1128,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f401])).
fof(f401,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f400])).
fof(f400,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : ! [X17] : s(X8,'const.combin.I_1'(s(X8,X17))) = s(X8,happ(s(fun(X8,X8),'const.combin.I_0'),s(X8,X17)))),
  file('thm.extreal.add_not_infty|split|0','ARITY_0_1_const.combin.I')).
fof(f70074,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK60(X8,'const.combin.I_0',bool))))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) )),
  inference(superposition,[],[f1145,f1462])).
fof(f1145,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK60(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f808])).
fof(f808,plain,(
  ! [X0,X1,X2] : (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK60(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK60(X2,X1,X0)))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK60])],[f628])).
fof(f628,plain,(
  ! [X0,X1,X2] : (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f627])).
fof(f627,plain,(
  ! [X0,X1,X2] : ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f423])).
fof(f423,plain,(
  ! [X0,X1,X2] : ((? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f422])).
fof(f422,plain,(
  ! [X0] : ! [X1,X2] : ((? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) => p(s(bool,happ(s(fun(X0,bool),X2),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f206])).
fof(f206,axiom,(
  ! [X8] : ! [X16,X62] : ((? [X6] : p(s(bool,happ(s(fun(X8,bool),X16),s(X8,X6)))) & ! [X6] : (p(s(bool,happ(s(fun(X8,bool),X16),s(X8,X6)))) => p(s(bool,happ(s(fun(X8,bool),X62),s(X8,X6)))))) => p(s(bool,happ(s(fun(X8,bool),X62),s(X8,'const.min.@_1'(s(fun(X8,bool),X16)))))))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.SELECT_ELIM_THM')).
fof(f67191,plain,(
  $false | ($spl41 | $spl48)),
  inference(subsumption_resolution,[],[f67187,f67102])).
fof(f67187,plain,(
  ~p(s(bool,sK78(sK14))) | $spl41),
  inference(unit_resulting_resolution,[],[f1710,f39721,f1351])).
fof(f1351,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f890])).
fof(f890,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP6(X2,X1,X0)) & (~sP6(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f716])).
fof(f716,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP6(X2,X1,X0))),
  inference(definition_folding,[],[f561,f715,f714,f713,f712])).
fof(f712,plain,(
  ! [X0,X1,X2] : (sP3(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f713,plain,(
  ! [X1,X0,X2] : (sP4(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f714,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f715,plain,(
  ! [X2,X1,X0] : (sP6(X2,X1,X0) <=> (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f561,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f560])).
fof(f560,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f81])).
fof(f81,axiom,(
  ! [X33,X43,X44] : ((p(s(bool,X44)) <=> s(bool,X33) = s(bool,X43)) <=> ((p(s(bool,X44)) | p(s(bool,X43)) | p(s(bool,X33))) & (p(s(bool,X44)) | ~p(s(bool,X33)) | ~p(s(bool,X43))) & (p(s(bool,X43)) | ~p(s(bool,X33)) | ~p(s(bool,X44))) & (p(s(bool,X33)) | ~p(s(bool,X43)) | ~p(s(bool,X44)))))),
  file('thm.extreal.add_not_infty|split|0','thm.sat.sat_2|49|_prove_ltb')).
fof(f39721,plain,(
  sP6(sK78(sK14),'F_0','T_0') | $spl41),
  inference(unit_resulting_resolution,[],[f1710,f39495,f1349])).
fof(f1349,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f890])).
fof(f39495,plain,(
  ~p(s(bool,sK78(sK14))) | $spl41),
  inference(unit_resulting_resolution,[],[f917,f38918,f1204])).
fof(f1204,plain,(
  ( ! [X0] : (~p(s(bool,sK78(X0))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f846])).
fof(f846,plain,(
  ! [X0] : ((~p(s(bool,sK78(X0))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | p(s(bool,sK78(X0)))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK78(X0)),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK78])],[f845])).
fof(f845,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | p(s(bool,X1))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(flattening,[],[f844])).
fof(f844,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))) & ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | p(s(bool,X1)))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(nnf_transformation,[],[f461])).
fof(f461,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X6] : ? [X45] : ((p(s(bool,X45)) <=> (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X6))) = s('type.realax.real','const.bool.COND_3'(s(bool,X45),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X6))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.real_def')).
fof(f67102,plain,(
  ( ! [X7] : (p(s(bool,X7))) ) | $spl48),
  inference(cnf_transformation,[],[f67102_D])).
fof(f67102_D,plain,(
  ( ! [X7] : (p(s(bool,X7))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f67103,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(X8,X10) = s(X8,'const.bool.COND_3'(s(bool,X11),s(X8,X12),s(X8,X9))) | s(bool,'F_0') != s(bool,X11) | s(X8,X9) != s(X8,X10)) ) | $spl50),
  inference(cnf_transformation,[],[f67103_D])).
fof(f67103_D,plain,(
  ( ! [X12,X10,X8,X11,X9] : (s(X8,X10) = s(X8,'const.bool.COND_3'(s(bool,X11),s(X8,X12),s(X8,X9))) | s(bool,'F_0') != s(bool,X11) | s(X8,X9) != s(X8,X10)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f67101,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(X8,X10) = s(X8,'const.bool.COND_3'(s(bool,X11),s(X8,X12),s(X8,X9))) | p(s(bool,X7)) | s(X8,X9) != s(X8,X10) | s(bool,'F_0') != s(bool,X11)) )),
  inference(forward_demodulation,[],[f67100,f1141])).
fof(f1141,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f415])).
fof(f415,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f414])).
fof(f414,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f193])).
fof(f193,axiom,(
  ! [X8] : ! [X19,X20] : s(X8,X20) = s(X8,'const.bool.COND_3'(s(bool,'F_0'),s(X8,X19),s(X8,X20)))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.COND_CLAUSES|split|1')).
fof(f67100,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,X7)) | s(X8,X9) != s(X8,X10) | s(X8,'const.bool.COND_3'(s(bool,'F_0'),s(X8,X13),s(X8,X10))) = s(X8,'const.bool.COND_3'(s(bool,X11),s(X8,X12),s(X8,X9))) | s(bool,'F_0') != s(bool,X11)) )),
  inference(subsumption_resolution,[],[f57579,f1038])).
fof(f57579,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (p(s(bool,X7)) | s(X8,X9) != s(X8,X10) | s(X8,'const.bool.COND_3'(s(bool,'F_0'),s(X8,X13),s(X8,X10))) = s(X8,'const.bool.COND_3'(s(bool,X11),s(X8,X12),s(X8,X9))) | s(bool,'F_0') != s(bool,X11) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1179,f992])).
fof(f992,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.extreal.add_not_infty|split|0',aHL_BOOL_CASES)).
fof(f1179,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,X2)) | s(X0,X5) != s(X0,X6) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f631])).
fof(f631,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f630])).
fof(f630,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6))) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(ennf_transformation,[],[f453])).
fof(f453,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f452])).
fof(f452,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(rectify,[],[f203])).
fof(f203,axiom,(
  ! [X8] : ! [X16,X62,X6,X41,X7,X42] : ((s(bool,X16) = s(bool,X62) & (p(s(bool,X62)) => s(X8,X6) = s(X8,X41)) & (~p(s(bool,X62)) => s(X8,X7) = s(X8,X42))) => s(X8,'const.bool.COND_3'(s(bool,X16),s(X8,X6),s(X8,X7))) = s(X8,'const.bool.COND_3'(s(bool,X62),s(X8,X41),s(X8,X42))))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.COND_CONG')).
fof(f41273,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f41272,f1710])).
fof(f41272,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl44),
  inference(forward_demodulation,[],[f41271,f957])).
fof(f41271,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14))) | $spl44),
  inference(forward_demodulation,[],[f41133,f924])).
fof(f924,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(cnf_transformation,[],[f88])).
fof(f88,axiom,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_ainv_def|split|1')).
fof(f41133,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal',sK14))) | $spl44),
  inference(backward_demodulation,[],[f40669,f30858])).
fof(f30858,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK13))),s('type.extreal.extreal',sK14)))),
  inference(forward_demodulation,[],[f30787,f3661])).
fof(f3661,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(superposition,[],[f967,f966])).
fof(f966,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X36] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X36))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X36)))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_of_num_def')).
fof(f967,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(rectify,[],[f137])).
fof(f137,axiom,(
  ! [X22] : s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X22))),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_le_def|split|5')).
fof(f30787,plain,(
  s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK13))),s('type.extreal.extreal',sK14)))),
  inference(superposition,[],[f1228,f918])).
fof(f1228,plain,(
  ( ! [X0,X1] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))) )),
  inference(cnf_transformation,[],[f482])).
fof(f482,plain,(
  ! [X0,X1] : s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X0))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X6,X7] : s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',X6))),s('type.extreal.extreal',X7))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.le_lneg')).
fof(f41270,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f40683,f6615])).
fof(f6615,plain,(
  ( ! [X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X1)))) )),
  inference(subsumption_resolution,[],[f6614,f1451])).
fof(f6614,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X1)))) )),
  inference(forward_demodulation,[],[f6547,f927])).
fof(f927,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f139])).
fof(f139,axiom,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_le_def|split|4')).
fof(f6547,plain,(
  ( ! [X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X1)))) )),
  inference(superposition,[],[f1059,f960])).
fof(f960,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X12] : s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X12)))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_add_def|split|1')).
fof(f1059,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f760])).
fof(f40683,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK14))) | $spl44),
  inference(backward_demodulation,[],[f40669,f918])).
fof(f40669,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK13) | $spl44),
  inference(cnf_transformation,[],[f40669_D])).
fof(f40669_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK13) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f40671,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK14))) | $spl46),
  inference(cnf_transformation,[],[f40671_D])).
fof(f40671_D,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK14))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f40667,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK14))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK13)),
  inference(forward_demodulation,[],[f40666,f6033])).
fof(f40666,plain,(
  s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK14))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK13)),
  inference(subsumption_resolution,[],[f40567,f916])).
fof(f40567,plain,(
  s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK14))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK13) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK13)),
  inference(superposition,[],[f1241,f918])).
fof(f1241,plain,(
  ( ! [X0,X1] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f656])).
fof(f656,plain,(
  ! [X0,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f655])).
fof(f655,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(ennf_transformation,[],[f499])).
fof(f499,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f498])).
fof(f498,plain,(
  ! [X0,X1] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X6,X7] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X7))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.le_addr')).
fof(f39483,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f39482,f1710])).
fof(f39482,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl40),
  inference(forward_demodulation,[],[f39372,f1732])).
fof(f1732,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f949,f1039])).
fof(f1039,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f750])).
fof(f949,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(rectify,[],[f132])).
fof(f132,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.le_infty|split|1')).
fof(f39372,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal',sK13))),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl40),
  inference(backward_demodulation,[],[f38919,f30858])).
fof(f38919,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK14) | $spl40),
  inference(cnf_transformation,[],[f38919_D])).
fof(f38919_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK14) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f38921,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | $spl42),
  inference(cnf_transformation,[],[f38921_D])).
fof(f38921_D,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f38917,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK14)),
  inference(forward_demodulation,[],[f38916,f1762])).
fof(f1762,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(unit_resulting_resolution,[],[f1124,f1043])).
fof(f1043,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f752])).
fof(f752,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f342])).
fof(f342,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X21] : (s(bool,'F_0') = s(bool,X21) <=> ~p(s(bool,X21)))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f1124,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(cnf_transformation,[],[f393])).
fof(f393,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(flattening,[],[f392])).
fof(f392,plain,(
  ! [X0] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(rectify,[],[f110])).
fof(f110,axiom,(
  ! [X6] : ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.lt_infty|split|3')).
fof(f38916,plain,(
  s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK14)),
  inference(subsumption_resolution,[],[f38827,f917])).
fof(f38827,plain,(
  s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK14) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK14)),
  inference(superposition,[],[f1240,f918])).
fof(f1240,plain,(
  ( ! [X0,X1] : (s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1)) )),
  inference(cnf_transformation,[],[f654])).
fof(f654,plain,(
  ! [X0,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f653])).
fof(f653,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1)) | s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(ennf_transformation,[],[f497])).
fof(f497,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X1) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X1)) => s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f496])).
fof(f496,plain,(
  ! [X0,X1] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X1) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X1)) => s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X0))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X6,X7] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X7) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X7)) => s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X6))) = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.lt_addl')).
fof(f38696,plain,(
  ( ! [X8] : (p(s(bool,sK70('const.combin.I_0',X8,bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK70('const.combin.I_0',X8,bool)))))) ) | $spl38),
  inference(cnf_transformation,[],[f38696_D])).
fof(f38696_D,plain,(
  ( ! [X8] : (p(s(bool,sK70('const.combin.I_0',X8,bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK70('const.combin.I_0',X8,bool)))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f38695,plain,(
  ( ! [X8,X7] : (p(s(bool,sK70('const.combin.I_0',X8,bool))) | ~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK70('const.combin.I_0',X8,bool)))))) )),
  inference(forward_demodulation,[],[f38671,f1462])).
fof(f38671,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK70('const.combin.I_0',X8,bool))))) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK70('const.combin.I_0',X8,bool)))))) )),
  inference(superposition,[],[f1166,f1462])).
fof(f1166,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK70(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK70(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f826])).
fof(f826,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK68(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK69(X2,X1,X0)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK70(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK70(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK68,sK69,sK70])],[f825])).
fof(f825,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ? [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f824])).
fof(f824,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f823])).
fof(f823,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f435])).
fof(f435,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f434])).
fof(f434,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f191])).
fof(f191,axiom,(
  ! [X8] : ! [X16,X62] : (? [X6] : (p(s(bool,happ(s(fun(X8,bool),X16),s(X8,X6)))) | p(s(bool,happ(s(fun(X8,bool),X62),s(X8,X6))))) <=> (? [X6] : p(s(bool,happ(s(fun(X8,bool),X16),s(X8,X6)))) | ? [X6] : p(s(bool,happ(s(fun(X8,bool),X62),s(X8,X6))))))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.EXISTS_OR_THM')).
fof(f38529,plain,(
  ( ! [X8] : (p(s(bool,sK70(X8,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK70(X8,'const.combin.I_0',bool)))))) ) | $spl36),
  inference(cnf_transformation,[],[f38529_D])).
fof(f38529_D,plain,(
  ( ! [X8] : (p(s(bool,sK70(X8,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK70(X8,'const.combin.I_0',bool)))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f38528,plain,(
  ( ! [X8,X7] : (p(s(bool,sK70(X8,'const.combin.I_0',bool))) | ~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK70(X8,'const.combin.I_0',bool)))))) )),
  inference(forward_demodulation,[],[f38504,f1462])).
fof(f38504,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK70(X8,'const.combin.I_0',bool))))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK70(X8,'const.combin.I_0',bool)))))) )),
  inference(superposition,[],[f1165,f1462])).
fof(f1165,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK70(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK70(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f826])).
fof(f38363,plain,(
  ( ! [X8] : (p(s(bool,sK69('const.combin.I_0',X8,bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK68('const.combin.I_0',X8,bool)))))) ) | $spl34),
  inference(cnf_transformation,[],[f38363_D])).
fof(f38363_D,plain,(
  ( ! [X8] : (p(s(bool,sK69('const.combin.I_0',X8,bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK68('const.combin.I_0',X8,bool)))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f38362,plain,(
  ( ! [X8,X7] : (p(s(bool,sK69('const.combin.I_0',X8,bool))) | ~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK68('const.combin.I_0',X8,bool)))))) )),
  inference(forward_demodulation,[],[f38338,f1462])).
fof(f38338,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK68('const.combin.I_0',X8,bool))))) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK69('const.combin.I_0',X8,bool)))))) )),
  inference(superposition,[],[f1164,f1462])).
fof(f1164,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK68(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK69(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f826])).
fof(f38197,plain,(
  ( ! [X8] : (p(s(bool,sK68(X8,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK69(X8,'const.combin.I_0',bool)))))) ) | $spl32),
  inference(cnf_transformation,[],[f38197_D])).
fof(f38197_D,plain,(
  ( ! [X8] : (p(s(bool,sK68(X8,'const.combin.I_0',bool))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK69(X8,'const.combin.I_0',bool)))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f38196,plain,(
  ( ! [X8,X7] : (p(s(bool,sK68(X8,'const.combin.I_0',bool))) | ~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK69(X8,'const.combin.I_0',bool)))))) )),
  inference(forward_demodulation,[],[f38172,f1462])).
fof(f38172,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK68(X8,'const.combin.I_0',bool))))) | p(s(bool,happ(s(fun(bool,bool),X8),s(bool,sK69(X8,'const.combin.I_0',bool)))))) )),
  inference(superposition,[],[f1163,f1462])).
fof(f1163,plain,(
  ( ! [X2,X0,X8,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK68(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK69(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f826])).
fof(f35327,plain,(
  ( ! [X59] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X59)))) ) | $spl26),
  inference(cnf_transformation,[],[f35327_D])).
fof(f35327_D,plain,(
  ( ! [X59] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X59)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f35334,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK37('const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))) | $spl30),
  inference(cnf_transformation,[],[f35334_D])).
fof(f35334_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK37('const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f35332,plain,(
  ( ! [X60] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK37('const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X60)))) )),
  inference(forward_demodulation,[],[f35331,f1454])).
fof(f1454,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(backward_demodulation,[],[f1453,f1449])).
fof(f1449,plain,(
  s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(backward_demodulation,[],[f937,f936])).
fof(f936,plain,(
  s('type.extreal.extreal','const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f128])).
fof(f128,axiom,(
  s('type.extreal.extreal','const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_exp_def|split|2')).
fof(f937,plain,(
  s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f126])).
fof(f126,axiom,(
  s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_inv_def|split|0')).
fof(f1453,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(backward_demodulation,[],[f966,f937])).
fof(f35331,plain,(
  ( ! [X60] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X60))) | s('type.extreal.extreal','const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK37('const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))))) )),
  inference(subsumption_resolution,[],[f35330,f6832])).
fof(f6832,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(subsumption_resolution,[],[f6747,f6660])).
fof(f6660,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))))) )),
  inference(unit_resulting_resolution,[],[f1119,f1060])).
fof(f1060,plain,(
  ( ! [X0] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f761])).
fof(f761,plain,(
  ! [X0] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0))))))),
  inference(nnf_transformation,[],[f353])).
fof(f353,plain,(
  ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) <=> s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f133])).
fof(f133,axiom,(
  ! [X6] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X6)))) <=> s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.le_infty|split|3')).
fof(f1119,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f383])).
fof(f383,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f382])).
fof(f382,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f167])).
fof(f167,axiom,(
  ! [X36] : ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X36)))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.num_not_infty|split|1')).
fof(f6747,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(superposition,[],[f1061,f1454])).
fof(f1061,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f761])).
fof(f35330,plain,(
  ( ! [X60] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X60))) | s('type.extreal.extreal','const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK37('const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(subsumption_resolution,[],[f35253,f6649])).
fof(f6649,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(subsumption_resolution,[],[f6648,f1451])).
fof(f6648,plain,(
  p(s(bool,'F_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(forward_demodulation,[],[f6565,f3661])).
fof(f6565,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(superposition,[],[f1059,f1454])).
fof(f35253,plain,(
  ( ! [X60] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X60))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK37('const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_exp_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(superposition,[],[f1005,f1454])).
fof(f1005,plain,(
  ( ! [X2,X0] : (s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK37(X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f742])).
fof(f742,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X2] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK37(X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK37])],[f741])).
fof(f741,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X2] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2)))) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f606])).
fof(f606,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X1] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) | ? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(flattening,[],[f605])).
fof(f605,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) & ! [X1] : s('type.extreal.extreal',X0) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) | (? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)))),
  inference(ennf_transformation,[],[f328])).
fof(f328,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | ? [X1] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1)))) => (? [X2] : s('type.extreal.extreal',X0) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X2))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X12] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X12) | ? [X6] : s('type.extreal.extreal',X12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))) => (? [X6] : s('type.extreal.extreal',X12) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X12) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X12)))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_46_prove_ltb')).
fof(f35329,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK37('const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))) | $spl28),
  inference(cnf_transformation,[],[f35329_D])).
fof(f35329_D,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK37('const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f35326,plain,(
  ( ! [X59] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK37('const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X59)))) )),
  inference(forward_demodulation,[],[f35325,f1453])).
fof(f35325,plain,(
  ( ! [X59] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X59))) | s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK37('const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))))) )),
  inference(subsumption_resolution,[],[f35324,f6831])).
fof(f6831,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(subsumption_resolution,[],[f6746,f6660])).
fof(f6746,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(superposition,[],[f1061,f1453])).
fof(f35324,plain,(
  ( ! [X59] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X59))) | s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK37('const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(subsumption_resolution,[],[f35252,f6647])).
fof(f6647,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(subsumption_resolution,[],[f6646,f1451])).
fof(f6646,plain,(
  p(s(bool,'F_0')) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(forward_demodulation,[],[f6564,f3661])).
fof(f6564,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(superposition,[],[f1059,f1453])).
fof(f35252,plain,(
  ( ! [X59] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X59))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) | s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK37('const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_inv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(superposition,[],[f1005,f1453])).
fof(f30072,plain,(
  ( ! [X8] : (p(s(bool,sK65(X8,'const.combin.I_0',bool))) | p(s(bool,X8))) ) | $spl24),
  inference(cnf_transformation,[],[f30072_D])).
fof(f30072_D,plain,(
  ( ! [X8] : (p(s(bool,sK65(X8,'const.combin.I_0',bool))) | p(s(bool,X8))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f30071,plain,(
  ( ! [X8,X7] : (p(s(bool,sK65(X8,'const.combin.I_0',bool))) | ~p(s(bool,X7)) | p(s(bool,X8))) )),
  inference(forward_demodulation,[],[f30052,f1462])).
fof(f30052,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK65(X8,'const.combin.I_0',bool))))) | p(s(bool,X8))) )),
  inference(superposition,[],[f1157,f1462])).
fof(f1157,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK65(X2,X1,X0))))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f818])).
fof(f818,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK64(X2,X1,X0))))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK65(X2,X1,X0))))) | p(s(bool,X2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK64,sK65])],[f817])).
fof(f817,plain,(
  ! [X0,X1,X2] : (((! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) & ~p(s(bool,X2))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(rectify,[],[f816])).
fof(f816,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))))),
  inference(flattening,[],[f815])).
fof(f815,plain,(
  ! [X0,X1,X2] : (((! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,X2))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2)))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ~p(s(bool,X2))) | (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f431])).
fof(f431,plain,(
  ! [X0,X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(flattening,[],[f430])).
fof(f430,plain,(
  ! [X0] : ! [X1,X2] : ((? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,X2))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,X2))))),
  inference(rectify,[],[f205])).
fof(f205,axiom,(
  ! [X8] : ! [X16,X62] : ((? [X6] : p(s(bool,happ(s(fun(X8,bool),X16),s(X8,X6)))) | p(s(bool,X62))) <=> ? [X6] : (p(s(bool,happ(s(fun(X8,bool),X16),s(X8,X6)))) | p(s(bool,X62))))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.LEFT_OR_EXISTS_THM')).
fof(f29933,plain,(
  ( ! [X8] : (p(s(bool,sK64(X8,'const.combin.I_0',bool))) | p(s(bool,X8))) ) | $spl22),
  inference(cnf_transformation,[],[f29933_D])).
fof(f29933_D,plain,(
  ( ! [X8] : (p(s(bool,sK64(X8,'const.combin.I_0',bool))) | p(s(bool,X8))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f29932,plain,(
  ( ! [X8,X7] : (p(s(bool,sK64(X8,'const.combin.I_0',bool))) | ~p(s(bool,X7)) | p(s(bool,X8))) )),
  inference(forward_demodulation,[],[f29913,f1462])).
fof(f29913,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK64(X8,'const.combin.I_0',bool))))) | p(s(bool,X8))) )),
  inference(superposition,[],[f1155,f1462])).
fof(f1155,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK64(X2,X1,X0))))) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f818])).
fof(f29797,plain,(
  ( ! [X8] : (p(s(bool,sK63('const.combin.I_0',X8,bool))) | p(s(bool,X8))) ) | $spl20),
  inference(cnf_transformation,[],[f29797_D])).
fof(f29797_D,plain,(
  ( ! [X8] : (p(s(bool,sK63('const.combin.I_0',X8,bool))) | p(s(bool,X8))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f29796,plain,(
  ( ! [X8,X7] : (p(s(bool,sK63('const.combin.I_0',X8,bool))) | ~p(s(bool,X7)) | p(s(bool,X8))) )),
  inference(forward_demodulation,[],[f29777,f1462])).
fof(f29777,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK63('const.combin.I_0',X8,bool)))))) )),
  inference(superposition,[],[f1154,f1462])).
fof(f1154,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK63(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f814])).
fof(f814,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK62(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK63(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK62,sK63])],[f813])).
fof(f813,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f812])).
fof(f812,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f811])).
fof(f811,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f429])).
fof(f429,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f428])).
fof(f428,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f156])).
fof(f156,axiom,(
  ! [X8] : ! [X16,X62] : ((p(s(bool,X16)) | ? [X6] : p(s(bool,happ(s(fun(X8,bool),X62),s(X8,X6))))) <=> ? [X6] : (p(s(bool,X16)) | p(s(bool,happ(s(fun(X8,bool),X62),s(X8,X6))))))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f29661,plain,(
  ( ! [X8] : (p(s(bool,sK62('const.combin.I_0',X8,bool))) | p(s(bool,X8))) ) | $spl18),
  inference(cnf_transformation,[],[f29661_D])).
fof(f29661_D,plain,(
  ( ! [X8] : (p(s(bool,sK62('const.combin.I_0',X8,bool))) | p(s(bool,X8))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f29660,plain,(
  ( ! [X8,X7] : (p(s(bool,sK62('const.combin.I_0',X8,bool))) | ~p(s(bool,X7)) | p(s(bool,X8))) )),
  inference(forward_demodulation,[],[f29641,f1462])).
fof(f29641,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK62('const.combin.I_0',X8,bool)))))) )),
  inference(superposition,[],[f1152,f1462])).
fof(f1152,plain,(
  ( ! [X6,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK62(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f814])).
fof(f27503,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl16),
  inference(cnf_transformation,[],[f27503_D])).
fof(f27503_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f27501,plain,(
  ( ! [X3] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f27486,f1462])).
fof(f27486,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) )),
  inference(superposition,[],[f1144,f1462])).
fof(f1144,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f626])).
fof(f626,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f421])).
fof(f421,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f420])).
fof(f420,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8] : ! [X16,X6] : (p(s(bool,happ(s(fun(X8,bool),X16),s(X8,X6)))) => p(s(bool,happ(s(fun(X8,bool),X16),s(X8,'const.min.@_1'(s(fun(X8,bool),X16)))))))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.SELECT_AX')).
fof(f25780,plain,(
  p(s(bool,sK58('const.combin.I_0',bool))) | $spl14),
  inference(cnf_transformation,[],[f25780_D])).
fof(f25780_D,plain,(
  p(s(bool,sK58('const.combin.I_0',bool))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f25778,plain,(
  ( ! [X3] : (p(s(bool,sK58('const.combin.I_0',bool))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f25763,f1462])).
fof(f25763,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK58('const.combin.I_0',bool)))))) )),
  inference(superposition,[],[f1135,f1462])).
fof(f1135,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK58(X1,X0)))))) )),
  inference(cnf_transformation,[],[f802])).
fof(f802,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK57(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK58(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK57,sK58])],[f801])).
fof(f801,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f800])).
fof(f800,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f625])).
fof(f625,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f409])).
fof(f409,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f408])).
fof(f408,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f155])).
fof(f155,axiom,(
  ! [X8] : ! [X16] : (~? [X6] : p(s(bool,happ(s(fun(X8,bool),X16),s(X8,X6)))) <=> ! [X6] : ~p(s(bool,happ(s(fun(X8,bool),X16),s(X8,X6)))))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.NOT_EXISTS_THM')).
fof(f25527,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17180,f24489])).
fof(f24489,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1323])).
fof(f1323,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f876])).
fof(f876,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f558])).
fof(f558,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f106])).
fof(f106,axiom,(
  ! [X52,X53,X54] : (p(s(bool,happ(s(fun(X52,bool),'=_1'(s(X52,X53))),s(X52,X54)))) <=> s(X52,X53) = s(X52,X54))),
  file('thm.extreal.add_not_infty|split|0',arity_1_HL_EQ)).
fof(f17180,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1043])).
fof(f25526,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17182,f24489])).
fof(f17182,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1451,f17173,f1242])).
fof(f25525,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17183,f24489])).
fof(f17183,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f17173,f1242])).
fof(f25524,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17186,f1451])).
fof(f17186,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1710,f17173,f1242])).
fof(f25523,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17188,f24489])).
fof(f17188,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1451,f17173,f1242])).
fof(f25522,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17189,f24489])).
fof(f17189,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f17173,f1242])).
fof(f25521,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17192,f1451])).
fof(f17192,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1710,f17173,f1242])).
fof(f25520,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17194,f1451])).
fof(f17194,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1789,f17173,f1294])).
fof(f1294,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f871])).
fof(f871,plain,(
  ! [X0,X1] : ((~sP2(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1)))),
  inference(rectify,[],[f870])).
fof(f870,plain,(
  ! [X1,X0] : ((~sP2(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP2(X1,X0)))),
  inference(flattening,[],[f869])).
fof(f869,plain,(
  ! [X1,X0] : ((~sP2(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP2(X1,X0)))),
  inference(nnf_transformation,[],[f710])).
fof(f710,plain,(
  ! [X1,X0] : (sP2(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f1789,plain,(
  sP2('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1451,f919,f1296])).
fof(f1296,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f871])).
fof(f25519,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17195,f17173])).
fof(f17195,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6039,f17173,f1294])).
fof(f6039,plain,(
  sP2('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(unit_resulting_resolution,[],[f919,f5976,f1296])).
fof(f25518,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17196,f17173])).
fof(f17196,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6853,f17173,f1294])).
fof(f6853,plain,(
  sP2('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(unit_resulting_resolution,[],[f919,f5977,f1296])).
fof(f5977,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(unit_resulting_resolution,[],[f917,f1058])).
fof(f25517,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17198,f17173])).
fof(f17198,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7681,f17173,f1294])).
fof(f7681,plain,(
  sP2('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),
  inference(unit_resulting_resolution,[],[f1451,f7614,f1299])).
fof(f1299,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f871])).
fof(f7614,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))),
  inference(unit_resulting_resolution,[],[f916,f1113])).
fof(f1113,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f795])).
fof(f795,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))) & (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0)))),
  inference(nnf_transformation,[],[f375])).
fof(f375,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f374])).
fof(f374,plain,(
  ! [X0] : (~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X6] : (~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.lt_infty|split|6')).
fof(f25516,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17199,f17173])).
fof(f17199,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14)))) | $spl10),
  inference(unit_resulting_resolution,[],[f8606,f17173,f1294])).
fof(f8606,plain,(
  sP2('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14)))),
  inference(unit_resulting_resolution,[],[f1451,f7615,f1299])).
fof(f7615,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14))))),
  inference(unit_resulting_resolution,[],[f917,f1113])).
fof(f25515,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17200,f17173])).
fof(f17200,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1883,f17173,f1294])).
fof(f1883,plain,(
  sP2('F_0',sK73(sK59('const.num.|48|_0','type.num.num')))),
  inference(unit_resulting_resolution,[],[f1451,f1871,f1299])).
fof(f1871,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num'))))),
  inference(unit_resulting_resolution,[],[f1140,f1190])).
fof(f1190,plain,(
  ( ! [X0] : (p(s(bool,sK73(X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f831])).
fof(f831,plain,(
  ! [X0] : ((~p(s(bool,sK73(X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,sK73(X0)))) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK73(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK73])],[f830])).
fof(f830,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(flattening,[],[f829])).
fof(f829,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f456])).
fof(f456,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(rectify,[],[f119])).
fof(f119,axiom,(
  ! [X36] : ? [X45] : ((p(s(bool,X45)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X36)) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.num.num',X36))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X45),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X36))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_pow_def|split|2')).
fof(f1140,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,sK59(X1,X0))) )),
  inference(cnf_transformation,[],[f807])).
fof(f807,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,sK59(X1,X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK59])],[f413])).
fof(f413,plain,(
  ! [X0,X1] : ? [X2] : s(X0,X1) = s(X0,X2)),
  inference(flattening,[],[f412])).
fof(f412,plain,(
  ! [X0] : ! [X1] : ? [X2] : s(X0,X1) = s(X0,X2)),
  inference(rectify,[],[f194])).
fof(f194,axiom,(
  ! [X8] : ! [X12] : ? [X6] : s(X8,X6) = s(X8,X12)),
  file('thm.extreal.add_not_infty|split|0','thm.bool.EXISTS_REFL')).
fof(f25514,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17201,f17173])).
fof(f17201,plain,(
  p(s(bool,sK73('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl10),
  inference(unit_resulting_resolution,[],[f4784,f17173,f1294])).
fof(f4784,plain,(
  sP2('F_0',sK73('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(unit_resulting_resolution,[],[f1451,f1870,f1299])).
fof(f1870,plain,(
  p(s(bool,sK73('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))),
  inference(unit_resulting_resolution,[],[f925,f1190])).
fof(f925,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f138])).
fof(f138,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_size_def|split|0')).
fof(f25513,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17202,f17173])).
fof(f17202,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1912,f17173,f1294])).
fof(f1912,plain,(
  sP2('F_0',sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),
  inference(unit_resulting_resolution,[],[f1451,f1899,f1299])).
fof(f1899,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))))),
  inference(unit_resulting_resolution,[],[f918,f1205])).
fof(f1205,plain,(
  ( ! [X0] : (p(s(bool,sK78(X0))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f846])).
fof(f25512,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17203,f17173])).
fof(f17203,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2103,f17173,f1294])).
fof(f2103,plain,(
  sP2('F_0',sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))),
  inference(unit_resulting_resolution,[],[f1451,f1901,f1299])).
fof(f1901,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal'))))),
  inference(unit_resulting_resolution,[],[f1140,f1205])).
fof(f25511,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17204,f17173])).
fof(f17204,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2227,f17173,f1294])).
fof(f2227,plain,(
  sP2('F_0',sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))),
  inference(unit_resulting_resolution,[],[f1451,f1930,f1299])).
fof(f1930,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal'))))),
  inference(unit_resulting_resolution,[],[f1140,f1206])).
fof(f1206,plain,(
  ( ! [X0] : (p(s(bool,sK78(X0))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f846])).
fof(f25510,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17207,f1451])).
fof(f17207,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7680,f17173,f1294])).
fof(f7680,plain,(
  sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0')),
  inference(unit_resulting_resolution,[],[f1451,f7614,f1296])).
fof(f25509,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17208,f1451])).
fof(f17208,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f8605,f17173,f1294])).
fof(f8605,plain,(
  sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14)),'F_0')),
  inference(unit_resulting_resolution,[],[f1451,f7615,f1296])).
fof(f25508,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17209,f1451])).
fof(f17209,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1882,f17173,f1294])).
fof(f1882,plain,(
  sP2(sK73(sK59('const.num.|48|_0','type.num.num')),'F_0')),
  inference(unit_resulting_resolution,[],[f1451,f1871,f1296])).
fof(f25507,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17210,f1451])).
fof(f17210,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f4783,f17173,f1294])).
fof(f4783,plain,(
  sP2(sK73('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1451,f1870,f1296])).
fof(f25506,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17211,f1451])).
fof(f17211,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1911,f17173,f1294])).
fof(f1911,plain,(
  sP2(sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),'F_0')),
  inference(unit_resulting_resolution,[],[f1451,f1899,f1296])).
fof(f25505,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17212,f1451])).
fof(f17212,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2102,f17173,f1294])).
fof(f2102,plain,(
  sP2(sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),'F_0')),
  inference(unit_resulting_resolution,[],[f1451,f1901,f1296])).
fof(f25504,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17213,f1451])).
fof(f17213,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2226,f17173,f1294])).
fof(f2226,plain,(
  sP2(sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),'F_0')),
  inference(unit_resulting_resolution,[],[f1451,f1930,f1296])).
fof(f25503,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17221,f1451])).
fof(f17221,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1804,f17173,f1294])).
fof(f1804,plain,(
  sP2('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1451,f919,f1299])).
fof(f25502,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17222,f1451])).
fof(f17222,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7681,f17173,f1294])).
fof(f25501,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17223,f1451])).
fof(f17223,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f8606,f17173,f1294])).
fof(f25500,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17224,f1451])).
fof(f17224,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1883,f17173,f1294])).
fof(f25499,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17225,f1451])).
fof(f17225,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f4784,f17173,f1294])).
fof(f25498,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17226,f1451])).
fof(f17226,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1912,f17173,f1294])).
fof(f25497,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17227,f1451])).
fof(f17227,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2103,f17173,f1294])).
fof(f25496,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17228,f1451])).
fof(f17228,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2227,f17173,f1294])).
fof(f25495,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17229,f17173])).
fof(f17229,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6040,f17173,f1294])).
fof(f6040,plain,(
  sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f919,f5976,f1299])).
fof(f25494,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17230,f17173])).
fof(f17230,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6854,f17173,f1294])).
fof(f6854,plain,(
  sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f919,f5977,f1299])).
fof(f25493,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17231,f17173])).
fof(f17231,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7680,f17173,f1294])).
fof(f25492,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17232,f17173])).
fof(f17232,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14)))) | $spl10),
  inference(unit_resulting_resolution,[],[f8605,f17173,f1294])).
fof(f25491,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17233,f17173])).
fof(f17233,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1882,f17173,f1294])).
fof(f25490,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17234,f17173])).
fof(f17234,plain,(
  p(s(bool,sK73('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl10),
  inference(unit_resulting_resolution,[],[f4783,f17173,f1294])).
fof(f25489,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17235,f17173])).
fof(f17235,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1911,f17173,f1294])).
fof(f25488,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17236,f17173])).
fof(f17236,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2102,f17173,f1294])).
fof(f25487,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17237,f17173])).
fof(f17237,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2226,f17173,f1294])).
fof(f25486,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17309,f19909])).
fof(f19909,plain,(
  ( ! [X0] : (sP6('T_0','T_0',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f1483,f1489,f17173,f1336])).
fof(f1336,plain,(
  ( ! [X2,X0,X1] : (~sP4(X1,X2,X0) | p(s(bool,X0)) | ~sP5(X0,X2,X1) | sP6(X0,X1,X2) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f881])).
fof(f881,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (sP3(X2,X1,X0) & sP4(X1,X2,X0) & sP5(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP3(X2,X1,X0) | ~sP4(X1,X2,X0) | ~sP5(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f880])).
fof(f880,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP6(X2,X1,X0)))),
  inference(flattening,[],[f879])).
fof(f879,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP6(X2,X1,X0)))),
  inference(nnf_transformation,[],[f715])).
fof(f1489,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f919,f1342])).
fof(f1342,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f887])).
fof(f887,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f886])).
fof(f886,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(flattening,[],[f885])).
fof(f885,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(nnf_transformation,[],[f713])).
fof(f1483,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f919,f1338])).
fof(f1338,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f884])).
fof(f884,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f883])).
fof(f883,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f882])).
fof(f882,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f714])).
fof(f1495,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f919,f1346])).
fof(f1346,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f889])).
fof(f889,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(flattening,[],[f888])).
fof(f888,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(nnf_transformation,[],[f712])).
fof(f17309,plain,(
  ( ! [X0] : (~sP6('T_0','T_0',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f919,f919,f17173,f1333])).
fof(f1333,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f881])).
fof(f25485,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17311,f22324])).
fof(f22324,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1344])).
fof(f1344,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f887])).
fof(f17311,plain,(
  ~sP4('T_0','F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1483,f17173,f1335])).
fof(f1335,plain,(
  ( ! [X2,X0,X1] : (~sP4(X1,X2,X0) | p(s(bool,X1)) | ~sP5(X0,X2,X1) | sP6(X0,X1,X2) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f881])).
fof(f3177,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f919,f1710,f1351])).
fof(f1497,plain,(
  ( ! [X0,X1] : (sP3(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f919,f1347])).
fof(f1347,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f889])).
fof(f25484,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17312,f22324])).
fof(f17312,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f1483,f17173,f1335])).
fof(f6041,plain,(
  ~sP6('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(unit_resulting_resolution,[],[f919,f919,f5976,f1333])).
fof(f25483,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17313,f22324])).
fof(f17313,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f1483,f17173,f1335])).
fof(f6855,plain,(
  ~sP6('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(unit_resulting_resolution,[],[f919,f919,f5977,f1333])).
fof(f25482,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17319,f22324])).
fof(f17319,plain,(
  ~sP4('T_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))) | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f1483,f17173,f1335])).
fof(f7779,plain,(
  ~sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1710,f7614,f1351])).
fof(f25481,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17321,f22324])).
fof(f17321,plain,(
  ( ! [X0] : (~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.combin.I_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))),'F_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1543,f3205,f1540,f17173,f1335])).
fof(f1540,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f947,f1338])).
fof(f947,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f270])).
fof(f270,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f134])).
fof(f134,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X6))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.le_refl')).
fof(f3205,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1451,f1127,f1352])).
fof(f1352,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f890])).
fof(f1543,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f947,f1347])).
fof(f25480,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17323,f22324])).
fof(f17323,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1530,f3205,f1527,f17173,f1335])).
fof(f1527,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f926,f1338])).
fof(f926,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(cnf_transformation,[],[f112])).
fof(f112,axiom,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.lt_infty|split|2')).
fof(f1530,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f926,f1347])).
fof(f25479,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17324,f22324])).
fof(f17324,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f7716,f17173,f1335])).
fof(f7716,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) )),
  inference(unit_resulting_resolution,[],[f7614,f1338])).
fof(f7684,plain,(
  ~sP6('T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0')),
  inference(unit_resulting_resolution,[],[f919,f1451,f7614,f1333])).
fof(f7745,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f7614,f1347])).
fof(f25478,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17326,f22324])).
fof(f17326,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f3205,f7716,f17173,f1335])).
fof(f25477,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17327,f22324])).
fof(f17327,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14)),'F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f8641,f17173,f1335])).
fof(f8641,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14)))) )),
  inference(unit_resulting_resolution,[],[f7615,f1338])).
fof(f8609,plain,(
  ~sP6('T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14)),'F_0')),
  inference(unit_resulting_resolution,[],[f919,f1451,f7615,f1333])).
fof(f25476,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17329,f22324])).
fof(f17329,plain,(
  ~sP4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1498,f3205,f1484,f17173,f1335])).
fof(f1484,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f921,f1338])).
fof(f921,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f228])).
fof(f228,axiom,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.extreal.add_not_infty|split|0','thm.numeral.numeral_distrib|split|35')).
fof(f1498,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f921,f1347])).
fof(f25475,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17331,f22324])).
fof(f17331,plain,(
  ~sP4(sK73(sK59('const.num.|48|_0','type.num.num')),'const.combin.I_1'(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1887,f3205,f1884,f17173,f1335])).
fof(f1884,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,sK73(sK59('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f1871,f1338])).
fof(f1887,plain,(
  ( ! [X0,X1] : (sP3(X0,sK73(sK59('const.num.|48|_0','type.num.num')),X1)) )),
  inference(unit_resulting_resolution,[],[f1871,f1347])).
fof(f25474,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17335,f22324])).
fof(f17335,plain,(
  ~sP4(sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),'const.combin.I_1'(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1916,f3205,f1913,f17173,f1335])).
fof(f1913,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) )),
  inference(unit_resulting_resolution,[],[f1899,f1338])).
fof(f1916,plain,(
  ( ! [X0,X1] : (sP3(X0,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),X1)) )),
  inference(unit_resulting_resolution,[],[f1899,f1347])).
fof(f25473,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17337,f22324])).
fof(f17337,plain,(
  ~sP4(sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2107,f3205,f2104,f17173,f1335])).
fof(f2104,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1901,f1338])).
fof(f2107,plain,(
  ( ! [X0,X1] : (sP3(X0,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),X1)) )),
  inference(unit_resulting_resolution,[],[f1901,f1347])).
fof(f25472,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17339,f22324])).
fof(f17339,plain,(
  ~sP4(sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2231,f3205,f2228,f17173,f1335])).
fof(f2228,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1930,f1338])).
fof(f2231,plain,(
  ( ! [X0,X1] : (sP3(X0,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),X1)) )),
  inference(unit_resulting_resolution,[],[f1930,f1347])).
fof(f25471,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17340,f22324])).
fof(f17340,plain,(
  ~sP4('T_0','F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1485,f17173,f1335])).
fof(f1485,plain,(
  ( ! [X0,X1] : (sP5(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1451,f1339])).
fof(f1339,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f884])).
fof(f25470,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17341,f22324])).
fof(f17341,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f1485,f17173,f1335])).
fof(f25469,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17342,f22324])).
fof(f17342,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14)),'F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f1485,f17173,f1335])).
fof(f25468,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17343,f22324])).
fof(f17343,plain,(
  ~sP4('T_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))) | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f1485,f17173,f1335])).
fof(f25467,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17344,f22324])).
fof(f17344,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f6051,f17173,f1335])).
fof(f6051,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f5976,f1339])).
fof(f25466,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17345,f22324])).
fof(f17345,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f6865,f17173,f1335])).
fof(f6865,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f5977,f1339])).
fof(f25465,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17356,f22324])).
fof(f17356,plain,(
  ( ! [X0] : (~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.combin.I_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))),'F_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1543,f3205,f1487,f17173,f1335])).
fof(f1487,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1451,f1340])).
fof(f1340,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f884])).
fof(f25464,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17357,f22324])).
fof(f17357,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1530,f3205,f1487,f17173,f1335])).
fof(f25463,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17358,f22324])).
fof(f17358,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f3205,f1487,f17173,f1335])).
fof(f25462,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17359,f22324])).
fof(f17359,plain,(
  ~sP4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1498,f3205,f1487,f17173,f1335])).
fof(f25461,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17360,f22324])).
fof(f17360,plain,(
  ~sP4(sK73(sK59('const.num.|48|_0','type.num.num')),'const.combin.I_1'(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1887,f3205,f1487,f17173,f1335])).
fof(f25460,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17361,f22324])).
fof(f17361,plain,(
  ~sP4(sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),'const.combin.I_1'(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1916,f3205,f1487,f17173,f1335])).
fof(f25459,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17362,f22324])).
fof(f17362,plain,(
  ~sP4(sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2107,f3205,f1487,f17173,f1335])).
fof(f25458,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17363,f22324])).
fof(f17363,plain,(
  ~sP4(sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2231,f3205,f1487,f17173,f1335])).
fof(f25457,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17368,f22324])).
fof(f17368,plain,(
  ( ! [X0] : (~sP4('const.combin.I_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1544,f3194,f1487,f17173,f1335])).
fof(f3194,plain,(
  ( ! [X0] : (~sP6('F_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f1451,f1127,f1352])).
fof(f1544,plain,(
  ( ! [X2,X0,X1] : (sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f947,f1348])).
fof(f1348,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f889])).
fof(f25456,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17369,f22324])).
fof(f17369,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1531,f3194,f1487,f17173,f1335])).
fof(f1531,plain,(
  ( ! [X0,X1] : (sP3('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f926,f1348])).
fof(f25455,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17370,f22324])).
fof(f17370,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7746,f3194,f1487,f17173,f1335])).
fof(f7746,plain,(
  ( ! [X0,X1] : (sP3('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f7614,f1348])).
fof(f25454,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17371,f22324])).
fof(f17371,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1500,f3194,f1487,f17173,f1335])).
fof(f1500,plain,(
  ( ! [X0,X1] : (sP3('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f921,f1348])).
fof(f25453,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17372,f22324])).
fof(f17372,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))),sK73(sK59('const.num.|48|_0','type.num.num')),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1888,f3194,f1487,f17173,f1335])).
fof(f1888,plain,(
  ( ! [X0,X1] : (sP3(sK73(sK59('const.num.|48|_0','type.num.num')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1871,f1348])).
fof(f25452,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17373,f22324])).
fof(f17373,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f3194,f1487,f17173,f1335])).
fof(f1917,plain,(
  ( ! [X0,X1] : (sP3(sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1899,f1348])).
fof(f25451,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17374,f22324])).
fof(f17374,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))),sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2108,f3194,f1487,f17173,f1335])).
fof(f2108,plain,(
  ( ! [X0,X1] : (sP3(sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1901,f1348])).
fof(f25450,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17375,f22324])).
fof(f17375,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))),sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2232,f3194,f1487,f17173,f1335])).
fof(f2232,plain,(
  ( ! [X0,X1] : (sP3(sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1930,f1348])).
fof(f25449,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17390,f22292])).
fof(f22292,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1340])).
fof(f17390,plain,(
  ~sP5('T_0','F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1489,f17173,f1335])).
fof(f25448,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17391,f22292])).
fof(f17391,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f1489,f17173,f1335])).
fof(f25447,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17392,f22292])).
fof(f17392,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f1489,f17173,f1335])).
fof(f25446,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17393,f22292])).
fof(f17393,plain,(
  ~sP5('T_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))) | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f1489,f17173,f1335])).
fof(f25445,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17394,f22292])).
fof(f17394,plain,(
  ~sP5('T_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14))) | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f1489,f17173,f1335])).
fof(f25444,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f17905,f22292])).
fof(f17905,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f7743,f17173,f1335])).
fof(f7743,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) )),
  inference(unit_resulting_resolution,[],[f7614,f1342])).
fof(f25443,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f18805,f22292])).
fof(f18805,plain,(
  ~sP5('T_0','F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1491,f17173,f1335])).
fof(f1491,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1451,f1343])).
fof(f1343,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f887])).
fof(f25442,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f18806,f22292])).
fof(f18806,plain,(
  ~sP5('T_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))) | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f1491,f17173,f1335])).
fof(f25441,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f18807,f22292])).
fof(f18807,plain,(
  ~sP5('T_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14))) | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f1491,f17173,f1335])).
fof(f25440,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f18808,f22292])).
fof(f18808,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f1491,f17173,f1335])).
fof(f25439,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f18986,f22292])).
fof(f18986,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f6058,f17173,f1335])).
fof(f6058,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f5976,f1343])).
fof(f25438,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19150,f22292])).
fof(f19150,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f6872,f17173,f1335])).
fof(f6872,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f5977,f1343])).
fof(f25437,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19781,f22324])).
fof(f19781,plain,(
  ~sP4('T_0','F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1483,f17173,f1336])).
fof(f25436,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19782,f22324])).
fof(f19782,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f1483,f17173,f1336])).
fof(f25435,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19783,f22324])).
fof(f19783,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f1483,f17173,f1336])).
fof(f25434,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19789,f22324])).
fof(f19789,plain,(
  ~sP4('T_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))) | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f1483,f17173,f1336])).
fof(f25433,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19791,f22324])).
fof(f19791,plain,(
  ( ! [X0] : (~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.combin.I_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))),'F_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1543,f3205,f1540,f17173,f1336])).
fof(f25432,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19793,f22324])).
fof(f19793,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1530,f3205,f1527,f17173,f1336])).
fof(f25431,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19794,f22324])).
fof(f19794,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f7716,f17173,f1336])).
fof(f25430,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19796,f22324])).
fof(f19796,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f3205,f7716,f17173,f1336])).
fof(f25429,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19797,f22324])).
fof(f19797,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14)),'F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f8641,f17173,f1336])).
fof(f25428,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19799,f22324])).
fof(f19799,plain,(
  ~sP4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1498,f3205,f1484,f17173,f1336])).
fof(f25427,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19801,f22324])).
fof(f19801,plain,(
  ~sP4(sK73(sK59('const.num.|48|_0','type.num.num')),'const.combin.I_1'(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1887,f3205,f1884,f17173,f1336])).
fof(f25426,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19805,f22324])).
fof(f19805,plain,(
  ~sP4(sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),'const.combin.I_1'(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1916,f3205,f1913,f17173,f1336])).
fof(f25425,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19807,f22324])).
fof(f19807,plain,(
  ~sP4(sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2107,f3205,f2104,f17173,f1336])).
fof(f25424,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19809,f22324])).
fof(f19809,plain,(
  ~sP4(sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2231,f3205,f2228,f17173,f1336])).
fof(f25423,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19810,f22324])).
fof(f19810,plain,(
  ~sP4('T_0','F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1485,f17173,f1336])).
fof(f25422,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19811,f22324])).
fof(f19811,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f1485,f17173,f1336])).
fof(f25421,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19812,f22324])).
fof(f19812,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14)),'F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f1485,f17173,f1336])).
fof(f25420,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19813,f22324])).
fof(f19813,plain,(
  ~sP4('T_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))) | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f1485,f17173,f1336])).
fof(f25419,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19814,f22324])).
fof(f19814,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f6051,f17173,f1336])).
fof(f25418,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19815,f22324])).
fof(f19815,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f6865,f17173,f1336])).
fof(f25417,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19826,f22324])).
fof(f19826,plain,(
  ( ! [X0] : (~sP4('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'const.combin.I_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))),'F_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1543,f3205,f1487,f17173,f1336])).
fof(f25416,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19827,f22324])).
fof(f19827,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1530,f3205,f1487,f17173,f1336])).
fof(f25415,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19828,f22324])).
fof(f19828,plain,(
  ~sP4('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f3205,f1487,f17173,f1336])).
fof(f25414,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19829,f22324])).
fof(f19829,plain,(
  ~sP4('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1498,f3205,f1487,f17173,f1336])).
fof(f25413,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19830,f22324])).
fof(f19830,plain,(
  ~sP4(sK73(sK59('const.num.|48|_0','type.num.num')),'const.combin.I_1'(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1887,f3205,f1487,f17173,f1336])).
fof(f25412,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19831,f22324])).
fof(f19831,plain,(
  ~sP4(sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),'const.combin.I_1'(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1916,f3205,f1487,f17173,f1336])).
fof(f25411,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19832,f22324])).
fof(f19832,plain,(
  ~sP4(sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2107,f3205,f1487,f17173,f1336])).
fof(f25410,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19833,f22324])).
fof(f19833,plain,(
  ~sP4(sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),'const.combin.I_1'(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2231,f3205,f1487,f17173,f1336])).
fof(f25409,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19838,f22324])).
fof(f19838,plain,(
  ( ! [X0] : (~sP4('const.combin.I_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'F_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1544,f3194,f1487,f17173,f1336])).
fof(f25408,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19839,f22324])).
fof(f19839,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1531,f3194,f1487,f17173,f1336])).
fof(f25407,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19840,f22324])).
fof(f19840,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7746,f3194,f1487,f17173,f1336])).
fof(f25406,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19841,f22324])).
fof(f19841,plain,(
  ~sP4('const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1500,f3194,f1487,f17173,f1336])).
fof(f25405,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19842,f22324])).
fof(f19842,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))),sK73(sK59('const.num.|48|_0','type.num.num')),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1888,f3194,f1487,f17173,f1336])).
fof(f25404,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19843,f22324])).
fof(f19843,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f3194,f1487,f17173,f1336])).
fof(f25403,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19844,f22324])).
fof(f19844,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))),sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2108,f3194,f1487,f17173,f1336])).
fof(f25402,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19845,f22324])).
fof(f19845,plain,(
  ~sP4('const.combin.I_1'(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))),sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2232,f3194,f1487,f17173,f1336])).
fof(f25401,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19860,f22292])).
fof(f19860,plain,(
  ~sP5('T_0','F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1489,f17173,f1336])).
fof(f25400,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19861,f22292])).
fof(f19861,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f1489,f17173,f1336])).
fof(f25399,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19862,f22292])).
fof(f19862,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f1489,f17173,f1336])).
fof(f25398,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19863,f22292])).
fof(f19863,plain,(
  ~sP5('T_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))) | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f1489,f17173,f1336])).
fof(f25397,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f19864,f22292])).
fof(f19864,plain,(
  ~sP5('T_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14))) | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f1489,f17173,f1336])).
fof(f25396,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f20085,f22705])).
fof(f22705,plain,(
  ( ! [X0] : (~sP6(X0,'T_0',sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1903,f17173,f1352])).
fof(f1903,plain,(
  s(bool,'T_0') = s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),
  inference(unit_resulting_resolution,[],[f1899,f1039])).
fof(f20085,plain,(
  ( ! [X0] : (sP6('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),'T_0',sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f1483,f1541,f17173,f1336])).
fof(f1541,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f947,f1342])).
fof(f25395,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f20234,f22705])).
fof(f20234,plain,(
  sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0',sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f1483,f1528,f17173,f1336])).
fof(f1528,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f926,f1342])).
fof(f25394,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f20375,f22292])).
fof(f20375,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f7743,f17173,f1336])).
fof(f25393,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f20538,f22705])).
fof(f20538,plain,(
  sP6('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0',sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f1483,f1490,f17173,f1336])).
fof(f1490,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f921,f1342])).
fof(f25392,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f20687,f22705])).
fof(f20687,plain,(
  sP6(sK73(sK59('const.num.|48|_0','type.num.num')),'T_0',sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f1483,f1885,f17173,f1336])).
fof(f1885,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK73(sK59('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f1871,f1342])).
fof(f25391,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f20836,f22705])).
fof(f20836,plain,(
  sP6(sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),'T_0',sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f1483,f1914,f17173,f1336])).
fof(f1914,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) )),
  inference(unit_resulting_resolution,[],[f1899,f1342])).
fof(f25390,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f20985,f22705])).
fof(f20985,plain,(
  sP6(sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),'T_0',sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f1483,f2105,f17173,f1336])).
fof(f2105,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1901,f1342])).
fof(f25389,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f21134,f22705])).
fof(f21134,plain,(
  sP6(sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),'T_0',sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f1483,f2229,f17173,f1336])).
fof(f2229,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1930,f1342])).
fof(f25388,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f21275,f22292])).
fof(f21275,plain,(
  ~sP5('T_0','F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1491,f17173,f1336])).
fof(f25387,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f21276,f22292])).
fof(f21276,plain,(
  ~sP5('T_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))) | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f1491,f17173,f1336])).
fof(f25386,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f21277,f22292])).
fof(f21277,plain,(
  ~sP5('T_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14))) | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f1491,f17173,f1336])).
fof(f25385,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f21278,f22292])).
fof(f21278,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f1491,f17173,f1336])).
fof(f25384,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f21343,f7684])).
fof(f21343,plain,(
  sP6('T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f7716,f1491,f17173,f1336])).
fof(f25383,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f21346,f7685])).
fof(f7685,plain,(
  ~sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0')),
  inference(unit_resulting_resolution,[],[f7614,f1451,f7614,f1333])).
fof(f21346,plain,(
  sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7744,f7716,f1491,f17173,f1336])).
fof(f7744,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) )),
  inference(unit_resulting_resolution,[],[f7614,f1346])).
fof(f25382,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f21354,f8609])).
fof(f21354,plain,(
  sP6('T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14)),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8641,f1491,f17173,f1336])).
fof(f25381,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f21456,f22292])).
fof(f21456,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f6058,f17173,f1336])).
fof(f25380,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f21480,f6041])).
fof(f21480,plain,(
  sP6('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f1483,f6058,f17173,f1336])).
fof(f25379,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f21620,f22292])).
fof(f21620,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f6872,f17173,f1336])).
fof(f25378,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f21644,f6855])).
fof(f21644,plain,(
  sP6('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f1483,f6872,f17173,f1336])).
fof(f25377,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22288,f22292])).
fof(f22288,plain,(
  ( ! [X0] : (~sP5('T_0','T_0',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f919,f919,f17173,f1337])).
fof(f1337,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f884])).
fof(f25376,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22320,f22324])).
fof(f22320,plain,(
  ( ! [X0] : (~sP4('T_0','T_0',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f919,f919,f17173,f1341])).
fof(f1341,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f887])).
fof(f25375,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22327,f17173])).
fof(f22327,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1543,f1451,f17173,f1345])).
fof(f1345,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f889])).
fof(f25374,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22328,f17173])).
fof(f22328,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1530,f1451,f17173,f1345])).
fof(f25373,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22329,f17173])).
fof(f22329,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f1451,f17173,f1345])).
fof(f25372,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22330,f17173])).
fof(f22330,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1498,f1451,f17173,f1345])).
fof(f25371,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22331,f17173])).
fof(f22331,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1887,f1451,f17173,f1345])).
fof(f25370,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22332,f17173])).
fof(f22332,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1916,f1451,f17173,f1345])).
fof(f25369,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22333,f17173])).
fof(f22333,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2107,f1451,f17173,f1345])).
fof(f25368,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22334,f17173])).
fof(f22334,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2231,f1451,f17173,f1345])).
fof(f25367,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22335,f17173])).
fof(f22335,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1499,f1451,f17173,f1345])).
fof(f1499,plain,(
  ( ! [X0,X1] : (sP3('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f919,f1348])).
fof(f25366,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22336,f17173])).
fof(f22336,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1544,f1451,f17173,f1345])).
fof(f25365,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22337,f17173])).
fof(f22337,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1531,f1451,f17173,f1345])).
fof(f25364,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22338,f17173])).
fof(f22338,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7746,f1451,f17173,f1345])).
fof(f25363,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22339,f17173])).
fof(f22339,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1500,f1451,f17173,f1345])).
fof(f25362,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22340,f17173])).
fof(f22340,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1888,f1451,f17173,f1345])).
fof(f25361,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22341,f17173])).
fof(f22341,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f1451,f17173,f1345])).
fof(f25360,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22342,f17173])).
fof(f22342,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2108,f1451,f17173,f1345])).
fof(f25359,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22343,f17173])).
fof(f22343,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2232,f1451,f17173,f1345])).
fof(f25358,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22344,f17173])).
fof(f22344,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f17173,f17173,f1345])).
fof(f25357,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22345,f17173])).
fof(f22345,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1542,f17173,f17173,f1345])).
fof(f1542,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f947,f1346])).
fof(f25356,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22346,f17173])).
fof(f22346,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1529,f17173,f17173,f1345])).
fof(f1529,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f926,f1346])).
fof(f25355,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22347,f17173])).
fof(f22347,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7744,f17173,f17173,f1345])).
fof(f25354,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22348,f17173])).
fof(f22348,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1496,f17173,f17173,f1345])).
fof(f1496,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f921,f1346])).
fof(f25353,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22349,f17173])).
fof(f22349,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1886,f17173,f17173,f1345])).
fof(f1886,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK73(sK59('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f1871,f1346])).
fof(f25352,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22350,f17173])).
fof(f22350,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1915,f17173,f17173,f1345])).
fof(f1915,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) )),
  inference(unit_resulting_resolution,[],[f1899,f1346])).
fof(f25351,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22351,f17173])).
fof(f22351,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2106,f17173,f17173,f1345])).
fof(f2106,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1901,f1346])).
fof(f25350,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22352,f17173])).
fof(f22352,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2230,f17173,f17173,f1345])).
fof(f2230,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1930,f1346])).
fof(f25349,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22354,f17173])).
fof(f22354,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1543,f17173,f17173,f1345])).
fof(f25348,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22355,f17173])).
fof(f22355,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1530,f17173,f17173,f1345])).
fof(f25347,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22356,f17173])).
fof(f22356,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f17173,f17173,f1345])).
fof(f25346,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22357,f17173])).
fof(f22357,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1498,f17173,f17173,f1345])).
fof(f25345,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22358,f17173])).
fof(f22358,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1887,f17173,f17173,f1345])).
fof(f25344,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22359,f17173])).
fof(f22359,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1916,f17173,f17173,f1345])).
fof(f25343,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22360,f17173])).
fof(f22360,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2107,f17173,f17173,f1345])).
fof(f25342,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22361,f17173])).
fof(f22361,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2231,f17173,f17173,f1345])).
fof(f25341,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22362,f17173])).
fof(f22362,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1499,f17173,f17173,f1345])).
fof(f25340,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22363,f17173])).
fof(f22363,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1544,f17173,f17173,f1345])).
fof(f25339,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22364,f17173])).
fof(f22364,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1531,f17173,f17173,f1345])).
fof(f25338,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22365,f17173])).
fof(f22365,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7746,f17173,f17173,f1345])).
fof(f25337,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22366,f17173])).
fof(f22366,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1500,f17173,f17173,f1345])).
fof(f25336,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22367,f17173])).
fof(f22367,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1888,f17173,f17173,f1345])).
fof(f25335,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22368,f17173])).
fof(f22368,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f17173,f17173,f1345])).
fof(f25334,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22369,f17173])).
fof(f22369,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2108,f17173,f17173,f1345])).
fof(f25333,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22370,f17173])).
fof(f22370,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2232,f17173,f17173,f1345])).
fof(f25332,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22372,f17173])).
fof(f22372,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1542,f1451,f17173,f1345])).
fof(f25331,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22373,f17173])).
fof(f22373,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1529,f1451,f17173,f1345])).
fof(f25330,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22374,f17173])).
fof(f22374,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7744,f1451,f17173,f1345])).
fof(f25329,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22375,f17173])).
fof(f22375,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1496,f1451,f17173,f1345])).
fof(f25328,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22376,f17173])).
fof(f22376,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1886,f1451,f17173,f1345])).
fof(f25327,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22377,f17173])).
fof(f22377,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1915,f1451,f17173,f1345])).
fof(f25326,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22378,f17173])).
fof(f22378,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2106,f1451,f17173,f1345])).
fof(f25325,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22379,f17173])).
fof(f22379,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2230,f1451,f17173,f1345])).
fof(f25324,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22380,f17173])).
fof(f22380,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1499,f1451,f17173,f1345])).
fof(f25323,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22381,f17173])).
fof(f22381,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1544,f1451,f17173,f1345])).
fof(f25322,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22382,f17173])).
fof(f22382,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1531,f1451,f17173,f1345])).
fof(f25321,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22383,f17173])).
fof(f22383,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7746,f1451,f17173,f1345])).
fof(f25320,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22384,f17173])).
fof(f22384,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1500,f1451,f17173,f1345])).
fof(f25319,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22385,f17173])).
fof(f22385,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1888,f1451,f17173,f1345])).
fof(f25318,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22386,f17173])).
fof(f22386,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f1451,f17173,f1345])).
fof(f25317,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22387,f17173])).
fof(f22387,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2108,f1451,f17173,f1345])).
fof(f25316,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22388,f17173])).
fof(f22388,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2232,f1451,f17173,f1345])).
fof(f25315,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22394,f17173])).
fof(f22394,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1542,f17173,f17173,f1345])).
fof(f25314,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22395,f17173])).
fof(f22395,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1529,f17173,f17173,f1345])).
fof(f25313,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22396,f17173])).
fof(f22396,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7744,f17173,f17173,f1345])).
fof(f25312,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22397,f17173])).
fof(f22397,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1496,f17173,f17173,f1345])).
fof(f25311,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22398,f17173])).
fof(f22398,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1886,f17173,f17173,f1345])).
fof(f25310,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22399,f17173])).
fof(f22399,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1915,f17173,f17173,f1345])).
fof(f25309,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22400,f17173])).
fof(f22400,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2106,f17173,f17173,f1345])).
fof(f25308,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22401,f17173])).
fof(f22401,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2230,f17173,f17173,f1345])).
fof(f25307,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22402,f17173])).
fof(f22402,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f17173,f17173,f1345])).
fof(f25306,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22403,f17173])).
fof(f22403,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1543,f17173,f17173,f1345])).
fof(f25305,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22404,f17173])).
fof(f22404,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1530,f17173,f17173,f1345])).
fof(f25304,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22405,f17173])).
fof(f22405,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f17173,f17173,f1345])).
fof(f25303,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22406,f17173])).
fof(f22406,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1498,f17173,f17173,f1345])).
fof(f25302,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22407,f17173])).
fof(f22407,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1887,f17173,f17173,f1345])).
fof(f25301,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22408,f17173])).
fof(f22408,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1916,f17173,f17173,f1345])).
fof(f25300,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22409,f17173])).
fof(f22409,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2107,f17173,f17173,f1345])).
fof(f25299,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22410,f17173])).
fof(f22410,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2231,f17173,f17173,f1345])).
fof(f25298,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22411,f17173])).
fof(f22411,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1499,f17173,f17173,f1345])).
fof(f25297,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22412,f17173])).
fof(f22412,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1544,f17173,f17173,f1345])).
fof(f25296,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22413,f17173])).
fof(f22413,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1531,f17173,f17173,f1345])).
fof(f25295,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22414,f17173])).
fof(f22414,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7746,f17173,f17173,f1345])).
fof(f25294,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22415,f17173])).
fof(f22415,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1500,f17173,f17173,f1345])).
fof(f25293,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22416,f17173])).
fof(f22416,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1888,f17173,f17173,f1345])).
fof(f25292,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22417,f17173])).
fof(f22417,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f17173,f17173,f1345])).
fof(f25291,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22418,f17173])).
fof(f22418,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2108,f17173,f17173,f1345])).
fof(f25290,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22419,f17173])).
fof(f22419,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2232,f17173,f17173,f1345])).
fof(f25289,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22425,f17173])).
fof(f22425,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1542,f1451,f17173,f1345])).
fof(f25288,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22426,f17173])).
fof(f22426,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1529,f1451,f17173,f1345])).
fof(f25287,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22427,f17173])).
fof(f22427,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7744,f1451,f17173,f1345])).
fof(f25286,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22428,f17173])).
fof(f22428,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1496,f1451,f17173,f1345])).
fof(f25285,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22429,f17173])).
fof(f22429,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1886,f1451,f17173,f1345])).
fof(f25284,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22430,f17173])).
fof(f22430,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1915,f1451,f17173,f1345])).
fof(f25283,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22431,f17173])).
fof(f22431,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2106,f1451,f17173,f1345])).
fof(f25282,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22432,f17173])).
fof(f22432,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2230,f1451,f17173,f1345])).
fof(f25281,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22433,f17173])).
fof(f22433,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f1451,f17173,f1345])).
fof(f25280,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22434,f17173])).
fof(f22434,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1543,f1451,f17173,f1345])).
fof(f25279,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22435,f17173])).
fof(f22435,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1530,f1451,f17173,f1345])).
fof(f25278,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22436,f17173])).
fof(f22436,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f1451,f17173,f1345])).
fof(f25277,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22437,f17173])).
fof(f22437,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1498,f1451,f17173,f1345])).
fof(f25276,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22438,f17173])).
fof(f22438,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1887,f1451,f17173,f1345])).
fof(f25275,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22439,f17173])).
fof(f22439,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1916,f1451,f17173,f1345])).
fof(f25274,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22440,f17173])).
fof(f22440,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2107,f1451,f17173,f1345])).
fof(f25273,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22441,f17173])).
fof(f22441,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2231,f1451,f17173,f1345])).
fof(f25272,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22443,f17173])).
fof(f22443,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1542,f17173,f17173,f1345])).
fof(f25271,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22444,f17173])).
fof(f22444,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1529,f17173,f17173,f1345])).
fof(f25270,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22445,f17173])).
fof(f22445,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7744,f17173,f17173,f1345])).
fof(f25269,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22446,f17173])).
fof(f22446,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1496,f17173,f17173,f1345])).
fof(f25268,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22447,f17173])).
fof(f22447,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1886,f17173,f17173,f1345])).
fof(f25267,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22448,f17173])).
fof(f22448,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1915,f17173,f17173,f1345])).
fof(f25266,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22449,f17173])).
fof(f22449,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2106,f17173,f17173,f1345])).
fof(f25265,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22450,f17173])).
fof(f22450,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2230,f17173,f17173,f1345])).
fof(f25264,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22451,f17173])).
fof(f22451,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f17173,f17173,f1345])).
fof(f25263,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22452,f17173])).
fof(f22452,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1543,f17173,f17173,f1345])).
fof(f25262,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22453,f17173])).
fof(f22453,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1530,f17173,f17173,f1345])).
fof(f25261,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22454,f17173])).
fof(f22454,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f17173,f17173,f1345])).
fof(f25260,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22455,f17173])).
fof(f22455,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1498,f17173,f17173,f1345])).
fof(f25259,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22456,f17173])).
fof(f22456,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1887,f17173,f17173,f1345])).
fof(f25258,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22457,f17173])).
fof(f22457,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1916,f17173,f17173,f1345])).
fof(f25257,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22458,f17173])).
fof(f22458,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2107,f17173,f17173,f1345])).
fof(f25256,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22459,f17173])).
fof(f22459,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2231,f17173,f17173,f1345])).
fof(f25255,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22460,f17173])).
fof(f22460,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1499,f17173,f17173,f1345])).
fof(f25254,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22461,f17173])).
fof(f22461,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1544,f17173,f17173,f1345])).
fof(f25253,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22462,f17173])).
fof(f22462,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1531,f17173,f17173,f1345])).
fof(f25252,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22463,f17173])).
fof(f22463,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7746,f17173,f17173,f1345])).
fof(f25251,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22464,f17173])).
fof(f22464,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1500,f17173,f17173,f1345])).
fof(f25250,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22465,f17173])).
fof(f22465,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1888,f17173,f17173,f1345])).
fof(f25249,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22466,f17173])).
fof(f22466,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f17173,f17173,f1345])).
fof(f25248,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22467,f17173])).
fof(f22467,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2108,f17173,f17173,f1345])).
fof(f25247,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22468,f17173])).
fof(f22468,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2232,f17173,f17173,f1345])).
fof(f25246,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22469,f17173])).
fof(f22469,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f1451,f17173,f1345])).
fof(f25245,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22470,f17173])).
fof(f22470,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1543,f1451,f17173,f1345])).
fof(f25244,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22471,f17173])).
fof(f22471,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1530,f1451,f17173,f1345])).
fof(f25243,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22472,f17173])).
fof(f22472,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f1451,f17173,f1345])).
fof(f25242,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22473,f17173])).
fof(f22473,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1498,f1451,f17173,f1345])).
fof(f25241,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22474,f17173])).
fof(f22474,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1887,f1451,f17173,f1345])).
fof(f25240,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22475,f17173])).
fof(f22475,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1916,f1451,f17173,f1345])).
fof(f25239,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22476,f17173])).
fof(f22476,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2107,f1451,f17173,f1345])).
fof(f25238,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22477,f17173])).
fof(f22477,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2231,f1451,f17173,f1345])).
fof(f25237,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22479,f17173])).
fof(f22479,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1544,f1451,f17173,f1345])).
fof(f25236,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22480,f17173])).
fof(f22480,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1531,f1451,f17173,f1345])).
fof(f25235,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22481,f17173])).
fof(f22481,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7746,f1451,f17173,f1345])).
fof(f25234,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22482,f17173])).
fof(f22482,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1500,f1451,f17173,f1345])).
fof(f25233,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22483,f17173])).
fof(f22483,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1888,f1451,f17173,f1345])).
fof(f25232,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22484,f17173])).
fof(f22484,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f1451,f17173,f1345])).
fof(f25231,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22485,f17173])).
fof(f22485,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2108,f1451,f17173,f1345])).
fof(f25230,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22486,f17173])).
fof(f22486,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2232,f1451,f17173,f1345])).
fof(f25229,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22491,f17173])).
fof(f22491,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f17173,f17173,f1345])).
fof(f25228,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22492,f17173])).
fof(f22492,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1542,f17173,f17173,f1345])).
fof(f25227,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22493,f17173])).
fof(f22493,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1529,f17173,f17173,f1345])).
fof(f25226,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22494,f17173])).
fof(f22494,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7744,f17173,f17173,f1345])).
fof(f25225,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22495,f17173])).
fof(f22495,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1496,f17173,f17173,f1345])).
fof(f25224,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22496,f17173])).
fof(f22496,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1886,f17173,f17173,f1345])).
fof(f25223,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22497,f17173])).
fof(f22497,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1915,f17173,f17173,f1345])).
fof(f25222,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22498,f17173])).
fof(f22498,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2106,f17173,f17173,f1345])).
fof(f25221,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22499,f17173])).
fof(f22499,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2230,f17173,f17173,f1345])).
fof(f25220,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22500,f17173])).
fof(f22500,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f17173,f17173,f1345])).
fof(f25219,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22501,f17173])).
fof(f22501,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1543,f17173,f17173,f1345])).
fof(f25218,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22502,f17173])).
fof(f22502,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1530,f17173,f17173,f1345])).
fof(f25217,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22503,f17173])).
fof(f22503,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f17173,f17173,f1345])).
fof(f25216,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22504,f17173])).
fof(f22504,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1498,f17173,f17173,f1345])).
fof(f25215,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22505,f17173])).
fof(f22505,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1887,f17173,f17173,f1345])).
fof(f25214,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22506,f17173])).
fof(f22506,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1916,f17173,f17173,f1345])).
fof(f25213,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22507,f17173])).
fof(f22507,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2107,f17173,f17173,f1345])).
fof(f25212,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22508,f17173])).
fof(f22508,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2231,f17173,f17173,f1345])).
fof(f25211,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22510,f17173])).
fof(f22510,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1544,f17173,f17173,f1345])).
fof(f25210,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22511,f17173])).
fof(f22511,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1531,f17173,f17173,f1345])).
fof(f25209,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22512,f17173])).
fof(f22512,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7746,f17173,f17173,f1345])).
fof(f25208,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22513,f17173])).
fof(f22513,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1500,f17173,f17173,f1345])).
fof(f25207,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22514,f17173])).
fof(f22514,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1888,f17173,f17173,f1345])).
fof(f25206,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22515,f17173])).
fof(f22515,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f17173,f17173,f1345])).
fof(f25205,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22516,f17173])).
fof(f22516,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2108,f17173,f17173,f1345])).
fof(f25204,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22517,f17173])).
fof(f22517,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2232,f17173,f17173,f1345])).
fof(f25203,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22522,f17173])).
fof(f22522,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f1451,f17173,f1345])).
fof(f25202,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22523,f17173])).
fof(f22523,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1542,f1451,f17173,f1345])).
fof(f25201,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22524,f17173])).
fof(f22524,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1529,f1451,f17173,f1345])).
fof(f25200,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22525,f17173])).
fof(f22525,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7744,f1451,f17173,f1345])).
fof(f25199,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22526,f17173])).
fof(f22526,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1496,f1451,f17173,f1345])).
fof(f25198,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22527,f17173])).
fof(f22527,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1886,f1451,f17173,f1345])).
fof(f25197,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22528,f17173])).
fof(f22528,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1915,f1451,f17173,f1345])).
fof(f25196,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22529,f17173])).
fof(f22529,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2106,f1451,f17173,f1345])).
fof(f25195,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22530,f17173])).
fof(f22530,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2230,f1451,f17173,f1345])).
fof(f25194,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22532,f17173])).
fof(f22532,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1543,f1451,f17173,f1345])).
fof(f25193,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22533,f17173])).
fof(f22533,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1530,f1451,f17173,f1345])).
fof(f25192,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22534,f17173])).
fof(f22534,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f1451,f17173,f1345])).
fof(f25191,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22535,f17173])).
fof(f22535,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1498,f1451,f17173,f1345])).
fof(f25190,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22536,f17173])).
fof(f22536,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1887,f1451,f17173,f1345])).
fof(f25189,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22537,f17173])).
fof(f22537,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1916,f1451,f17173,f1345])).
fof(f25188,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22538,f17173])).
fof(f22538,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2107,f1451,f17173,f1345])).
fof(f25187,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22539,f17173])).
fof(f22539,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2231,f1451,f17173,f1345])).
fof(f25186,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22540,f17173])).
fof(f22540,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f17173,f17173,f1345])).
fof(f25185,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22541,f17173])).
fof(f22541,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1542,f17173,f17173,f1345])).
fof(f25184,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22542,f17173])).
fof(f22542,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1529,f17173,f17173,f1345])).
fof(f25183,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22543,f17173])).
fof(f22543,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7744,f17173,f17173,f1345])).
fof(f25182,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22544,f17173])).
fof(f22544,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1496,f17173,f17173,f1345])).
fof(f25181,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22545,f17173])).
fof(f22545,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1886,f17173,f17173,f1345])).
fof(f25180,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22546,f17173])).
fof(f22546,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1915,f17173,f17173,f1345])).
fof(f25179,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22547,f17173])).
fof(f22547,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2106,f17173,f17173,f1345])).
fof(f25178,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22548,f17173])).
fof(f22548,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2230,f17173,f17173,f1345])).
fof(f25177,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22550,f17173])).
fof(f22550,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1543,f17173,f17173,f1345])).
fof(f25176,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22551,f17173])).
fof(f22551,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1530,f17173,f17173,f1345])).
fof(f25175,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22552,f17173])).
fof(f22552,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f17173,f17173,f1345])).
fof(f25174,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22553,f17173])).
fof(f22553,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1498,f17173,f17173,f1345])).
fof(f25173,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22554,f17173])).
fof(f22554,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1887,f17173,f17173,f1345])).
fof(f25172,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22555,f17173])).
fof(f22555,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1916,f17173,f17173,f1345])).
fof(f25171,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22556,f17173])).
fof(f22556,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2107,f17173,f17173,f1345])).
fof(f25170,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22557,f17173])).
fof(f22557,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2231,f17173,f17173,f1345])).
fof(f25169,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22558,f17173])).
fof(f22558,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1499,f17173,f17173,f1345])).
fof(f25168,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22559,f17173])).
fof(f22559,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1544,f17173,f17173,f1345])).
fof(f25167,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22560,f17173])).
fof(f22560,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1531,f17173,f17173,f1345])).
fof(f25166,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22561,f17173])).
fof(f22561,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7746,f17173,f17173,f1345])).
fof(f25165,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22562,f17173])).
fof(f22562,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1500,f17173,f17173,f1345])).
fof(f25164,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22563,f17173])).
fof(f22563,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1888,f17173,f17173,f1345])).
fof(f25163,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22564,f17173])).
fof(f22564,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f17173,f17173,f1345])).
fof(f25162,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22565,f17173])).
fof(f22565,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2108,f17173,f17173,f1345])).
fof(f25161,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22566,f17173])).
fof(f22566,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2232,f17173,f17173,f1345])).
fof(f25160,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22567,f17173])).
fof(f22567,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f1451,f17173,f1345])).
fof(f25159,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22568,f17173])).
fof(f22568,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1542,f1451,f17173,f1345])).
fof(f25158,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22569,f17173])).
fof(f22569,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1529,f1451,f17173,f1345])).
fof(f25157,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22570,f17173])).
fof(f22570,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7744,f1451,f17173,f1345])).
fof(f25156,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22571,f17173])).
fof(f22571,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1496,f1451,f17173,f1345])).
fof(f25155,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22572,f17173])).
fof(f22572,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1886,f1451,f17173,f1345])).
fof(f25154,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22573,f17173])).
fof(f22573,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1915,f1451,f17173,f1345])).
fof(f25153,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22574,f17173])).
fof(f22574,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2106,f1451,f17173,f1345])).
fof(f25152,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22575,f17173])).
fof(f22575,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2230,f1451,f17173,f1345])).
fof(f25151,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22577,f17173])).
fof(f22577,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1544,f1451,f17173,f1345])).
fof(f25150,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22578,f17173])).
fof(f22578,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1531,f1451,f17173,f1345])).
fof(f25149,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22579,f17173])).
fof(f22579,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7746,f1451,f17173,f1345])).
fof(f25148,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22580,f17173])).
fof(f22580,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1500,f1451,f17173,f1345])).
fof(f25147,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22581,f17173])).
fof(f22581,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1888,f1451,f17173,f1345])).
fof(f25146,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22582,f17173])).
fof(f22582,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f1451,f17173,f1345])).
fof(f25145,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22583,f17173])).
fof(f22583,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2108,f1451,f17173,f1345])).
fof(f25144,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22584,f17173])).
fof(f22584,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2232,f1451,f17173,f1345])).
fof(f25143,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22589,f17173])).
fof(f22589,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f17173,f17173,f1345])).
fof(f25142,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22590,f17173])).
fof(f22590,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1542,f17173,f17173,f1345])).
fof(f25141,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22591,f17173])).
fof(f22591,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1529,f17173,f17173,f1345])).
fof(f25140,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22592,f17173])).
fof(f22592,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7744,f17173,f17173,f1345])).
fof(f25139,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22593,f17173])).
fof(f22593,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1496,f17173,f17173,f1345])).
fof(f25138,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22594,f17173])).
fof(f22594,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1886,f17173,f17173,f1345])).
fof(f25137,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22595,f17173])).
fof(f22595,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1915,f17173,f17173,f1345])).
fof(f25136,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22596,f17173])).
fof(f22596,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2106,f17173,f17173,f1345])).
fof(f25135,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22597,f17173])).
fof(f22597,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2230,f17173,f17173,f1345])).
fof(f25134,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22598,f17173])).
fof(f22598,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f17173,f17173,f1345])).
fof(f25133,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22599,f17173])).
fof(f22599,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1543,f17173,f17173,f1345])).
fof(f25132,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22600,f17173])).
fof(f22600,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1530,f17173,f17173,f1345])).
fof(f25131,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22601,f17173])).
fof(f22601,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f17173,f17173,f1345])).
fof(f25130,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22602,f17173])).
fof(f22602,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1498,f17173,f17173,f1345])).
fof(f25129,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22603,f17173])).
fof(f22603,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1887,f17173,f17173,f1345])).
fof(f25128,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22604,f17173])).
fof(f22604,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1916,f17173,f17173,f1345])).
fof(f25127,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22605,f17173])).
fof(f22605,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2107,f17173,f17173,f1345])).
fof(f25126,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22606,f17173])).
fof(f22606,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2231,f17173,f17173,f1345])).
fof(f25125,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22608,f17173])).
fof(f22608,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1544,f17173,f17173,f1345])).
fof(f25124,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22609,f17173])).
fof(f22609,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1531,f17173,f17173,f1345])).
fof(f25123,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22610,f17173])).
fof(f22610,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7746,f17173,f17173,f1345])).
fof(f25122,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22611,f17173])).
fof(f22611,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1500,f17173,f17173,f1345])).
fof(f25121,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22612,f17173])).
fof(f22612,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1888,f17173,f17173,f1345])).
fof(f25120,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22613,f17173])).
fof(f22613,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1917,f17173,f17173,f1345])).
fof(f25119,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22614,f17173])).
fof(f22614,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2108,f17173,f17173,f1345])).
fof(f25118,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22615,f17173])).
fof(f22615,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2232,f17173,f17173,f1345])).
fof(f25117,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22620,f1710])).
fof(f22620,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f3175,f17173,f1349])).
fof(f3175,plain,(
  ~sP6('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f919,f1710,f1351])).
fof(f25116,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22621,f1710])).
fof(f22621,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f3177,f17173,f1349])).
fof(f25115,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22622,f6334])).
fof(f6334,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(unit_resulting_resolution,[],[f919,f5976,f1435])).
fof(f1435,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f1434])).
fof(f1434,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f1415])).
fof(f1415,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f915])).
fof(f915,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f914])).
fof(f914,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f691])).
fof(f691,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f690])).
fof(f690,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f576])).
fof(f576,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X6,X41,X7,X42] : ((s(bool,X6) = s(bool,X41) & (p(s(bool,X41)) => s(bool,X7) = s(bool,X42))) => ((p(s(bool,X6)) => p(s(bool,X7))) <=> (p(s(bool,X41)) => p(s(bool,X42)))))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.IMP_CONG')).
fof(f22622,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl10),
  inference(unit_resulting_resolution,[],[f6041,f17173,f1349])).
fof(f25114,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22623,f7148])).
fof(f7148,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(unit_resulting_resolution,[],[f919,f5977,f1435])).
fof(f22623,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0'))) | $spl10),
  inference(unit_resulting_resolution,[],[f6855,f17173,f1349])).
fof(f25113,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22624,f24489])).
fof(f22624,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))) | $spl10),
  inference(unit_resulting_resolution,[],[f7684,f17173,f1349])).
fof(f25112,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22625,f24489])).
fof(f22625,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14))) | $spl10),
  inference(unit_resulting_resolution,[],[f8609,f17173,f1349])).
fof(f25111,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22633,f24489])).
fof(f22633,plain,(
  s(bool,'const.combin.I_1'(s(bool,'T_0'))) = s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) | $spl10),
  inference(unit_resulting_resolution,[],[f3648,f17173,f1349])).
fof(f3648,plain,(
  ~sP6('F_0','const.combin.I_1'(s(bool,'T_0')),sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),
  inference(superposition,[],[f3194,f1903])).
fof(f25110,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22637,f24489])).
fof(f22637,plain,(
  s(bool,'const.combin.I_1'(s(bool,'T_0'))) = s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) | $spl10),
  inference(unit_resulting_resolution,[],[f3743,f17173,f1349])).
fof(f3743,plain,(
  ~sP6('F_0',sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),'const.combin.I_1'(s(bool,'T_0')))),
  inference(superposition,[],[f3205,f1903])).
fof(f25109,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22638,f1710])).
fof(f22638,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7779,f17173,f1349])).
fof(f25108,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22639,f1710])).
fof(f22639,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7780,f17173,f1349])).
fof(f7780,plain,(
  ~sP6('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1710,f7614,f1351])).
fof(f25107,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22759,f24363])).
fof(f24363,plain,(
  ( ! [X0,X1] : (sP8('T_0',X0,X1)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1501,f17173,f1438])).
fof(f1438,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f1359])).
fof(f1359,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f893])).
fof(f893,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (sP7(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP7(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f892])).
fof(f892,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f891])).
fof(f891,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f718])).
fof(f718,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f1501,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f919,f1361])).
fof(f1361,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f895])).
fof(f895,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(flattening,[],[f894])).
fof(f894,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(nnf_transformation,[],[f717])).
fof(f717,plain,(
  ! [X0,X1,X2] : (sP7(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f22759,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f919,f17173,f1354])).
fof(f1354,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f893])).
fof(f25106,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22784,f24363])).
fof(f22784,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f919,f17173,f1355])).
fof(f1355,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f893])).
fof(f25105,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22812,f22816])).
fof(f22816,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1363])).
fof(f1363,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f895])).
fof(f22812,plain,(
  ( ! [X0] : (~sP7('T_0','T_0',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f919,f919,f17173,f1360])).
fof(f1360,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f895])).
fof(f25104,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22824,f1451])).
fof(f22824,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f4642,f17173,f1364])).
fof(f1364,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f897])).
fof(f897,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f896])).
fof(f896,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f719])).
fof(f719,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP8(X2,X0,X1))),
  inference(definition_folding,[],[f563,f718,f717])).
fof(f563,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f562])).
fof(f562,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X33,X43,X44] : ((p(s(bool,X44)) <=> (p(s(bool,X43)) & p(s(bool,X33)))) <=> ((p(s(bool,X44)) | ~p(s(bool,X43)) | ~p(s(bool,X33))) & (p(s(bool,X43)) | ~p(s(bool,X44))) & (p(s(bool,X33)) | ~p(s(bool,X44)))))),
  file('thm.extreal.add_not_infty|split|0','thm.sat.sat_23_prove_ltb')).
fof(f4642,plain,(
  ~sP8('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f4627,f1353])).
fof(f1353,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f893])).
fof(f4627,plain,(
  ~sP7('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f919,f1451,f919,f1360])).
fof(f25103,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22825,f1451])).
fof(f22825,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7862,f17173,f1364])).
fof(f7862,plain,(
  ~sP8('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'T_0')),
  inference(unit_resulting_resolution,[],[f1451,f919,f7614,f1369])).
fof(f1369,plain,(
  ( ! [X2,X0,X1] : (~sP8(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f897])).
fof(f25102,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22826,f1451])).
fof(f22826,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7866,f17173,f1364])).
fof(f7866,plain,(
  ~sP8('F_0','T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),
  inference(unit_resulting_resolution,[],[f1451,f919,f7614,f1369])).
fof(f25101,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22827,f17173])).
fof(f22827,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6187,f17173,f1364])).
fof(f6187,plain,(
  ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f919,f919,f5976,f1369])).
fof(f25100,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22828,f17173])).
fof(f22828,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7001,f17173,f1364])).
fof(f7001,plain,(
  ~sP8('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f919,f919,f5977,f1369])).
fof(f25099,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22829,f17173])).
fof(f22829,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7781,f17173,f1364])).
fof(f7781,plain,(
  ( ! [X0] : (~sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1451,f7614,f1354])).
fof(f25098,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22830,f17173])).
fof(f22830,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7782,f17173,f1364])).
fof(f7782,plain,(
  ( ! [X0] : (~sP8('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1451,f7614,f1355])).
fof(f25097,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22831,f17173])).
fof(f22831,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2109,f17173,f1364])).
fof(f2109,plain,(
  ( ! [X0] : (~sP8(sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1451,f1901,f1354])).
fof(f25096,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22832,f17173])).
fof(f22832,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2110,f17173,f1364])).
fof(f2110,plain,(
  ( ! [X0] : (~sP8(sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1451,f1901,f1355])).
fof(f25095,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22833,f17173])).
fof(f22833,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2233,f17173,f1364])).
fof(f2233,plain,(
  ( ! [X0] : (~sP8(sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1451,f1930,f1354])).
fof(f25094,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22834,f17173])).
fof(f22834,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2234,f17173,f1364])).
fof(f2234,plain,(
  ( ! [X0] : (~sP8(sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1451,f1930,f1355])).
fof(f25093,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22839,f1451])).
fof(f22839,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1936,f17173,f1364])).
fof(f1936,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f919,f1451,f1354])).
fof(f25092,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22840,f17173])).
fof(f22840,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6160,f17173,f1364])).
fof(f6160,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(unit_resulting_resolution,[],[f919,f5976,f1354])).
fof(f25091,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22841,f17173])).
fof(f22841,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6974,f17173,f1364])).
fof(f6974,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(unit_resulting_resolution,[],[f919,f5977,f1354])).
fof(f25090,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22842,f17173])).
fof(f22842,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1939,f17173,f1364])).
fof(f1939,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f919,f1451,f1355])).
fof(f25089,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22843,f17173])).
fof(f22843,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6163,f17173,f1364])).
fof(f6163,plain,(
  ( ! [X0] : (~sP8('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f919,f5976,f1355])).
fof(f25088,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22844,f17173])).
fof(f22844,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6977,f17173,f1364])).
fof(f6977,plain,(
  ( ! [X0] : (~sP8('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f919,f5977,f1355])).
fof(f25087,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22847,f17173])).
fof(f22847,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7866,f17173,f1364])).
fof(f25086,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22850,f1451])).
fof(f22850,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7781,f17173,f1364])).
fof(f25085,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22851,f17173])).
fof(f22851,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7782,f17173,f1364])).
fof(f25084,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22852,f1451])).
fof(f22852,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2109,f17173,f1364])).
fof(f25083,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22853,f17173])).
fof(f22853,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2110,f17173,f1364])).
fof(f25082,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22854,f1451])).
fof(f22854,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2233,f17173,f1364])).
fof(f25081,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22855,f17173])).
fof(f22855,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2234,f17173,f1364])).
fof(f25080,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22866,f1451])).
fof(f22866,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f4642,f17173,f1365])).
fof(f1365,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f897])).
fof(f25079,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22867,f1451])).
fof(f22867,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7862,f17173,f1365])).
fof(f25078,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22868,f1451])).
fof(f22868,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7866,f17173,f1365])).
fof(f25077,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22869,f17173])).
fof(f22869,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6187,f17173,f1365])).
fof(f25076,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22870,f17173])).
fof(f22870,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7001,f17173,f1365])).
fof(f25075,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22871,f17173])).
fof(f22871,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7781,f17173,f1365])).
fof(f25074,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22872,f17173])).
fof(f22872,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7782,f17173,f1365])).
fof(f25073,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22873,f17173])).
fof(f22873,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2109,f17173,f1365])).
fof(f25072,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22874,f17173])).
fof(f22874,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2110,f17173,f1365])).
fof(f25071,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22875,f17173])).
fof(f22875,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2233,f17173,f1365])).
fof(f25070,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22876,f17173])).
fof(f22876,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2234,f17173,f1365])).
fof(f25069,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22881,f17173])).
fof(f22881,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1936,f17173,f1365])).
fof(f25068,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22882,f17173])).
fof(f22882,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6160,f17173,f1365])).
fof(f25067,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22883,f17173])).
fof(f22883,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6974,f17173,f1365])).
fof(f25066,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22884,f1451])).
fof(f22884,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1939,f17173,f1365])).
fof(f25065,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22885,f17173])).
fof(f22885,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6163,f17173,f1365])).
fof(f25064,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22886,f17173])).
fof(f22886,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6977,f17173,f1365])).
fof(f25063,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22888,f17173])).
fof(f22888,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7862,f17173,f1365])).
fof(f25062,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22892,f17173])).
fof(f22892,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7781,f17173,f1365])).
fof(f25061,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22893,f1451])).
fof(f22893,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7782,f17173,f1365])).
fof(f25060,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22894,f17173])).
fof(f22894,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2109,f17173,f1365])).
fof(f25059,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22895,f1451])).
fof(f22895,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2110,f17173,f1365])).
fof(f25058,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22896,f17173])).
fof(f22896,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2233,f17173,f1365])).
fof(f25057,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22897,f1451])).
fof(f22897,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2234,f17173,f1365])).
fof(f25056,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f22968,f23796])).
fof(f23796,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f17173,f17173,f1381])).
fof(f1381,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f905])).
fof(f905,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f904])).
fof(f904,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f722])).
fof(f722,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP10(X2,X0,X1))),
  inference(definition_folding,[],[f565,f721,f720])).
fof(f720,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f721,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f565,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f564])).
fof(f564,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X33,X43,X44] : ((p(s(bool,X44)) <=> (p(s(bool,X43)) | p(s(bool,X33)))) <=> ((p(s(bool,X44)) | ~p(s(bool,X43))) & (p(s(bool,X44)) | ~p(s(bool,X33))) & (p(s(bool,X43)) | p(s(bool,X33)) | ~p(s(bool,X44)))))),
  file('thm.extreal.add_not_infty|split|0','thm.sat.sat_25_prove_ltb')).
fof(f22968,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1,'T_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f919,f17173,f1370])).
fof(f1370,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f900])).
fof(f900,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP9(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f899])).
fof(f899,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f898])).
fof(f898,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1)) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f721])).
fof(f25055,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23004,f23796])).
fof(f23004,plain,(
  ( ! [X0,X1] : (~sP10(X0,'T_0',X1)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f919,f17173,f1371])).
fof(f1371,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f900])).
fof(f25054,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23180,f17173])).
fof(f23180,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f2112,f17173,f1376])).
fof(f1376,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f900])).
fof(f2112,plain,(
  ( ! [X0] : (~sP10('F_0',X0,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1451,f1901,f1370])).
fof(f1511,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1451,f1380])).
fof(f1380,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f903])).
fof(f903,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f902])).
fof(f902,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f901])).
fof(f901,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f720])).
fof(f25053,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23181,f17173])).
fof(f23181,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f2236,f17173,f1376])).
fof(f2236,plain,(
  ( ! [X0] : (~sP10('F_0',X0,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1451,f1930,f1370])).
fof(f25052,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23182,f17173])).
fof(f23182,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f7870,f17173,f1376])).
fof(f7870,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) )),
  inference(unit_resulting_resolution,[],[f1451,f7614,f1370])).
fof(f25051,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23183,f17173])).
fof(f23183,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f1957,f17173,f1376])).
fof(f1957,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f919,f1451,f1371])).
fof(f25050,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23184,f17173])).
fof(f23184,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1509,f1957,f17173,f1376])).
fof(f1509,plain,(
  ( ! [X0,X1] : (sP9(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f919,f1379])).
fof(f1379,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f903])).
fof(f25049,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23186,f17173])).
fof(f23186,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1546,f1957,f17173,f1376])).
fof(f1546,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal',X2)))) )),
  inference(unit_resulting_resolution,[],[f947,f1378])).
fof(f1378,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f903])).
fof(f25048,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23187,f17173])).
fof(f23187,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1533,f1957,f17173,f1376])).
fof(f1533,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f926,f1378])).
fof(f25047,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23188,f17173])).
fof(f23188,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7925,f1957,f17173,f1376])).
fof(f7925,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) )),
  inference(unit_resulting_resolution,[],[f7614,f1378])).
fof(f25046,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23189,f17173])).
fof(f23189,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1508,f1957,f17173,f1376])).
fof(f1508,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f921,f1378])).
fof(f25045,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23190,f17173])).
fof(f23190,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1890,f1957,f17173,f1376])).
fof(f1890,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,sK73(sK59('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f1871,f1378])).
fof(f25044,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23191,f17173])).
fof(f23191,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1919,f1957,f17173,f1376])).
fof(f1919,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) )),
  inference(unit_resulting_resolution,[],[f1899,f1378])).
fof(f25043,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23192,f17173])).
fof(f23192,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2118,f1957,f17173,f1376])).
fof(f2118,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1901,f1378])).
fof(f25042,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23193,f17173])).
fof(f23193,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2242,f1957,f17173,f1376])).
fof(f2242,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1930,f1378])).
fof(f25041,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23194,f17173])).
fof(f23194,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f7882,f17173,f1376])).
fof(f7882,plain,(
  ( ! [X0] : (~sP10('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f1451,f7614,f1371])).
fof(f25040,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23195,f17173])).
fof(f23195,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7926,f7882,f17173,f1376])).
fof(f7926,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f7614,f1379])).
fof(f25039,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23197,f17173])).
fof(f23197,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1546,f7882,f17173,f1376])).
fof(f25038,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23198,f17173])).
fof(f23198,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1533,f7882,f17173,f1376])).
fof(f25037,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23199,f17173])).
fof(f23199,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7925,f7882,f17173,f1376])).
fof(f25036,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23200,f17173])).
fof(f23200,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1508,f7882,f17173,f1376])).
fof(f25035,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23201,f17173])).
fof(f23201,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1890,f7882,f17173,f1376])).
fof(f25034,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23202,f17173])).
fof(f23202,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1919,f7882,f17173,f1376])).
fof(f25033,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23203,f17173])).
fof(f23203,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2118,f7882,f17173,f1376])).
fof(f25032,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23204,f17173])).
fof(f23204,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2242,f7882,f17173,f1376])).
fof(f25031,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23205,f17173])).
fof(f23205,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f2115,f17173,f1376])).
fof(f2115,plain,(
  ( ! [X0] : (~sP10('F_0',sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f1451,f1901,f1371])).
fof(f25030,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23206,f17173])).
fof(f23206,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2119,f2115,f17173,f1376])).
fof(f2119,plain,(
  ( ! [X0,X1] : (sP9(X0,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),X1)) )),
  inference(unit_resulting_resolution,[],[f1901,f1379])).
fof(f25029,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23208,f17173])).
fof(f23208,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1546,f2115,f17173,f1376])).
fof(f25028,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23209,f17173])).
fof(f23209,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1533,f2115,f17173,f1376])).
fof(f25027,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23210,f17173])).
fof(f23210,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7925,f2115,f17173,f1376])).
fof(f25026,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23211,f17173])).
fof(f23211,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1508,f2115,f17173,f1376])).
fof(f25025,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23212,f17173])).
fof(f23212,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1890,f2115,f17173,f1376])).
fof(f25024,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23213,f17173])).
fof(f23213,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1919,f2115,f17173,f1376])).
fof(f25023,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23214,f17173])).
fof(f23214,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2118,f2115,f17173,f1376])).
fof(f25022,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23215,f17173])).
fof(f23215,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2242,f2115,f17173,f1376])).
fof(f25021,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23216,f17173])).
fof(f23216,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f2239,f17173,f1376])).
fof(f2239,plain,(
  ( ! [X0] : (~sP10('F_0',sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f1451,f1930,f1371])).
fof(f25020,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23217,f17173])).
fof(f23217,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2243,f2239,f17173,f1376])).
fof(f2243,plain,(
  ( ! [X0,X1] : (sP9(X0,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),X1)) )),
  inference(unit_resulting_resolution,[],[f1930,f1379])).
fof(f25019,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23219,f17173])).
fof(f23219,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1546,f2239,f17173,f1376])).
fof(f25018,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23220,f17173])).
fof(f23220,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1533,f2239,f17173,f1376])).
fof(f25017,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23221,f17173])).
fof(f23221,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7925,f2239,f17173,f1376])).
fof(f25016,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23222,f17173])).
fof(f23222,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1508,f2239,f17173,f1376])).
fof(f25015,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23223,f17173])).
fof(f23223,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1890,f2239,f17173,f1376])).
fof(f25014,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23224,f17173])).
fof(f23224,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1919,f2239,f17173,f1376])).
fof(f25013,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23225,f17173])).
fof(f23225,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2118,f2239,f17173,f1376])).
fof(f25012,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23226,f17173])).
fof(f23226,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2242,f2239,f17173,f1376])).
fof(f25011,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23228,f17173])).
fof(f23228,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6256,f6189,f17173,f1376])).
fof(f6189,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f919,f5976,f1371])).
fof(f6256,plain,(
  ( ! [X0,X1] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f5976,f1380])).
fof(f25010,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23229,f17173])).
fof(f23229,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1509,f6189,f17173,f1376])).
fof(f25009,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23231,f17173])).
fof(f23231,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1546,f6189,f17173,f1376])).
fof(f25008,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23232,f17173])).
fof(f23232,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1533,f6189,f17173,f1376])).
fof(f25007,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23233,f17173])).
fof(f23233,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7925,f6189,f17173,f1376])).
fof(f25006,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23234,f17173])).
fof(f23234,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1508,f6189,f17173,f1376])).
fof(f25005,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23235,f17173])).
fof(f23235,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1890,f6189,f17173,f1376])).
fof(f25004,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23236,f17173])).
fof(f23236,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1919,f6189,f17173,f1376])).
fof(f25003,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23237,f17173])).
fof(f23237,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2118,f6189,f17173,f1376])).
fof(f25002,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23238,f17173])).
fof(f23238,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2242,f6189,f17173,f1376])).
fof(f25001,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23240,f17173])).
fof(f23240,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7070,f7003,f17173,f1376])).
fof(f7003,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f919,f5977,f1371])).
fof(f7070,plain,(
  ( ! [X0,X1] : (sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f5977,f1380])).
fof(f25000,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23241,f17173])).
fof(f23241,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1509,f7003,f17173,f1376])).
fof(f24999,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23243,f17173])).
fof(f23243,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1546,f7003,f17173,f1376])).
fof(f24998,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23244,f17173])).
fof(f23244,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1533,f7003,f17173,f1376])).
fof(f24997,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23245,f17173])).
fof(f23245,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7925,f7003,f17173,f1376])).
fof(f24996,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23246,f17173])).
fof(f23246,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1508,f7003,f17173,f1376])).
fof(f24995,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23247,f17173])).
fof(f23247,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1890,f7003,f17173,f1376])).
fof(f24994,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23248,f17173])).
fof(f23248,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1919,f7003,f17173,f1376])).
fof(f24993,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23249,f17173])).
fof(f23249,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2118,f7003,f17173,f1376])).
fof(f24992,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23250,f17173])).
fof(f23250,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2242,f7003,f17173,f1376])).
fof(f24991,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23266,f23608])).
fof(f23608,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1380])).
fof(f23266,plain,(
  ~sP9('F_0','T_0','F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f1451,f17173,f1376])).
fof(f24990,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23268,f23608])).
fof(f23268,plain,(
  ~sP9('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f1451,f17173,f1376])).
fof(f24989,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23270,f23608])).
fof(f23270,plain,(
  ~sP9('F_0',sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f1451,f17173,f1376])).
fof(f24988,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23272,f23608])).
fof(f23272,plain,(
  ~sP9('F_0',sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f1451,f17173,f1376])).
fof(f24987,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23274,f23608])).
fof(f23274,plain,(
  ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0','F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f6189,f1451,f17173,f1376])).
fof(f24986,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23276,f23608])).
fof(f23276,plain,(
  ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0','F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7003,f1451,f17173,f1376])).
fof(f24985,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23305,f23608])).
fof(f23305,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f17173,f17173,f1376])).
fof(f1956,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f919,f1451,f1370])).
fof(f24984,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23307,f23608])).
fof(f23307,plain,(
  ( ! [X0] : (~sP9('F_0',X0,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f17173,f17173,f1376])).
fof(f24983,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23309,f23608])).
fof(f23309,plain,(
  ( ! [X0] : (~sP9('F_0',X0,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f17173,f17173,f1376])).
fof(f24982,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23311,f23608])).
fof(f23311,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f17173,f17173,f1376])).
fof(f24981,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23313,f23608])).
fof(f23313,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f17173,f17173,f1376])).
fof(f24980,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23315,f23608])).
fof(f23315,plain,(
  ( ! [X0] : (~sP9('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f17173,f17173,f1376])).
fof(f24979,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23317,f23608])).
fof(f23317,plain,(
  ( ! [X0] : (~sP9('F_0',sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f17173,f17173,f1376])).
fof(f24978,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23319,f23608])).
fof(f23319,plain,(
  ( ! [X0] : (~sP9('F_0',sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f17173,f17173,f1376])).
fof(f24977,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23321,f23608])).
fof(f23321,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,'T_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6188,f17173,f17173,f1376])).
fof(f6188,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f919,f5976,f1370])).
fof(f24976,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23323,f23608])).
fof(f23323,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6189,f17173,f17173,f1376])).
fof(f24975,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23325,f23608])).
fof(f23325,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,'T_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7002,f17173,f17173,f1376])).
fof(f7002,plain,(
  ( ! [X0] : (~sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f919,f5977,f1370])).
fof(f24974,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23327,f23608])).
fof(f23327,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7003,f17173,f17173,f1376])).
fof(f24973,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23330,f17173])).
fof(f23330,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1507,f17173,f1376])).
fof(f1507,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f919,f1378])).
fof(f24972,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23332,f17173])).
fof(f23332,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f1507,f17173,f1376])).
fof(f24971,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23333,f17173])).
fof(f23333,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f1507,f17173,f1376])).
fof(f24970,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23334,f17173])).
fof(f23334,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f1507,f17173,f1376])).
fof(f24969,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23335,f17173])).
fof(f23335,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6188,f1507,f17173,f1376])).
fof(f24968,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23337,f17173])).
fof(f23337,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7002,f1507,f17173,f1376])).
fof(f24967,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23340,f17173])).
fof(f23340,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f1546,f17173,f1376])).
fof(f24966,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23341,f17173])).
fof(f23341,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f1546,f17173,f1376])).
fof(f24965,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23342,f17173])).
fof(f23342,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f1546,f17173,f1376])).
fof(f24964,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23346,f17173])).
fof(f23346,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f1533,f17173,f1376])).
fof(f24963,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23347,f17173])).
fof(f23347,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f1533,f17173,f1376])).
fof(f24962,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23348,f17173])).
fof(f23348,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f1533,f17173,f1376])).
fof(f24961,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23351,f17173])).
fof(f23351,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f7925,f17173,f1376])).
fof(f24960,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23353,f17173])).
fof(f23353,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f7925,f17173,f1376])).
fof(f24959,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23354,f17173])).
fof(f23354,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f7925,f17173,f1376])).
fof(f24958,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23355,f17173])).
fof(f23355,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f7925,f17173,f1376])).
fof(f24957,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23359,f17173])).
fof(f23359,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f1508,f17173,f1376])).
fof(f24956,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23360,f17173])).
fof(f23360,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f1508,f17173,f1376])).
fof(f24955,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23361,f17173])).
fof(f23361,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f1508,f17173,f1376])).
fof(f24954,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23365,f17173])).
fof(f23365,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f1890,f17173,f1376])).
fof(f24953,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23366,f17173])).
fof(f23366,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f1890,f17173,f1376])).
fof(f24952,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23367,f17173])).
fof(f23367,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f1890,f17173,f1376])).
fof(f24951,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23371,f17173])).
fof(f23371,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f1919,f17173,f1376])).
fof(f24950,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23372,f17173])).
fof(f23372,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f1919,f17173,f1376])).
fof(f24949,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23373,f17173])).
fof(f23373,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f1919,f17173,f1376])).
fof(f24948,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23376,f17173])).
fof(f23376,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f2118,f17173,f1376])).
fof(f24947,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23378,f17173])).
fof(f23378,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f2118,f17173,f1376])).
fof(f24946,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23379,f17173])).
fof(f23379,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f2118,f17173,f1376])).
fof(f24945,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23380,f17173])).
fof(f23380,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f2118,f17173,f1376])).
fof(f24944,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23383,f17173])).
fof(f23383,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f2242,f17173,f1376])).
fof(f24943,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23385,f17173])).
fof(f23385,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f2242,f17173,f1376])).
fof(f24942,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23386,f17173])).
fof(f23386,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f2242,f17173,f1376])).
fof(f24941,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23387,f17173])).
fof(f23387,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f2242,f17173,f1376])).
fof(f24940,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23391,f17173])).
fof(f23391,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1534,f17173,f1376])).
fof(f1534,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f926,f1379])).
fof(f24939,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23392,f17173])).
fof(f23392,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f7926,f17173,f1376])).
fof(f24938,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23393,f17173])).
fof(f23393,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1510,f17173,f1376])).
fof(f1510,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f921,f1379])).
fof(f24937,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23394,f17173])).
fof(f23394,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1891,f17173,f1376])).
fof(f1891,plain,(
  ( ! [X0,X1] : (sP9(X0,sK73(sK59('const.num.|48|_0','type.num.num')),X1)) )),
  inference(unit_resulting_resolution,[],[f1871,f1379])).
fof(f24936,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23395,f17173])).
fof(f23395,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1920,f17173,f1376])).
fof(f1920,plain,(
  ( ! [X0,X1] : (sP9(X0,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),X1)) )),
  inference(unit_resulting_resolution,[],[f1899,f1379])).
fof(f24935,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23396,f17173])).
fof(f23396,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f2119,f17173,f1376])).
fof(f24934,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23397,f17173])).
fof(f23397,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f2243,f17173,f1376])).
fof(f24933,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23398,f17173])).
fof(f23398,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1511,f17173,f1376])).
fof(f24932,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23399,f17173])).
fof(f23399,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f1511,f17173,f1376])).
fof(f24931,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23400,f17173])).
fof(f23400,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f1511,f17173,f1376])).
fof(f24930,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23401,f17173])).
fof(f23401,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f1511,f17173,f1376])).
fof(f24929,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23403,f17173])).
fof(f23403,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f1511,f17173,f1376])).
fof(f24928,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23404,f17173])).
fof(f23404,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f1511,f17173,f1376])).
fof(f24927,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23405,f17173])).
fof(f23405,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f1511,f17173,f1376])).
fof(f24926,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23406,f17173])).
fof(f23406,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6188,f6256,f17173,f1376])).
fof(f24925,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23408,f17173])).
fof(f23408,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7002,f7070,f17173,f1376])).
fof(f24924,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23413,f23608])).
fof(f23413,plain,(
  ~sP9('F_0','F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1451,f17173,f1376])).
fof(f24923,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23414,f23608])).
fof(f23414,plain,(
  ~sP9('F_0','F_0',sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal'))) | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f1451,f17173,f1376])).
fof(f24922,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23415,f23608])).
fof(f23415,plain,(
  ~sP9('F_0','F_0',sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal'))) | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f1451,f17173,f1376])).
fof(f24921,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23416,f23608])).
fof(f23416,plain,(
  ~sP9('F_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))) | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f1451,f17173,f1376])).
fof(f24920,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23417,f23608])).
fof(f23417,plain,(
  ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f6188,f1451,f17173,f1376])).
fof(f24919,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23418,f23608])).
fof(f23418,plain,(
  ~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'F_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7002,f1451,f17173,f1376])).
fof(f24918,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23445,f23608])).
fof(f23445,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f17173,f17173,f1376])).
fof(f24917,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23446,f23608])).
fof(f23446,plain,(
  ( ! [X0] : (~sP9('F_0',X0,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f17173,f17173,f1376])).
fof(f24916,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23447,f23608])).
fof(f23447,plain,(
  ( ! [X0] : (~sP9('F_0',X0,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f17173,f17173,f1376])).
fof(f24915,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23448,f23608])).
fof(f23448,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f17173,f17173,f1376])).
fof(f24914,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23449,f23608])).
fof(f23449,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f17173,f17173,f1376])).
fof(f24913,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23450,f23608])).
fof(f23450,plain,(
  ( ! [X0] : (~sP9('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f17173,f17173,f1376])).
fof(f24912,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23451,f23608])).
fof(f23451,plain,(
  ( ! [X0] : (~sP9('F_0',sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f17173,f17173,f1376])).
fof(f24911,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23452,f23608])).
fof(f23452,plain,(
  ( ! [X0] : (~sP9('F_0',sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f17173,f17173,f1376])).
fof(f24910,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23453,f23608])).
fof(f23453,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,'T_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6188,f17173,f17173,f1376])).
fof(f24909,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23454,f23608])).
fof(f23454,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6189,f17173,f17173,f1376])).
fof(f24908,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23455,f23608])).
fof(f23455,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,'T_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7002,f17173,f17173,f1376])).
fof(f24907,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23456,f23608])).
fof(f23456,plain,(
  ( ! [X0] : (~sP9('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7003,f17173,f17173,f1376])).
fof(f24906,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23500,f17173])).
fof(f23500,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1546,f919,f17173,f1377])).
fof(f1377,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f903])).
fof(f24905,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23501,f17173])).
fof(f23501,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1533,f919,f17173,f1377])).
fof(f24904,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23502,f17173])).
fof(f23502,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7925,f919,f17173,f1377])).
fof(f24903,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23503,f17173])).
fof(f23503,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1508,f919,f17173,f1377])).
fof(f24902,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23504,f17173])).
fof(f23504,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1890,f919,f17173,f1377])).
fof(f24901,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23505,f17173])).
fof(f23505,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1919,f919,f17173,f1377])).
fof(f24900,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23506,f17173])).
fof(f23506,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2118,f919,f17173,f1377])).
fof(f24899,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23507,f17173])).
fof(f23507,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2242,f919,f17173,f1377])).
fof(f24898,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23508,f17173])).
fof(f23508,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1509,f919,f17173,f1377])).
fof(f24897,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23509,f17173])).
fof(f23509,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1534,f919,f17173,f1377])).
fof(f24896,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23510,f17173])).
fof(f23510,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7926,f919,f17173,f1377])).
fof(f24895,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23511,f17173])).
fof(f23511,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1510,f919,f17173,f1377])).
fof(f24894,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23512,f17173])).
fof(f23512,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1891,f919,f17173,f1377])).
fof(f24893,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23513,f17173])).
fof(f23513,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1920,f919,f17173,f1377])).
fof(f24892,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23514,f17173])).
fof(f23514,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2119,f919,f17173,f1377])).
fof(f24891,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23515,f17173])).
fof(f23515,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2243,f919,f17173,f1377])).
fof(f24890,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23528,f23608])).
fof(f23528,plain,(
  ( ! [X0] : (~sP9('T_0',X0,'F_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f919,f1451,f17173,f1377])).
fof(f24889,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23551,f23608])).
fof(f23551,plain,(
  ( ! [X0,X1] : (~sP9('T_0',X0,X1)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f919,f17173,f17173,f1377])).
fof(f24888,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23553,f17173])).
fof(f23553,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1507,f919,f17173,f1377])).
fof(f24887,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23554,f17173])).
fof(f23554,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1546,f919,f17173,f1377])).
fof(f24886,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23555,f17173])).
fof(f23555,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1533,f919,f17173,f1377])).
fof(f24885,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23556,f17173])).
fof(f23556,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7925,f919,f17173,f1377])).
fof(f24884,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23557,f17173])).
fof(f23557,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1508,f919,f17173,f1377])).
fof(f24883,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23558,f17173])).
fof(f23558,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1890,f919,f17173,f1377])).
fof(f24882,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23559,f17173])).
fof(f23559,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1919,f919,f17173,f1377])).
fof(f24881,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23560,f17173])).
fof(f23560,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2118,f919,f17173,f1377])).
fof(f24880,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23561,f17173])).
fof(f23561,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2242,f919,f17173,f1377])).
fof(f24879,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23563,f17173])).
fof(f23563,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1534,f919,f17173,f1377])).
fof(f24878,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23564,f17173])).
fof(f23564,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7926,f919,f17173,f1377])).
fof(f24877,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23565,f17173])).
fof(f23565,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1510,f919,f17173,f1377])).
fof(f24876,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23566,f17173])).
fof(f23566,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1891,f919,f17173,f1377])).
fof(f24875,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23567,f17173])).
fof(f23567,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1920,f919,f17173,f1377])).
fof(f24874,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23568,f17173])).
fof(f23568,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2119,f919,f17173,f1377])).
fof(f24873,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23569,f17173])).
fof(f23569,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2243,f919,f17173,f1377])).
fof(f24872,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23583,f23608])).
fof(f23583,plain,(
  ( ! [X0] : (~sP9('T_0','F_0',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f919,f1451,f17173,f1377])).
fof(f24871,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23606,f23608])).
fof(f23606,plain,(
  ( ! [X0,X1] : (~sP9('T_0',X0,X1)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f919,f17173,f17173,f1377])).
fof(f24870,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23611,f17173])).
fof(f23611,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f1451,f17173,f1381])).
fof(f24869,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23612,f17173])).
fof(f23612,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f1451,f17173,f1381])).
fof(f24868,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23613,f17173])).
fof(f23613,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f1451,f17173,f1381])).
fof(f24867,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23614,f17173])).
fof(f23614,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f1451,f17173,f1381])).
fof(f24866,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23615,f17173])).
fof(f23615,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f1451,f17173,f1381])).
fof(f24865,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23616,f17173])).
fof(f23616,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f1451,f17173,f1381])).
fof(f24864,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23617,f17173])).
fof(f23617,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f1451,f17173,f1381])).
fof(f24863,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23618,f1451])).
fof(f23618,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f4751,f17173,f17173,f1381])).
fof(f4751,plain,(
  ~sP10('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f4731,f1372])).
fof(f1372,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f900])).
fof(f4731,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f919,f1451,f1451,f1377])).
fof(f24862,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23619,f1451])).
fof(f23619,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f6275,f17173,f17173,f1381])).
fof(f6275,plain,(
  ~sP10('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f1451,f919,f5976,f1384])).
fof(f1384,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f905])).
fof(f24861,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23620,f1451])).
fof(f23620,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7089,f17173,f17173,f1381])).
fof(f7089,plain,(
  ~sP10('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f1451,f919,f5977,f1384])).
fof(f24860,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23621,f17173])).
fof(f23621,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6283,f17173,f17173,f1381])).
fof(f6283,plain,(
  ~sP10('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(unit_resulting_resolution,[],[f919,f1451,f5976,f1384])).
fof(f24859,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23622,f17173])).
fof(f23622,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7097,f17173,f17173,f1381])).
fof(f7097,plain,(
  ~sP10('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(unit_resulting_resolution,[],[f919,f1451,f5977,f1384])).
fof(f24858,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23624,f17173])).
fof(f23624,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f17173,f17173,f1381])).
fof(f24857,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23625,f17173])).
fof(f23625,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f17173,f17173,f1381])).
fof(f24856,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23626,f17173])).
fof(f23626,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f17173,f17173,f1381])).
fof(f24855,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23627,f17173])).
fof(f23627,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f17173,f17173,f1381])).
fof(f24854,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23628,f17173])).
fof(f23628,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f17173,f17173,f1381])).
fof(f24853,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23629,f17173])).
fof(f23629,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f17173,f17173,f1381])).
fof(f24852,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23630,f17173])).
fof(f23630,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f17173,f17173,f1381])).
fof(f24851,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23632,f17173])).
fof(f23632,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6189,f17173,f17173,f1381])).
fof(f24850,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23634,f17173])).
fof(f23634,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7003,f17173,f17173,f1381])).
fof(f24849,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23635,f1451])).
fof(f23635,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7941,f17173,f17173,f1381])).
fof(f7941,plain,(
  ~sP10('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1451,f1451,f7614,f1384])).
fof(f24848,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23639,f1451])).
fof(f23639,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f1451,f17173,f1381])).
fof(f24847,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23640,f1451])).
fof(f23640,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f1451,f17173,f1381])).
fof(f24846,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23641,f1451])).
fof(f23641,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f1451,f17173,f1381])).
fof(f24845,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23642,f1451])).
fof(f23642,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f1451,f17173,f1381])).
fof(f24844,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23643,f17173])).
fof(f23643,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6189,f1451,f17173,f1381])).
fof(f24843,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23644,f17173])).
fof(f23644,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7003,f1451,f17173,f1381])).
fof(f24842,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23645,f17173])).
fof(f23645,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7941,f1451,f17173,f1381])).
fof(f24841,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23655,f1451])).
fof(f23655,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f17173,f17173,f1381])).
fof(f24840,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23656,f1451])).
fof(f23656,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f17173,f17173,f1381])).
fof(f24839,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23657,f1451])).
fof(f23657,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f17173,f17173,f1381])).
fof(f24838,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23658,f1451])).
fof(f23658,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f17173,f17173,f1381])).
fof(f24837,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23659,f1451])).
fof(f23659,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f17173,f17173,f1381])).
fof(f24836,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23660,f1451])).
fof(f23660,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f17173,f17173,f1381])).
fof(f24835,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23661,f1451])).
fof(f23661,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f17173,f17173,f1381])).
fof(f24834,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23662,f1451])).
fof(f23662,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f17173,f17173,f1381])).
fof(f24833,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23663,f17173])).
fof(f23663,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6188,f17173,f17173,f1381])).
fof(f24832,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23664,f17173])).
fof(f23664,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6189,f17173,f17173,f1381])).
fof(f24831,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23665,f17173])).
fof(f23665,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7002,f17173,f17173,f1381])).
fof(f24830,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23666,f17173])).
fof(f23666,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7003,f17173,f17173,f1381])).
fof(f24829,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23667,f17173])).
fof(f23667,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7941,f17173,f17173,f1381])).
fof(f24828,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23672,f17173])).
fof(f23672,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1451,f17173,f1381])).
fof(f24827,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23673,f17173])).
fof(f23673,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f1451,f17173,f1381])).
fof(f24826,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23674,f17173])).
fof(f23674,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f1451,f17173,f1381])).
fof(f24825,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23675,f17173])).
fof(f23675,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f1451,f17173,f1381])).
fof(f24824,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23677,f17173])).
fof(f23677,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f1451,f17173,f1381])).
fof(f24823,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23678,f17173])).
fof(f23678,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f1451,f17173,f1381])).
fof(f24822,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23679,f17173])).
fof(f23679,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f1451,f17173,f1381])).
fof(f24821,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23680,f1451])).
fof(f23680,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f4751,f17173,f17173,f1381])).
fof(f24820,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23681,f17173])).
fof(f23681,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6275,f17173,f17173,f1381])).
fof(f24819,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23682,f17173])).
fof(f23682,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7089,f17173,f17173,f1381])).
fof(f24818,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23683,f1451])).
fof(f23683,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f6283,f17173,f17173,f1381])).
fof(f24817,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23684,f1451])).
fof(f23684,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7097,f17173,f17173,f1381])).
fof(f24816,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23685,f17173])).
fof(f23685,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f17173,f17173,f1381])).
fof(f24815,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23686,f17173])).
fof(f23686,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f17173,f17173,f1381])).
fof(f24814,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23687,f17173])).
fof(f23687,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f17173,f17173,f1381])).
fof(f24813,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23688,f17173])).
fof(f23688,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f17173,f17173,f1381])).
fof(f24812,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23690,f17173])).
fof(f23690,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f17173,f17173,f1381])).
fof(f24811,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23691,f17173])).
fof(f23691,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f17173,f17173,f1381])).
fof(f24810,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23692,f17173])).
fof(f23692,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f17173,f17173,f1381])).
fof(f24809,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23693,f17173])).
fof(f23693,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6188,f17173,f17173,f1381])).
fof(f24808,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23695,f17173])).
fof(f23695,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7002,f17173,f17173,f1381])).
fof(f24807,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23697,f1451])).
fof(f23697,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7941,f17173,f17173,f1381])).
fof(f24806,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23701,f1451])).
fof(f23701,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1451,f17173,f1381])).
fof(f24805,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23702,f1451])).
fof(f23702,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f1451,f17173,f1381])).
fof(f24804,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23703,f1451])).
fof(f23703,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f1451,f17173,f1381])).
fof(f24803,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23704,f1451])).
fof(f23704,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f1451,f17173,f1381])).
fof(f24802,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23705,f17173])).
fof(f23705,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6188,f1451,f17173,f1381])).
fof(f24801,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23706,f17173])).
fof(f23706,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7002,f1451,f17173,f1381])).
fof(f24800,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23707,f17173])).
fof(f23707,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7941,f1451,f17173,f1381])).
fof(f24799,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23717,f1451])).
fof(f23717,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f17173,f17173,f1381])).
fof(f24798,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23718,f1451])).
fof(f23718,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f17173,f17173,f1381])).
fof(f24797,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23719,f1451])).
fof(f23719,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f17173,f17173,f1381])).
fof(f24796,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23720,f1451])).
fof(f23720,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f17173,f17173,f1381])).
fof(f24795,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23721,f1451])).
fof(f23721,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f17173,f17173,f1381])).
fof(f24794,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23722,f1451])).
fof(f23722,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f17173,f17173,f1381])).
fof(f24793,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23723,f1451])).
fof(f23723,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f17173,f17173,f1381])).
fof(f24792,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23724,f1451])).
fof(f23724,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f17173,f17173,f1381])).
fof(f24791,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23725,f17173])).
fof(f23725,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6188,f17173,f17173,f1381])).
fof(f24790,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23726,f17173])).
fof(f23726,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6189,f17173,f17173,f1381])).
fof(f24789,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23727,f17173])).
fof(f23727,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7002,f17173,f17173,f1381])).
fof(f24788,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23728,f17173])).
fof(f23728,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7003,f17173,f17173,f1381])).
fof(f24787,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23729,f17173])).
fof(f23729,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7941,f17173,f17173,f1381])).
fof(f24786,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23734,f1451])).
fof(f23734,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f4751,f1451,f17173,f1381])).
fof(f24785,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23735,f17173])).
fof(f23735,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6275,f1451,f17173,f1381])).
fof(f24784,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23736,f17173])).
fof(f23736,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7089,f1451,f17173,f1381])).
fof(f24783,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23738,f17173])).
fof(f23738,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f1451,f17173,f1381])).
fof(f24782,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23739,f17173])).
fof(f23739,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f1451,f17173,f1381])).
fof(f24781,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23740,f17173])).
fof(f23740,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f1451,f17173,f1381])).
fof(f24780,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23743,f1451])).
fof(f23743,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7941,f1451,f17173,f1381])).
fof(f24779,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23744,f1451])).
fof(f23744,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f4751,f17173,f17173,f1381])).
fof(f24778,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23745,f17173])).
fof(f23745,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6275,f17173,f17173,f1381])).
fof(f24777,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23746,f17173])).
fof(f23746,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7089,f17173,f17173,f1381])).
fof(f24776,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23747,f1451])).
fof(f23747,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f6283,f17173,f17173,f1381])).
fof(f24775,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23748,f1451])).
fof(f23748,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7097,f17173,f17173,f1381])).
fof(f24774,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23749,f17173])).
fof(f23749,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f17173,f17173,f1381])).
fof(f24773,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23750,f17173])).
fof(f23750,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f17173,f17173,f1381])).
fof(f24772,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23751,f17173])).
fof(f23751,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f17173,f17173,f1381])).
fof(f24771,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23752,f17173])).
fof(f23752,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f17173,f17173,f1381])).
fof(f24770,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23754,f17173])).
fof(f23754,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f17173,f17173,f1381])).
fof(f24769,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23755,f17173])).
fof(f23755,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f17173,f17173,f1381])).
fof(f24768,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23756,f17173])).
fof(f23756,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f17173,f17173,f1381])).
fof(f24767,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23757,f17173])).
fof(f23757,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6188,f17173,f17173,f1381])).
fof(f24766,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23759,f17173])).
fof(f23759,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7002,f17173,f17173,f1381])).
fof(f24765,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23761,f1451])).
fof(f23761,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7941,f17173,f17173,f1381])).
fof(f24764,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23762,f1451])).
fof(f23762,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f4751,f1451,f17173,f1381])).
fof(f24763,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23763,f17173])).
fof(f23763,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6283,f1451,f17173,f1381])).
fof(f24762,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23764,f17173])).
fof(f23764,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7097,f1451,f17173,f1381])).
fof(f24761,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23766,f17173])).
fof(f23766,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f1451,f17173,f1381])).
fof(f24760,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23767,f17173])).
fof(f23767,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f1451,f17173,f1381])).
fof(f24759,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23768,f17173])).
fof(f23768,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f1451,f17173,f1381])).
fof(f24758,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23771,f1451])).
fof(f23771,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7941,f1451,f17173,f1381])).
fof(f24757,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23776,f1451])).
fof(f23776,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f4751,f17173,f17173,f1381])).
fof(f24756,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23777,f1451])).
fof(f23777,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f6275,f17173,f17173,f1381])).
fof(f24755,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23778,f1451])).
fof(f23778,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7089,f17173,f17173,f1381])).
fof(f24754,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23779,f17173])).
fof(f23779,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6283,f17173,f17173,f1381])).
fof(f24753,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23780,f17173])).
fof(f23780,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7097,f17173,f17173,f1381])).
fof(f24752,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23782,f17173])).
fof(f23782,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f17173,f17173,f1381])).
fof(f24751,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23783,f17173])).
fof(f23783,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f17173,f17173,f1381])).
fof(f24750,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23784,f17173])).
fof(f23784,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f17173,f17173,f1381])).
fof(f24749,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23785,f17173])).
fof(f23785,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f17173,f17173,f1381])).
fof(f24748,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23786,f17173])).
fof(f23786,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f17173,f17173,f1381])).
fof(f24747,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23787,f17173])).
fof(f23787,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f17173,f17173,f1381])).
fof(f24746,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23788,f17173])).
fof(f23788,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f17173,f17173,f1381])).
fof(f24745,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23790,f17173])).
fof(f23790,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6189,f17173,f17173,f1381])).
fof(f24744,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23792,f17173])).
fof(f23792,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7003,f17173,f17173,f1381])).
fof(f24743,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23793,f1451])).
fof(f23793,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7941,f17173,f17173,f1381])).
fof(f24742,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23799,f17173])).
fof(f23799,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7929,f919,f17173,f1384])).
fof(f7929,plain,(
  ( ! [X0] : (sP10('T_0',X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) )),
  inference(unit_resulting_resolution,[],[f919,f7614,f1382])).
fof(f1382,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f905])).
fof(f24741,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23800,f17173])).
fof(f23800,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2019,f919,f17173,f1384])).
fof(f2019,plain,(
  ( ! [X0] : (sP10('T_0',X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f1508,f919,f1437])).
fof(f1437,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f1373])).
fof(f1373,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f900])).
fof(f24740,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23801,f17173])).
fof(f23801,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2450,f919,f17173,f1384])).
fof(f2450,plain,(
  ( ! [X0] : (sP10('T_0',X0,sK73(sK59('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f919,f1890,f1437])).
fof(f24739,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23802,f17173])).
fof(f23802,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f2020,f919,f17173,f1384])).
fof(f2020,plain,(
  ( ! [X0] : (sP10('T_0',X0,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) )),
  inference(unit_resulting_resolution,[],[f1919,f919,f1437])).
fof(f24738,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23803,f17173])).
fof(f23803,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2122,f919,f17173,f1384])).
fof(f2122,plain,(
  ( ! [X0] : (sP10('T_0',X0,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f919,f1901,f1382])).
fof(f24737,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23804,f17173])).
fof(f23804,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2246,f919,f17173,f1384])).
fof(f2246,plain,(
  ( ! [X0] : (sP10('T_0',X0,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f919,f1930,f1382])).
fof(f24736,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23805,f17173])).
fof(f23805,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1966,f919,f17173,f1384])).
fof(f1966,plain,(
  ( ! [X0] : (sP10('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f919,f919,f1383])).
fof(f1383,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f905])).
fof(f24735,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23806,f17173])).
fof(f23806,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7934,f919,f17173,f1384])).
fof(f7934,plain,(
  ( ! [X0] : (sP10('T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f919,f7614,f1383])).
fof(f24734,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23807,f17173])).
fof(f23807,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2023,f919,f17173,f1384])).
fof(f2023,plain,(
  ( ! [X0] : (sP10('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f1510,f919,f1437])).
fof(f24733,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23808,f17173])).
fof(f23808,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2469,f919,f17173,f1384])).
fof(f2469,plain,(
  ( ! [X0] : (sP10('T_0',sK73(sK59('const.num.|48|_0','type.num.num')),X0)) )),
  inference(unit_resulting_resolution,[],[f919,f1891,f1437])).
fof(f24732,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23809,f17173])).
fof(f23809,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2024,f919,f17173,f1384])).
fof(f2024,plain,(
  ( ! [X0] : (sP10('T_0',sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),X0)) )),
  inference(unit_resulting_resolution,[],[f1920,f919,f1437])).
fof(f24731,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23810,f17173])).
fof(f23810,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2125,f919,f17173,f1384])).
fof(f2125,plain,(
  ( ! [X0] : (sP10('T_0',sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f919,f1901,f1383])).
fof(f24730,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23811,f17173])).
fof(f23811,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2249,f919,f17173,f1384])).
fof(f2249,plain,(
  ( ! [X0] : (sP10('T_0',sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f919,f1930,f1383])).
fof(f24729,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23922,f17173])).
fof(f23922,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1960,f919,f17173,f1384])).
fof(f1960,plain,(
  ( ! [X0] : (sP10('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f919,f919,f1382])).
fof(f24728,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23923,f17173])).
fof(f23923,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7929,f919,f17173,f1384])).
fof(f24727,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23924,f17173])).
fof(f23924,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2019,f919,f17173,f1384])).
fof(f24726,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23925,f17173])).
fof(f23925,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2450,f919,f17173,f1384])).
fof(f24725,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23926,f17173])).
fof(f23926,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2020,f919,f17173,f1384])).
fof(f24724,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23927,f17173])).
fof(f23927,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2122,f919,f17173,f1384])).
fof(f24723,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23928,f17173])).
fof(f23928,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2246,f919,f17173,f1384])).
fof(f24722,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23930,f17173])).
fof(f23930,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7934,f919,f17173,f1384])).
fof(f24721,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23931,f17173])).
fof(f23931,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2023,f919,f17173,f1384])).
fof(f24720,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23932,f17173])).
fof(f23932,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2469,f919,f17173,f1384])).
fof(f24719,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23933,f17173])).
fof(f23933,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f2024,f919,f17173,f1384])).
fof(f24718,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23934,f17173])).
fof(f23934,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2125,f919,f17173,f1384])).
fof(f24717,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23935,f17173])).
fof(f23935,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2249,f919,f17173,f1384])).
fof(f24716,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23940,f1451])).
fof(f23940,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1974,f17173,f1387])).
fof(f1387,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f908])).
fof(f908,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP11(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f907])).
fof(f907,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1) | sP12(X2,X0,X1)))),
  inference(flattening,[],[f906])).
fof(f906,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1)) | sP12(X2,X0,X1)))),
  inference(nnf_transformation,[],[f724])).
fof(f724,plain,(
  ! [X2,X0,X1] : (sP12(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f1974,plain,(
  ( ! [X0] : (sP12('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f919,f1451,f1399])).
fof(f1399,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f913])).
fof(f913,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f912])).
fof(f912,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f725])).
fof(f725,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP12(X2,X0,X1))),
  inference(definition_folding,[],[f689,f724,f723])).
fof(f723,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f689,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f567])).
fof(f567,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f566])).
fof(f566,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X33,X43,X44] : ((p(s(bool,X44)) <=> (p(s(bool,X43)) => p(s(bool,X33)))) <=> ((p(s(bool,X44)) | p(s(bool,X43))) & (p(s(bool,X44)) | ~p(s(bool,X33))) & (~p(s(bool,X43)) | p(s(bool,X33)) | ~p(s(bool,X44)))))),
  file('thm.extreal.add_not_infty|split|0','thm.sat.sat_27_prove_ltb')).
fof(f24715,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23941,f17173])).
fof(f23941,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1994,f17173,f1387])).
fof(f1994,plain,(
  ( ! [X0] : (sP12('T_0',X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f1514,f919,f1436])).
fof(f1436,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f1390])).
fof(f1390,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f908])).
fof(f1514,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f946,f1395])).
fof(f1395,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f911])).
fof(f911,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f910])).
fof(f910,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f909])).
fof(f909,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f723])).
fof(f946,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f231])).
fof(f231,axiom,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.extreal.add_not_infty|split|0','thm.numeral.numeral_distrib|split|34')).
fof(f24714,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23942,f17173])).
fof(f23942,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6327,f17173,f1387])).
fof(f6327,plain,(
  ( ! [X0] : (sP12('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(unit_resulting_resolution,[],[f919,f5976,f1399])).
fof(f24713,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23943,f17173])).
fof(f23943,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7141,f17173,f1387])).
fof(f7141,plain,(
  ( ! [X0] : (sP12('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(unit_resulting_resolution,[],[f919,f5977,f1399])).
fof(f24712,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23944,f17173])).
fof(f23944,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1980,f17173,f1387])).
fof(f1980,plain,(
  ( ! [X0] : (sP12('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f919,f919,f1400])).
fof(f1400,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f913])).
fof(f24711,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23945,f17173])).
fof(f23945,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f8017,f17173,f1387])).
fof(f8017,plain,(
  ( ! [X0] : (sP12('T_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f919,f7614,f1400])).
fof(f24710,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23946,f17173])).
fof(f23946,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1997,f17173,f1387])).
fof(f1997,plain,(
  ( ! [X0] : (sP12('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f1516,f919,f1436])).
fof(f1516,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f921,f1396])).
fof(f1396,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f911])).
fof(f24709,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23947,f17173])).
fof(f23947,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2471,f17173,f1387])).
fof(f2471,plain,(
  ( ! [X0] : (sP12('T_0',sK73(sK59('const.num.|48|_0','type.num.num')),X0)) )),
  inference(unit_resulting_resolution,[],[f919,f1892,f1436])).
fof(f1892,plain,(
  ( ! [X0,X1] : (sP11(X0,sK73(sK59('const.num.|48|_0','type.num.num')),X1)) )),
  inference(unit_resulting_resolution,[],[f1871,f1396])).
fof(f24708,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23948,f17173])).
fof(f23948,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1998,f17173,f1387])).
fof(f1998,plain,(
  ( ! [X0] : (sP12('T_0',sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),X0)) )),
  inference(unit_resulting_resolution,[],[f1921,f919,f1436])).
fof(f1921,plain,(
  ( ! [X0,X1] : (sP11(X0,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),X1)) )),
  inference(unit_resulting_resolution,[],[f1899,f1396])).
fof(f24707,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23949,f17173])).
fof(f23949,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2134,f17173,f1387])).
fof(f2134,plain,(
  ( ! [X0] : (sP12('T_0',sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f919,f1901,f1400])).
fof(f24706,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23950,f17173])).
fof(f23950,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2258,f17173,f1387])).
fof(f2258,plain,(
  ( ! [X0] : (sP12('T_0',sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f919,f1930,f1400])).
fof(f24705,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23954,f17173])).
fof(f23954,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f8012,f17173,f1387])).
fof(f8012,plain,(
  sP12('F_0','F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),
  inference(unit_resulting_resolution,[],[f1451,f1451,f7614,f1398])).
fof(f1398,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f913])).
fof(f24704,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23957,f1451])).
fof(f23957,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f8038,f17173,f1387])).
fof(f8038,plain,(
  ( ! [X0] : (sP12('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1513,f7614,f1436])).
fof(f1513,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1451,f1395])).
fof(f24703,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23958,f17173])).
fof(f23958,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f8040,f17173,f1387])).
fof(f8040,plain,(
  ( ! [X0] : (sP12('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1515,f7614,f1436])).
fof(f1515,plain,(
  ( ! [X0,X1] : (sP11(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f919,f1396])).
fof(f24702,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23960,f17173])).
fof(f23960,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f5066,f17173,f1387])).
fof(f5066,plain,(
  ( ! [X0] : (sP12(sK73('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),X0)) )),
  inference(unit_resulting_resolution,[],[f1921,f1870,f1436])).
fof(f24701,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23961,f1451])).
fof(f23961,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2139,f17173,f1387])).
fof(f2139,plain,(
  ( ! [X0] : (sP12(sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1513,f1901,f1436])).
fof(f24700,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23962,f17173])).
fof(f23962,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2141,f17173,f1387])).
fof(f2141,plain,(
  ( ! [X0] : (sP12(sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1515,f1901,f1436])).
fof(f24699,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23963,f17173])).
fof(f23963,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2144,f17173,f1387])).
fof(f2144,plain,(
  ( ! [X0] : (sP12(sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),X0)) )),
  inference(unit_resulting_resolution,[],[f1921,f1901,f1436])).
fof(f24698,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23964,f1451])).
fof(f23964,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2263,f17173,f1387])).
fof(f2263,plain,(
  ( ! [X0] : (sP12(sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1513,f1930,f1436])).
fof(f24697,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23965,f17173])).
fof(f23965,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2265,f17173,f1387])).
fof(f2265,plain,(
  ( ! [X0] : (sP12(sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1515,f1930,f1436])).
fof(f24696,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23966,f17173])).
fof(f23966,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2268,f17173,f1387])).
fof(f2268,plain,(
  ( ! [X0] : (sP12(sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),X0)) )),
  inference(unit_resulting_resolution,[],[f1921,f1930,f1436])).
fof(f24695,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23982,f1451])).
fof(f23982,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f5261,f17173,f1387])).
fof(f5261,plain,(
  sP12('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1517,f919,f1451,f1393])).
fof(f1393,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f908])).
fof(f1517,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1451,f1397])).
fof(f1397,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f911])).
fof(f24694,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23983,f1451])).
fof(f23983,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f6322,f17173,f1387])).
fof(f6322,plain,(
  sP12('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f1451,f919,f5976,f1398])).
fof(f24693,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23984,f1451])).
fof(f23984,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7136,f17173,f1387])).
fof(f7136,plain,(
  sP12('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f1451,f919,f5977,f1398])).
fof(f24692,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23985,f1451])).
fof(f23985,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f8012,f17173,f1387])).
fof(f24691,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23986,f17173])).
fof(f23986,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6325,f17173,f1387])).
fof(f6325,plain,(
  sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1451,f919,f5976,f1398])).
fof(f24690,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23987,f17173])).
fof(f23987,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7139,f17173,f1387])).
fof(f7139,plain,(
  sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1451,f919,f5977,f1398])).
fof(f24689,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23988,f17173])).
fof(f23988,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f8038,f17173,f1387])).
fof(f24688,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23989,f17173])).
fof(f23989,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f8040,f17173,f1387])).
fof(f24687,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23990,f17173])).
fof(f23990,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f5262,f17173,f1387])).
fof(f5262,plain,(
  sP12('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1518,f919,f1451,f1393])).
fof(f1518,plain,(
  ( ! [X0,X1] : (sP11('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f946,f1397])).
fof(f24686,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23991,f17173])).
fof(f23991,plain,(
  p(s(bool,sK73('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl10),
  inference(unit_resulting_resolution,[],[f5066,f17173,f1387])).
fof(f24685,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23992,f17173])).
fof(f23992,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2139,f17173,f1387])).
fof(f24684,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23993,f17173])).
fof(f23993,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2141,f17173,f1387])).
fof(f24683,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23994,f17173])).
fof(f23994,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2144,f17173,f1387])).
fof(f24682,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23995,f17173])).
fof(f23995,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2263,f17173,f1387])).
fof(f24681,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23996,f17173])).
fof(f23996,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2265,f17173,f1387])).
fof(f24680,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f23997,f17173])).
fof(f23997,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2268,f17173,f1387])).
fof(f24679,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24114,f24155])).
fof(f24155,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1397])).
fof(f24114,plain,(
  ~sP11('F_0','T_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f1971,f919,f17173,f1393])).
fof(f1971,plain,(
  ( ! [X0] : (~sP12('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f919,f1451,f1388])).
fof(f1388,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f908])).
fof(f24678,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24116,f24155])).
fof(f24116,plain,(
  ~sP11('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7942,f919,f17173,f1393])).
fof(f7942,plain,(
  ( ! [X0] : (~sP12('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f1451,f7614,f1388])).
fof(f24677,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24118,f24155])).
fof(f24118,plain,(
  ~sP11('F_0',sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2130,f919,f17173,f1393])).
fof(f2130,plain,(
  ( ! [X0] : (~sP12('F_0',sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f1451,f1901,f1388])).
fof(f24676,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24120,f24155])).
fof(f24120,plain,(
  ~sP11('F_0',sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f2254,f919,f17173,f1393])).
fof(f2254,plain,(
  ( ! [X0] : (~sP12('F_0',sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f1451,f1930,f1388])).
fof(f24675,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24122,f24155])).
fof(f24122,plain,(
  ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f6306,f919,f17173,f1393])).
fof(f6306,plain,(
  ( ! [X0] : (~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f919,f5976,f1388])).
fof(f24674,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24124,f24155])).
fof(f24124,plain,(
  ~sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0','T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f7120,f919,f17173,f1393])).
fof(f7120,plain,(
  ( ! [X0] : (~sP12('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f919,f5977,f1388])).
fof(f24673,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24151,f24155])).
fof(f24151,plain,(
  ( ! [X0] : (~sP11('T_0',X0,'T_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f919,f919,f17173,f1394])).
fof(f1394,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f911])).
fof(f24672,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24183,f1451])).
fof(f24183,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f1971,f919,f17173,f1398])).
fof(f24671,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24184,f1451])).
fof(f24184,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f7942,f919,f17173,f1398])).
fof(f24670,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24185,f1451])).
fof(f24185,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2130,f919,f17173,f1398])).
fof(f24669,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24186,f1451])).
fof(f24186,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f2254,f919,f17173,f1398])).
fof(f24668,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24187,f17173])).
fof(f24187,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6306,f919,f17173,f1398])).
fof(f24667,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24188,f17173])).
fof(f24188,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7120,f919,f17173,f1398])).
fof(f24666,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24189,f17173])).
fof(f24189,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f8022,f919,f17173,f1398])).
fof(f8022,plain,(
  ~sP12('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1451,f919,f7614,f1401])).
fof(f1401,plain,(
  ( ! [X2,X0,X1] : (~sP12(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f913])).
fof(f24665,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24194,f1451])).
fof(f24194,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f5292,f919,f17173,f1398])).
fof(f5292,plain,(
  ~sP12('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f5278,f1389])).
fof(f1389,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f908])).
fof(f5278,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f919,f919,f1451,f1394])).
fof(f24664,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24195,f17173])).
fof(f24195,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f6331,f919,f17173,f1398])).
fof(f6331,plain,(
  ~sP12('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f919,f919,f5976,f1401])).
fof(f24663,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24196,f17173])).
fof(f24196,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f7145,f919,f17173,f1398])).
fof(f7145,plain,(
  ~sP12('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f919,f919,f5977,f1401])).
fof(f24662,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24198,f17173])).
fof(f24198,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7942,f919,f17173,f1398])).
fof(f24661,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24199,f17173])).
fof(f24199,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2130,f919,f17173,f1398])).
fof(f24660,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24200,f17173])).
fof(f24200,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2254,f919,f17173,f1398])).
fof(f24659,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24203,f1451])).
fof(f24203,plain,(
  p(s(bool,'F_0')) | $spl10),
  inference(unit_resulting_resolution,[],[f8022,f919,f17173,f1398])).
fof(f24658,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24380,f1501])).
fof(f24380,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1936,f17173,f1438])).
fof(f24657,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24382,f1501])).
fof(f24382,plain,(
  ( ! [X0] : (~sP7(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6160,f17173,f1438])).
fof(f24656,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24384,f1501])).
fof(f24384,plain,(
  ( ! [X0] : (~sP7(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),'T_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6974,f17173,f1438])).
fof(f24655,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24386,f1501])).
fof(f24386,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1939,f17173,f1438])).
fof(f24654,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24388,f1501])).
fof(f24388,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,'T_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6163,f17173,f1438])).
fof(f24653,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24390,f1501])).
fof(f24390,plain,(
  ( ! [X0] : (~sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,'T_0')) ) | $spl10),
  inference(unit_resulting_resolution,[],[f6977,f17173,f1438])).
fof(f24652,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24397,f7847])).
fof(f7847,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) )),
  inference(unit_resulting_resolution,[],[f7614,f1361])).
fof(f24397,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7781,f17173,f1438])).
fof(f24651,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24399,f7847])).
fof(f24399,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7782,f17173,f1438])).
fof(f24650,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24401,f2111])).
fof(f2111,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1901,f1361])).
fof(f24401,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2109,f17173,f1438])).
fof(f24649,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24403,f2111])).
fof(f24403,plain,(
  ( ! [X0] : (~sP7('F_0',X0,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2110,f17173,f1438])).
fof(f24648,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24405,f2235])).
fof(f2235,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1930,f1361])).
fof(f24405,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2233,f17173,f1438])).
fof(f24647,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24407,f2235])).
fof(f24407,plain,(
  ( ! [X0] : (~sP7('F_0',X0,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2234,f17173,f1438])).
fof(f24646,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24410,f17173])).
fof(f24410,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1538,f17173,f1474])).
fof(f1474,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1279,f1451])).
fof(f1279,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f864])).
fof(f864,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP1(X0,X1)))),
  inference(flattening,[],[f863])).
fof(f863,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP1(X0,X1)))),
  inference(nnf_transformation,[],[f708])).
fof(f708,plain,(
  ! [X0,X1] : (sP1(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f1538,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f947,f1280])).
fof(f1280,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f864])).
fof(f24645,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24411,f17173])).
fof(f24411,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1566,f17173,f1474])).
fof(f1566,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f949,f1280])).
fof(f24644,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24412,f17173])).
fof(f24412,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1525,f17173,f1474])).
fof(f1525,plain,(
  ( ! [X0] : (sP1(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f926,f1280])).
fof(f24643,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24413,f17173])).
fof(f24413,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7674,f17173,f1474])).
fof(f7674,plain,(
  ( ! [X0] : (sP1(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) )),
  inference(unit_resulting_resolution,[],[f7614,f1280])).
fof(f24642,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24414,f17173])).
fof(f24414,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14)))) | $spl10),
  inference(unit_resulting_resolution,[],[f8599,f17173,f1474])).
fof(f8599,plain,(
  ( ! [X0] : (sP1(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14)))) )),
  inference(unit_resulting_resolution,[],[f7615,f1280])).
fof(f24641,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24415,f17173])).
fof(f24415,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1480,f17173,f1474])).
fof(f1480,plain,(
  ( ! [X0] : (sP1(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f921,f1280])).
fof(f24640,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24416,f17173])).
fof(f24416,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.DATATYPE_1'(s(X0,X1))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1651,f17173,f1474])).
fof(f1651,plain,(
  ( ! [X4,X5,X3] : (sP1(X5,'const.bool.DATATYPE_1'(s(X3,X4)))) )),
  inference(subsumption_resolution,[],[f1649,f919])).
fof(f1649,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'T_0')) | sP1(X5,'const.bool.DATATYPE_1'(s(X3,X4)))) )),
  inference(superposition,[],[f1280,f1126])).
fof(f1126,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f397])).
fof(f397,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(flattening,[],[f396])).
fof(f396,plain,(
  ! [X0] : ! [X1] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X0,X1)))),
  inference(rectify,[],[f147])).
fof(f147,axiom,(
  ! [X8] : ! [X6] : s(bool,'T_0') = s(bool,'const.bool.DATATYPE_1'(s(X8,X6)))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.DATATYPE_TAG_THM')).
fof(f24639,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24417,f17173])).
fof(f24417,plain,(
  p(s(bool,sK59('T_0',bool))) | $spl10),
  inference(unit_resulting_resolution,[],[f1523,f17173,f1474])).
fof(f1523,plain,(
  ( ! [X0] : (sP1(X0,sK59('T_0',bool))) )),
  inference(unit_resulting_resolution,[],[f919,f1520])).
fof(f1520,plain,(
  ( ! [X2,X3] : (sP1(X3,sK59(X2,bool)) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f1280,f1140])).
fof(f24638,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24418,f17173])).
fof(f24418,plain,(
  p(s(bool,sK59('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),bool))) | $spl10),
  inference(unit_resulting_resolution,[],[f1524,f17173,f1474])).
fof(f1524,plain,(
  ( ! [X0] : (sP1(X0,sK59('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),bool))) )),
  inference(unit_resulting_resolution,[],[f921,f1520])).
fof(f24637,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24419,f17173])).
fof(f24419,plain,(
  p(s(bool,sK59(sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),bool))) | $spl10),
  inference(unit_resulting_resolution,[],[f1923,f17173,f1474])).
fof(f1923,plain,(
  ( ! [X0] : (sP1(X0,sK59(sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),bool))) )),
  inference(unit_resulting_resolution,[],[f1899,f1520])).
fof(f24636,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24420,f17173])).
fof(f24420,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1876,f17173,f1474])).
fof(f1876,plain,(
  ( ! [X0] : (sP1(X0,sK73(sK59('const.num.|48|_0','type.num.num')))) )),
  inference(unit_resulting_resolution,[],[f1871,f1280])).
fof(f24635,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24421,f17173])).
fof(f24421,plain,(
  p(s(bool,sK73('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl10),
  inference(unit_resulting_resolution,[],[f4777,f17173,f1474])).
fof(f4777,plain,(
  ( ! [X0] : (sP1(X0,sK73('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1870,f1280])).
fof(f24634,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24422,f17173])).
fof(f24422,plain,(
  p(s(bool,sK73('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl10),
  inference(unit_resulting_resolution,[],[f5419,f17173,f1474])).
fof(f5419,plain,(
  ( ! [X0] : (sP1(X0,sK73('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f1872,f1280])).
fof(f1872,plain,(
  p(s(bool,sK73('const.combin.I_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f1127,f1190])).
fof(f24633,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24423,f17173])).
fof(f24423,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1905,f17173,f1474])).
fof(f1905,plain,(
  ( ! [X0] : (sP1(X0,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) )),
  inference(unit_resulting_resolution,[],[f1899,f1280])).
fof(f24632,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24424,f17173])).
fof(f24424,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2096,f17173,f1474])).
fof(f2096,plain,(
  ( ! [X0] : (sP1(X0,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1901,f1280])).
fof(f24631,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24425,f17173])).
fof(f24425,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2220,f17173,f1474])).
fof(f2220,plain,(
  ( ! [X0] : (sP1(X0,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) )),
  inference(unit_resulting_resolution,[],[f1930,f1280])).
fof(f24630,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24426,f17173])).
fof(f24426,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1481,f17173,f1474])).
fof(f1481,plain,(
  ( ! [X0] : (sP1('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f919,f1281])).
fof(f1281,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f864])).
fof(f24629,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24427,f17173])).
fof(f24427,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1539,f17173,f1474])).
fof(f1539,plain,(
  ( ! [X0,X1] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f947,f1281])).
fof(f24628,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24428,f17173])).
fof(f24428,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1567,f17173,f1474])).
fof(f1567,plain,(
  ( ! [X0,X1] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f949,f1281])).
fof(f24627,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24429,f17173])).
fof(f24429,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1526,f17173,f1474])).
fof(f1526,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f926,f1281])).
fof(f24626,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24430,f17173])).
fof(f24430,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7675,f17173,f1474])).
fof(f7675,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f7614,f1281])).
fof(f24625,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24431,f17173])).
fof(f24431,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f8600,f17173,f1474])).
fof(f8600,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14)),X0)) )),
  inference(unit_resulting_resolution,[],[f7615,f1281])).
fof(f24624,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24432,f17173])).
fof(f24432,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1482,f17173,f1474])).
fof(f1482,plain,(
  ( ! [X0] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f921,f1281])).
fof(f24623,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24433,f17173])).
fof(f24433,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1650,f17173,f1474])).
fof(f1650,plain,(
  ( ! [X2,X0,X1] : (sP1('const.bool.DATATYPE_1'(s(X0,X1)),X2)) )),
  inference(subsumption_resolution,[],[f1648,f919])).
fof(f1648,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | sP1('const.bool.DATATYPE_1'(s(X0,X1)),X2)) )),
  inference(superposition,[],[f1281,f1126])).
fof(f24622,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24434,f17173])).
fof(f24434,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1521,f17173,f1474])).
fof(f1521,plain,(
  ( ! [X0] : (sP1(sK59('T_0',bool),X0)) )),
  inference(unit_resulting_resolution,[],[f919,f1519])).
fof(f1519,plain,(
  ( ! [X0,X1] : (sP1(sK59(X0,bool),X1) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f1281,f1140])).
fof(f24621,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24435,f17173])).
fof(f24435,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1522,f17173,f1474])).
fof(f1522,plain,(
  ( ! [X0] : (sP1(sK59('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),bool),X0)) )),
  inference(unit_resulting_resolution,[],[f921,f1519])).
fof(f24620,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24436,f17173])).
fof(f24436,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1922,f17173,f1474])).
fof(f1922,plain,(
  ( ! [X0] : (sP1(sK59(sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),bool),X0)) )),
  inference(unit_resulting_resolution,[],[f1899,f1519])).
fof(f24619,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24437,f17173])).
fof(f24437,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1877,f17173,f1474])).
fof(f1877,plain,(
  ( ! [X0] : (sP1(sK73(sK59('const.num.|48|_0','type.num.num')),X0)) )),
  inference(unit_resulting_resolution,[],[f1871,f1281])).
fof(f24618,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24438,f17173])).
fof(f24438,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f4778,f17173,f1474])).
fof(f4778,plain,(
  ( ! [X0] : (sP1(sK73('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f1870,f1281])).
fof(f24617,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24439,f17173])).
fof(f24439,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f5420,f17173,f1474])).
fof(f5420,plain,(
  ( ! [X0] : (sP1(sK73('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f1872,f1281])).
fof(f24616,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24440,f17173])).
fof(f24440,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1906,f17173,f1474])).
fof(f1906,plain,(
  ( ! [X0] : (sP1(sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),X0)) )),
  inference(unit_resulting_resolution,[],[f1899,f1281])).
fof(f24615,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24441,f17173])).
fof(f24441,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2097,f17173,f1474])).
fof(f2097,plain,(
  ( ! [X0] : (sP1(sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f1901,f1281])).
fof(f24614,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24442,f17173])).
fof(f24442,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2221,f17173,f1474])).
fof(f2221,plain,(
  ( ! [X0] : (sP1(sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')),X0)) )),
  inference(unit_resulting_resolution,[],[f1930,f1281])).
fof(f24613,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24447,f17173])).
fof(f24447,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1479,f17173,f1474])).
fof(f1479,plain,(
  ( ! [X0] : (sP1(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f919,f1280])).
fof(f24612,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24448,f17173])).
fof(f24448,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1538,f17173,f1474])).
fof(f24611,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24449,f17173])).
fof(f24449,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1566,f17173,f1474])).
fof(f24610,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24450,f17173])).
fof(f24450,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1525,f17173,f1474])).
fof(f24609,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24451,f17173])).
fof(f24451,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f7674,f17173,f1474])).
fof(f24608,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24452,f17173])).
fof(f24452,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f8599,f17173,f1474])).
fof(f24607,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24453,f17173])).
fof(f24453,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1480,f17173,f1474])).
fof(f24606,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24454,f17173])).
fof(f24454,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1651,f17173,f1474])).
fof(f24605,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24455,f17173])).
fof(f24455,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1523,f17173,f1474])).
fof(f24604,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24456,f17173])).
fof(f24456,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1524,f17173,f1474])).
fof(f24603,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24457,f17173])).
fof(f24457,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1923,f17173,f1474])).
fof(f24602,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24458,f17173])).
fof(f24458,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1876,f17173,f1474])).
fof(f24601,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24459,f17173])).
fof(f24459,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f4777,f17173,f1474])).
fof(f24600,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24460,f17173])).
fof(f24460,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f5419,f17173,f1474])).
fof(f24599,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24461,f17173])).
fof(f24461,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1905,f17173,f1474])).
fof(f24598,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24462,f17173])).
fof(f24462,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2096,f17173,f1474])).
fof(f24597,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24463,f17173])).
fof(f24463,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f2220,f17173,f1474])).
fof(f24596,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24465,f17173])).
fof(f24465,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1539,f17173,f1474])).
fof(f24595,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24466,f17173])).
fof(f24466,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1567,f17173,f1474])).
fof(f24594,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24467,f17173])).
fof(f24467,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1526,f17173,f1474])).
fof(f24593,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24468,f17173])).
fof(f24468,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))) | $spl10),
  inference(unit_resulting_resolution,[],[f7675,f17173,f1474])).
fof(f24592,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24469,f17173])).
fof(f24469,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK14)))) | $spl10),
  inference(unit_resulting_resolution,[],[f8600,f17173,f1474])).
fof(f24591,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24470,f17173])).
fof(f24470,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1482,f17173,f1474])).
fof(f24590,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24471,f17173])).
fof(f24471,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.DATATYPE_1'(s(X0,X1))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f1650,f17173,f1474])).
fof(f24589,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24472,f17173])).
fof(f24472,plain,(
  p(s(bool,sK59('T_0',bool))) | $spl10),
  inference(unit_resulting_resolution,[],[f1521,f17173,f1474])).
fof(f24588,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24473,f17173])).
fof(f24473,plain,(
  p(s(bool,sK59('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),bool))) | $spl10),
  inference(unit_resulting_resolution,[],[f1522,f17173,f1474])).
fof(f24587,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24474,f17173])).
fof(f24474,plain,(
  p(s(bool,sK59(sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),bool))) | $spl10),
  inference(unit_resulting_resolution,[],[f1922,f17173,f1474])).
fof(f24586,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24475,f17173])).
fof(f24475,plain,(
  p(s(bool,sK73(sK59('const.num.|48|_0','type.num.num')))) | $spl10),
  inference(unit_resulting_resolution,[],[f1877,f17173,f1474])).
fof(f24585,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24476,f17173])).
fof(f24476,plain,(
  p(s(bool,sK73('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl10),
  inference(unit_resulting_resolution,[],[f4778,f17173,f1474])).
fof(f24584,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24477,f17173])).
fof(f24477,plain,(
  p(s(bool,sK73('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl10),
  inference(unit_resulting_resolution,[],[f5420,f17173,f1474])).
fof(f24583,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24478,f17173])).
fof(f24478,plain,(
  p(s(bool,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) | $spl10),
  inference(unit_resulting_resolution,[],[f1906,f17173,f1474])).
fof(f24582,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24479,f17173])).
fof(f24479,plain,(
  p(s(bool,sK78(sK59('const.extreal.NegInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2097,f17173,f1474])).
fof(f24581,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24480,f17173])).
fof(f24480,plain,(
  p(s(bool,sK78(sK59('const.extreal.PosInf_0','type.extreal.extreal')))) | $spl10),
  inference(unit_resulting_resolution,[],[f2221,f17173,f1474])).
fof(f24567,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24496,f17173])).
fof(f24496,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1251])).
fof(f1251,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f847])).
fof(f847,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) & (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(nnf_transformation,[],[f509])).
fof(f509,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f508])).
fof(f508,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f84])).
fof(f84,axiom,(
  ! [X6,X7] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_lt_def')).
fof(f24566,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24500,f24522])).
fof(f24522,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1206])).
fof(f24500,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1111])).
fof(f1111,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f794])).
fof(f794,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) & (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)))),
  inference(nnf_transformation,[],[f373])).
fof(f373,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(flattening,[],[f372])).
fof(f372,plain,(
  ! [X0] : (~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(rectify,[],[f160])).
fof(f160,axiom,(
  ! [X6] : (~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.lt_infty|split|5')).
fof(f24565,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24502,f17173])).
fof(f24502,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1251])).
fof(f24564,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24504,f24524])).
fof(f24524,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1205])).
fof(f24504,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1113])).
fof(f24563,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24506,f17173])).
fof(f24506,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1253])).
fof(f1253,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f848])).
fof(f848,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) & (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))))),
  inference(nnf_transformation,[],[f511])).
fof(f511,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) <=> ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(flattening,[],[f510])).
fof(f510,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) <=> ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(rectify,[],[f113])).
fof(f113,axiom,(
  ! [X7,X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) <=> ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X6)))))),
  file('thm.extreal.add_not_infty|split|0','thm.real.real_lt')).
fof(f24562,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f24508,f17173])).
fof(f24508,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1253])).
fof(f24525,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f918,f17173,f1205])).
fof(f24523,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f931,f17173,f1206])).
fof(f931,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(cnf_transformation,[],[f34])).
fof(f34,axiom,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_add_def|split|2')).
fof(f24521,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1456,f17173,f1202])).
fof(f1202,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK77(X0)))) )),
  inference(cnf_transformation,[],[f843])).
fof(f843,plain,(
  ! [X0] : ((~p(s(bool,sK77(X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK77(X0)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK77(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK77])],[f842])).
fof(f842,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(flattening,[],[f841])).
fof(f841,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f460])).
fof(f460,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X6] : ? [X45] : ((p(s(bool,X45)) <=> s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X45),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_mul_def|split|6')).
fof(f1456,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.realax.real',X0)))) )),
  inference(forward_demodulation,[],[f984,f935])).
fof(f935,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f251])).
fof(f251,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.extreal.add_not_infty|split|0','thm.numeral.numeral_distrib|split|13')).
fof(f984,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))),
  inference(rectify,[],[f221])).
fof(f221,axiom,(
  ! [X6] : s('type.realax.real',X6) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X6)))),
  file('thm.extreal.add_not_infty|split|0','thm.real.REAL_MUL_LID')).
fof(f24519,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1456,f17173,f1199])).
fof(f1199,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK76(X0)))) )),
  inference(cnf_transformation,[],[f840])).
fof(f840,plain,(
  ! [X0] : ((~p(s(bool,sK76(X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK76(X0)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK76(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK76])],[f839])).
fof(f839,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(flattening,[],[f838])).
fof(f838,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(nnf_transformation,[],[f459])).
fof(f459,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(rectify,[],[f120])).
fof(f120,axiom,(
  ! [X7] : ? [X45] : ((p(s(bool,X45)) <=> s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X45),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_mul_def|split|5')).
fof(f24517,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1456,f17173,f1196])).
fof(f1196,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK75(X0)))) )),
  inference(cnf_transformation,[],[f837])).
fof(f837,plain,(
  ! [X0] : ((~p(s(bool,sK75(X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK75(X0)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK75(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK75])],[f836])).
fof(f836,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(flattening,[],[f835])).
fof(f835,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(nnf_transformation,[],[f458])).
fof(f458,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(rectify,[],[f118])).
fof(f118,axiom,(
  ! [X6] : ? [X45] : ((p(s(bool,X45)) <=> s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X45),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_mul_def|split|4')).
fof(f24515,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1456,f17173,f1193])).
fof(f1193,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK74(X0)))) )),
  inference(cnf_transformation,[],[f834])).
fof(f834,plain,(
  ! [X0] : ((~p(s(bool,sK74(X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK74(X0)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK74(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK74])],[f833])).
fof(f833,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(flattening,[],[f832])).
fof(f832,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f457])).
fof(f457,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(rectify,[],[f117])).
fof(f117,axiom,(
  ! [X7] : ? [X45] : ((p(s(bool,X45)) <=> s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X45),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_mul_def|split|7')).
fof(f24513,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f925,f17173,f1190])).
fof(f24511,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f925,f17173,f1063])).
fof(f1063,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f762])).
fof(f762,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f354])).
fof(f354,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f244])).
fof(f244,axiom,(
  ! [X36] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X36)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X36))),
  file('thm.extreal.add_not_infty|split|0','thm.numeral.numeral_distrib|split|30')).
fof(f24509,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f17173,f1253])).
fof(f24507,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f17173,f1253])).
fof(f24505,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f945,f17173,f1113])).
fof(f945,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0')),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.PosInf_0')),
  inference(flattening,[],[f127])).
fof(f127,axiom,(
  ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.PosInf_0')),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_distinct|split|0')).
fof(f24503,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f17173,f1251])).
fof(f24501,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1119,f17173,f1111])).
fof(f24499,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f931,f17173,f1061])).
fof(f24497,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f17173,f1251])).
fof(f24495,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f918,f17173,f1059])).
fof(f24490,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1167,f17173,f1323])).
fof(f1167,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f437])).
fof(f437,plain,(
  ! [X0,X1,X2,X3] : s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(flattening,[],[f436])).
fof(f436,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X8] : ! [X9,X10,X11] : s(X8,X9) = s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X8,X9),s(X8,X10),s(fun('type.realax.real',X8),X11)))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_case_def|split|0')).
fof(f24488,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1909,f17173,f2390])).
fof(f2390,plain,(
  ( ! [X3] : (sP2(X3,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14)))) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f2379,f919])).
fof(f2379,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | sP2(X3,sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))) )),
  inference(superposition,[],[f1299,f1903])).
fof(f1909,plain,(
  ~sP2('T_0',sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),
  inference(unit_resulting_resolution,[],[f919,f1899,f1295])).
fof(f1295,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f871])).
fof(f24486,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1907,f17173,f2389])).
fof(f2389,plain,(
  ( ! [X2] : (sP2(sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),X2) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f2378,f919])).
fof(f2378,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP2(sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),X2)) )),
  inference(superposition,[],[f1296,f1903])).
fof(f1907,plain,(
  ~sP2(sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),'T_0')),
  inference(unit_resulting_resolution,[],[f919,f1899,f1295])).
fof(f24484,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1479,f17173,f17173,f1474])).
fof(f24482,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1479,f1451,f17173,f1474])).
fof(f24446,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1479,f17173,f17173,f1474])).
fof(f24444,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1481,f1451,f17173,f1474])).
fof(f24408,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1505,f2234,f17173,f1438])).
fof(f1505,plain,(
  ( ! [X0,X1] : (sP7('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1451,f1363])).
fof(f24406,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1503,f2233,f17173,f1438])).
fof(f1503,plain,(
  ( ! [X0,X1] : (sP7(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1451,f1362])).
fof(f1362,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f895])).
fof(f24404,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1505,f2110,f17173,f1438])).
fof(f24402,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1503,f2109,f17173,f1438])).
fof(f24400,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1505,f7782,f17173,f1438])).
fof(f24398,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1503,f7781,f17173,f1438])).
fof(f24391,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f6986,f6977,f17173,f1438])).
fof(f6986,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f5977,f1363])).
fof(f24389,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f6172,f6163,f17173,f1438])).
fof(f6172,plain,(
  ( ! [X0,X1] : (sP7('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f5976,f1363])).
fof(f24387,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1505,f1939,f17173,f1438])).
fof(f24385,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f6985,f6974,f17173,f1438])).
fof(f6985,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f5977,f1362])).
fof(f24383,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f6171,f6160,f17173,f1438])).
fof(f6171,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f5976,f1362])).
fof(f24381,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1503,f1936,f17173,f1438])).
fof(f24330,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f1140,f17173,f1435])).
fof(f24328,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f1142,f17173,f1435])).
fof(f1142,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f417])).
fof(f417,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f416])).
fof(f416,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f192])).
fof(f192,axiom,(
  ! [X8] : ! [X19,X20] : s(X8,X19) = s(X8,'const.bool.COND_3'(s(bool,'T_0'),s(X8,X19),s(X8,X20)))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.COND_CLAUSES|split|0')).
fof(f24326,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f1141,f17173,f1435])).
fof(f24324,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f1143,f17173,f1435])).
fof(f1143,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f419])).
fof(f419,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f418])).
fof(f418,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X8] : ! [X12,X11] : s(X8,X12) = s(X8,'const.ind_type.FCONS_3'(s(X8,X12),s(fun('type.num.num',X8),X11),s('type.num.num','const.num.|48|_0')))),
  file('thm.extreal.add_not_infty|split|0','thm.ind_type.FCONS|split|0')).
fof(f24322,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f1127,f17173,f1435])).
fof(f24320,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f1173,f17173,f1435])).
fof(f1173,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK71(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(cnf_transformation,[],[f827])).
fof(f827,plain,(
  ! [X0,X1,X2,X3] : (s(X0,X1) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK71(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s(X0,X2) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK71(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.PosInf_0'))) & ! [X5] : s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X5))) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK71(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK71])],[f447])).
fof(f447,plain,(
  ! [X0,X1,X2,X3] : ? [X4] : (s(X0,X1) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s(X0,X2) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.PosInf_0'))) & ! [X5] : s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X5))) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))))),
  inference(flattening,[],[f446])).
fof(f446,plain,(
  ! [X0] : ! [X1,X2,X3] : ? [X4] : (s(X0,X1) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s(X0,X2) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.PosInf_0'))) & ! [X5] : s(X0,happ(s(fun('type.realax.real',X0),X3),s('type.realax.real',X5))) = s(X0,happ(s(fun('type.extreal.extreal',X0),X4),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X5))))))),
  inference(rectify,[],[f135])).
fof(f135,axiom,(
  ! [X8] : ! [X56,X57,X58] : ? [X59] : (s(X8,X56) = s(X8,happ(s(fun('type.extreal.extreal',X8),X59),s('type.extreal.extreal','const.extreal.NegInf_0'))) & s(X8,X57) = s(X8,happ(s(fun('type.extreal.extreal',X8),X59),s('type.extreal.extreal','const.extreal.PosInf_0'))) & ! [X12] : s(X8,happ(s(fun('type.realax.real',X8),X58),s('type.realax.real',X12))) = s(X8,happ(s(fun('type.extreal.extreal',X8),X59),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X12))))))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_Axiom')).
fof(f24318,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f1172,f17173,f1435])).
fof(f1172,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,happ(s(fun('type.extreal.extreal',X0),sK71(X3,X2,X1,X0)),s('type.extreal.extreal','const.extreal.NegInf_0')))) )),
  inference(cnf_transformation,[],[f827])).
fof(f24316,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f1306,f17173,f1435])).
fof(f1306,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f535])).
fof(f535,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f534])).
fof(f534,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f199])).
fof(f199,axiom,(
  ! [X26,X24] : ! [X27,X29] : s(X24,X27) = s(X24,happ(s(fun(X26,X24),'_dst_x0_1'(s(X24,X27))),s(X26,X29)))),
  file('thm.extreal.add_not_infty|split|0','_dst_x0_ax')).
fof(f24314,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f1462,f17173,f1435])).
fof(f24312,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f1168,f17173,f1435])).
fof(f1168,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f439])).
fof(f439,plain,(
  ! [X0,X1,X2,X3] : s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(flattening,[],[f438])).
fof(f438,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8] : ! [X9,X10,X11] : s(X8,X10) = s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X8,X9),s(X8,X10),s(fun('type.realax.real',X8),X11)))),
  file('thm.extreal.add_not_infty|split|0','thm.extreal.extreal_case_def|split|1')).
fof(f24310,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f1167,f17173,f1435])).
fof(f24292,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f1903,f17173,f1435])).
fof(f24290,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f952,f17173,f1435])).
fof(f952,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f237])).
fof(f237,axiom,(
  ! [X36] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X36),s('type.num.num','const.num.|48|_0')))),
  file('thm.extreal.add_not_infty|split|0','thm.numeral.numeral_distrib|split|29')).
fof(f24288,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f954,f17173,f1435])).
fof(f954,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f235])).
fof(f235,axiom,(
  ! [X36] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X36)))),
  file('thm.extreal.add_not_infty|split|0','thm.numeral.numeral_distrib|split|26')).
fof(f24286,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f1126,f17173,f1435])).
fof(f24284,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f957,f17173,f1435])).
fof(f24281,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1980,f919,f919,f17173,f1401])).
fof(f24249,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f5292,f919,f17173,f1399])).
fof(f24231,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f5292,f919,f17173,f17173,f1398])).
fof(f24213,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f5292,f919,f1451,f17173,f1398])).
fof(f24193,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f5292,f17173,f919,f17173,f1398])).
fof(f24191,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1971,f1451,f919,f17173,f1398])).
fof(f24156,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f5278,f17173,f1397])).
fof(f24154,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f5278,f17173,f1395])).
fof(f24152,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1515,f919,f919,f17173,f1394])).
fof(f24125,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7135,f7120,f919,f17173,f1393])).
fof(f7135,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK14),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f5977,f1397])).
fof(f24123,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f6321,f6306,f919,f17173,f1393])).
fof(f6321,plain,(
  ( ! [X0,X1] : (sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.NegInf_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f5976,f1397])).
fof(f24121,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1517,f2254,f919,f17173,f1393])).
fof(f24119,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1517,f2130,f919,f17173,f1393])).
fof(f24117,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1517,f7942,f919,f17173,f1393])).
fof(f24115,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1517,f1971,f919,f17173,f1393])).
fof(f24030,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1980,f919,f17173,f1388])).
fof(f24001,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1974,f17173,f17173,f1387])).
fof(f23999,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f5261,f1451,f17173,f1387])).
fof(f23970,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1974,f17173,f17173,f1387])).
fof(f23968,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1974,f1451,f17173,f1387])).
fof(f23939,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1960,f17173,f919,f17173,f1384])).
fof(f23937,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1960,f1451,f919,f17173,f1384])).
fof(f23868,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1960,f919,f17173,f17173,f1384])).
fof(f23832,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1966,f919,f1451,f17173,f1384])).
fof(f23797,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4751,f17173,f17173,f17173,f1381])).
fof(f23795,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4751,f1451,f17173,f17173,f1381])).
fof(f23775,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4751,f17173,f1451,f17173,f1381])).
fof(f23773,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4751,f1451,f1451,f17173,f1381])).
fof(f23733,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4751,f17173,f17173,f17173,f1381])).
fof(f23731,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1451,f17173,f17173,f1381])).
fof(f23711,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4751,f17173,f1451,f17173,f1381])).
fof(f23709,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1451,f1451,f17173,f1381])).
fof(f23671,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4751,f17173,f17173,f17173,f1381])).
fof(f23669,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1451,f17173,f17173,f1381])).
fof(f23649,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4751,f17173,f1451,f17173,f1381])).
fof(f23647,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f1451,f1451,f17173,f1381])).
fof(f23609,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4731,f17173,f1380])).
fof(f23607,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1507,f919,f17173,f17173,f1377])).
fof(f23584,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1507,f919,f1451,f17173,f1377])).
fof(f23552,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1507,f919,f17173,f17173,f1377])).
fof(f23529,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1509,f919,f1451,f17173,f1377])).
fof(f23497,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7002,f7070,f17173,f17173,f1376])).
fof(f23495,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f6188,f6256,f17173,f17173,f1376])).
fof(f23493,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1511,f17173,f17173,f1376])).
fof(f23491,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2239,f2243,f17173,f17173,f1376])).
fof(f23489,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2115,f2119,f17173,f17173,f1376])).
fof(f23487,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1920,f17173,f17173,f1376])).
fof(f23485,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1891,f17173,f17173,f1376])).
fof(f23483,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1510,f17173,f17173,f1376])).
fof(f23481,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7882,f7926,f17173,f17173,f1376])).
fof(f23479,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1534,f17173,f17173,f1376])).
fof(f23477,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f1509,f17173,f17173,f1376])).
fof(f23475,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f2242,f17173,f17173,f1376])).
fof(f23473,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f2118,f17173,f17173,f1376])).
fof(f23471,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f1919,f17173,f17173,f1376])).
fof(f23469,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f1890,f17173,f17173,f1376])).
fof(f23467,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f1508,f17173,f17173,f1376])).
fof(f23465,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f7925,f17173,f17173,f1376])).
fof(f23463,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f1533,f17173,f17173,f1376])).
fof(f23461,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1957,f1546,f17173,f17173,f1376])).
fof(f23459,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1507,f17173,f17173,f1376])).
fof(f23438,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7002,f7070,f1451,f17173,f1376])).
fof(f23436,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f6188,f6256,f1451,f17173,f1376])).
fof(f23434,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1511,f1451,f17173,f1376])).
fof(f23432,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2236,f2242,f1451,f17173,f1376])).
fof(f23430,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f2112,f2118,f1451,f17173,f1376])).
fof(f23425,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7870,f7925,f1451,f17173,f1376])).
fof(f23421,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1956,f1507,f1451,f17173,f1376])).
fof(f23328,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7070,f7003,f17173,f17173,f1376])).
fof(f23326,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7070,f7002,f17173,f17173,f1376])).
fof(f23324,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f6256,f6189,f17173,f17173,f1376])).
fof(f23322,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f6256,f6188,f17173,f17173,f1376])).
fof(f23320,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f2239,f17173,f17173,f1376])).
fof(f23318,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f2115,f17173,f17173,f1376])).
fof(f23316,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f7882,f17173,f17173,f1376])).
fof(f23314,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f1957,f17173,f17173,f1376])).
fof(f23312,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f7870,f17173,f17173,f1376])).
fof(f23310,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f2236,f17173,f17173,f1376])).
fof(f23308,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f2112,f17173,f17173,f1376])).
fof(f23306,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f1956,f17173,f17173,f1376])).
fof(f23277,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7070,f7003,f1451,f17173,f1376])).
fof(f23275,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f6256,f6189,f1451,f17173,f1376])).
fof(f23273,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f2239,f1451,f17173,f1376])).
fof(f23271,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f2115,f1451,f17173,f1376])).
fof(f23269,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f7882,f1451,f17173,f1376])).
fof(f23267,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1511,f1957,f1451,f17173,f1376])).
fof(f23005,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1966,f919,f17173,f1371])).
fof(f22969,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1960,f919,f17173,f1370])).
fof(f22933,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4605,f919,f919,f17173,f1369])).
fof(f4605,plain,(
  sP8('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1501,f919,f919,f1356])).
fof(f1356,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f893])).
fof(f22901,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1936,f17173,f17173,f1365])).
fof(f22899,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1939,f1451,f17173,f1365])).
fof(f22880,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1936,f17173,f17173,f1365])).
fof(f22878,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4642,f1451,f17173,f1365])).
fof(f22859,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1936,f17173,f17173,f1364])).
fof(f22857,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1936,f1451,f17173,f1364])).
fof(f22838,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1936,f17173,f17173,f1364])).
fof(f22836,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4642,f1451,f17173,f1364])).
fof(f22817,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4627,f17173,f1363])).
fof(f22815,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4627,f17173,f1362])).
fof(f22813,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1501,f919,f919,f17173,f1360])).
fof(f22785,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4605,f919,f17173,f1355])).
fof(f22760,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4605,f919,f17173,f1354])).
fof(f22735,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3164,f1140,f17173,f1352])).
fof(f3164,plain,(
  ( ! [X0] : (sP6('T_0',X0,sK59(X0,bool))) )),
  inference(unit_resulting_resolution,[],[f919,f1140,f1350])).
fof(f1350,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f890])).
fof(f22730,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3165,f1127,f17173,f1352])).
fof(f3165,plain,(
  ( ! [X0] : (sP6('T_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f919,f1127,f1350])).
fof(f22706,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3160,f1903,f17173,f1352])).
fof(f3160,plain,(
  sP6('T_0','T_0',sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))))),
  inference(unit_resulting_resolution,[],[f919,f1903,f1350])).
fof(f22702,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3157,f1126,f17173,f1352])).
fof(f3157,plain,(
  ( ! [X0,X1] : (sP6('T_0','T_0','const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f919,f1126,f1350])).
fof(f22699,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3154,f1140,f17173,f1352])).
fof(f3154,plain,(
  ( ! [X0] : (sP6('T_0',sK59(X0,bool),X0)) )),
  inference(unit_resulting_resolution,[],[f919,f1140,f1350])).
fof(f22694,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3155,f1127,f17173,f1352])).
fof(f3155,plain,(
  ( ! [X0] : (sP6('T_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f919,f1127,f1350])).
fof(f22670,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3150,f1903,f17173,f1352])).
fof(f3150,plain,(
  sP6('T_0',sK78('const.extreal.extreal_add_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal',sK14))),'T_0')),
  inference(unit_resulting_resolution,[],[f919,f1903,f1350])).
fof(f22666,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3147,f1126,f17173,f1352])).
fof(f3147,plain,(
  ( ! [X0,X1] : (sP6('T_0','const.bool.DATATYPE_1'(s(X0,X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f919,f1126,f1350])).
fof(f24561,plain,(
  $false | $spl10),
  inference(trivial_inequality_removal,[],[f22649])).
fof(f22649,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl10),
  inference(unit_resulting_resolution,[],[f12977,f17173,f1352])).
fof(f12977,plain,(
  sP6('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1495,f1485,f1491,f1451,f1335])).
fof(f24560,plain,(
  $false | $spl10),
  inference(trivial_inequality_removal,[],[f22648])).
fof(f22648,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl10),
  inference(unit_resulting_resolution,[],[f12730,f17173,f1352])).
fof(f12730,plain,(
  sP6('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1499,f1483,f1489,f919,f1334])).
fof(f1334,plain,(
  ( ! [X2,X0,X1] : (~sP4(X1,X2,X0) | ~p(s(bool,X2)) | ~sP5(X0,X2,X1) | sP6(X0,X1,X2) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f881])).
fof(f22643,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3177,f1710,f17173,f1349])).
fof(f22641,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3175,f1710,f17173,f1349])).
fof(f22619,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f17173,f17173,f17173,f1345])).
fof(f22617,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f1451,f17173,f17173,f1345])).
fof(f22588,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f17173,f1451,f17173,f1345])).
fof(f22586,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f1451,f1451,f17173,f1345])).
fof(f22521,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f17173,f17173,f17173,f1345])).
fof(f22519,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f1451,f17173,f17173,f1345])).
fof(f22490,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f17173,f1451,f17173,f1345])).
fof(f22488,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f1451,f1451,f17173,f1345])).
fof(f22423,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f17173,f17173,f17173,f1345])).
fof(f22421,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f1451,f17173,f17173,f1345])).
fof(f22392,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f17173,f1451,f17173,f1345])).
fof(f22390,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1499,f1451,f1451,f17173,f1345])).
fof(f22325,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4531,f17173,f1344])).
fof(f4531,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f919,f1451,f919,f1341])).
fof(f22323,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4531,f17173,f1343])).
fof(f22321,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1489,f919,f919,f17173,f1341])).
fof(f22293,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4498,f17173,f1340])).
fof(f4498,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f919,f1451,f919,f1337])).
fof(f22291,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f4498,f17173,f1339])).
fof(f22289,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3428,f919,f919,f17173,f1337])).
fof(f3428,plain,(
  ( ! [X0,X1] : (sP5('T_0','T_0','const.bool.DATATYPE_1'(s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f3147,f1332])).
fof(f1332,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f881])).
fof(f21654,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f1483,f6872,f17173,f1336])).
fof(f21640,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f6865,f6872,f17173,f1336])).
fof(f21490,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f1483,f6058,f17173,f1336])).
fof(f21476,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f6051,f6058,f17173,f1336])).
fof(f21364,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f8641,f1491,f17173,f1336])).
fof(f21353,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f7716,f1491,f17173,f1336])).
fof(f21321,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f1483,f1491,f17173,f1336])).
fof(f21319,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1483,f1491,f17173,f1336])).
fof(f21307,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f1485,f1491,f17173,f1336])).
fof(f21302,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f1485,f1491,f17173,f1336])).
fof(f21300,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f1485,f1491,f17173,f1336])).
fof(f21298,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1485,f1491,f17173,f1336])).
fof(f20499,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f1485,f7743,f17173,f1336])).
fof(f20389,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f1483,f7743,f17173,f1336])).
fof(f20066,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f6865,f1489,f17173,f1336])).
fof(f20054,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f6051,f1489,f17173,f1336])).
fof(f20042,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f1485,f1489,f17173,f1336])).
fof(f20040,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f1485,f1489,f17173,f1336])).
fof(f20038,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1485,f1489,f17173,f1336])).
fof(f19971,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f8641,f1489,f17173,f1336])).
fof(f19959,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f7716,f1489,f17173,f1336])).
fof(f19924,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f1483,f1489,f17173,f1336])).
fof(f19922,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f1483,f1489,f17173,f1336])).
fof(f19920,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1483,f1489,f17173,f1336])).
fof(f19184,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f1483,f6872,f17173,f1335])).
fof(f19170,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f6865,f6872,f17173,f1335])).
fof(f19020,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f1483,f6058,f17173,f1335])).
fof(f19006,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f6051,f6058,f17173,f1335])).
fof(f18894,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f8641,f1491,f17173,f1335])).
fof(f18883,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f7716,f1491,f17173,f1335])).
fof(f18851,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f1483,f1491,f17173,f1335])).
fof(f18849,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1483,f1491,f17173,f1335])).
fof(f18837,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f1485,f1491,f17173,f1335])).
fof(f18832,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f1485,f1491,f17173,f1335])).
fof(f18830,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f1485,f1491,f17173,f1335])).
fof(f18828,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1485,f1491,f17173,f1335])).
fof(f18029,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f1485,f7743,f17173,f1335])).
fof(f17919,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f7779,f1483,f7743,f17173,f1335])).
fof(f17596,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f6865,f1489,f17173,f1335])).
fof(f17584,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f6051,f1489,f17173,f1335])).
fof(f17572,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f1485,f1489,f17173,f1335])).
fof(f17570,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f1485,f1489,f17173,f1335])).
fof(f17568,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1485,f1489,f17173,f1335])).
fof(f17501,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1495,f8609,f8641,f1489,f17173,f1335])).
fof(f17489,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f7745,f7684,f7716,f1489,f17173,f1335])).
fof(f17454,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6855,f1483,f1489,f17173,f1335])).
fof(f17452,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f6041,f1483,f1489,f17173,f1335])).
fof(f17450,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1497,f3177,f1483,f1489,f17173,f1335])).
fof(f17310,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f3160,f919,f919,f17173,f1333])).
fof(f17287,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1780,f919,f17173,f1299])).
fof(f1780,plain,(
  ~sP2('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f919,f919,f1295])).
fof(f17264,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1780,f919,f17173,f1296])).
fof(f17241,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1789,f17173,f17173,f1294])).
fof(f17239,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1804,f1451,f17173,f1294])).
fof(f17217,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1789,f17173,f17173,f1294])).
fof(f17215,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1789,f1451,f17173,f1294])).
fof(f17193,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1451,f1710,f17173,f1242])).
fof(f17191,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1710,f17173,f1242])).
fof(f17187,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1451,f1710,f17173,f1242])).
fof(f17185,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f17173,f1710,f17173,f1242])).
fof(f17181,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f1710,f17173,f1043])).
fof(f17179,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f957,f17173,f1038])).
fof(f17177,plain,(
  $false | $spl10),
  inference(unit_resulting_resolution,[],[f919,f17173])).
fof(f17175,plain,(
  p(s(bool,sK57('const.combin.I_0',bool))) | $spl12),
  inference(cnf_transformation,[],[f17175_D])).
fof(f17175_D,plain,(
  p(s(bool,sK57('const.combin.I_0',bool))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f17172,plain,(
  ( ! [X3] : (p(s(bool,sK57('const.combin.I_0',bool))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f17157,f1462])).
fof(f17157,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,sK57('const.combin.I_0',bool)))))) )),
  inference(superposition,[],[f1134,f1462])).
fof(f1134,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK57(X1,X0)))))) )),
  inference(cnf_transformation,[],[f802])).
fof(f15954,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl6),
  inference(cnf_transformation,[],[f15954_D])).
fof(f15954_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f15955,plain,(
  ( ! [X3] : (~p(s(bool,sK61(X3,'const.combin.I_0',bool)))) ) | $spl8),
  inference(cnf_transformation,[],[f15955_D])).
fof(f15955_D,plain,(
  ( ! [X3] : (~p(s(bool,sK61(X3,'const.combin.I_0',bool)))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f15940,plain,(
  ( ! [X3] : (~p(s(bool,sK61(X3,'const.combin.I_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))))) )),
  inference(superposition,[],[f1148,f1462])).
fof(f1148,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK61(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f809])).
fof(f809,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK61(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK61(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK61])],[f629])).
fof(f629,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f425])).
fof(f425,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f424])).
fof(f424,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8] : ! [X16,X21] : (! [X6] : (s(X8,X6) = s(X8,X21) => p(s(bool,happ(s(fun(X8,bool),X16),s(X8,X6))))) => p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X16)))))),
  file('thm.extreal.add_not_infty|split|0','thm.bool.JRH_INDUCT_UTIL')).
fof(f5406,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(cnf_transformation,[],[f5406_D])).
fof(f5406_D,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f5403,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(cnf_transformation,[],[f5403_D])).
fof(f5403_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f5405,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f5388,f978])).
fof(f978,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f258])).
fof(f258,axiom,(
  ! [X36] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36)))))),
  file('thm.extreal.add_not_infty|split|0','thm.numeral.numeral_distrib|split|14')).
fof(f5388,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1054,f1455])).
fof(f1455,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f980,f935])).
fof(f980,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f248])).
fof(f248,axiom,(
  ! [X36] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X36),s('type.num.num','const.num.|48|_0')))),
  file('thm.extreal.add_not_infty|split|0','thm.numeral.numeral_distrib|split|11')).
fof(f1054,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f758])).
fof(f758,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f350])).
fof(f350,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f243])).
fof(f243,axiom,(
  ! [X36] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X36))),
  file('thm.extreal.add_not_infty|split|0','thm.numeral.numeral_distrib|split|17')).
fof(f5401,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl0),
  inference(cnf_transformation,[],[f5401_D])).
fof(f5401_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f5399,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f5386,f978])).
fof(f5386,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f1054,f935])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 905
% Passive clauses: 9569
% Generated clauses: 71744
% Final active clauses: 586
% Final passive clauses: 5126
% Input formulas: 261
% Initial clauses: 508
% 
% Duplicate literals: 8
% Trivial inequalities: 25
% Fw subsumption resolutions: 1384
% Fw demodulations: 27516
% Bw demodulations: 1327
% Condensations: 19
% Global subsumptions: 3
% 
% Simple tautologies: 102
% Equational tautologies: 50
% Forward subsumptions: 31414
% Fw demodulations to eq. taut.: 431
% Bw demodulations to eq. taut.: 12
% 
% Unit resulting resolution: 23094
% Factoring: 1
% Forward superposition: 9266
% Backward superposition: 7266
% Self superposition: 96
% Equality factoring: 2
% Equality resolution: 23
% 
% Split clauses: 20
% Split components: 40
% SAT solver clauses: 10817
% SAT solver unit clauses: 4966
% SAT solver binary clauses: 4523
% 
% Sat splits: 20
% Sat splitting refutations: 1203
% 
% Memory used [KB]: 46950
% Time elapsed: 1.886 s
% ------------------------------
% ------------------------------
% Success in time 3.22 s
