Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.option
% (76211)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 156
% Passive clauses: 2477
% Generated clauses: 22513
% Final active clauses: 139
% Final passive clauses: 2011
% Input formulas: 20
% Initial clauses: 33
% 
% Duplicate literals: 1216
% Trivial inequalities: 224
% Fw subsumption resolutions: 2455
% Fw demodulations: 1807
% Bw demodulations: 10
% 
% Simple tautologies: 344
% Equational tautologies: 79
% Forward subsumptions: 12597
% 
% Binary resolution: 35
% Forward superposition: 8364
% Backward superposition: 7575
% Self superposition: 371
% Equality factoring: 204
% Equality resolution: 21
% 
% Split clauses: 1071
% Split components: 2176
% SAT solver clauses: 1514
% SAT solver unit clauses: 17
% SAT solver binary clauses: 134
% 
% Sat splits: 1071
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 1488
% TWLsolver calls for satisfiability: 121
% 
% Memory used [KB]: 14200
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.option
% (76224)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 95
% Passive clauses: 453
% Generated clauses: 11817
% Final active clauses: 87
% Final passive clauses: 272
% Input formulas: 20
% Initial clauses: 33
% 
% Duplicate literals: 1464
% Trivial inequalities: 64
% Fw subsumption resolutions: 540
% Fw demodulations: 465
% Bw demodulations: 80
% Condensations: 1918
% Global subsumptions: 4
% 
% Simple tautologies: 160
% Equational tautologies: 62
% Forward subsumptions: 5379
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 7
% Forward superposition: 3448
% Backward superposition: 2952
% Self superposition: 220
% Equality factoring: 157
% Equality resolution: 5
% 
% Split clauses: 115
% Split components: 235
% SAT solver clauses: 1100
% SAT solver unit clauses: 41
% SAT solver binary clauses: 75
% SAT solver learnt clauses: 8
% 
% Sat splits: 115
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1084
% TWLsolver calls for satisfiability: 138
% 
% Memory used [KB]: 9850
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.option
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.option
% SZS output start Proof for thm.option
fof(f1572,plain,(
  $false),
  inference(subsumption_resolution,[],[f1571,f735])).
fof(f735,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.option.option'(X2),happ(s(fun(X0,'type.option.option'(X2)),X3),s(X0,X1))) = s('type.option.option'(X2),'const.option.OPTION_MCOMP_3'(s(fun(X0,'type.option.option'(X2)),X3),s(fun(X0,'type.option.option'(X0)),'const.option.SOME_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f718,f97])).
fof(f97,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.option.option'(X0),happ(s(fun(X1,'type.option.option'(X0)),X3),s(X1,X2))) = s('type.option.option'(X0),'const.option.OPTION_BIND_2'(s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X2))),s(fun(X1,'type.option.option'(X0)),X3)))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3] : s('type.option.option'(X0),happ(s(fun(X1,'type.option.option'(X0)),X3),s(X1,X2))) = s('type.option.option'(X0),'const.option.OPTION_BIND_2'(s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X2))),s(fun(X1,'type.option.option'(X0)),X3)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0,X1] : ! [X2,X3] : s('type.option.option'(X0),happ(s(fun(X1,'type.option.option'(X0)),X3),s(X1,X2))) = s('type.option.option'(X0),'const.option.OPTION_BIND_2'(s('type.option.option'(X1),'const.option.SOME_1'(s(X1,X2))),s(fun(X1,'type.option.option'(X0)),X3)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X9,X12] : ! [X11,X13] : s('type.option.option'(X9),happ(s(fun(X12,'type.option.option'(X9)),X13),s(X12,X11))) = s('type.option.option'(X9),'const.option.OPTION_BIND_2'(s('type.option.option'(X12),'const.option.SOME_1'(s(X12,X11))),s(fun(X12,'type.option.option'(X9)),X13)))),
  file('thm.option.OPTION_MCOMP_ID|split|0','thm.option.OPTION_BIND_def|split|1')).
fof(f718,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.option.option'(X2),'const.option.OPTION_BIND_2'(s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))),s(fun(X0,'type.option.option'(X2)),X3))) = s('type.option.option'(X2),'const.option.OPTION_MCOMP_3'(s(fun(X0,'type.option.option'(X2)),X3),s(fun(X0,'type.option.option'(X0)),'const.option.SOME_0'),s(X0,X1)))) )),
  inference(superposition,[],[f101,f90])).
fof(f90,plain,(
  ( ! [X0,X1] : (s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),happ(s(fun(X0,'type.option.option'(X0)),'const.option.SOME_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),happ(s(fun(X0,'type.option.option'(X0)),'const.option.SOME_0'),s(X0,X1)))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : ! [X1] : s('type.option.option'(X0),'const.option.SOME_1'(s(X0,X1))) = s('type.option.option'(X0),happ(s(fun(X0,'type.option.option'(X0)),'const.option.SOME_0'),s(X0,X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X9] : ! [X16] : s('type.option.option'(X9),'const.option.SOME_1'(s(X9,X16))) = s('type.option.option'(X9),happ(s(fun(X9,'type.option.option'(X9)),'const.option.SOME_0'),s(X9,X16)))),
  file('thm.option.OPTION_MCOMP_ID|split|0','ARITY_0_1_const.option.SOME')).
fof(f101,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.option.option'(X2),'const.option.OPTION_MCOMP_3'(s(fun(X1,'type.option.option'(X2)),X3),s(fun(X0,'type.option.option'(X1)),X4),s(X0,X5))) = s('type.option.option'(X2),'const.option.OPTION_BIND_2'(s('type.option.option'(X1),happ(s(fun(X0,'type.option.option'(X1)),X4),s(X0,X5))),s(fun(X1,'type.option.option'(X2)),X3)))) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s('type.option.option'(X2),'const.option.OPTION_MCOMP_3'(s(fun(X1,'type.option.option'(X2)),X3),s(fun(X0,'type.option.option'(X1)),X4),s(X0,X5))) = s('type.option.option'(X2),'const.option.OPTION_BIND_2'(s('type.option.option'(X1),happ(s(fun(X0,'type.option.option'(X1)),X4),s(X0,X5))),s(fun(X1,'type.option.option'(X2)),X3)))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s('type.option.option'(X2),'const.option.OPTION_MCOMP_3'(s(fun(X1,'type.option.option'(X2)),X3),s(fun(X0,'type.option.option'(X1)),X4),s(X0,X5))) = s('type.option.option'(X2),'const.option.OPTION_BIND_2'(s('type.option.option'(X1),happ(s(fun(X0,'type.option.option'(X1)),X4),s(X0,X5))),s(fun(X1,'type.option.option'(X2)),X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6,X12,X9] : ! [X8,X13,X15] : s('type.option.option'(X9),'const.option.OPTION_MCOMP_3'(s(fun(X12,'type.option.option'(X9)),X8),s(fun(X6,'type.option.option'(X12)),X13),s(X6,X15))) = s('type.option.option'(X9),'const.option.OPTION_BIND_2'(s('type.option.option'(X12),happ(s(fun(X6,'type.option.option'(X12)),X13),s(X6,X15))),s(fun(X12,'type.option.option'(X9)),X8)))),
  file('thm.option.OPTION_MCOMP_ID|split|0','thm.option.OPTION_MCOMP_def')).
fof(f1571,plain,(
  s('type.option.option'(sK1),happ(s(fun(sK0,'type.option.option'(sK1)),sK2),s(sK0,sK5('const.option.OPTION_MCOMP_2'(s(fun(sK0,'type.option.option'(sK1)),sK2),s(fun(sK0,'type.option.option'(sK0)),'const.option.SOME_0')),sK2,sK0,'type.option.option'(sK1))))) != s('type.option.option'(sK1),'const.option.OPTION_MCOMP_3'(s(fun(sK0,'type.option.option'(sK1)),sK2),s(fun(sK0,'type.option.option'(sK0)),'const.option.SOME_0'),s(sK0,sK5('const.option.OPTION_MCOMP_2'(s(fun(sK0,'type.option.option'(sK1)),sK2),s(fun(sK0,'type.option.option'(sK0)),'const.option.SOME_0')),sK2,sK0,'type.option.option'(sK1)))))),
  inference(superposition,[],[f541,f102])).
fof(f102,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.option.option'(X0),'const.option.OPTION_MCOMP_3'(s(fun(X1,'type.option.option'(X0)),X3),s(fun(X2,'type.option.option'(X1)),X4),s(X2,X5))) = s('type.option.option'(X0),happ(s(fun(X2,'type.option.option'(X0)),'const.option.OPTION_MCOMP_2'(s(fun(X1,'type.option.option'(X0)),X3),s(fun(X2,'type.option.option'(X1)),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s('type.option.option'(X0),'const.option.OPTION_MCOMP_3'(s(fun(X1,'type.option.option'(X0)),X3),s(fun(X2,'type.option.option'(X1)),X4),s(X2,X5))) = s('type.option.option'(X0),happ(s(fun(X2,'type.option.option'(X0)),'const.option.OPTION_MCOMP_2'(s(fun(X1,'type.option.option'(X0)),X3),s(fun(X2,'type.option.option'(X1)),X4))),s(X2,X5)))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s('type.option.option'(X0),'const.option.OPTION_MCOMP_3'(s(fun(X1,'type.option.option'(X0)),X3),s(fun(X2,'type.option.option'(X1)),X4),s(X2,X5))) = s('type.option.option'(X0),happ(s(fun(X2,'type.option.option'(X0)),'const.option.OPTION_MCOMP_2'(s(fun(X1,'type.option.option'(X0)),X3),s(fun(X2,'type.option.option'(X1)),X4))),s(X2,X5)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X9,X12,X6] : ! [X16,X17,X18] : s('type.option.option'(X9),'const.option.OPTION_MCOMP_3'(s(fun(X12,'type.option.option'(X9)),X16),s(fun(X6,'type.option.option'(X12)),X17),s(X6,X18))) = s('type.option.option'(X9),happ(s(fun(X6,'type.option.option'(X9)),'const.option.OPTION_MCOMP_2'(s(fun(X12,'type.option.option'(X9)),X16),s(fun(X6,'type.option.option'(X12)),X17))),s(X6,X18)))),
  file('thm.option.OPTION_MCOMP_ID|split|0','ARITY_2_3_const.option.OPTION_MCOMP')).
fof(f541,plain,(
  s('type.option.option'(sK1),happ(s(fun(sK0,'type.option.option'(sK1)),sK2),s(sK0,sK5('const.option.OPTION_MCOMP_2'(s(fun(sK0,'type.option.option'(sK1)),sK2),s(fun(sK0,'type.option.option'(sK0)),'const.option.SOME_0')),sK2,sK0,'type.option.option'(sK1))))) != s('type.option.option'(sK1),happ(s(fun(sK0,'type.option.option'(sK1)),'const.option.OPTION_MCOMP_2'(s(fun(sK0,'type.option.option'(sK1)),sK2),s(fun(sK0,'type.option.option'(sK0)),'const.option.SOME_0'))),s(sK0,sK5('const.option.OPTION_MCOMP_2'(s(fun(sK0,'type.option.option'(sK1)),sK2),s(fun(sK0,'type.option.option'(sK0)),'const.option.SOME_0')),sK2,sK0,'type.option.option'(sK1)))))),
  inference(unit_resulting_resolution,[],[f70,f100])).
fof(f100,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK5(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK5(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK5(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK5(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X12,X9] : ! [X13,X8] : (s(fun(X9,X12),X8) = s(fun(X9,X12),X13) <=> ! [X11] : s(X12,happ(s(fun(X9,X12),X8),s(X9,X11))) = s(X12,happ(s(fun(X9,X12),X13),s(X9,X11))))),
  file('thm.option.OPTION_MCOMP_ID|split|0','thm.bool.FUN_EQ_THM')).
fof(f70,plain,(
  s(fun(sK0,'type.option.option'(sK1)),sK2) != s(fun(sK0,'type.option.option'(sK1)),'const.option.OPTION_MCOMP_2'(s(fun(sK0,'type.option.option'(sK1)),sK2),s(fun(sK0,'type.option.option'(sK0)),'const.option.SOME_0')))),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  s(fun(sK0,'type.option.option'(sK1)),sK2) != s(fun(sK0,'type.option.option'(sK1)),'const.option.OPTION_MCOMP_2'(s(fun(sK0,'type.option.option'(sK1)),sK2),s(fun(sK0,'type.option.option'(sK0)),'const.option.SOME_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f49])).
fof(f49,plain,(
  ? [X0,X1,X2] : s(fun(X0,'type.option.option'(X1)),X2) != s(fun(X0,'type.option.option'(X1)),'const.option.OPTION_MCOMP_2'(s(fun(X0,'type.option.option'(X1)),X2),s(fun(X0,'type.option.option'(X0)),'const.option.SOME_0')))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1,X2] : s(fun(X0,'type.option.option'(X1)),X2) = s(fun(X0,'type.option.option'(X1)),'const.option.OPTION_MCOMP_2'(s(fun(X0,'type.option.option'(X1)),X2),s(fun(X0,'type.option.option'(X0)),'const.option.SOME_0')))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1] : ! [X2] : s(fun(X0,'type.option.option'(X1)),X2) = s(fun(X0,'type.option.option'(X1)),'const.option.OPTION_MCOMP_2'(s(fun(X0,'type.option.option'(X1)),X2),s(fun(X0,'type.option.option'(X0)),'const.option.SOME_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8] : s(fun(X6,'type.option.option'(X7)),X8) = s(fun(X6,'type.option.option'(X7)),'const.option.OPTION_MCOMP_2'(s(fun(X6,'type.option.option'(X7)),X8),s(fun(X6,'type.option.option'(X6)),'const.option.SOME_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8] : s(fun(X6,'type.option.option'(X7)),X8) = s(fun(X6,'type.option.option'(X7)),'const.option.OPTION_MCOMP_2'(s(fun(X6,'type.option.option'(X7)),X8),s(fun(X6,'type.option.option'(X6)),'const.option.SOME_0')))),
  file('thm.option.OPTION_MCOMP_ID|split|0','thm.option.OPTION_MCOMP_ID|split|0')).
% SZS output end Proof for thm.option
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 41
% Passive clauses: 179
% Generated clauses: 1487
% Final active clauses: 25
% Final passive clauses: 44
% Input formulas: 20
% Initial clauses: 33
% 
% Duplicate literals: 175
% Trivial inequalities: 10
% Fw subsumption resolutions: 21
% Fw demodulations: 88
% Bw demodulations: 34
% Condensations: 141
% Global subsumptions: 2
% 
% Simple tautologies: 43
% Equational tautologies: 33
% Forward subsumptions: 761
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 32
% Forward superposition: 428
% Backward superposition: 365
% Self superposition: 57
% Equality factoring: 44
% Equality resolution: 2
% 
% Split clauses: 37
% Split components: 74
% SAT solver clauses: 257
% SAT solver unit clauses: 14
% SAT solver binary clauses: 26
% 
% Sat splits: 37
% 
% Memory used [KB]: 2174
% Time elapsed: 0.057 s
% ------------------------------
% ------------------------------
% Success in time 1.371 s
