Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.numpair
% (64391)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 97
% Passive clauses: 2237
% Generated clauses: 17610
% Final active clauses: 92
% Final passive clauses: 1977
% Input formulas: 10
% Initial clauses: 16
% 
% Duplicate literals: 1096
% Trivial inequalities: 223
% Fw subsumption resolutions: 2186
% Fw demodulations: 1269
% Bw demodulations: 1
% 
% Simple tautologies: 351
% Equational tautologies: 103
% Forward subsumptions: 8014
% 
% Binary resolution: 46
% Forward superposition: 7688
% Backward superposition: 4275
% Self superposition: 370
% Equality factoring: 266
% Equality resolution: 23
% 
% Split clauses: 1011
% Split components: 2065
% SAT solver clauses: 2278
% SAT solver unit clauses: 269
% SAT solver binary clauses: 620
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 1011
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2245
% TWLsolver calls for satisfiability: 113
% 
% Memory used [KB]: 7547
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.numpair
% (64411)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 47
% Passive clauses: 179
% Generated clauses: 6363
% Final active clauses: 47
% Final passive clauses: 131
% Input formulas: 10
% Initial clauses: 16
% 
% Duplicate literals: 1044
% Trivial inequalities: 76
% Fw subsumption resolutions: 313
% Fw demodulations: 146
% Condensations: 1600
% 
% Simple tautologies: 120
% Equational tautologies: 68
% Forward subsumptions: 2389
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 11
% Forward superposition: 1544
% Backward superposition: 1084
% Self superposition: 228
% Equality factoring: 197
% Equality resolution: 18
% 
% Split clauses: 74
% Split components: 150
% SAT solver clauses: 528
% SAT solver unit clauses: 18
% SAT solver binary clauses: 83
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 74
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 527
% TWLsolver calls for satisfiability: 108
% 
% Memory used [KB]: 4349
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.numpair
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numpair
% SZS output start Proof for thm.numpair
fof(f1324,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f33,f53,f51,f49,f32,f52,f31,f50,f48,f56,f63,f66,f64,f800,f814,f1052,f1055,f1054,f1030,f1031,f1057,f1058,f1032,f1033,f1060,f1034,f1062,f1059,f1311])).
fof(f1311,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(unit_resulting_resolution,[],[f815,f1304])).
fof(f1304,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',X0))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK3),s('type.num.num',X1)))) ) | $spl3),
  inference(superposition,[],[f987,f40])).
fof(f40,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numpair.npair_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(cnf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0,X1] : s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numpair.npair_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X9,X8] : s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X9),s('type.num.num',X8))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numpair.npair_2'(s('type.num.num',X9),s('type.num.num',X8))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.numpair.ncons_|49||49|','thm.numpair.ncons_def')).
fof(f987,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK3),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numpair.npair_2'(s('type.num.num',sK1),s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f820,f359])).
fof(f359,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X0),s('type.num.num',X1))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.num.num',X2) = s('type.num.num','const.numpair.npair_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(superposition,[],[f42,f40])).
fof(f42,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2] : ((s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))),
  inference(nnf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1,X2] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10,X11,X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X12))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X12))) <=> s('type.num.num',X10) = s('type.num.num',X11))),
  file('thm.numpair.ncons_|49||49|','thm.arithmetic.EQ_MONO_ADD_EQ')).
fof(f820,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.npair_2'(s('type.num.num',sK1),s('type.num.num',X0))) != s('type.num.num','const.numpair.npair_2'(s('type.num.num',sK3),s('type.num.num',X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f51,f44])).
fof(f44,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.numpair.npair_2'(s('type.num.num',X2),s('type.num.num',X0))) != s('type.num.num','const.numpair.npair_2'(s('type.num.num',X3),s('type.num.num',X1))) | s('type.num.num',X2) = s('type.num.num',X3)) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2,X3] : ((s('type.num.num','const.numpair.npair_2'(s('type.num.num',X2),s('type.num.num',X0))) != s('type.num.num','const.numpair.npair_2'(s('type.num.num',X3),s('type.num.num',X1))) | (s('type.num.num',X2) = s('type.num.num',X3) & s('type.num.num',X0) = s('type.num.num',X1))) & (s('type.num.num',X2) != s('type.num.num',X3) | s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.numpair.npair_2'(s('type.num.num',X2),s('type.num.num',X0))) = s('type.num.num','const.numpair.npair_2'(s('type.num.num',X3),s('type.num.num',X1)))))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2,X3] : ((s('type.num.num','const.numpair.npair_2'(s('type.num.num',X2),s('type.num.num',X0))) != s('type.num.num','const.numpair.npair_2'(s('type.num.num',X3),s('type.num.num',X1))) | (s('type.num.num',X2) = s('type.num.num',X3) & s('type.num.num',X0) = s('type.num.num',X1))) & ((s('type.num.num',X2) != s('type.num.num',X3) | s('type.num.num',X0) != s('type.num.num',X1)) | s('type.num.num','const.numpair.npair_2'(s('type.num.num',X2),s('type.num.num',X0))) = s('type.num.num','const.numpair.npair_2'(s('type.num.num',X3),s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0,X1,X2,X3] : (s('type.num.num','const.numpair.npair_2'(s('type.num.num',X2),s('type.num.num',X0))) = s('type.num.num','const.numpair.npair_2'(s('type.num.num',X3),s('type.num.num',X1))) <=> (s('type.num.num',X2) = s('type.num.num',X3) & s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X14,X15,X16,X17] : (s('type.num.num','const.numpair.npair_2'(s('type.num.num',X16),s('type.num.num',X14))) = s('type.num.num','const.numpair.npair_2'(s('type.num.num',X17),s('type.num.num',X15))) <=> (s('type.num.num',X16) = s('type.num.num',X17) & s('type.num.num',X14) = s('type.num.num',X15)))),
  file('thm.numpair.ncons_|49||49|','thm.numpair.npair_|49||49|')).
fof(f815,plain,(
  s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK3),s('type.num.num',sK0))) | ($spl0 | $spl4)),
  inference(backward_demodulation,[],[f52,f48])).
fof(f1059,plain,(
  ( ! [X8,X9] : (s(fun(bool,bool),X8) = s(fun(bool,bool),X9) | s(bool,'T_0') = s(bool,sK4(X8,X9,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,sK4(X8,X9,bool,bool))))) ) | $spl18),
  inference(cnf_transformation,[],[f1059_D])).
fof(f1059_D,plain,(
  ( ! [X8,X9] : (s(fun(bool,bool),X8) = s(fun(bool,bool),X9) | s(bool,'T_0') = s(bool,sK4(X8,X9,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,sK4(X8,X9,bool,bool))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1062,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl22),
  inference(cnf_transformation,[],[f1062_D])).
fof(f1062_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1034,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'F_0') != s(bool,f(s(bool,sK4(X14,X15,bool,bool)))) | s(bool,'T_0') = s(bool,sK4(X14,X15,bool,bool))) )),
  inference(superposition,[],[f432,f37])).
fof(f37,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.numpair.ncons_|49||49|',aHL_BOOL_CASES)).
fof(f432,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,g(s(X16,sK4(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'F_0') != s(bool,f(s(X16,sK4(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f41,f37])).
fof(f41,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.numpair.ncons_|49||49|',aHL_EXT)).
fof(f1060,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,g(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) ) | $spl20),
  inference(cnf_transformation,[],[f1060_D])).
fof(f1060_D,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,g(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1033,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') = s(bool,g(s(bool,X13))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'F_0') != s(bool,f(s(bool,sK4(X11,X12,bool,bool)))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,sK4(X11,X12,bool,bool))) )),
  inference(superposition,[],[f432,f58])).
fof(f58,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f37,f37])).
fof(f1032,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,g(s(bool,X10))) | s(fun(bool,bool),X8) = s(fun(bool,bool),X9) | s(bool,'F_0') != s(bool,f(s(bool,sK4(X8,X9,bool,bool)))) | s(bool,'T_0') = s(bool,sK4(X8,X9,bool,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f432,f58])).
fof(f1058,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK4(X2,X3,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,sK4(X2,X3,bool,bool))))) ) | $spl16),
  inference(cnf_transformation,[],[f1058_D])).
fof(f1058_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK4(X2,X3,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,sK4(X2,X3,bool,bool))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1057,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl14),
  inference(cnf_transformation,[],[f1057_D])).
fof(f1057_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1031,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'F_0') != s(bool,f(s(bool,sK4(X5,X6,bool,bool)))) | s(bool,X7) = s(bool,sK4(X5,X6,bool,bool)) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f432,f58])).
fof(f1030,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'F_0') != s(bool,f(s(bool,sK4(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK4(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f432,f58])).
fof(f1054,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0'))) | $spl11),
  inference(cnf_transformation,[],[f1054_D])).
fof(f1054_D,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0'))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f1055,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK4(X0,X1,bool,bool))))) ) | $spl12),
  inference(cnf_transformation,[],[f1055_D])).
fof(f1055_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK4(X0,X1,bool,bool))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1052,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK4(X0,X1,bool,bool)))) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f1029])).
fof(f1029,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK4(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f432,f427])).
fof(f427,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,sK4(X16,X17,X18,bool)) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f41,f37])).
fof(f814,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f48,f758])).
fof(f758,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X0),s('type.num.num',sK2))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X1),s('type.num.num',sK0)))) ) | $spl5),
  inference(superposition,[],[f651,f40])).
fof(f651,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X0),s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numpair.npair_2'(s('type.num.num',X1),s('type.num.num',sK2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl5),
  inference(superposition,[],[f332,f40])).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numpair.npair_2'(s('type.num.num',X0),s('type.num.num',sK0))),s('type.num.num',X1))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numpair.npair_2'(s('type.num.num',X2),s('type.num.num',sK2))),s('type.num.num',X1)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f309,f42])).
fof(f309,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numpair.npair_2'(s('type.num.num',X0),s('type.num.num',sK2))) != s('type.num.num','const.numpair.npair_2'(s('type.num.num',X1),s('type.num.num',sK0)))) ) | $spl5),
  inference(unit_resulting_resolution,[],[f53,f45])).
fof(f45,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num','const.numpair.npair_2'(s('type.num.num',X2),s('type.num.num',X0))) != s('type.num.num','const.numpair.npair_2'(s('type.num.num',X3),s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f30])).
fof(f800,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(unit_resulting_resolution,[],[f57,f758])).
fof(f57,plain,(
  s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f48,f50])).
fof(f64,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(cnf_transformation,[],[f64_D])).
fof(f64_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f66,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl9),
  inference(cnf_transformation,[],[f66_D])).
fof(f66_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f63,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f37])).
fof(f56,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f55])).
fof(f55,plain,(
  s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f54,f50])).
fof(f54,plain,(
  s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK3),s('type.num.num',sK0))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f49,f52])).
fof(f48,plain,(
  s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK3),s('type.num.num',sK2))) | $spl0),
  inference(cnf_transformation,[],[f48_D])).
fof(f48_D,plain,(
  s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK3),s('type.num.num',sK2))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f50,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK3) | $spl2),
  inference(cnf_transformation,[],[f50_D])).
fof(f50_D,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK3) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f31,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK3) | s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK3),s('type.num.num',sK2)))),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK3),s('type.num.num',sK2))) | (s('type.num.num',sK1) = s('type.num.num',sK3) & s('type.num.num',sK0) = s('type.num.num',sK2))) & (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK3),s('type.num.num',sK2))) | s('type.num.num',sK1) != s('type.num.num',sK3) | s('type.num.num',sK0) != s('type.num.num',sK2))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f24])).
fof(f24,plain,(
  ? [X0,X1,X2,X3] : ((s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X2))) | (s('type.num.num',X1) = s('type.num.num',X3) & s('type.num.num',X0) = s('type.num.num',X2))) & (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X2))) | s('type.num.num',X1) != s('type.num.num',X3) | s('type.num.num',X0) != s('type.num.num',X2)))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ? [X0,X1,X2,X3] : ((s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X2))) | (s('type.num.num',X1) = s('type.num.num',X3) & s('type.num.num',X0) = s('type.num.num',X2))) & (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X2))) | (s('type.num.num',X1) != s('type.num.num',X3) | s('type.num.num',X0) != s('type.num.num',X2))))),
  inference(nnf_transformation,[],[f21])).
fof(f21,plain,(
  ? [X0,X1,X2,X3] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X2))) <~> (s('type.num.num',X1) = s('type.num.num',X3) & s('type.num.num',X0) = s('type.num.num',X2)))),
  inference(ennf_transformation,[],[f12])).
fof(f12,plain,(
  ~! [X0,X1,X2,X3] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X3),s('type.num.num',X2))) <=> (s('type.num.num',X1) = s('type.num.num',X3) & s('type.num.num',X0) = s('type.num.num',X2)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X7),s('type.num.num',X6))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X9),s('type.num.num',X8))) <=> (s('type.num.num',X7) = s('type.num.num',X9) & s('type.num.num',X6) = s('type.num.num',X8)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9] : (s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X7),s('type.num.num',X6))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',X9),s('type.num.num',X8))) <=> (s('type.num.num',X7) = s('type.num.num',X9) & s('type.num.num',X6) = s('type.num.num',X8)))),
  file('thm.numpair.ncons_|49||49|','thm.numpair.ncons_|49||49|')).
fof(f52,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK2) | $spl4),
  inference(cnf_transformation,[],[f52_D])).
fof(f52_D,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK2) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f32,plain,(
  s('type.num.num',sK0) = s('type.num.num',sK2) | s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK3),s('type.num.num',sK2)))),
  inference(cnf_transformation,[],[f25])).
fof(f49,plain,(
  s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK3),s('type.num.num',sK2))) | $spl1),
  inference(cnf_transformation,[],[f49_D])).
fof(f49_D,plain,(
  s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK3),s('type.num.num',sK2))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f51,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK3) | $spl3),
  inference(cnf_transformation,[],[f51_D])).
fof(f51_D,plain,(
  s('type.num.num',sK1) != s('type.num.num',sK3) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f53,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK2) | $spl5),
  inference(cnf_transformation,[],[f53_D])).
fof(f53_D,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK2) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f33,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK2) | s('type.num.num',sK1) != s('type.num.num',sK3) | s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK1),s('type.num.num',sK0))) != s('type.num.num','const.numpair.ncons_2'(s('type.num.num',sK3),s('type.num.num',sK2)))),
  inference(cnf_transformation,[],[f25])).
% SZS output end Proof for thm.numpair
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 55
% Passive clauses: 149
% Generated clauses: 1272
% Final active clauses: 35
% Final passive clauses: 45
% Input formulas: 10
% Initial clauses: 16
% 
% Duplicate literals: 52
% Trivial inequalities: 4
% Fw subsumption resolutions: 1
% Fw demodulations: 188
% Bw demodulations: 1
% Condensations: 40
% 
% Simple tautologies: 16
% Equational tautologies: 26
% Forward subsumptions: 686
% Fw demodulations to eq. taut.: 96
% 
% Unit resulting resolution: 453
% Forward superposition: 333
% Backward superposition: 113
% Self superposition: 24
% Equality factoring: 26
% Equality resolution: 9
% 
% Split clauses: 10
% Split components: 22
% SAT solver clauses: 175
% SAT solver unit clauses: 3
% SAT solver binary clauses: 137
% 
% Sat splits: 10
% Sat splitting refutations: 4
% 
% Memory used [KB]: 1791
% Time elapsed: 0.043 s
% ------------------------------
% ------------------------------
% Success in time 1.36 s
