Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.util_prob
% (54576)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 143
% Passive clauses: 3408
% Generated clauses: 24323
% Final active clauses: 124
% Final passive clauses: 3038
% Input formulas: 25
% Initial clauses: 37
% 
% Duplicate literals: 829
% Trivial inequalities: 397
% Fw subsumption resolutions: 2578
% Fw demodulations: 3944
% Bw demodulations: 15
% 
% Simple tautologies: 414
% Equational tautologies: 42
% Forward subsumptions: 12239
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 275
% Factoring: 4
% Forward superposition: 8601
% Backward superposition: 7308
% Self superposition: 162
% Equality factoring: 75
% Equality resolution: 31
% 
% Split clauses: 197
% Split components: 409
% SAT solver clauses: 612
% SAT solver unit clauses: 14
% SAT solver binary clauses: 132
% 
% Sat splits: 197
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 609
% TWLsolver calls for satisfiability: 128
% 
% Memory used [KB]: 12920
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.util_prob
% (54592)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 107
% Passive clauses: 1006
% Generated clauses: 13390
% Final active clauses: 93
% Final passive clauses: 880
% Input formulas: 25
% Initial clauses: 37
% 
% Duplicate literals: 531
% Trivial inequalities: 169
% Fw subsumption resolutions: 816
% Fw demodulations: 1609
% Bw demodulations: 23
% Condensations: 1210
% Global subsumptions: 5
% 
% Simple tautologies: 165
% Equational tautologies: 39
% Forward subsumptions: 6049
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 169
% Factoring: 2
% Forward superposition: 4375
% Backward superposition: 3168
% Self superposition: 129
% Equality factoring: 74
% Equality resolution: 19
% 
% Split clauses: 86
% Split components: 182
% SAT solver clauses: 1835
% SAT solver unit clauses: 56
% SAT solver binary clauses: 39
% SAT solver learnt clauses: 10
% 
% Sat splits: 86
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1814
% TWLsolver calls for satisfiability: 219
% 
% Memory used [KB]: 12025
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.util_prob
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.util_prob
% SZS output start Proof for thm.util_prob
fof(f7647,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f383,f387,f642,f684,f1017,f385,f1034,f1045,f1037,f1035,f3761,f3762,f3763,f3764,f3780,f3773,f3765,f3779,f3777,f3766,f3767,f3768,f3784,f3783,f3769,f3776,f3774,f3781,f4515,f4601,f4892,f4893,f4894,f4898,f4900,f4903,f4906,f4909,f4912,f4926,f4927,f4928,f4929,f4930,f4931,f4933,f4935,f4940,f4942,f4944,f4946,f4948,f4950,f4952,f4954,f4956,f4958,f7229,f7307,f4600,f7375,f7381,f7382,f7384,f7385,f7386,f7426,f7428,f7445,f7447,f7594,f7595,f7596,f7600,f7602,f7605,f7608,f7611,f7614,f7618,f7619,f7621,f7623,f7628,f7630,f7632,f7634,f7636,f7638,f7640,f7642,f7644,f7646])).
fof(f7646,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599,f7601,f7604,f7607,f7620,f7622,f7627,f7629,f7631,f7633,f7635,f7637,f7639,f7641,f7643,f7645])).
fof(f7645,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7583,f93])).
fof(f7583,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | ($spl2 | $spl4)),
  inference(condensation,[],[f7582])).
fof(f7582,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl2 | $spl4)),
  inference(condensation,[],[f7574])).
fof(f7574,plain,(
  ( ! [X94,X93,X91] : (s(bool,'T_0') = s(bool,X94) | p(s(bool,X93)) | p(s(bool,X94)) | s(bool,'T_0') = s(bool,X91)) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f2513])).
fof(f7643,plain,(
  ( ! [X17,X15,X13] : (s(fun(X13,bool),X15) = s(fun(X13,bool),X17)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7569,f7449])).
fof(f7569,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X13,sK4(happ(s(fun(X14,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X14,X16)),X17,X13)),s(fun(X13,bool),X15))) | s(fun(X13,bool),X15) = s(fun(X13,bool),X17)) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f1089])).
fof(f7641,plain,(
  ( ! [X4,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X4)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7568,f7449])).
fof(f7568,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK4('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X4,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X4)) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f1086])).
fof(f7639,plain,(
  ( ! [X6,X8,X5] : (s(fun(X5,bool),X6) = s(fun(X5,bool),X8)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7567,f7449])).
fof(f7567,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X5,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5)),s(fun(X5,bool),X6))) | s(fun(X5,bool),X6) = s(fun(X5,bool),X8)) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f1087])).
fof(f7637,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(fun(X5,bool),X9) = s(fun(X5,bool),'const.util_prob.PREIMAGE_2'(s(fun(X5,X6),X7),s(fun(X6,bool),X8)))) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7565,f7449])).
fof(f7565,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(X5,X6),X7),s(X5,sK4('const.util_prob.PREIMAGE_2'(s(fun(X5,X6),X7),s(fun(X6,bool),X8)),X9,X5)))),s(fun(X6,bool),X8))) | s(fun(X5,bool),X9) = s(fun(X5,bool),'const.util_prob.PREIMAGE_2'(s(fun(X5,X6),X7),s(fun(X6,bool),X8)))) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f1263])).
fof(f7635,plain,(
  ( ! [X17,X15,X13] : (s(fun(X13,bool),X15) = s(fun(X13,bool),X17)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7564,f7449])).
fof(f7564,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X13,sK4(X17,happ(s(fun(X14,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X14,X16)),X13)),s(fun(X13,bool),X15))) | s(fun(X13,bool),X15) = s(fun(X13,bool),X17)) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f1085])).
fof(f7633,plain,(
  ( ! [X4,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X4)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7563,f7449])).
fof(f7563,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK4(X4,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X4)) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f1082])).
fof(f7631,plain,(
  ( ! [X6,X8,X5] : (s(fun(X5,bool),X6) = s(fun(X5,bool),X8)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7562,f7449])).
fof(f7562,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X5,sK4(X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X5)),s(fun(X5,bool),X6))) | s(fun(X5,bool),X6) = s(fun(X5,bool),X8)) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f1083])).
fof(f7629,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X2),X3),s(fun(X2,bool),X4)))) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7560,f7449])).
fof(f7560,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X2,happ(s(fun(X0,X2),X3),s(X0,sK4(X1,'const.util_prob.PREIMAGE_2'(s(fun(X0,X2),X3),s(fun(X2,bool),X4)),X0)))),s(fun(X2,bool),X4))) | s(fun(X0,bool),X1) = s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X2),X3),s(fun(X2,bool),X4)))) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f1262])).
fof(f7627,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7559,f7449])).
fof(f7559,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f110])).
fof(f7622,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(fun(X10,bool),X14) = s(fun(X10,bool),'const.util_prob.PREIMAGE_2'(s(fun(X10,X11),'const.combin.K_1'(s(X11,X12))),s(fun(X11,bool),X13)))) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7543,f7449])).
fof(f7543,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X10,sK4('const.util_prob.PREIMAGE_2'(s(fun(X10,X11),'const.combin.K_1'(s(X11,X12))),s(fun(X11,bool),X13)),X14,X10)),s(fun(X10,bool),X14))) | s(fun(X10,bool),X14) = s(fun(X10,bool),'const.util_prob.PREIMAGE_2'(s(fun(X10,X11),'const.combin.K_1'(s(X11,X12))),s(fun(X11,bool),X13)))) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f4142])).
fof(f7620,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(fun(X5,bool),X6) = s(fun(X5,bool),'const.util_prob.PREIMAGE_2'(s(fun(X5,X7),'const.combin.K_1'(s(X7,X8))),s(fun(X7,bool),X9)))) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7542,f7449])).
fof(f7542,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X5,sK4(X6,'const.util_prob.PREIMAGE_2'(s(fun(X5,X7),'const.combin.K_1'(s(X7,X8))),s(fun(X7,bool),X9)),X5)),s(fun(X5,bool),X6))) | s(fun(X5,bool),X6) = s(fun(X5,bool),'const.util_prob.PREIMAGE_2'(s(fun(X5,X7),'const.combin.K_1'(s(X7,X8))),s(fun(X7,bool),X9)))) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f4141])).
fof(f7607,plain,(
  ( ! [X12,X11,X9] : (s(fun(X9,bool),X11) = s(fun(X9,bool),X12)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7606,f7449])).
fof(f7606,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X9,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X12,X9)),s(fun(X9,bool),X11))) | s(fun(X9,bool),X11) = s(fun(X9,bool),X12)) ) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f7478,f7449])).
fof(f7478,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'const.bool.IN_2'(s(X9,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X12,X9)),s(fun(X9,bool),X11))) != s(bool,'const.bool.IN_2'(s(X9,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X12,X9)),s(fun(X9,bool),X12))) | s(fun(X9,bool),X11) = s(fun(X9,bool),X12)) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f1088])).
fof(f7604,plain,(
  ( ! [X12,X11,X9] : (s(fun(X9,bool),X11) = s(fun(X9,bool),X12)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7603,f7449])).
fof(f7603,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X9,sK4(X12,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X9)),s(fun(X9,bool),X11))) | s(fun(X9,bool),X11) = s(fun(X9,bool),X12)) ) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f7477,f7449])).
fof(f7477,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'const.bool.IN_2'(s(X9,sK4(X12,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X9)),s(fun(X9,bool),X11))) != s(bool,'const.bool.IN_2'(s(X9,sK4(X12,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X9)),s(fun(X9,bool),X12))) | s(fun(X9,bool),X11) = s(fun(X9,bool),X12)) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f1084])).
fof(f7601,plain,(
  ( ! [X2,X3] : (s(fun(X2,bool),'const.pred_set.EMPTY_0') = s(fun(X2,bool),X3)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7475,f7449])).
fof(f7475,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X2,sK4('const.pred_set.EMPTY_0',X3,X2)),s(fun(X2,bool),X3))) | s(fun(X2,bool),'const.pred_set.EMPTY_0') = s(fun(X2,bool),X3)) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f1073])).
fof(f7599,plain,(
  ( ! [X2,X3] : (s(fun(X2,bool),'const.pred_set.EMPTY_0') = s(fun(X2,bool),X3)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7472,f7449])).
fof(f7472,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X2,sK4(X3,'const.pred_set.EMPTY_0',X2)),s(fun(X2,bool),X3))) | s(fun(X2,bool),'const.pred_set.EMPTY_0') = s(fun(X2,bool),X3)) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f1057])).
fof(f7452,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f108])).
fof(f7593,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7450,f7449])).
fof(f7450,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f97])).
fof(f7449,plain,(
  ( ! [X60] : (s(bool,'T_0') = s(bool,X60)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7448,f281])).
fof(f7448,plain,(
  ( ! [X60] : (s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),'const.bool.COND_3'(s(bool,X60),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X60)) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f386,f1017])).
fof(f386,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) | $spl2),
  inference(cnf_transformation,[],[f386_D])).
fof(f386_D,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') = s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f7396,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3)))),
  inference(superposition,[],[f82,f281])).
fof(f82,plain,(
  s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) != s(fun(sK1,bool),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) != s(fun(sK1,bool),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f63])).
fof(f63,plain,(
  ? [X0,X1,X2,X3] : s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(fun(X0,bool),X3))) != s(fun(X1,bool),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X3))),s(fun(X1,bool),'const.pred_set.UNIV_0'),s(fun(X1,bool),'const.pred_set.EMPTY_0')))),
  inference(ennf_transformation,[],[f28])).
fof(f28,plain,(
  ~! [X0,X1,X2,X3] : s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(fun(X0,bool),X3))) = s(fun(X1,bool),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X3))),s(fun(X1,bool),'const.pred_set.UNIV_0'),s(fun(X1,bool),'const.pred_set.EMPTY_0')))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ~! [X0,X1] : ! [X2,X3] : s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(fun(X0,bool),X3))) = s(fun(X1,bool),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X3))),s(fun(X1,bool),'const.pred_set.UNIV_0'),s(fun(X1,bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : s(fun(X7,bool),'const.util_prob.PREIMAGE_2'(s(fun(X7,X6),'const.combin.K_1'(s(X6,X8))),s(fun(X6,bool),X9))) = s(fun(X7,bool),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),X9))),s(fun(X7,bool),'const.pred_set.UNIV_0'),s(fun(X7,bool),'const.pred_set.EMPTY_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : s(fun(X7,bool),'const.util_prob.PREIMAGE_2'(s(fun(X7,X6),'const.combin.K_1'(s(X6,X8))),s(fun(X6,bool),X9))) = s(fun(X7,bool),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),X9))),s(fun(X7,bool),'const.pred_set.UNIV_0'),s(fun(X7,bool),'const.pred_set.EMPTY_0')))),
  file('thm.util_prob.PREIMAGE_K','thm.util_prob.PREIMAGE_K')).
fof(f431,plain,(
  s(bool,f(s(sK1,sK5('const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),bool,sK1)))) != s(bool,g(s(sK1,sK5('const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),bool,sK1))))),
  inference(unit_resulting_resolution,[],[f82,f117])).
fof(f433,plain,(
  s(bool,f(s(sK1,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),bool,sK1)))) != s(bool,g(s(sK1,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),bool,sK1))))),
  inference(unit_resulting_resolution,[],[f82,f117])).
fof(f7398,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),sK1)),s(fun(sK1,bool),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))))),
  inference(forward_demodulation,[],[f7397,f120])).
fof(f7397,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,happ(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(sK1,sK4('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),sK1)))),s(fun(sK0,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),sK1)),s(fun(sK1,bool),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))))),
  inference(forward_demodulation,[],[f7390,f118])).
fof(f7390,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK4('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),sK1)),s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),sK1)),s(fun(sK1,bool),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))))),
  inference(unit_resulting_resolution,[],[f82,f110])).
fof(f7400,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),sK1)),s(fun(sK1,bool),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))))),
  inference(forward_demodulation,[],[f7399,f120])).
fof(f7399,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,happ(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(sK1,sK4('const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),sK1)))),s(fun(sK0,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),sK1)),s(fun(sK1,bool),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))))),
  inference(forward_demodulation,[],[f7388,f118])).
fof(f7388,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK4('const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),sK1)),s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))))) != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),sK1)),s(fun(sK1,bool),'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))))),
  inference(unit_resulting_resolution,[],[f82,f110])).
fof(f7343,plain,(
  ( ! [X39,X41,X40] : (s(bool,X41) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40))) | s(bool,'F_0') = s(bool,X41) | s(bool,'F_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f2158,f83])).
fof(f2158,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'T_0')) | s(bool,X41) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40))) | s(bool,'F_0') = s(bool,X41) | s(bool,'F_0') = s(bool,X39)) )),
  inference(superposition,[],[f503,f210])).
fof(f7342,plain,(
  ( ! [X52,X50,X51] : (s(bool,X52) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X50),s(bool,X51))) | s(bool,'F_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f2834,f83])).
fof(f2834,plain,(
  ( ! [X52,X50,X51] : (~p(s(bool,'T_0')) | s(bool,X52) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X50),s(bool,X51))) | s(bool,'F_0') = s(bool,X52) | s(bool,'F_0') = s(bool,X51)) )),
  inference(superposition,[],[f503,f297])).
fof(f3329,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,'T_0') = s(bool,happ(s(fun(X115,bool),'const.combin.K_1'(s(bool,X113))),s(X115,X116))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X113),s(bool,X114)))) )),
  inference(superposition,[],[f355,f107])).
fof(f3330,plain,(
  ( ! [X118,X120,X119,X117] : (s(bool,'T_0') = s(bool,happ(s(fun(X119,bool),'const.combin.K_1'(s(bool,X118))),s(X119,X120))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X117),s(bool,X118)))) )),
  inference(superposition,[],[f355,f108])).
fof(f7341,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X79) = s(bool,happ(s(fun(X76,bool),'const.combin.K_1'(s(bool,X77))),s(X76,X78))) | s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f3392,f83])).
fof(f3392,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,'T_0')) | s(bool,X79) = s(bool,happ(s(fun(X76,bool),'const.combin.K_1'(s(bool,X77))),s(X76,X78))) | s(bool,'F_0') = s(bool,X79) | s(bool,'F_0') = s(bool,X77)) )),
  inference(superposition,[],[f503,f355])).
fof(f97,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.util_prob.PREIMAGE_K','thm.bool.EQ_CLAUSES|split|3')).
fof(f98,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f73])).
fof(f108,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X10,X11] : s(X6,X11) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X10),s(X6,X11)))),
  file('thm.util_prob.PREIMAGE_K','thm.bool.COND_CLAUSES|split|1')).
fof(f119,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f86,f83])).
fof(f86,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f29])).
fof(f29,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f13])).
fof(f13,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.util_prob.PREIMAGE_K','thm.bool.NOT_CLAUSES|split|1')).
fof(f132,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f106,f98])).
fof(f153,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X2),s(X3,X4))) | s(bool,'F_0') = s(bool,X2)) )),
  inference(superposition,[],[f91,f115])).
fof(f210,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X3),s(bool,X4))) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f107,f91])).
fof(f297,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,X4))) | s(bool,'F_0') = s(bool,X4)) )),
  inference(superposition,[],[f108,f91])).
fof(f355,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),'const.combin.K_1'(s(bool,X5))),s(X4,X6))) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f120,f91])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,X2,bool))) )),
  inference(superposition,[],[f117,f91])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK5(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK5(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f117,f91])).
fof(f462,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X19),s(X22,X23))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X19),s(X20,X21)))) )),
  inference(superposition,[],[f153,f115])).
fof(f463,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X24),s(X26,X27))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X24),s(bool,X25)))) )),
  inference(superposition,[],[f153,f107])).
fof(f464,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X29),s(X30,X31))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X28),s(bool,X29)))) )),
  inference(superposition,[],[f153,f108])).
fof(f465,plain,(
  ( ! [X35,X33,X36,X34,X32] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X33),s(X35,X36))) | s(bool,'F_0') = s(bool,happ(s(fun(X32,bool),'const.combin.K_1'(s(bool,X33))),s(X32,X34)))) )),
  inference(superposition,[],[f153,f120])).
fof(f503,plain,(
  ( ! [X26,X23] : (~p(s(bool,X26)) | s(bool,X23) = s(bool,X26) | s(bool,'F_0') = s(bool,X23)) )),
  inference(forward_demodulation,[],[f502,f115])).
fof(f502,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,X26)) | s(bool,X26) = s(bool,'const.combin.K_2'(s(bool,X23),s(X24,X25))) | s(bool,'F_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f484,f83])).
fof(f484,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,X26)) | s(bool,X26) = s(bool,'const.combin.K_2'(s(bool,X23),s(X24,X25))) | s(bool,'F_0') = s(bool,X23)) )),
  inference(superposition,[],[f114,f153])).
fof(f504,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X30) != s(bool,'const.combin.K_2'(s(bool,X27),s(X28,X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f485,f83])).
fof(f485,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | s(bool,X30) != s(bool,'const.combin.K_2'(s(bool,X27),s(X28,X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X27)) )),
  inference(superposition,[],[f154,f153])).
fof(f546,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.UNIV_0'))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f537,f83])).
fof(f537,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.UNIV_0'))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f503,f129])).
fof(f547,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,X13) = s(bool,'const.combin.K_2'(s(bool,X10),s(X11,X12))) | s(bool,'F_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f539,f83])).
fof(f539,plain,(
  ( ! [X12,X10,X13,X11] : (~p(s(bool,'T_0')) | s(bool,X13) = s(bool,'const.combin.K_2'(s(bool,X10),s(X11,X12))) | s(bool,'F_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X10)) )),
  inference(superposition,[],[f503,f153])).
fof(f573,plain,(
  ( ! [X39,X37,X38] : (s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X37),s(bool,X38)))) )),
  inference(superposition,[],[f151,f108])).
fof(f603,plain,(
  ( ! [X30,X31,X29] : (s(bool,'T_0') = s(bool,X30) | s(bool,X31) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X29),s(bool,X30))) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f151,f108])).
fof(f727,plain,(
  ( ! [X76,X74,X75] : (s(bool,X74) = s(bool,'const.combin.K_2'(s(bool,X75),s(bool,X74))) | s(bool,'F_0') = s(bool,X75) | s(bool,X74) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f153,f151])).
fof(f1057,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X2,sK4(X3,'const.pred_set.EMPTY_0',X2)),s(fun(X2,bool),X3))) | s(fun(X2,bool),'const.pred_set.EMPTY_0') = s(fun(X2,bool),X3)) )),
  inference(superposition,[],[f110,f132])).
fof(f1062,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X20,sK4(X21,X22,X20)),s(fun(X20,bool),X21))) | s(fun(X20,bool),X21) = s(fun(X20,bool),X22) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X20,sK4(X21,X22,X20)),s(fun(X20,bool),X22)))) )),
  inference(superposition,[],[f110,f91])).
fof(f1067,plain,(
  ( ! [X12,X13] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X12))) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X13))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK4(X12,X13,bool))) )),
  inference(superposition,[],[f110,f91])).
fof(f1073,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X2,sK4('const.pred_set.EMPTY_0',X3,X2)),s(fun(X2,bool),X3))) | s(fun(X2,bool),'const.pred_set.EMPTY_0') = s(fun(X2,bool),X3)) )),
  inference(superposition,[],[f110,f132])).
fof(f1078,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X20,sK4(X21,X22,X20)),s(fun(X20,bool),X22))) | s(fun(X20,bool),X21) = s(fun(X20,bool),X22) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X20,sK4(X21,X22,X20)),s(fun(X20,bool),X21)))) )),
  inference(superposition,[],[f110,f91])).
fof(f1084,plain,(
  ( ! [X12,X10,X11,X9] : (s(fun(X9,bool),X11) = s(fun(X9,bool),X12) | s(bool,'const.bool.IN_2'(s(X9,sK4(X12,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X9)),s(fun(X9,bool),X11))) != s(bool,'const.bool.IN_2'(s(X9,sK4(X12,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X9)),s(fun(X9,bool),X12)))) )),
  inference(forward_demodulation,[],[f1054,f108])).
fof(f1054,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'const.bool.IN_2'(s(X9,sK4(X12,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X9)),s(fun(X9,bool),X11))) != s(bool,'const.bool.IN_2'(s(X9,sK4(X12,'const.bool.COND_3'(s(bool,'F_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X9)),s(fun(X9,bool),X12))) | s(fun(X9,bool),X12) = s(fun(X9,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)))) )),
  inference(superposition,[],[f110,f108])).
fof(f1088,plain,(
  ( ! [X12,X10,X11,X9] : (s(fun(X9,bool),X11) = s(fun(X9,bool),X12) | s(bool,'const.bool.IN_2'(s(X9,sK4('const.bool.COND_3'(s(bool,'F_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X12,X9)),s(fun(X9,bool),X11))) != s(bool,'const.bool.IN_2'(s(X9,sK4('const.bool.COND_3'(s(bool,'F_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X12,X9)),s(fun(X9,bool),X12)))) )),
  inference(forward_demodulation,[],[f1070,f108])).
fof(f1070,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'const.bool.IN_2'(s(X9,sK4('const.bool.COND_3'(s(bool,'F_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X12,X9)),s(fun(X9,bool),X11))) != s(bool,'const.bool.IN_2'(s(X9,sK4('const.bool.COND_3'(s(bool,'F_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X12,X9)),s(fun(X9,bool),X12))) | s(fun(X9,bool),X12) = s(fun(X9,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)))) )),
  inference(superposition,[],[f110,f108])).
fof(f1191,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X10,X12),s(fun(X10,bool),'const.util_prob.PREIMAGE_2'(s(fun(X10,X9),X11),s(fun(X9,bool),'const.pred_set.EMPTY_0')))))) )),
  inference(superposition,[],[f118,f132])).
fof(f1197,plain,(
  ( ! [X47,X45,X48,X46,X44] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X45,X47),s(fun(X45,bool),'const.util_prob.PREIMAGE_2'(s(fun(X45,X44),X46),s(fun(X44,bool),X48))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X44,happ(s(fun(X45,X44),X46),s(X45,X47))),s(fun(X44,bool),X48)))) )),
  inference(superposition,[],[f118,f91])).
fof(f1229,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X38,happ(s(fun(X36,X38),X39),s(X36,X37))),s(fun(X38,bool),X40))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),'const.util_prob.PREIMAGE_2'(s(fun(X36,X38),X39),s(fun(X38,bool),X40)))))) )),
  inference(superposition,[],[f118,f91])).
fof(f1236,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X28,X30),s(fun(X28,bool),'const.util_prob.PREIMAGE_2'(s(fun(X28,X27),X29),s(fun(X27,bool),X31))))) | ~p(s(bool,'const.bool.IN_2'(s(X27,happ(s(fun(X28,X27),X29),s(X28,X30))),s(fun(X27,bool),X31))))) )),
  inference(superposition,[],[f97,f118])).
fof(f1237,plain,(
  ( ! [X35,X33,X36,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(X33,X35),s(fun(X33,bool),'const.util_prob.PREIMAGE_2'(s(fun(X33,X32),X34),s(fun(X32,bool),X36)))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X32,happ(s(fun(X33,X32),X34),s(X33,X35))),s(fun(X32,bool),X36)))) )),
  inference(superposition,[],[f98,f118])).
fof(f1245,plain,(
  ( ! [X80,X85,X83,X81,X79,X84,X82] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.bool.IN_2'(s(X80,X82),s(fun(X80,bool),'const.util_prob.PREIMAGE_2'(s(fun(X80,X79),X81),s(fun(X79,bool),X83))))),s(X84,X85))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X79,happ(s(fun(X80,X79),X81),s(X80,X82))),s(fun(X79,bool),X83)))) )),
  inference(superposition,[],[f153,f118])).
fof(f1248,plain,(
  ( ! [X103,X101,X99,X102,X100,X98] : (~p(s(bool,'const.bool.IN_2'(s(X99,X101),s(fun(X99,bool),'const.util_prob.PREIMAGE_2'(s(fun(X99,X98),X100),s(fun(X98,bool),X102)))))) | s(bool,X103) = s(bool,'const.bool.IN_2'(s(X98,happ(s(fun(X99,X98),X100),s(X99,X101))),s(fun(X98,bool),X102))) | s(bool,'F_0') = s(bool,X103)) )),
  inference(superposition,[],[f503,f118])).
fof(f1267,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X27,happ(s(fun(X25,X27),X28),s(X25,X26))),s(fun(X27,bool),X29))) | ~p(s(bool,'const.bool.IN_2'(s(X25,X26),s(fun(X25,bool),'const.util_prob.PREIMAGE_2'(s(fun(X25,X27),X28),s(fun(X27,bool),X29))))))) )),
  inference(superposition,[],[f97,f118])).
fof(f1268,plain,(
  ( ! [X30,X33,X31,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(X32,happ(s(fun(X30,X32),X33),s(X30,X31))),s(fun(X32,bool),X34)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),'const.util_prob.PREIMAGE_2'(s(fun(X30,X32),X33),s(fun(X32,bool),X34)))))) )),
  inference(superposition,[],[f98,f118])).
fof(f1276,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.bool.IN_2'(s(X79,happ(s(fun(X77,X79),X80),s(X77,X78))),s(fun(X79,bool),X81))),s(X82,X83))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X77,X78),s(fun(X77,bool),'const.util_prob.PREIMAGE_2'(s(fun(X77,X79),X80),s(fun(X79,bool),X81)))))) )),
  inference(superposition,[],[f153,f118])).
fof(f1279,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (~p(s(bool,'const.bool.IN_2'(s(X98,happ(s(fun(X96,X98),X99),s(X96,X97))),s(fun(X98,bool),X100)))) | s(bool,X101) = s(bool,'const.bool.IN_2'(s(X96,X97),s(fun(X96,bool),'const.util_prob.PREIMAGE_2'(s(fun(X96,X98),X99),s(fun(X98,bool),X100))))) | s(bool,'F_0') = s(bool,X101)) )),
  inference(superposition,[],[f503,f118])).
fof(f2074,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f210,f151])).
fof(f2099,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.IN_2'(s(X48,happ(s(fun(X46,X48),X49),s(X46,X47))),s(fun(X48,bool),X50))),s(bool,X51))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),'const.util_prob.PREIMAGE_2'(s(fun(X46,X48),X49),s(fun(X48,bool),X50)))))) )),
  inference(superposition,[],[f210,f118])).
fof(f2100,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.IN_2'(s(X53,X55),s(fun(X53,bool),'const.util_prob.PREIMAGE_2'(s(fun(X53,X52),X54),s(fun(X52,bool),X56))))),s(bool,X57))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X52,happ(s(fun(X53,X52),X54),s(X53,X55))),s(fun(X52,bool),X56)))) )),
  inference(superposition,[],[f210,f118])).
fof(f2105,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X68),s(bool,X71))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X68),s(X69,X70)))) )),
  inference(superposition,[],[f210,f115])).
fof(f2106,plain,(
  ( ! [X74,X72,X73] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X72),s(bool,X74))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X72),s(bool,X73)))) )),
  inference(superposition,[],[f210,f107])).
fof(f2107,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X76),s(bool,X77))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X75),s(bool,X76)))) )),
  inference(superposition,[],[f210,f108])).
fof(f2108,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X79),s(bool,X81))) | s(bool,'F_0') = s(bool,happ(s(fun(X78,bool),'const.combin.K_1'(s(bool,X79))),s(X78,X80)))) )),
  inference(superposition,[],[f210,f120])).
fof(f2221,plain,(
  ( ! [X35,X33,X34] : (s(bool,X35) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X33),s(bool,X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f2156,f83])).
fof(f2156,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | s(bool,X35) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X33),s(bool,X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f154,f210])).
fof(f2422,plain,(
  ( ! [X127,X125,X128,X126,X129] : (s(X127,X128) = s(X127,'const.bool.COND_3'(s(bool,X126),s(X127,X129),s(X127,X128))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X125),s(bool,X126)))) )),
  inference(superposition,[],[f281,f108])).
fof(f2483,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X12),s(bool,X13),s(bool,X14)))) )),
  inference(superposition,[],[f281,f91])).
fof(f2511,plain,(
  ( ! [X87,X85,X86] : (s(bool,'F_0') != s(bool,X87) | ~p(s(bool,'const.bool.COND_3'(s(bool,X85),s(bool,X86),s(bool,X87)))) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f97,f281])).
fof(f2512,plain,(
  ( ! [X90,X88,X89] : (p(s(bool,X90)) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X88),s(bool,X89),s(bool,X90))) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f98,f281])).
fof(f2520,plain,(
  ( ! [X123,X121,X122,X120,X119] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X121),s(X122,X123))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X119),s(bool,X120),s(bool,X121))) | s(bool,'T_0') = s(bool,X119)) )),
  inference(superposition,[],[f153,f281])).
fof(f2523,plain,(
  ( ! [X134,X132,X135,X133] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X134),s(bool,X135))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X132),s(bool,X133),s(bool,X134))) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f210,f281])).
fof(f2525,plain,(
  ( ! [X142,X140,X143,X141] : (~p(s(bool,X142)) | s(bool,X143) = s(bool,'const.bool.COND_3'(s(bool,X140),s(bool,X141),s(bool,X142))) | s(bool,'F_0') = s(bool,X143) | s(bool,'T_0') = s(bool,X140)) )),
  inference(superposition,[],[f503,f281])).
fof(f2760,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X28),s(bool,X29))) | s(bool,'F_0') = s(bool,X29) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f297,f151])).
fof(f2796,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X51),s(bool,'const.bool.IN_2'(s(X48,happ(s(fun(X46,X48),X49),s(X46,X47))),s(fun(X48,bool),X50))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),'const.util_prob.PREIMAGE_2'(s(fun(X46,X48),X49),s(fun(X48,bool),X50)))))) )),
  inference(superposition,[],[f297,f118])).
fof(f2797,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X57),s(bool,'const.bool.IN_2'(s(X53,X55),s(fun(X53,bool),'const.util_prob.PREIMAGE_2'(s(fun(X53,X52),X54),s(fun(X52,bool),X56))))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X52,happ(s(fun(X53,X52),X54),s(X53,X55))),s(fun(X52,bool),X56)))) )),
  inference(superposition,[],[f297,f118])).
fof(f2803,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X74),s(bool,X71))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X71),s(X72,X73)))) )),
  inference(superposition,[],[f297,f115])).
fof(f2804,plain,(
  ( ! [X76,X77,X75] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X77),s(bool,X75))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X75),s(bool,X76)))) )),
  inference(superposition,[],[f297,f107])).
fof(f2805,plain,(
  ( ! [X80,X78,X79] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X80),s(bool,X79))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X78),s(bool,X79)))) )),
  inference(superposition,[],[f297,f108])).
fof(f2806,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X84),s(bool,X83))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X81),s(bool,X82),s(bool,X83))) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f297,f281])).
fof(f2807,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X88),s(bool,X86))) | s(bool,'F_0') = s(bool,happ(s(fun(X85,bool),'const.combin.K_1'(s(bool,X86))),s(X85,X87)))) )),
  inference(superposition,[],[f297,f120])).
fof(f2911,plain,(
  ( ! [X35,X33,X34] : (s(bool,X35) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X33),s(bool,X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f2829,f83])).
fof(f2829,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'T_0')) | s(bool,X35) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X33),s(bool,X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X34)) )),
  inference(superposition,[],[f154,f297])).
fof(f3318,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(X68,bool),'const.combin.K_1'(s(bool,'const.bool.IN_2'(s(X65,happ(s(fun(X63,X65),X66),s(X63,X64))),s(fun(X65,bool),X67))))),s(X68,X69))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X63,X64),s(fun(X63,bool),'const.util_prob.PREIMAGE_2'(s(fun(X63,X65),X66),s(fun(X65,bool),X67)))))) )),
  inference(superposition,[],[f355,f118])).
fof(f3319,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (s(bool,'T_0') = s(bool,happ(s(fun(X75,bool),'const.combin.K_1'(s(bool,'const.bool.IN_2'(s(X71,X73),s(fun(X71,bool),'const.util_prob.PREIMAGE_2'(s(fun(X71,X70),X72),s(fun(X70,bool),X74))))))),s(X75,X76))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X70,happ(s(fun(X71,X70),X72),s(X71,X73))),s(fun(X70,bool),X74)))) )),
  inference(superposition,[],[f355,f118])).
fof(f3328,plain,(
  ( ! [X111,X109,X112,X110,X108] : (s(bool,'T_0') = s(bool,happ(s(fun(X111,bool),'const.combin.K_1'(s(bool,X108))),s(X111,X112))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X108),s(X109,X110)))) )),
  inference(superposition,[],[f355,f115])).
fof(f3331,plain,(
  ( ! [X125,X123,X121,X124,X122] : (s(bool,'T_0') = s(bool,happ(s(fun(X124,bool),'const.combin.K_1'(s(bool,X123))),s(X124,X125))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X121),s(bool,X122),s(bool,X123))) | s(bool,'T_0') = s(bool,X121)) )),
  inference(superposition,[],[f355,f281])).
fof(f3332,plain,(
  ( ! [X127,X130,X128,X126,X129] : (s(bool,'T_0') = s(bool,happ(s(fun(X129,bool),'const.combin.K_1'(s(bool,X127))),s(X129,X130))) | s(bool,'F_0') = s(bool,happ(s(fun(X126,bool),'const.combin.K_1'(s(bool,X127))),s(X126,X128)))) )),
  inference(superposition,[],[f355,f120])).
fof(f3348,plain,(
  ( ! [X52,X53,X51] : (s(bool,X51) = s(bool,happ(s(fun(bool,bool),'const.combin.K_1'(s(bool,X52))),s(bool,X51))) | s(bool,'F_0') = s(bool,X52) | s(bool,X51) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f355,f151])).
fof(f3474,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,X49) != s(bool,happ(s(fun(X46,bool),'const.combin.K_1'(s(bool,X47))),s(X46,X48))) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f3385,f83])).
fof(f3385,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'T_0')) | s(bool,X49) != s(bool,happ(s(fun(X46,bool),'const.combin.K_1'(s(bool,X47))),s(X46,X48))) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X47)) )),
  inference(superposition,[],[f154,f355])).
fof(f4130,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X19,X20),s(fun(X19,bool),X21))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X17,X18),s(fun(X17,bool),'const.util_prob.PREIMAGE_2'(s(fun(X17,X19),'const.combin.K_1'(s(X19,X20))),s(fun(X19,bool),X21)))))) )),
  inference(superposition,[],[f1189,f91])).
fof(f4156,plain,(
  ( ! [X94,X92,X90,X95,X93,X91,X89] : (s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,'const.bool.IN_2'(s(X91,X92),s(fun(X91,bool),X93))),s(X94,X95))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X89,X90),s(fun(X89,bool),'const.util_prob.PREIMAGE_2'(s(fun(X89,X91),'const.combin.K_1'(s(X91,X92))),s(fun(X91,bool),X93)))))) )),
  inference(superposition,[],[f153,f1189])).
fof(f4159,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.IN_2'(s(X110,X111),s(fun(X110,bool),X112))),s(bool,X113))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X108,X109),s(fun(X108,bool),'const.util_prob.PREIMAGE_2'(s(fun(X108,X110),'const.combin.K_1'(s(X110,X111))),s(fun(X110,bool),X112)))))) )),
  inference(superposition,[],[f210,f1189])).
fof(f4163,plain,(
  ( ! [X134,X138,X136,X135,X139,X137] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X139),s(bool,'const.bool.IN_2'(s(X136,X137),s(fun(X136,bool),X138))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X134,X135),s(fun(X134,bool),'const.util_prob.PREIMAGE_2'(s(fun(X134,X136),'const.combin.K_1'(s(X136,X137))),s(fun(X136,bool),X138)))))) )),
  inference(superposition,[],[f297,f1189])).
fof(f4164,plain,(
  ( ! [X146,X144,X142,X140,X145,X143,X141] : (s(bool,'T_0') = s(bool,happ(s(fun(X145,bool),'const.combin.K_1'(s(bool,'const.bool.IN_2'(s(X142,X143),s(fun(X142,bool),X144))))),s(X145,X146))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X140,X141),s(fun(X140,bool),'const.util_prob.PREIMAGE_2'(s(fun(X140,X142),'const.combin.K_1'(s(X142,X143))),s(fun(X142,bool),X144)))))) )),
  inference(superposition,[],[f355,f1189])).
fof(f762,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'T_0') != s(bool,f(s(X101,sK5(X102,X103,bool,X101)))) | s(fun(X101,bool),X102) = s(fun(X101,bool),X103) | s(bool,X104) = s(bool,g(s(X101,sK5(X102,X103,bool,X101)))) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f117,f151])).
fof(f656,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,X109) != s(bool,f(s(X106,sK5(X107,X108,bool,X106)))) | s(fun(X106,bool),X107) = s(fun(X106,bool),X108) | s(bool,'T_0') = s(bool,X109) | s(bool,'T_0') = s(bool,g(s(X106,sK5(X107,X108,bool,X106))))) )),
  inference(superposition,[],[f117,f151])).
fof(f109,plain,(
  ( ! [X4,X2,X0,X1] : (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X9,X12] : (s(fun(X6,bool),X9) = s(fun(X6,bool),X12) <=> ! [X8] : s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),X9))) = s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),X12))))),
  file('thm.util_prob.PREIMAGE_K','thm.pred_set.EXTENSION')).
fof(f1189,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X4))) = s(bool,'const.bool.IN_2'(s(X1,X3),s(fun(X1,bool),'const.util_prob.PREIMAGE_2'(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(fun(X0,bool),X4)))))) )),
  inference(superposition,[],[f118,f120])).
fof(f2538,plain,(
  ( ! [X198,X196,X200,X199,X197,X195] : (s(fun(X195,bool),X198) != s(fun(X195,bool),X199) | s(bool,'const.bool.IN_2'(s(X195,X200),s(fun(X195,bool),X199))) = s(bool,'const.bool.IN_2'(s(X195,X200),s(fun(X195,bool),'const.bool.COND_3'(s(bool,X196),s(fun(X195,bool),X197),s(fun(X195,bool),X198))))) | s(bool,'T_0') = s(bool,X196)) )),
  inference(superposition,[],[f109,f281])).
fof(f4129,plain,(
  ( ! [X14,X12,X10,X15,X13,X11,X16] : (s(bool,'const.bool.IN_2'(s(X14,X15),s(fun(X14,bool),X16))) = s(bool,'const.bool.IN_2'(s(X11,X13),s(fun(X11,bool),'const.util_prob.PREIMAGE_2'(s(fun(X11,X10),X12),s(fun(X10,bool),'const.util_prob.PREIMAGE_2'(s(fun(X10,X14),'const.combin.K_1'(s(X14,X15))),s(fun(X14,bool),X16)))))))) )),
  inference(superposition,[],[f1189,f118])).
fof(f4131,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X24,X25),s(fun(X24,bool),X26))) | s(bool,X27) = s(bool,'const.bool.IN_2'(s(X22,X23),s(fun(X22,bool),'const.util_prob.PREIMAGE_2'(s(fun(X22,X24),'const.combin.K_1'(s(X24,X25))),s(fun(X24,bool),X26))))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f1189,f151])).
fof(f4132,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (s(bool,X33) = s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32))) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X28,X29),s(fun(X28,bool),'const.util_prob.PREIMAGE_2'(s(fun(X28,X30),'const.combin.K_1'(s(X30,X31))),s(fun(X30,bool),X32)))))) )),
  inference(superposition,[],[f1189,f151])).
fof(f4141,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),X9))) != s(bool,'const.bool.IN_2'(s(X5,sK4(X6,'const.util_prob.PREIMAGE_2'(s(fun(X5,X7),'const.combin.K_1'(s(X7,X8))),s(fun(X7,bool),X9)),X5)),s(fun(X5,bool),X6))) | s(fun(X5,bool),X6) = s(fun(X5,bool),'const.util_prob.PREIMAGE_2'(s(fun(X5,X7),'const.combin.K_1'(s(X7,X8))),s(fun(X7,bool),X9)))) )),
  inference(superposition,[],[f110,f1189])).
fof(f4142,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,'const.bool.IN_2'(s(X11,X12),s(fun(X11,bool),X13))) != s(bool,'const.bool.IN_2'(s(X10,sK4('const.util_prob.PREIMAGE_2'(s(fun(X10,X11),'const.combin.K_1'(s(X11,X12))),s(fun(X11,bool),X13)),X14,X10)),s(fun(X10,bool),X14))) | s(fun(X10,bool),X14) = s(fun(X10,bool),'const.util_prob.PREIMAGE_2'(s(fun(X10,X11),'const.combin.K_1'(s(X11,X12))),s(fun(X11,bool),X13)))) )),
  inference(superposition,[],[f110,f1189])).
fof(f4161,plain,(
  ( ! [X127,X125,X123,X121,X126,X124,X122,X120] : (s(X125,X126) = s(X125,'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X122,X123),s(fun(X122,bool),X124))),s(X125,X127),s(X125,X126))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X120,X121),s(fun(X120,bool),'const.util_prob.PREIMAGE_2'(s(fun(X120,X122),'const.combin.K_1'(s(X122,X123))),s(fun(X122,bool),X124)))))) )),
  inference(superposition,[],[f281,f1189])).
fof(f118,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X4))),s(fun(X1,bool),X3))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X4))),s(fun(X1,bool),X3))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X4))),s(fun(X1,bool),X3))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X7] : ! [X14,X9,X8] : s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X14),s(X6,X8))),s(fun(X7,bool),X9))) = s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X14),s(fun(X7,bool),X9)))))),
  file('thm.util_prob.PREIMAGE_K','thm.util_prob.IN_PREIMAGE')).
fof(f1234,plain,(
  ( ! [X21,X19,X17,X20,X18] : (p(s(bool,'const.bool.IN_2'(s(X18,X20),s(fun(X18,bool),'const.util_prob.PREIMAGE_2'(s(fun(X18,X17),X19),s(fun(X17,bool),X21)))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X17,happ(s(fun(X18,X17),X19),s(X18,X20))),s(fun(X17,bool),X21)))) )),
  inference(superposition,[],[f93,f118])).
fof(f1235,plain,(
  ( ! [X26,X24,X23,X25,X22] : (~p(s(bool,'const.bool.IN_2'(s(X23,X25),s(fun(X23,bool),'const.util_prob.PREIMAGE_2'(s(fun(X23,X22),X24),s(fun(X22,bool),X26)))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X22,happ(s(fun(X23,X22),X24),s(X23,X25))),s(fun(X22,bool),X26)))) )),
  inference(superposition,[],[f94,f118])).
fof(f1238,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (p(s(bool,'const.bool.IN_2'(s(X38,X40),s(fun(X38,bool),'const.util_prob.PREIMAGE_2'(s(fun(X38,X37),X39),s(fun(X37,bool),X41)))))) | p(s(bool,X42)) | s(bool,X42) = s(bool,'const.bool.IN_2'(s(X37,happ(s(fun(X38,X37),X39),s(X38,X40))),s(fun(X37,bool),X41)))) )),
  inference(superposition,[],[f111,f118])).
fof(f1240,plain,(
  ( ! [X54,X52,X50,X53,X51,X49] : (~p(s(bool,'const.bool.IN_2'(s(X50,X52),s(fun(X50,bool),'const.util_prob.PREIMAGE_2'(s(fun(X50,X49),X51),s(fun(X49,bool),X53)))))) | ~p(s(bool,X54)) | s(bool,X54) = s(bool,'const.bool.IN_2'(s(X49,happ(s(fun(X50,X49),X51),s(X50,X52))),s(fun(X49,bool),X53)))) )),
  inference(superposition,[],[f114,f118])).
fof(f1246,plain,(
  ( ! [X90,X88,X87,X91,X89,X86] : (~p(s(bool,'const.bool.IN_2'(s(X87,X89),s(fun(X87,bool),'const.util_prob.PREIMAGE_2'(s(fun(X87,X86),X88),s(fun(X86,bool),X90)))))) | s(bool,X91) != s(bool,'const.bool.IN_2'(s(X86,happ(s(fun(X87,X86),X88),s(X87,X89))),s(fun(X86,bool),X90))) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f154,f118])).
fof(f1247,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (p(s(bool,'const.bool.IN_2'(s(X93,X95),s(fun(X93,bool),'const.util_prob.PREIMAGE_2'(s(fun(X93,X92),X94),s(fun(X92,bool),X96)))))) | s(bool,X97) = s(bool,'const.bool.IN_2'(s(X92,happ(s(fun(X93,X92),X94),s(X93,X95))),s(fun(X92,bool),X96))) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f188,f118])).
fof(f1265,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(X17,happ(s(fun(X15,X17),X18),s(X15,X16))),s(fun(X17,bool),X19)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X15,X16),s(fun(X15,bool),'const.util_prob.PREIMAGE_2'(s(fun(X15,X17),X18),s(fun(X17,bool),X19)))))) )),
  inference(superposition,[],[f93,f118])).
fof(f1266,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(X22,happ(s(fun(X20,X22),X23),s(X20,X21))),s(fun(X22,bool),X24)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X20,X21),s(fun(X20,bool),'const.util_prob.PREIMAGE_2'(s(fun(X20,X22),X23),s(fun(X22,bool),X24)))))) )),
  inference(superposition,[],[f94,f118])).
fof(f1269,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (p(s(bool,'const.bool.IN_2'(s(X37,happ(s(fun(X35,X37),X38),s(X35,X36))),s(fun(X37,bool),X39)))) | p(s(bool,X40)) | s(bool,X40) = s(bool,'const.bool.IN_2'(s(X35,X36),s(fun(X35,bool),'const.util_prob.PREIMAGE_2'(s(fun(X35,X37),X38),s(fun(X37,bool),X39)))))) )),
  inference(superposition,[],[f111,f118])).
fof(f1271,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (~p(s(bool,'const.bool.IN_2'(s(X49,happ(s(fun(X47,X49),X50),s(X47,X48))),s(fun(X49,bool),X51)))) | ~p(s(bool,X52)) | s(bool,X52) = s(bool,'const.bool.IN_2'(s(X47,X48),s(fun(X47,bool),'const.util_prob.PREIMAGE_2'(s(fun(X47,X49),X50),s(fun(X49,bool),X51)))))) )),
  inference(superposition,[],[f114,f118])).
fof(f1277,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (~p(s(bool,'const.bool.IN_2'(s(X86,happ(s(fun(X84,X86),X87),s(X84,X85))),s(fun(X86,bool),X88)))) | s(bool,X89) != s(bool,'const.bool.IN_2'(s(X84,X85),s(fun(X84,bool),'const.util_prob.PREIMAGE_2'(s(fun(X84,X86),X87),s(fun(X86,bool),X88))))) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f154,f118])).
fof(f1278,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (p(s(bool,'const.bool.IN_2'(s(X92,happ(s(fun(X90,X92),X93),s(X90,X91))),s(fun(X92,bool),X94)))) | s(bool,X95) = s(bool,'const.bool.IN_2'(s(X90,X91),s(fun(X90,bool),'const.util_prob.PREIMAGE_2'(s(fun(X90,X92),X93),s(fun(X92,bool),X94))))) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f188,f118])).
fof(f2415,plain,(
  ( ! [X94,X92,X90,X88,X95,X93,X91,X89] : (s(X93,X94) = s(X93,'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X89,X91),s(fun(X89,bool),'const.util_prob.PREIMAGE_2'(s(fun(X89,X88),X90),s(fun(X88,bool),X92))))),s(X93,X95),s(X93,X94))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X88,happ(s(fun(X89,X88),X90),s(X89,X91))),s(fun(X88,bool),X92)))) )),
  inference(superposition,[],[f281,f118])).
fof(f1192,plain,(
  ( ! [X14,X19,X17,X15,X13,X18,X16] : (s(bool,'const.bool.IN_2'(s(X17,happ(s(fun(X13,X17),X18),s(X13,happ(s(fun(X14,X13),X15),s(X14,X16))))),s(fun(X17,bool),X19))) = s(bool,'const.bool.IN_2'(s(X14,X16),s(fun(X14,bool),'const.util_prob.PREIMAGE_2'(s(fun(X14,X13),X15),s(fun(X13,bool),'const.util_prob.PREIMAGE_2'(s(fun(X13,X17),X18),s(fun(X17,bool),X19)))))))) )),
  inference(superposition,[],[f118,f118])).
fof(f110,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f80])).
fof(f1262,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK4(X1,'const.util_prob.PREIMAGE_2'(s(fun(X0,X2),X3),s(fun(X2,bool),X4)),X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X2,happ(s(fun(X0,X2),X3),s(X0,sK4(X1,'const.util_prob.PREIMAGE_2'(s(fun(X0,X2),X3),s(fun(X2,bool),X4)),X0)))),s(fun(X2,bool),X4))) | s(fun(X0,bool),X1) = s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X2),X3),s(fun(X2,bool),X4)))) )),
  inference(superposition,[],[f110,f118])).
fof(f1083,plain,(
  ( ! [X6,X8,X7,X5] : (s(fun(X5,bool),X6) = s(fun(X5,bool),X8) | s(bool,'const.bool.IN_2'(s(X5,sK4(X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X5)),s(fun(X5,bool),X6))) != s(bool,'const.bool.IN_2'(s(X5,sK4(X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X5)),s(fun(X5,bool),X8)))) )),
  inference(forward_demodulation,[],[f1053,f107])).
fof(f1053,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'const.bool.IN_2'(s(X5,sK4(X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X5)),s(fun(X5,bool),X6))) != s(bool,'const.bool.IN_2'(s(X5,sK4(X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X5)),s(fun(X5,bool),X8))) | s(fun(X5,bool),X8) = s(fun(X5,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)))) )),
  inference(superposition,[],[f110,f107])).
fof(f2537,plain,(
  ( ! [X191,X194,X192,X190,X193] : (s(bool,'const.bool.IN_2'(s(X190,sK4(X194,'const.bool.COND_3'(s(bool,X191),s(fun(X190,bool),X192),s(fun(X190,bool),X193)),X190)),s(fun(X190,bool),X193))) != s(bool,'const.bool.IN_2'(s(X190,sK4(X194,'const.bool.COND_3'(s(bool,X191),s(fun(X190,bool),X192),s(fun(X190,bool),X193)),X190)),s(fun(X190,bool),X194))) | s(fun(X190,bool),X194) = s(fun(X190,bool),'const.bool.COND_3'(s(bool,X191),s(fun(X190,bool),X192),s(fun(X190,bool),X193))) | s(bool,'T_0') = s(bool,X191)) )),
  inference(superposition,[],[f110,f281])).
fof(f1082,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X4) | s(bool,'const.bool.IN_2'(s(X0,sK4(X4,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X4,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X0)),s(fun(X0,bool),X4)))) )),
  inference(forward_demodulation,[],[f1052,f115])).
fof(f1052,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK4(X4,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4(X4,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X0)),s(fun(X0,bool),X4))) | s(fun(X0,bool),X4) = s(fun(X0,bool),'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)))) )),
  inference(superposition,[],[f110,f115])).
fof(f1085,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(fun(X13,bool),X15) = s(fun(X13,bool),X17) | s(bool,'const.bool.IN_2'(s(X13,sK4(X17,happ(s(fun(X14,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X14,X16)),X13)),s(fun(X13,bool),X15))) != s(bool,'const.bool.IN_2'(s(X13,sK4(X17,happ(s(fun(X14,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X14,X16)),X13)),s(fun(X13,bool),X17)))) )),
  inference(forward_demodulation,[],[f1055,f120])).
fof(f1055,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,'const.bool.IN_2'(s(X13,sK4(X17,happ(s(fun(X14,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X14,X16)),X13)),s(fun(X13,bool),X15))) != s(bool,'const.bool.IN_2'(s(X13,sK4(X17,happ(s(fun(X14,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X14,X16)),X13)),s(fun(X13,bool),X17))) | s(fun(X13,bool),X17) = s(fun(X13,bool),happ(s(fun(X14,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X14,X16)))) )),
  inference(superposition,[],[f110,f120])).
fof(f1263,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'const.bool.IN_2'(s(X5,sK4('const.util_prob.PREIMAGE_2'(s(fun(X5,X6),X7),s(fun(X6,bool),X8)),X9,X5)),s(fun(X5,bool),X9))) != s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(X5,X6),X7),s(X5,sK4('const.util_prob.PREIMAGE_2'(s(fun(X5,X6),X7),s(fun(X6,bool),X8)),X9,X5)))),s(fun(X6,bool),X8))) | s(fun(X5,bool),X9) = s(fun(X5,bool),'const.util_prob.PREIMAGE_2'(s(fun(X5,X6),X7),s(fun(X6,bool),X8)))) )),
  inference(superposition,[],[f110,f118])).
fof(f1087,plain,(
  ( ! [X6,X8,X7,X5] : (s(fun(X5,bool),X6) = s(fun(X5,bool),X8) | s(bool,'const.bool.IN_2'(s(X5,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5)),s(fun(X5,bool),X6))) != s(bool,'const.bool.IN_2'(s(X5,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5)),s(fun(X5,bool),X8)))) )),
  inference(forward_demodulation,[],[f1069,f107])).
fof(f1069,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'const.bool.IN_2'(s(X5,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5)),s(fun(X5,bool),X6))) != s(bool,'const.bool.IN_2'(s(X5,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5)),s(fun(X5,bool),X8))) | s(fun(X5,bool),X8) = s(fun(X5,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)))) )),
  inference(superposition,[],[f110,f107])).
fof(f2536,plain,(
  ( ! [X189,X187,X185,X188,X186] : (s(bool,'const.bool.IN_2'(s(X185,sK4('const.bool.COND_3'(s(bool,X186),s(fun(X185,bool),X187),s(fun(X185,bool),X188)),X189,X185)),s(fun(X185,bool),X188))) != s(bool,'const.bool.IN_2'(s(X185,sK4('const.bool.COND_3'(s(bool,X186),s(fun(X185,bool),X187),s(fun(X185,bool),X188)),X189,X185)),s(fun(X185,bool),X189))) | s(fun(X185,bool),X189) = s(fun(X185,bool),'const.bool.COND_3'(s(bool,X186),s(fun(X185,bool),X187),s(fun(X185,bool),X188))) | s(bool,'T_0') = s(bool,X186)) )),
  inference(superposition,[],[f110,f281])).
fof(f1086,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X4) | s(bool,'const.bool.IN_2'(s(X0,sK4('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X4,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X4,X0)),s(fun(X0,bool),X4)))) )),
  inference(forward_demodulation,[],[f1068,f115])).
fof(f1068,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK4('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X4,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK4('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X4,X0)),s(fun(X0,bool),X4))) | s(fun(X0,bool),X4) = s(fun(X0,bool),'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)))) )),
  inference(superposition,[],[f110,f115])).
fof(f1089,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(fun(X13,bool),X15) = s(fun(X13,bool),X17) | s(bool,'const.bool.IN_2'(s(X13,sK4(happ(s(fun(X14,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X14,X16)),X17,X13)),s(fun(X13,bool),X15))) != s(bool,'const.bool.IN_2'(s(X13,sK4(happ(s(fun(X14,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X14,X16)),X17,X13)),s(fun(X13,bool),X17)))) )),
  inference(forward_demodulation,[],[f1071,f120])).
fof(f1071,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,'const.bool.IN_2'(s(X13,sK4(happ(s(fun(X14,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X14,X16)),X17,X13)),s(fun(X13,bool),X15))) != s(bool,'const.bool.IN_2'(s(X13,sK4(happ(s(fun(X14,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X14,X16)),X17,X13)),s(fun(X13,bool),X17))) | s(fun(X13,bool),X17) = s(fun(X13,bool),happ(s(fun(X14,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X14,X16)))) )),
  inference(superposition,[],[f110,f120])).
fof(f2414,plain,(
  ( ! [X80,X87,X85,X83,X81,X86,X84,X82] : (s(X85,X86) = s(X85,'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X82,happ(s(fun(X80,X82),X83),s(X80,X81))),s(fun(X82,bool),X84))),s(X85,X87),s(X85,X86))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X80,X81),s(fun(X80,bool),'const.util_prob.PREIMAGE_2'(s(fun(X80,X82),X83),s(fun(X82,bool),X84)))))) )),
  inference(superposition,[],[f281,f118])).
fof(f1065,plain,(
  ( ! [X6,X8,X7] : (s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X6))) != s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),X7))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK4(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f110,f151])).
fof(f1063,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X0))) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,X2) = s(bool,sK4(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f110,f151])).
fof(f2509,plain,(
  ( ! [X80,X81,X79] : (p(s(bool,X81)) | s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X79),s(bool,X80),s(bool,X81))) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f93,f281])).
fof(f2510,plain,(
  ( ! [X83,X84,X82] : (~p(s(bool,X84)) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X82),s(bool,X83),s(bool,X84))) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f94,f281])).
fof(f2513,plain,(
  ( ! [X94,X92,X93,X91] : (p(s(bool,X93)) | p(s(bool,X94)) | s(bool,X94) = s(bool,'const.bool.COND_3'(s(bool,X91),s(bool,X92),s(bool,X93))) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f111,f281])).
fof(f2515,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,X101)) | ~p(s(bool,X102)) | s(bool,X102) = s(bool,'const.bool.COND_3'(s(bool,X99),s(bool,X100),s(bool,X101))) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f114,f281])).
fof(f2521,plain,(
  ( ! [X127,X125,X126,X124] : (~p(s(bool,X126)) | s(bool,X127) != s(bool,'const.bool.COND_3'(s(bool,X124),s(bool,X125),s(bool,X126))) | s(bool,'T_0') = s(bool,X127) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f154,f281])).
fof(f2522,plain,(
  ( ! [X130,X128,X131,X129] : (p(s(bool,X130)) | s(bool,X131) = s(bool,'const.bool.COND_3'(s(bool,X128),s(bool,X129),s(bool,X130))) | s(bool,'T_0') = s(bool,X131) | s(bool,'T_0') = s(bool,X128)) )),
  inference(superposition,[],[f188,f281])).
fof(f2485,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X21) = s(bool,X22) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X19),s(bool,X20),s(bool,X21)))) )),
  inference(superposition,[],[f281,f151])).
fof(f2484,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X15) | s(bool,X18) = s(bool,'const.bool.COND_3'(s(bool,X15),s(bool,X16),s(bool,X17))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f281,f151])).
fof(f2424,plain,(
  ( ! [X140,X138,X136,X141,X139,X137] : (s(X139,X140) = s(X139,'const.bool.COND_3'(s(bool,X137),s(X139,X141),s(X139,X140))) | s(bool,'T_0') = s(bool,happ(s(fun(X136,bool),'const.combin.K_1'(s(bool,X137))),s(X136,X138)))) )),
  inference(superposition,[],[f281,f120])).
fof(f2423,plain,(
  ( ! [X134,X132,X130,X135,X133,X131] : (s(X133,X134) = s(X133,'const.bool.COND_3'(s(bool,X132),s(X133,X135),s(X133,X134))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X130),s(bool,X131),s(bool,X132))) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f281,f281])).
fof(f2421,plain,(
  ( ! [X123,X121,X124,X122,X120] : (s(X122,X123) = s(X122,'const.bool.COND_3'(s(bool,X120),s(X122,X124),s(X122,X123))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X120),s(bool,X121)))) )),
  inference(superposition,[],[f281,f107])).
fof(f2420,plain,(
  ( ! [X118,X116,X114,X119,X117,X115] : (s(X117,X118) = s(X117,'const.bool.COND_3'(s(bool,X114),s(X117,X119),s(X117,X118))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X114),s(X115,X116)))) )),
  inference(superposition,[],[f281,f115])).
fof(f1231,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X5,X7),s(fun(X5,bool),'const.util_prob.PREIMAGE_2'(s(fun(X5,X4),X6),s(fun(X4,bool),'const.pred_set.EMPTY_0')))))) )),
  inference(superposition,[],[f133,f118])).
fof(f1195,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (s(bool,X37) = s(bool,'const.bool.IN_2'(s(X33,X35),s(fun(X33,bool),'const.util_prob.PREIMAGE_2'(s(fun(X33,X32),X34),s(fun(X32,bool),X36))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X32,happ(s(fun(X33,X32),X34),s(X33,X35))),s(fun(X32,bool),X36))) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f118,f151])).
fof(f1193,plain,(
  ( ! [X24,X23,X21,X25,X22,X20] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X21,X23),s(fun(X21,bool),'const.util_prob.PREIMAGE_2'(s(fun(X21,X20),X22),s(fun(X20,bool),X24))))) | s(bool,X25) = s(bool,'const.bool.IN_2'(s(X20,happ(s(fun(X21,X20),X22),s(X21,X23))),s(fun(X20,bool),X24))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f118,f151])).
fof(f1190,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X5),X7),s(fun(X5,bool),'const.pred_set.UNIV_0')))))) )),
  inference(superposition,[],[f118,f129])).
fof(f1188,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X2),X3),s(fun(X2,bool),'const.pred_set.EMPTY_0'))))))) )),
  inference(unit_resulting_resolution,[],[f133,f118,f154])).
fof(f1076,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X15) != s(bool,'const.bool.IN_2'(s(X12,sK4(X13,X14,X12)),s(fun(X12,bool),X14))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X12,sK4(X13,X14,X12)),s(fun(X12,bool),X13))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f110,f151])).
fof(f1074,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X4,sK4(X5,X6,X4)),s(fun(X4,bool),X6))) | s(fun(X4,bool),X5) = s(fun(X4,bool),X6) | s(bool,X7) = s(bool,'const.bool.IN_2'(s(X4,sK4(X5,X6,X4)),s(fun(X4,bool),X5))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f110,f151])).
fof(f1072,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK4('const.pred_set.UNIV_0',X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.UNIV_0') = s(fun(X0,bool),X1)) )),
  inference(superposition,[],[f110,f129])).
fof(f1060,plain,(
  ( ! [X14,X12,X15,X13] : (s(bool,X15) != s(bool,'const.bool.IN_2'(s(X12,sK4(X13,X14,X12)),s(fun(X12,bool),X13))) | s(fun(X12,bool),X13) = s(fun(X12,bool),X14) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X12,sK4(X13,X14,X12)),s(fun(X12,bool),X14))) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f110,f151])).
fof(f1058,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X4,sK4(X5,X6,X4)),s(fun(X4,bool),X5))) | s(fun(X4,bool),X5) = s(fun(X4,bool),X6) | s(bool,X7) = s(bool,'const.bool.IN_2'(s(X4,sK4(X5,X6,X4)),s(fun(X4,bool),X6))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f110,f151])).
fof(f1056,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK4(X1,'const.pred_set.UNIV_0',X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.UNIV_0') = s(fun(X0,bool),X1)) )),
  inference(superposition,[],[f110,f129])).
fof(f767,plain,(
  ( ! [X121,X122,X120,X119] : (s(X121,f(s(bool,'T_0'))) != s(X121,g(s(bool,'T_0'))) | s(fun(bool,X121),X119) = s(fun(bool,X121),X120) | s(bool,X122) = s(bool,sK5(X119,X120,X121,bool)) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f117,f151])).
fof(f1032,plain,(
  ( ! [X30,X31,X29] : (s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f739,f83])).
fof(f739,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | s(bool,X29) = s(bool,X30) | s(bool,'F_0') = s(bool,X30) | s(bool,X29) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f503,f151])).
fof(f1031,plain,(
  ( ! [X24,X23,X25] : (s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(subsumption_resolution,[],[f737,f83])).
fof(f737,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | s(bool,X23) != s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f154,f151])).
fof(f1030,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f735,f83])).
fof(f735,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f114,f151])).
fof(f726,plain,(
  ( ! [X72,X73] : (s(bool,X72) != s(bool,'const.bool.IN_2'(s(bool,X72),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,X72) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f133,f151])).
fof(f724,plain,(
  ( ! [X68,X69] : (s(bool,X68) = s(bool,'const.bool.IN_2'(s(bool,X68),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X68) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f129,f151])).
fof(f704,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(X4,X5) = s(X4,'const.bool.COND_3'(s(bool,X3),s(X4,X5),s(X4,X6))) | s(bool,X3) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f107,f151])).
fof(f661,plain,(
  ( ! [X127,X125,X126,X124] : (s(X126,f(s(bool,X127))) != s(X126,g(s(bool,X127))) | s(fun(bool,X126),X124) = s(fun(bool,X126),X125) | s(bool,'T_0') = s(bool,X127) | s(bool,'T_0') = s(bool,sK5(X124,X125,X126,bool))) )),
  inference(superposition,[],[f117,f151])).
fof(f604,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,'T_0') = s(bool,X33) | s(bool,X35) = s(bool,happ(s(fun(X32,bool),'const.combin.K_1'(s(bool,X33))),s(X32,X34))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f151,f120])).
fof(f602,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,X26) | s(bool,X28) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X26),s(bool,X27))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f151,f107])).
fof(f601,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X25) = s(bool,'const.combin.K_2'(s(bool,X22),s(X23,X24))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f151,f115])).
fof(f590,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f151,f151])).
fof(f574,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43) | s(bool,'T_0') = s(bool,happ(s(fun(X40,bool),'const.combin.K_1'(s(bool,X41))),s(X40,X42)))) )),
  inference(superposition,[],[f151,f120])).
fof(f572,plain,(
  ( ! [X35,X36,X34] : (s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X34),s(bool,X35)))) )),
  inference(superposition,[],[f151,f107])).
fof(f571,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X30),s(X31,X32)))) )),
  inference(superposition,[],[f151,f115])).
fof(f451,plain,(
  ( ! [X8,X7,X9] : (s(bool,X9) = s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f445,f119])).
fof(f445,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,'F_0')) | s(bool,X9) = s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f188,f132])).
fof(f426,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) != s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f407,f83])).
fof(f407,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f154,f129])).
fof(f281,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.bool.COND_3'(s(bool,X0),s(X1,X3),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f108,f91])).
fof(f256,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f132,f91])).
fof(f188,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f183,f119])).
fof(f183,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f111,f91])).
fof(f154,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f97,f91])).
fof(f151,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f91,f91])).
fof(f133,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f106,f93])).
fof(f129,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) )),
  inference(unit_resulting_resolution,[],[f103,f94])).
fof(f121,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f119,f93])).
fof(f117,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.util_prob.PREIMAGE_K',aHL_EXT)).
fof(f120,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(forward_demodulation,[],[f116,f115])).
fof(f116,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.combin.K_2'(s(X0,X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.combin.K_1'(s(X0,X2))),s(X1,X3)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6,X7] : ! [X15,X16] : s(X6,'const.combin.K_2'(s(X6,X15),s(X7,X16))) = s(X6,happ(s(fun(X7,X6),'const.combin.K_1'(s(X6,X15))),s(X7,X16)))),
  file('thm.util_prob.PREIMAGE_K','ARITY_1_2_const.combin.K')).
fof(f115,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.combin.K_2'(s(X1,X2),s(X0,X3)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X7,X6] : ! [X8,X13] : s(X6,X8) = s(X6,'const.combin.K_2'(s(X6,X8),s(X7,X13)))),
  file('thm.util_prob.PREIMAGE_K','thm.combin.K_THM')).
fof(f111,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.util_prob.PREIMAGE_K','thm.bool.IMP_ANTISYM_AX')).
fof(f114,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f66])).
fof(f107,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X10,X11] : s(X6,X10) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X10),s(X6,X11)))),
  file('thm.util_prob.PREIMAGE_K','thm.bool.COND_CLAUSES|split|0')).
fof(f106,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X8] : ~p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.util_prob.PREIMAGE_K','thm.pred_set.NOT_IN_EMPTY')).
fof(f103,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))) )),
  inference(cnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X8] : p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),'const.pred_set.UNIV_0'))))),
  file('thm.util_prob.PREIMAGE_K','thm.pred_set.IN_UNIV')).
fof(f93,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.util_prob.PREIMAGE_K','thm.bool.EQ_CLAUSES|split|1')).
fof(f94,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f71])).
fof(f91,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.util_prob.PREIMAGE_K',aHL_BOOL_CASES)).
fof(f83,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.util_prob.PREIMAGE_K','thm.bool.TRUTH')).
fof(f7644,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599,f7601,f7604,f7607,f7620,f7622,f7627,f7629,f7631,f7633,f7635,f7637,f7639,f7641,f7643])).
fof(f7642,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599,f7601,f7604,f7607,f7620,f7622,f7627,f7629,f7631,f7633,f7635,f7637,f7639,f7641])).
fof(f7640,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599,f7601,f7604,f7607,f7620,f7622,f7627,f7629,f7631,f7633,f7635,f7637,f7639])).
fof(f7638,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599,f7601,f7604,f7607,f7620,f7622,f7627,f7629,f7631,f7633,f7635,f7637])).
fof(f7636,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599,f7601,f7604,f7607,f7620,f7622,f7627,f7629,f7631,f7633,f7635])).
fof(f7634,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599,f7601,f7604,f7607,f7620,f7622,f7627,f7629,f7631,f7633])).
fof(f7632,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599,f7601,f7604,f7607,f7620,f7622,f7627,f7629,f7631])).
fof(f7630,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599,f7601,f7604,f7607,f7620,f7622,f7627,f7629])).
fof(f7628,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599,f7601,f7604,f7607,f7620,f7622,f7627])).
fof(f7623,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599,f7601,f7604,f7607,f7620,f7622])).
fof(f7621,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599,f7601,f7604,f7607,f7620])).
fof(f7619,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7536,f7449])).
fof(f7536,plain,(
  s(bool,'T_0') != s(bool,g(s(sK1,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),bool,sK1)))) | ($spl0 | $spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f7440])).
fof(f7440,plain,(
  s(bool,f(s(sK1,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),bool,sK1)))) != s(bool,g(s(sK1,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),bool,sK1)))) | $spl0),
  inference(backward_demodulation,[],[f7437,f433])).
fof(f7437,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl0),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f385,f1056,f1058,f1060,f1072,f1074,f1076,f1132,f1150,f1177,f1179,f1188,f1190,f1193,f1195,f1231,f1298,f1299,f1315,f1430,f1510,f1688,f2319,f2334,f2420,f2421,f2423,f2424,f2484,f2485,f2577,f2585,f3093,f3107,f3482,f3517,f3518,f3520,f3521,f3761,f3762,f3763,f3764,f3765,f3766,f3767,f3768,f3769,f3791,f3969,f3989,f4055,f4515,f2619,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f3507,f2611,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f3508,f2610,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4233,f4181,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f1046,f1047,f4283,f4275,f4274,f4273,f4272,f4271,f4270,f4222,f4184,f4180,f4179,f4178,f4177,f4174,f4164,f4163,f4159,f4156,f4137,f4136,f4130,f4014,f3965,f3906,f3904,f3864,f3863,f3862,f3861,f3860,f3851,f3850,f3811,f3789,f3749,f3710,f3643,f3596,f3595,f3519,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f3290,f3178,f3177,f3166,f3165,f3105,f2968,f2955,f2954,f2944,f2943,f2918,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2738,f2698,f2690,f2689,f2664,f2576,f2534,f2533,f2532,f2525,f2523,f2520,f2512,f2511,f2490,f2489,f2483,f2422,f2368,f2367,f2366,f2365,f2360,f2359,f2294,f2293,f2292,f2291,f2286,f2285,f2267,f2266,f2265,f2264,f2259,f2258,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f2054,f1768,f1767,f1762,f1743,f1742,f1714,f1686,f1685,f1625,f1624,f1619,f1600,f1599,f1538,f1537,f1536,f1483,f1482,f1481,f1480,f1476,f1475,f1442,f1441,f1437,f1435,f1434,f1416,f1415,f1414,f1311,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1178,f1133,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f6636,f6637,f6638,f6639,f6640,f6659,f6671,f6692,f6730,f6731,f6732,f6733,f6734,f6754,f6781,f6893,f6894,f6948,f6949,f7009,f7010,f7011,f7012,f7013,f7028,f7045,f7229,f7341,f3330,f3329,f7342,f7343,f7400,f7398,f433,f431,f82,f7396])).
fof(f7045,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X2),s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f4014,f1315])).
fof(f7028,plain,(
  ( ! [X41,X42,X40] : (s(bool,X41) = s(bool,'const.bool.COND_3'(s(bool,X40),s(bool,X40),s(bool,X40))) | s(bool,'T_0') = s(bool,X41) | s(bool,X40) = s(bool,X42) | s(bool,'F_0') = s(bool,X42)) ) | $spl0),
  inference(superposition,[],[f4014,f1315])).
fof(f7013,plain,(
  ( ! [X92,X90,X93,X91] : (s(bool,X91) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X93),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(X90,bool),'const.combin.K_1'(s(bool,X91))),s(X90,X92)))) ) | $spl0),
  inference(superposition,[],[f4014,f120])).
fof(f7012,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,X88) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X89),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X86),s(bool,X87),s(bool,X88))) | s(bool,'T_0') = s(bool,X86)) ) | $spl0),
  inference(superposition,[],[f4014,f281])).
fof(f7011,plain,(
  ( ! [X85,X83,X84] : (s(bool,X84) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X85),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X83),s(bool,X84)))) ) | $spl0),
  inference(superposition,[],[f4014,f108])).
fof(f7010,plain,(
  ( ! [X80,X81,X82] : (s(bool,X80) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X82),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X80),s(bool,X81)))) ) | $spl0),
  inference(superposition,[],[f4014,f107])).
fof(f7009,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X76) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X79),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X76),s(X77,X78)))) ) | $spl0),
  inference(superposition,[],[f4014,f115])).
fof(f6949,plain,(
  ( ! [X241,X239,X237,X240,X238,X236] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X238,happ(s(fun(X236,X238),X239),s(X236,X237))),s(fun(X238,bool),X240))),s(bool,X241),s(bool,'const.bool.IN_2'(s(X238,happ(s(fun(X236,X238),X239),s(X236,X237))),s(fun(X238,bool),X240))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X236,X237),s(fun(X236,bool),'const.util_prob.PREIMAGE_2'(s(fun(X236,X238),X239),s(fun(X238,bool),X240)))))) ) | $spl0),
  inference(superposition,[],[f3989,f118])).
fof(f6948,plain,(
  ( ! [X231,X235,X233,X234,X232] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X233,happ(s(fun(X231,X233),X234),s(X231,X232))),s(fun(X233,bool),X235))),s(bool,'const.bool.IN_2'(s(X233,happ(s(fun(X231,X233),X234),s(X231,X232))),s(fun(X233,bool),X235))),s(bool,'const.bool.IN_2'(s(X233,happ(s(fun(X231,X233),X234),s(X231,X232))),s(fun(X233,bool),X235))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X231,X232),s(fun(X231,bool),'const.util_prob.PREIMAGE_2'(s(fun(X231,X233),X234),s(fun(X233,bool),X235)))))) ) | $spl0),
  inference(superposition,[],[f3969,f118])).
fof(f6894,plain,(
  ( ! [X239,X237,X235,X238,X236,X234] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X235,X237),s(fun(X235,bool),'const.util_prob.PREIMAGE_2'(s(fun(X235,X234),X236),s(fun(X234,bool),X238))))),s(bool,X239),s(bool,'const.bool.IN_2'(s(X235,X237),s(fun(X235,bool),'const.util_prob.PREIMAGE_2'(s(fun(X235,X234),X236),s(fun(X234,bool),X238))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X234,happ(s(fun(X235,X234),X236),s(X235,X237))),s(fun(X234,bool),X238)))) ) | $spl0),
  inference(superposition,[],[f3989,f118])).
fof(f6893,plain,(
  ( ! [X231,X229,X233,X230,X232] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X230,X232),s(fun(X230,bool),'const.util_prob.PREIMAGE_2'(s(fun(X230,X229),X231),s(fun(X229,bool),X233))))),s(bool,'const.bool.IN_2'(s(X230,X232),s(fun(X230,bool),'const.util_prob.PREIMAGE_2'(s(fun(X230,X229),X231),s(fun(X229,bool),X233))))),s(bool,'const.bool.IN_2'(s(X230,X232),s(fun(X230,bool),'const.util_prob.PREIMAGE_2'(s(fun(X230,X229),X231),s(fun(X229,bool),X233))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X229,happ(s(fun(X230,X229),X231),s(X230,X232))),s(fun(X229,bool),X233)))) ) | $spl0),
  inference(superposition,[],[f3969,f118])).
fof(f6781,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,X17) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X17),s(bool,X18),s(bool,X17)))) ) | $spl0),
  inference(subsumption_resolution,[],[f6743,f385])).
fof(f6743,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X17) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X17),s(bool,X18),s(bool,X17)))) ) | $spl0),
  inference(superposition,[],[f3989,f1133])).
fof(f6754,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,X7) = s(bool,'const.bool.COND_3'(s(bool,X5),s(bool,X6),s(bool,X5))) | s(bool,'T_0') = s(bool,X7)) ) | $spl0),
  inference(trivial_inequality_removal,[],[f6739])).
fof(f6739,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,X5) | s(bool,X7) = s(bool,'const.bool.COND_3'(s(bool,X5),s(bool,X6),s(bool,X5))) | s(bool,'T_0') = s(bool,X7)) ) | $spl0),
  inference(superposition,[],[f3989,f151])).
fof(f6734,plain,(
  ( ! [X74,X72,X75,X73] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X73),s(bool,X75),s(bool,X73))) | s(bool,'T_0') = s(bool,happ(s(fun(X72,bool),'const.combin.K_1'(s(bool,X73))),s(X72,X74)))) ) | $spl0),
  inference(superposition,[],[f3989,f120])).
fof(f6733,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X70))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X68),s(bool,X69),s(bool,X70))) | s(bool,'T_0') = s(bool,X68)) ) | $spl0),
  inference(superposition,[],[f3989,f281])).
fof(f6732,plain,(
  ( ! [X66,X67,X65] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X66),s(bool,X67),s(bool,X66))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X65),s(bool,X66)))) ) | $spl0),
  inference(superposition,[],[f3989,f108])).
fof(f6731,plain,(
  ( ! [X64,X62,X63] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X62),s(bool,X64),s(bool,X62))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X62),s(bool,X63)))) ) | $spl0),
  inference(superposition,[],[f3989,f107])).
fof(f6730,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X58),s(bool,X61),s(bool,X58))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X58),s(X59,X60)))) ) | $spl0),
  inference(superposition,[],[f3989,f115])).
fof(f6692,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) != s(bool,'const.bool.COND_3'(s(bool,X32),s(bool,X31),s(bool,X32))) | s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl0),
  inference(superposition,[],[f3989,f151])).
fof(f6671,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X9),s(bool,X9),s(bool,X9)))) ) | $spl0),
  inference(subsumption_resolution,[],[f6648,f385])).
fof(f6648,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X9) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X9),s(bool,X9),s(bool,X9)))) ) | $spl0),
  inference(superposition,[],[f3969,f1133])).
fof(f6659,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X1),s(bool,X1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(trivial_inequality_removal,[],[f6644])).
fof(f6644,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,'const.bool.COND_3'(s(bool,X1),s(bool,X1),s(bool,X1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(superposition,[],[f3969,f151])).
fof(f6640,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X47),s(bool,X47),s(bool,X47))) | s(bool,'T_0') = s(bool,happ(s(fun(X46,bool),'const.combin.K_1'(s(bool,X47))),s(X46,X48)))) ) | $spl0),
  inference(superposition,[],[f3969,f120])).
fof(f6639,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X45),s(bool,X45),s(bool,X45))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X43),s(bool,X44),s(bool,X45))) | s(bool,'T_0') = s(bool,X43)) ) | $spl0),
  inference(superposition,[],[f3969,f281])).
fof(f6638,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X42),s(bool,X42),s(bool,X42))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X41),s(bool,X42)))) ) | $spl0),
  inference(superposition,[],[f3969,f108])).
fof(f6637,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X39),s(bool,X39),s(bool,X39))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X39),s(bool,X40)))) ) | $spl0),
  inference(superposition,[],[f3969,f107])).
fof(f6636,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X36),s(bool,X36),s(bool,X36))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X36),s(X37,X38)))) ) | $spl0),
  inference(superposition,[],[f3969,f115])).
fof(f1133,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'F_0') = s(bool,X9)) ) | $spl0),
  inference(subsumption_resolution,[],[f1108,f83])).
fof(f1108,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | s(bool,X9) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'F_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f503,f385])).
fof(f1178,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X24),s(bool,X25)))) ) | $spl0),
  inference(forward_demodulation,[],[f1152,f385])).
fof(f1152,plain,(
  ( ! [X24,X25] : (s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X24),s(bool,X25)))) ) | $spl0),
  inference(superposition,[],[f1132,f108])).
fof(f1311,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X31),s(X32,X33)))) ) | $spl0),
  inference(superposition,[],[f1133,f115])).
fof(f1414,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X34),s(bool,X35))) | s(bool,'T_0') = s(bool,X34)) ) | $spl0),
  inference(forward_demodulation,[],[f1312,f385])).
fof(f1312,plain,(
  ( ! [X35,X34] : (s(bool,X34) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X34),s(bool,X35)))) ) | $spl0),
  inference(superposition,[],[f1133,f107])).
fof(f1415,plain,(
  ( ! [X37,X36] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X36),s(bool,X37))) | s(bool,'T_0') = s(bool,X37)) ) | $spl0),
  inference(forward_demodulation,[],[f1313,f385])).
fof(f1313,plain,(
  ( ! [X37,X36] : (s(bool,X37) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X36),s(bool,X37)))) ) | $spl0),
  inference(superposition,[],[f1133,f108])).
fof(f1416,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') = s(bool,X39) | s(bool,'F_0') = s(bool,happ(s(fun(X38,bool),'const.combin.K_1'(s(bool,X39))),s(X38,X40)))) ) | $spl0),
  inference(forward_demodulation,[],[f1314,f385])).
fof(f1314,plain,(
  ( ! [X39,X38,X40] : (s(bool,X39) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'F_0') = s(bool,happ(s(fun(X38,bool),'const.combin.K_1'(s(bool,X39))),s(X38,X40)))) ) | $spl0),
  inference(superposition,[],[f1133,f120])).
fof(f1434,plain,(
  ( ! [X41,X42,X40] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X40,sK4(X41,X42,X40)),s(fun(X40,bool),X41))) | s(fun(X40,bool),X41) = s(fun(X40,bool),X42) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X40,sK4(X41,X42,X40)),s(fun(X40,bool),X42)))) ) | $spl0),
  inference(forward_demodulation,[],[f1344,f385])).
fof(f1344,plain,(
  ( ! [X41,X42,X40] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) != s(bool,'const.bool.IN_2'(s(X40,sK4(X41,X42,X40)),s(fun(X40,bool),X41))) | s(fun(X40,bool),X41) = s(fun(X40,bool),X42) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X40,sK4(X41,X42,X40)),s(fun(X40,bool),X42)))) ) | $spl0),
  inference(superposition,[],[f110,f1133])).
fof(f1435,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X43,sK4(X44,X45,X43)),s(fun(X43,bool),X45))) | s(fun(X43,bool),X44) = s(fun(X43,bool),X45) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X43,sK4(X44,X45,X43)),s(fun(X43,bool),X44)))) ) | $spl0),
  inference(forward_demodulation,[],[f1345,f385])).
fof(f1345,plain,(
  ( ! [X45,X43,X44] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) != s(bool,'const.bool.IN_2'(s(X43,sK4(X44,X45,X43)),s(fun(X43,bool),X45))) | s(fun(X43,bool),X44) = s(fun(X43,bool),X45) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X43,sK4(X44,X45,X43)),s(fun(X43,bool),X44)))) ) | $spl0),
  inference(superposition,[],[f110,f1133])).
fof(f1437,plain,(
  ( ! [X88,X87,X89] : (s(bool,'T_0') != s(bool,f(s(X87,sK5(X88,X89,bool,X87)))) | s(fun(X87,bool),X88) = s(fun(X87,bool),X89) | s(bool,'F_0') = s(bool,g(s(X87,sK5(X88,X89,bool,X87))))) ) | $spl0),
  inference(forward_demodulation,[],[f1363,f385])).
fof(f1363,plain,(
  ( ! [X88,X87,X89] : (s(bool,f(s(X87,sK5(X88,X89,bool,X87)))) != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(fun(X87,bool),X88) = s(fun(X87,bool),X89) | s(bool,'F_0') = s(bool,g(s(X87,sK5(X88,X89,bool,X87))))) ) | $spl0),
  inference(superposition,[],[f117,f1133])).
fof(f1441,plain,(
  ( ! [X105,X104] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X104))) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X105))) | s(fun(bool,bool),X104) = s(fun(bool,bool),X105) | s(bool,'F_0') = s(bool,sK4(X104,X105,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f1368,f385])).
fof(f1368,plain,(
  ( ! [X105,X104] : (s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(bool,bool),X104))) != s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))),s(fun(bool,bool),X105))) | s(fun(bool,bool),X104) = s(fun(bool,bool),X105) | s(bool,'F_0') = s(bool,sK4(X104,X105,bool))) ) | $spl0),
  inference(superposition,[],[f110,f1133])).
fof(f1442,plain,(
  ( ! [X107,X108,X106] : (s(X108,f(s(bool,'T_0'))) != s(X108,g(s(bool,'T_0'))) | s(fun(bool,X108),X106) = s(fun(bool,X108),X107) | s(bool,'F_0') = s(bool,sK5(X106,X107,X108,bool))) ) | $spl0),
  inference(forward_demodulation,[],[f1369,f385])).
fof(f1369,plain,(
  ( ! [X107,X108,X106] : (s(X108,f(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))))) != s(X108,g(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))))) | s(fun(bool,X108),X106) = s(fun(bool,X108),X107) | s(bool,'F_0') = s(bool,sK5(X106,X107,X108,bool))) ) | $spl0),
  inference(superposition,[],[f117,f1133])).
fof(f1475,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (s(bool,X51) = s(bool,'const.bool.IN_2'(s(X48,happ(s(fun(X46,X48),X49),s(X46,X47))),s(fun(X48,bool),X50))) | s(bool,'F_0') = s(bool,X51) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),'const.util_prob.PREIMAGE_2'(s(fun(X46,X48),X49),s(fun(X48,bool),X50)))))) ) | $spl0),
  inference(superposition,[],[f1315,f118])).
fof(f1476,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,X57) = s(bool,'const.bool.IN_2'(s(X53,X55),s(fun(X53,bool),'const.util_prob.PREIMAGE_2'(s(fun(X53,X52),X54),s(fun(X52,bool),X56))))) | s(bool,'F_0') = s(bool,X57) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X52,happ(s(fun(X53,X52),X54),s(X53,X55))),s(fun(X52,bool),X56)))) ) | $spl0),
  inference(superposition,[],[f1315,f118])).
fof(f1480,plain,(
  ( ! [X68,X66,X67,X65] : (s(bool,X65) = s(bool,X68) | s(bool,'F_0') = s(bool,X68) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X65),s(X66,X67)))) ) | $spl0),
  inference(superposition,[],[f1315,f115])).
fof(f1481,plain,(
  ( ! [X70,X71,X69] : (s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X69),s(bool,X70)))) ) | $spl0),
  inference(superposition,[],[f1315,f107])).
fof(f1482,plain,(
  ( ! [X74,X72,X73] : (s(bool,X73) = s(bool,X74) | s(bool,'F_0') = s(bool,X74) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X72),s(bool,X73)))) ) | $spl0),
  inference(superposition,[],[f1315,f108])).
fof(f1483,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,X76) = s(bool,X78) | s(bool,'F_0') = s(bool,X78) | s(bool,'F_0') = s(bool,happ(s(fun(X75,bool),'const.combin.K_1'(s(bool,X76))),s(X75,X77)))) ) | $spl0),
  inference(superposition,[],[f1315,f120])).
fof(f1536,plain,(
  ( ! [X57,X58,X56] : (s(bool,'F_0') = s(bool,X56) | s(bool,X58) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X56),s(bool,X57))) | s(bool,'F_0') = s(bool,X58)) ) | $spl0),
  inference(superposition,[],[f1315,f107])).
fof(f1537,plain,(
  ( ! [X61,X59,X60] : (s(bool,'F_0') = s(bool,X60) | s(bool,X61) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X59),s(bool,X60))) | s(bool,'F_0') = s(bool,X61)) ) | $spl0),
  inference(superposition,[],[f1315,f108])).
fof(f1538,plain,(
  ( ! [X64,X62,X65,X63] : (s(bool,'F_0') = s(bool,X63) | s(bool,X65) = s(bool,happ(s(fun(X62,bool),'const.combin.K_1'(s(bool,X63))),s(X62,X64))) | s(bool,'F_0') = s(bool,X65)) ) | $spl0),
  inference(superposition,[],[f1315,f120])).
fof(f1599,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,X91) != s(bool,'const.bool.IN_2'(s(X88,sK4(X89,X90,X88)),s(fun(X88,bool),X89))) | s(fun(X88,bool),X89) = s(fun(X88,bool),X90) | s(bool,'F_0') = s(bool,X91) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X88,sK4(X89,X90,X88)),s(fun(X88,bool),X90)))) ) | $spl0),
  inference(superposition,[],[f110,f1315])).
fof(f1600,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,X95) != s(bool,'const.bool.IN_2'(s(X92,sK4(X93,X94,X92)),s(fun(X92,bool),X94))) | s(fun(X92,bool),X93) = s(fun(X92,bool),X94) | s(bool,'F_0') = s(bool,X95) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X92,sK4(X93,X94,X92)),s(fun(X92,bool),X93)))) ) | $spl0),
  inference(superposition,[],[f110,f1315])).
fof(f1619,plain,(
  ( ! [X158,X159,X157,X160] : (s(bool,X160) != s(bool,f(s(X157,sK5(X158,X159,bool,X157)))) | s(fun(X157,bool),X158) = s(fun(X157,bool),X159) | s(bool,'F_0') = s(bool,X160) | s(bool,'F_0') = s(bool,g(s(X157,sK5(X158,X159,bool,X157))))) ) | $spl0),
  inference(superposition,[],[f117,f1315])).
fof(f1624,plain,(
  ( ! [X182,X180,X181] : (s(bool,'const.bool.IN_2'(s(bool,X182),s(fun(bool,bool),X180))) != s(bool,'const.bool.IN_2'(s(bool,X182),s(fun(bool,bool),X181))) | s(fun(bool,bool),X180) = s(fun(bool,bool),X181) | s(bool,'F_0') = s(bool,X182) | s(bool,'F_0') = s(bool,sK4(X180,X181,bool))) ) | $spl0),
  inference(superposition,[],[f110,f1315])).
fof(f1625,plain,(
  ( ! [X185,X186,X184,X183] : (s(X185,f(s(bool,X186))) != s(X185,g(s(bool,X186))) | s(fun(bool,X185),X183) = s(fun(bool,X185),X184) | s(bool,'F_0') = s(bool,X186) | s(bool,'F_0') = s(bool,sK5(X183,X184,X185,bool))) ) | $spl0),
  inference(superposition,[],[f117,f1315])).
fof(f1685,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | ~p(s(bool,X4)) | s(bool,X3) = s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f97,f1315])).
fof(f1686,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(X7,X8) = s(X7,'const.bool.COND_3'(s(bool,X6),s(X7,X9),s(X7,X8))) | s(bool,X6) = s(bool,X10) | s(bool,'F_0') = s(bool,X10)) ) | $spl0),
  inference(superposition,[],[f108,f1315])).
fof(f1714,plain,(
  ( ! [X94,X93] : (s(bool,X93) = s(bool,'const.bool.IN_2'(s(bool,X93),s(fun(bool,bool),'const.pred_set.EMPTY_0'))) | s(bool,X93) = s(bool,X94) | s(bool,'F_0') = s(bool,X94)) ) | $spl0),
  inference(superposition,[],[f132,f1315])).
fof(f1742,plain,(
  ( ! [X74,X72,X71,X73] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X71,sK4(X72,X73,X71)),s(fun(X71,bool),X72))) | s(fun(X71,bool),X72) = s(fun(X71,bool),X73) | s(bool,X74) = s(bool,'const.bool.IN_2'(s(X71,sK4(X72,X73,X71)),s(fun(X71,bool),X73))) | s(bool,'F_0') = s(bool,X74)) ) | $spl0),
  inference(superposition,[],[f110,f1315])).
fof(f1743,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X75,sK4(X76,X77,X75)),s(fun(X75,bool),X77))) | s(fun(X75,bool),X76) = s(fun(X75,bool),X77) | s(bool,X78) = s(bool,'const.bool.IN_2'(s(X75,sK4(X76,X77,X75)),s(fun(X75,bool),X76))) | s(bool,'F_0') = s(bool,X78)) ) | $spl0),
  inference(superposition,[],[f110,f1315])).
fof(f1762,plain,(
  ( ! [X142,X140,X141,X139] : (s(bool,'F_0') != s(bool,f(s(X139,sK5(X140,X141,bool,X139)))) | s(fun(X139,bool),X140) = s(fun(X139,bool),X141) | s(bool,X142) = s(bool,g(s(X139,sK5(X140,X141,bool,X139)))) | s(bool,'F_0') = s(bool,X142)) ) | $spl0),
  inference(superposition,[],[f117,f1315])).
fof(f1767,plain,(
  ( ! [X163,X161,X162] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X161))) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X162))) | s(fun(bool,bool),X161) = s(fun(bool,bool),X162) | s(bool,X163) = s(bool,sK4(X161,X162,bool)) | s(bool,'F_0') = s(bool,X163)) ) | $spl0),
  inference(superposition,[],[f110,f1315])).
fof(f1768,plain,(
  ( ! [X167,X165,X166,X164] : (s(X166,f(s(bool,'F_0'))) != s(X166,g(s(bool,'F_0'))) | s(fun(bool,X166),X164) = s(fun(bool,X166),X165) | s(bool,X167) = s(bool,sK5(X164,X165,X166,bool)) | s(bool,'F_0') = s(bool,X167)) ) | $spl0),
  inference(superposition,[],[f117,f1315])).
fof(f2054,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl0),
  inference(subsumption_resolution,[],[f1721,f119])).
fof(f1721,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'F_0')) | p(s(bool,X10)) | s(bool,X9) = s(bool,X10) | s(bool,X9) = s(bool,X11) | s(bool,'F_0') = s(bool,X11)) ) | $spl0),
  inference(superposition,[],[f111,f1315])).
fof(f2258,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (s(bool,X51) != s(bool,'const.bool.IN_2'(s(X48,happ(s(fun(X46,X48),X49),s(X46,X47))),s(fun(X48,bool),X50))) | s(bool,'T_0') = s(bool,X51) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),'const.util_prob.PREIMAGE_2'(s(fun(X46,X48),X49),s(fun(X48,bool),X50)))))) ) | $spl0),
  inference(superposition,[],[f1430,f118])).
fof(f2259,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,X57) != s(bool,'const.bool.IN_2'(s(X53,X55),s(fun(X53,bool),'const.util_prob.PREIMAGE_2'(s(fun(X53,X52),X54),s(fun(X52,bool),X56))))) | s(bool,'T_0') = s(bool,X57) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X52,happ(s(fun(X53,X52),X54),s(X53,X55))),s(fun(X52,bool),X56)))) ) | $spl0),
  inference(superposition,[],[f1430,f118])).
fof(f2264,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X71) | s(bool,'T_0') = s(bool,X71) | s(bool,'F_0') = s(bool,'const.combin.K_2'(s(bool,X68),s(X69,X70)))) ) | $spl0),
  inference(superposition,[],[f1430,f115])).
fof(f2265,plain,(
  ( ! [X74,X72,X73] : (s(bool,X72) != s(bool,X74) | s(bool,'T_0') = s(bool,X74) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X72),s(bool,X73)))) ) | $spl0),
  inference(superposition,[],[f1430,f107])).
fof(f2266,plain,(
  ( ! [X76,X77,X75] : (s(bool,X76) != s(bool,X77) | s(bool,'T_0') = s(bool,X77) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X75),s(bool,X76)))) ) | $spl0),
  inference(superposition,[],[f1430,f108])).
fof(f2267,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,X79) != s(bool,X81) | s(bool,'T_0') = s(bool,X81) | s(bool,'F_0') = s(bool,happ(s(fun(X78,bool),'const.combin.K_1'(s(bool,X79))),s(X78,X80)))) ) | $spl0),
  inference(superposition,[],[f1430,f120])).
fof(f2285,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (s(bool,X51) != s(bool,'const.bool.IN_2'(s(X48,happ(s(fun(X46,X48),X49),s(X46,X47))),s(fun(X48,bool),X50))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),'const.util_prob.PREIMAGE_2'(s(fun(X46,X48),X49),s(fun(X48,bool),X50))))) | s(bool,'F_0') = s(bool,X51)) ) | $spl0),
  inference(superposition,[],[f1430,f118])).
fof(f2286,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,X57) != s(bool,'const.bool.IN_2'(s(X53,X55),s(fun(X53,bool),'const.util_prob.PREIMAGE_2'(s(fun(X53,X52),X54),s(fun(X52,bool),X56))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X52,happ(s(fun(X53,X52),X54),s(X53,X55))),s(fun(X52,bool),X56))) | s(bool,'F_0') = s(bool,X57)) ) | $spl0),
  inference(superposition,[],[f1430,f118])).
fof(f2291,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,X68) != s(bool,X71) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X68),s(X69,X70))) | s(bool,'F_0') = s(bool,X71)) ) | $spl0),
  inference(superposition,[],[f1430,f115])).
fof(f2292,plain,(
  ( ! [X74,X72,X73] : (s(bool,X72) != s(bool,X74) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X72),s(bool,X73))) | s(bool,'F_0') = s(bool,X74)) ) | $spl0),
  inference(superposition,[],[f1430,f107])).
fof(f2293,plain,(
  ( ! [X76,X77,X75] : (s(bool,X76) != s(bool,X77) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X75),s(bool,X76))) | s(bool,'F_0') = s(bool,X77)) ) | $spl0),
  inference(superposition,[],[f1430,f108])).
fof(f2294,plain,(
  ( ! [X80,X78,X81,X79] : (s(bool,X79) != s(bool,X81) | s(bool,'T_0') = s(bool,happ(s(fun(X78,bool),'const.combin.K_1'(s(bool,X79))),s(X78,X80))) | s(bool,'F_0') = s(bool,X81)) ) | $spl0),
  inference(superposition,[],[f1430,f120])).
fof(f2359,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X50,happ(s(fun(X48,X50),X51),s(X48,X49))),s(fun(X50,bool),X52))) | s(bool,X53) = s(bool,'const.bool.IN_2'(s(X48,X49),s(fun(X48,bool),'const.util_prob.PREIMAGE_2'(s(fun(X48,X50),X51),s(fun(X50,bool),X52))))) | s(bool,'F_0') = s(bool,X53)) ) | $spl0),
  inference(superposition,[],[f1688,f118])).
fof(f2360,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X55,X57),s(fun(X55,bool),'const.util_prob.PREIMAGE_2'(s(fun(X55,X54),X56),s(fun(X54,bool),X58))))) | s(bool,X59) = s(bool,'const.bool.IN_2'(s(X54,happ(s(fun(X55,X54),X56),s(X55,X57))),s(fun(X54,bool),X58))) | s(bool,'F_0') = s(bool,X59)) ) | $spl0),
  inference(superposition,[],[f1688,f118])).
fof(f2365,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') != s(bool,X70) | s(bool,X73) = s(bool,'const.combin.K_2'(s(bool,X70),s(X71,X72))) | s(bool,'F_0') = s(bool,X73)) ) | $spl0),
  inference(superposition,[],[f1688,f115])).
fof(f2366,plain,(
  ( ! [X76,X74,X75] : (s(bool,'T_0') != s(bool,X74) | s(bool,X76) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X74),s(bool,X75))) | s(bool,'F_0') = s(bool,X76)) ) | $spl0),
  inference(superposition,[],[f1688,f107])).
fof(f2367,plain,(
  ( ! [X78,X79,X77] : (s(bool,'T_0') != s(bool,X78) | s(bool,X79) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X77),s(bool,X78))) | s(bool,'F_0') = s(bool,X79)) ) | $spl0),
  inference(superposition,[],[f1688,f108])).
fof(f2368,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X81) | s(bool,X83) = s(bool,happ(s(fun(X80,bool),'const.combin.K_1'(s(bool,X81))),s(X80,X82))) | s(bool,'F_0') = s(bool,X83)) ) | $spl0),
  inference(superposition,[],[f1688,f120])).
fof(f2489,plain,(
  ( ! [X37,X35,X36,X34] : (s(bool,'F_0') = s(bool,X36) | s(bool,'T_0') = s(bool,X34) | s(bool,X37) = s(bool,'const.bool.COND_3'(s(bool,X34),s(bool,X35),s(bool,X36))) | s(bool,'F_0') = s(bool,X37)) ) | $spl0),
  inference(superposition,[],[f281,f1315])).
fof(f2490,plain,(
  ( ! [X39,X41,X38,X40] : (s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X41) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X38),s(bool,X39),s(bool,X40)))) ) | $spl0),
  inference(superposition,[],[f281,f1315])).
fof(f2532,plain,(
  ( ! [X167,X169,X166,X168] : (s(bool,X168) != s(bool,X169) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X166),s(bool,X167),s(bool,X168))) | s(bool,'F_0') = s(bool,X169) | s(bool,'T_0') = s(bool,X166)) ) | $spl0),
  inference(superposition,[],[f1430,f281])).
fof(f2533,plain,(
  ( ! [X173,X171,X172,X170] : (s(bool,X172) != s(bool,X173) | s(bool,'T_0') = s(bool,X173) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X170),s(bool,X171),s(bool,X172))) | s(bool,'T_0') = s(bool,X170)) ) | $spl0),
  inference(superposition,[],[f1430,f281])).
fof(f2534,plain,(
  ( ! [X177,X175,X176,X174] : (s(bool,'T_0') != s(bool,X176) | s(bool,X177) = s(bool,'const.bool.COND_3'(s(bool,X174),s(bool,X175),s(bool,X176))) | s(bool,'F_0') = s(bool,X177) | s(bool,'T_0') = s(bool,X174)) ) | $spl0),
  inference(superposition,[],[f1688,f281])).
fof(f2576,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X31),s(bool,X32),s(bool,X33)))) ) | $spl0),
  inference(forward_demodulation,[],[f2488,f385])).
fof(f2488,plain,(
  ( ! [X33,X31,X32] : (s(bool,X33) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X31),s(bool,X32),s(bool,X33)))) ) | $spl0),
  inference(superposition,[],[f281,f1133])).
fof(f2664,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X2),s(bool,X1))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f1178,f151])).
fof(f2689,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X50,happ(s(fun(X48,X50),X51),s(X48,X49))),s(fun(X50,bool),X52))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X53),s(bool,'const.bool.IN_2'(s(X48,X49),s(fun(X48,bool),'const.util_prob.PREIMAGE_2'(s(fun(X48,X50),X51),s(fun(X50,bool),X52)))))))) ) | $spl0),
  inference(superposition,[],[f1178,f118])).
fof(f2690,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X55,X57),s(fun(X55,bool),'const.util_prob.PREIMAGE_2'(s(fun(X55,X54),X56),s(fun(X54,bool),X58))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X59),s(bool,'const.bool.IN_2'(s(X54,happ(s(fun(X55,X54),X56),s(X55,X57))),s(fun(X54,bool),X58)))))) ) | $spl0),
  inference(superposition,[],[f1178,f118])).
fof(f2698,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X82) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X83),s(bool,'const.bool.COND_3'(s(bool,X80),s(bool,X81),s(bool,X82))))) | s(bool,'T_0') = s(bool,X80)) ) | $spl0),
  inference(superposition,[],[f1178,f281])).
fof(f2738,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f297,f1315])).
fof(f2918,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f1414,f151])).
fof(f2943,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.IN_2'(s(X48,happ(s(fun(X46,X48),X49),s(X46,X47))),s(fun(X48,bool),X50))),s(bool,X51))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),'const.util_prob.PREIMAGE_2'(s(fun(X46,X48),X49),s(fun(X48,bool),X50)))))) ) | $spl0),
  inference(superposition,[],[f1414,f118])).
fof(f2944,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.IN_2'(s(X53,X55),s(fun(X53,bool),'const.util_prob.PREIMAGE_2'(s(fun(X53,X52),X54),s(fun(X52,bool),X56))))),s(bool,X57))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X52,happ(s(fun(X53,X52),X54),s(X53,X55))),s(fun(X52,bool),X56)))) ) | $spl0),
  inference(superposition,[],[f1414,f118])).
fof(f2954,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X86),s(bool,X87))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X84),s(bool,X85),s(bool,X86))) | s(bool,'T_0') = s(bool,X84)) ) | $spl0),
  inference(superposition,[],[f1414,f281])).
fof(f2955,plain,(
  ( ! [X90,X88,X91,X89] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X89),s(bool,X91))) | s(bool,'T_0') = s(bool,happ(s(fun(X88,bool),'const.combin.K_1'(s(bool,X89))),s(X88,X90)))) ) | $spl0),
  inference(superposition,[],[f1414,f120])).
fof(f2968,plain,(
  ( ! [X35,X36,X34] : (s(bool,X34) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X35),s(bool,X34))) | s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,X36) | s(bool,'F_0') = s(bool,X36)) ) | $spl0),
  inference(superposition,[],[f1414,f1315])).
fof(f3105,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f1415,f1315])).
fof(f3165,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X51),s(bool,'const.bool.IN_2'(s(X48,happ(s(fun(X46,X48),X49),s(X46,X47))),s(fun(X48,bool),X50))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),'const.util_prob.PREIMAGE_2'(s(fun(X46,X48),X49),s(fun(X48,bool),X50)))))) ) | $spl0),
  inference(superposition,[],[f1415,f118])).
fof(f3166,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X57),s(bool,'const.bool.IN_2'(s(X53,X55),s(fun(X53,bool),'const.util_prob.PREIMAGE_2'(s(fun(X53,X52),X54),s(fun(X52,bool),X56))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X52,happ(s(fun(X53,X52),X54),s(X53,X55))),s(fun(X52,bool),X56)))) ) | $spl0),
  inference(superposition,[],[f1415,f118])).
fof(f3177,plain,(
  ( ! [X90,X88,X87,X89] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X90),s(bool,X89))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X87),s(bool,X88),s(bool,X89))) | s(bool,'T_0') = s(bool,X87)) ) | $spl0),
  inference(superposition,[],[f1415,f281])).
fof(f3178,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X94),s(bool,X92))) | s(bool,'T_0') = s(bool,happ(s(fun(X91,bool),'const.combin.K_1'(s(bool,X92))),s(X91,X93)))) ) | $spl0),
  inference(superposition,[],[f1415,f120])).
fof(f3290,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | $spl0),
  inference(subsumption_resolution,[],[f3190,f119])).
fof(f3190,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X4),s(bool,X5))) | s(bool,'T_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f93,f1415])).
fof(f3519,plain,(
  ( ! [X90,X88,X89] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X89),s(bool,X90))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X88),s(bool,X89)))) ) | $spl0),
  inference(superposition,[],[f3093,f108])).
fof(f3595,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X2)))) ) | $spl0),
  inference(forward_demodulation,[],[f3480,f132])).
fof(f3480,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))),s(bool,X2)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f133,f3093])).
fof(f3596,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X32),s(bool,X33))) | s(bool,X32) = s(bool,X34) | s(bool,'F_0') = s(bool,X34)) ) | $spl0),
  inference(subsumption_resolution,[],[f3502,f121])).
fof(f3502,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X32),s(bool,X33))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X32) = s(bool,X34) | s(bool,'F_0') = s(bool,X34)) ) | $spl0),
  inference(superposition,[],[f3093,f1315])).
fof(f3643,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'F_0'),s(bool,X1))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(superposition,[],[f3595,f151])).
fof(f3710,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) ) | $spl0),
  inference(trivial_inequality_removal,[],[f3698])).
fof(f3698,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X3))) | s(bool,'T_0') = s(bool,X4)) ) | $spl0),
  inference(superposition,[],[f3595,f151])).
fof(f3749,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X1,f(s(bool,X0))) != s(X1,g(s(bool,X0))) | s(fun(bool,X1),X2) = s(fun(bool,X1),X3) | s(bool,'T_0') = s(bool,sK5(X2,X3,X1,bool)) | s(bool,X0) = s(bool,X4) | s(bool,'F_0') = s(bool,X4)) ) | $spl0),
  inference(superposition,[],[f435,f1315])).
fof(f3789,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f3290,f1315])).
fof(f3811,plain,(
  ( ! [X30,X28,X29] : (s(bool,X28) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X28),s(bool,X29))) | s(bool,'T_0') = s(bool,X29) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | $spl0),
  inference(superposition,[],[f3290,f151])).
fof(f3850,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X51),s(bool,'const.bool.IN_2'(s(X48,happ(s(fun(X46,X48),X49),s(X46,X47))),s(fun(X48,bool),X50))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),'const.util_prob.PREIMAGE_2'(s(fun(X46,X48),X49),s(fun(X48,bool),X50)))))) ) | $spl0),
  inference(superposition,[],[f3290,f118])).
fof(f3851,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X57),s(bool,'const.bool.IN_2'(s(X53,X55),s(fun(X53,bool),'const.util_prob.PREIMAGE_2'(s(fun(X53,X52),X54),s(fun(X52,bool),X56))))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X52,happ(s(fun(X53,X52),X54),s(X53,X55))),s(fun(X52,bool),X56)))) ) | $spl0),
  inference(superposition,[],[f3290,f118])).
fof(f3860,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X84),s(bool,X81))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X81),s(X82,X83)))) ) | $spl0),
  inference(superposition,[],[f3290,f115])).
fof(f3861,plain,(
  ( ! [X87,X85,X86] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X87),s(bool,X85))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X85),s(bool,X86)))) ) | $spl0),
  inference(superposition,[],[f3290,f107])).
fof(f3862,plain,(
  ( ! [X90,X88,X89] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X90),s(bool,X89))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X88),s(bool,X89)))) ) | $spl0),
  inference(superposition,[],[f3290,f108])).
fof(f3863,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X94),s(bool,X93))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X91),s(bool,X92),s(bool,X93))) | s(bool,'T_0') = s(bool,X91)) ) | $spl0),
  inference(superposition,[],[f3290,f281])).
fof(f3864,plain,(
  ( ! [X97,X95,X98,X96] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X98),s(bool,X96))) | s(bool,'T_0') = s(bool,happ(s(fun(X95,bool),'const.combin.K_1'(s(bool,X96))),s(X95,X97)))) ) | $spl0),
  inference(superposition,[],[f3290,f120])).
fof(f3904,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'F_0')))) ) | $spl0),
  inference(forward_demodulation,[],[f3787,f107])).
fof(f3787,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X0),s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X1)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f3595,f3290])).
fof(f3906,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X32),s(bool,X33),s(bool,X32))) | s(bool,X32) = s(bool,X34) | s(bool,'F_0') = s(bool,X34)) ) | $spl0),
  inference(subsumption_resolution,[],[f3845,f121])).
fof(f3845,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X32),s(bool,X33),s(bool,X32))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X32) = s(bool,X34) | s(bool,'F_0') = s(bool,X34)) ) | $spl0),
  inference(superposition,[],[f3290,f1315])).
fof(f3965,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X18),s(bool,'F_0'))) | s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | $spl0),
  inference(superposition,[],[f3904,f151])).
fof(f4014,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl0),
  inference(trivial_inequality_removal,[],[f4002])).
fof(f4002,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X4) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl0),
  inference(superposition,[],[f3904,f151])).
fof(f4136,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X53,X54),s(fun(X53,bool),X55))) | s(bool,X56) = s(bool,'const.bool.IN_2'(s(X51,X52),s(fun(X51,bool),'const.util_prob.PREIMAGE_2'(s(fun(X51,X53),'const.combin.K_1'(s(X53,X54))),s(fun(X53,bool),X55))))) | s(bool,'F_0') = s(bool,X56)) ) | $spl0),
  inference(superposition,[],[f1189,f1315])).
fof(f4137,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (s(bool,X62) = s(bool,'const.bool.IN_2'(s(X59,X60),s(fun(X59,bool),X61))) | s(bool,'F_0') = s(bool,X62) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X57,X58),s(fun(X57,bool),'const.util_prob.PREIMAGE_2'(s(fun(X57,X59),'const.combin.K_1'(s(X59,X60))),s(fun(X59,bool),X61)))))) ) | $spl0),
  inference(superposition,[],[f1189,f1315])).
fof(f4174,plain,(
  ( ! [X204,X202,X200,X199,X203,X201] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.IN_2'(s(X201,X202),s(fun(X201,bool),X203))),s(bool,X204))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X199,X200),s(fun(X199,bool),'const.util_prob.PREIMAGE_2'(s(fun(X199,X201),'const.combin.K_1'(s(X201,X202))),s(fun(X201,bool),X203)))))) ) | $spl0),
  inference(superposition,[],[f1414,f1189])).
fof(f4177,plain,(
  ( ! [X218,X222,X220,X219,X217,X221] : (s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X222),s(bool,'const.bool.IN_2'(s(X219,X220),s(fun(X219,bool),X221))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X217,X218),s(fun(X217,bool),'const.util_prob.PREIMAGE_2'(s(fun(X217,X219),'const.combin.K_1'(s(X219,X220))),s(fun(X219,bool),X221)))))) ) | $spl0),
  inference(superposition,[],[f1415,f1189])).
fof(f4178,plain,(
  ( ! [X227,X225,X223,X228,X226,X224] : (s(bool,X228) != s(bool,'const.bool.IN_2'(s(X225,X226),s(fun(X225,bool),X227))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X223,X224),s(fun(X223,bool),'const.util_prob.PREIMAGE_2'(s(fun(X223,X225),'const.combin.K_1'(s(X225,X226))),s(fun(X225,bool),X227))))) | s(bool,'F_0') = s(bool,X228)) ) | $spl0),
  inference(superposition,[],[f1430,f1189])).
fof(f4179,plain,(
  ( ! [X231,X229,X233,X230,X234,X232] : (s(bool,X234) != s(bool,'const.bool.IN_2'(s(X231,X232),s(fun(X231,bool),X233))) | s(bool,'T_0') = s(bool,X234) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X229,X230),s(fun(X229,bool),'const.util_prob.PREIMAGE_2'(s(fun(X229,X231),'const.combin.K_1'(s(X231,X232))),s(fun(X231,bool),X233)))))) ) | $spl0),
  inference(superposition,[],[f1430,f1189])).
fof(f4180,plain,(
  ( ! [X239,X237,X235,X240,X238,X236] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X237,X238),s(fun(X237,bool),X239))) | s(bool,X240) = s(bool,'const.bool.IN_2'(s(X235,X236),s(fun(X235,bool),'const.util_prob.PREIMAGE_2'(s(fun(X235,X237),'const.combin.K_1'(s(X237,X238))),s(fun(X237,bool),X239))))) | s(bool,'F_0') = s(bool,X240)) ) | $spl0),
  inference(superposition,[],[f1688,f1189])).
fof(f4184,plain,(
  ( ! [X261,X263,X259,X260,X262,X264] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X264),s(bool,'const.bool.IN_2'(s(X261,X262),s(fun(X261,bool),X263))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X259,X260),s(fun(X259,bool),'const.util_prob.PREIMAGE_2'(s(fun(X259,X261),'const.combin.K_1'(s(X261,X262))),s(fun(X261,bool),X263)))))) ) | $spl0),
  inference(superposition,[],[f3290,f1189])).
fof(f4222,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X48,X49),s(fun(X48,bool),X50))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),'const.util_prob.PREIMAGE_2'(s(fun(X46,X48),'const.combin.K_1'(s(X48,X49))),s(fun(X48,bool),X50)))))) ) | $spl0),
  inference(forward_demodulation,[],[f4135,f385])).
fof(f4135,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) = s(bool,'const.bool.IN_2'(s(X48,X49),s(fun(X48,bool),X50))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),'const.util_prob.PREIMAGE_2'(s(fun(X46,X48),'const.combin.K_1'(s(X48,X49))),s(fun(X48,bool),X50)))))) ) | $spl0),
  inference(superposition,[],[f1189,f1133])).
fof(f4270,plain,(
  ( ! [X94,X95,X93,X96] : (s(bool,X93) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X96))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X93),s(X94,X95)))) ) | $spl0),
  inference(superposition,[],[f3710,f115])).
fof(f4271,plain,(
  ( ! [X99,X97,X98] : (s(bool,X97) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X99))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X97),s(bool,X98)))) ) | $spl0),
  inference(superposition,[],[f3710,f107])).
fof(f4272,plain,(
  ( ! [X101,X102,X100] : (s(bool,X101) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X102))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X100),s(bool,X101)))) ) | $spl0),
  inference(superposition,[],[f3710,f108])).
fof(f4273,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,X105) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X106))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X103),s(bool,X104),s(bool,X105))) | s(bool,'T_0') = s(bool,X103)) ) | $spl0),
  inference(superposition,[],[f3710,f281])).
fof(f4274,plain,(
  ( ! [X109,X107,X110,X108] : (s(bool,X108) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X110))) | s(bool,'T_0') = s(bool,happ(s(fun(X107,bool),'const.combin.K_1'(s(bool,X108))),s(X107,X109)))) ) | $spl0),
  inference(superposition,[],[f3710,f120])).
fof(f4275,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,'F_0'),s(bool,X2))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f3710,f151])).
fof(f4283,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,X2))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f3710,f1315])).
fof(f1047,plain,(
  s(bool,f(s(sK1,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),bool,sK1)))) != s(bool,g(s(sK1,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),bool,sK1)))) | $spl0),
  inference(backward_demodulation,[],[f385,f433])).
fof(f1046,plain,(
  s(bool,f(s(sK1,sK5('const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),bool,sK1)))) != s(bool,g(s(sK1,sK5('const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),bool,sK1)))) | $spl0),
  inference(backward_demodulation,[],[f385,f431])).
fof(f4181,plain,(
  ( ! [X243,X241,X246,X244,X242,X245] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.IN_2'(s(X243,X244),s(fun(X243,bool),X245))),s(bool,X246))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X241,X242),s(fun(X241,bool),'const.util_prob.PREIMAGE_2'(s(fun(X241,X243),'const.combin.K_1'(s(X243,X244))),s(fun(X243,bool),X245)))))) ) | $spl0),
  inference(superposition,[],[f3093,f1189])).
fof(f4233,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X155,X156),s(fun(X155,bool),X157))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X153,X154),s(fun(X153,bool),'const.util_prob.PREIMAGE_2'(s(fun(X153,X155),'const.combin.K_1'(s(X155,X156))),s(fun(X155,bool),X157)))))) ) | $spl0),
  inference(forward_demodulation,[],[f4166,f385])).
fof(f4166,plain,(
  ( ! [X156,X154,X157,X155,X153] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) != s(bool,'const.bool.IN_2'(s(X155,X156),s(fun(X155,bool),X157))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X153,X154),s(fun(X153,bool),'const.util_prob.PREIMAGE_2'(s(fun(X153,X155),'const.combin.K_1'(s(X155,X156))),s(fun(X155,bool),X157)))))) ) | $spl0),
  inference(superposition,[],[f1132,f1189])).
fof(f2610,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X50,happ(s(fun(X48,X50),X51),s(X48,X49))),s(fun(X50,bool),X52))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.IN_2'(s(X48,X49),s(fun(X48,bool),'const.util_prob.PREIMAGE_2'(s(fun(X48,X50),X51),s(fun(X50,bool),X52))))),s(bool,X53)))) ) | $spl0),
  inference(superposition,[],[f1177,f118])).
fof(f3508,plain,(
  ( ! [X57,X54,X52,X56,X55,X53] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.IN_2'(s(X53,X55),s(fun(X53,bool),'const.util_prob.PREIMAGE_2'(s(fun(X53,X52),X54),s(fun(X52,bool),X56))))),s(bool,X57))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X52,happ(s(fun(X53,X52),X54),s(X53,X55))),s(fun(X52,bool),X56)))) ) | $spl0),
  inference(superposition,[],[f3093,f118])).
fof(f2611,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X55,X57),s(fun(X55,bool),'const.util_prob.PREIMAGE_2'(s(fun(X55,X54),X56),s(fun(X54,bool),X58))))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.IN_2'(s(X54,happ(s(fun(X55,X54),X56),s(X55,X57))),s(fun(X54,bool),X58))),s(bool,X59)))) ) | $spl0),
  inference(superposition,[],[f1177,f118])).
fof(f3507,plain,(
  ( ! [X47,X50,X48,X46,X51,X49] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.IN_2'(s(X48,happ(s(fun(X46,X48),X49),s(X46,X47))),s(fun(X48,bool),X50))),s(bool,X51))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),'const.util_prob.PREIMAGE_2'(s(fun(X46,X48),X49),s(fun(X48,bool),X50)))))) ) | $spl0),
  inference(superposition,[],[f3093,f118])).
fof(f2619,plain,(
  ( ! [X80,X83,X81,X82] : (s(bool,'T_0') != s(bool,X82) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'const.bool.COND_3'(s(bool,X80),s(bool,X81),s(bool,X82))),s(bool,X83))) | s(bool,'T_0') = s(bool,X80)) ) | $spl0),
  inference(superposition,[],[f1177,f281])).
fof(f4055,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f1189,f1132])).
fof(f3989,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X6),s(bool,X7),s(bool,X6))) | s(bool,'T_0') = s(bool,X6)) ) | $spl0),
  inference(superposition,[],[f3904,f91])).
fof(f3969,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X26),s(bool,X26),s(bool,X26))) | s(bool,'T_0') = s(bool,X26)) ) | $spl0),
  inference(superposition,[],[f3904,f91])).
fof(f3791,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X8),s(bool,X9),s(bool,X10))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X8)) ) | $spl0),
  inference(superposition,[],[f3290,f91])).
fof(f3521,plain,(
  ( ! [X97,X95,X98,X96] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X96),s(bool,X98))) | s(bool,'T_0') = s(bool,happ(s(fun(X95,bool),'const.combin.K_1'(s(bool,X96))),s(X95,X97)))) ) | $spl0),
  inference(superposition,[],[f3093,f120])).
fof(f3520,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X93),s(bool,X94))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X91),s(bool,X92),s(bool,X93))) | s(bool,'T_0') = s(bool,X91)) ) | $spl0),
  inference(superposition,[],[f3093,f281])).
fof(f3518,plain,(
  ( ! [X87,X85,X86] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X85),s(bool,X87))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X85),s(bool,X86)))) ) | $spl0),
  inference(superposition,[],[f3093,f107])).
fof(f3517,plain,(
  ( ! [X83,X81,X84,X82] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X81),s(bool,X84))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X81),s(X82,X83)))) ) | $spl0),
  inference(superposition,[],[f3093,f115])).
fof(f3482,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,'const.bool.COND_3'(s(bool,X0),s(bool,X1),s(bool,X2))) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f3093,f151])).
fof(f3107,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) = s(bool,'const.bool.COND_3'(s(bool,X8),s(bool,X9),s(bool,X10))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X8)) ) | $spl0),
  inference(superposition,[],[f1415,f91])).
fof(f3093,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5))) | s(bool,'T_0') = s(bool,X4)) ) | $spl0),
  inference(subsumption_resolution,[],[f2997,f119])).
fof(f2997,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X4),s(bool,X5))) | s(bool,'T_0') = s(bool,X4)) ) | $spl0),
  inference(superposition,[],[f93,f1414])).
fof(f2585,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X1),s(bool,X2))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f1177,f151])).
fof(f2577,plain,(
  ( ! [X146,X144,X145] : (s(bool,'T_0') != s(bool,X146) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X144),s(bool,X145),s(bool,X146))) | s(bool,'T_0') = s(bool,X144)) ) | $spl0),
  inference(forward_demodulation,[],[f2526,f385])).
fof(f2526,plain,(
  ( ! [X146,X144,X145] : (s(bool,X146) != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X144),s(bool,X145),s(bool,X146))) | s(bool,'T_0') = s(bool,X144)) ) | $spl0),
  inference(superposition,[],[f1132,f281])).
fof(f2334,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f1688,f151])).
fof(f2319,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) != s(bool,X33) | s(bool,'F_0') = s(bool,X33) | s(bool,X32) = s(bool,X34) | s(bool,'F_0') = s(bool,X34)) ) | $spl0),
  inference(subsumption_resolution,[],[f2280,f121])).
fof(f2280,plain,(
  ( ! [X33,X34,X32] : (s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X33) | s(bool,X32) = s(bool,X34) | s(bool,'F_0') = s(bool,X34)) ) | $spl0),
  inference(superposition,[],[f1430,f1315])).
fof(f1688,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') != s(bool,X13) | s(bool,X13) = s(bool,X14) | s(bool,'F_0') = s(bool,X14)) ) | $spl0),
  inference(superposition,[],[f121,f1315])).
fof(f1510,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) ) | $spl0),
  inference(superposition,[],[f1315,f1315])).
fof(f1430,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) ) | $spl0),
  inference(subsumption_resolution,[],[f1429,f83])).
fof(f1429,plain,(
  ( ! [X23,X22] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) ) | $spl0),
  inference(forward_demodulation,[],[f1336,f385])).
fof(f1336,plain,(
  ( ! [X23,X22] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3)))) | s(bool,X22) != s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X22)) ) | $spl0),
  inference(superposition,[],[f154,f1133])).
fof(f1315,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f1133,f1133])).
fof(f1299,plain,(
  ( ! [X103,X105,X102,X106,X104] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X104,happ(s(fun(X102,X104),X105),s(X102,X103))),s(fun(X104,bool),X106))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X102,X103),s(fun(X102,bool),'const.util_prob.PREIMAGE_2'(s(fun(X102,X104),X105),s(fun(X104,bool),X106)))))) ) | $spl0),
  inference(forward_demodulation,[],[f1280,f385])).
fof(f1280,plain,(
  ( ! [X103,X105,X102,X106,X104] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) != s(bool,'const.bool.IN_2'(s(X104,happ(s(fun(X102,X104),X105),s(X102,X103))),s(fun(X104,bool),X106))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X102,X103),s(fun(X102,bool),'const.util_prob.PREIMAGE_2'(s(fun(X102,X104),X105),s(fun(X104,bool),X106)))))) ) | $spl0),
  inference(superposition,[],[f1132,f118])).
fof(f1298,plain,(
  ( ! [X107,X105,X108,X106,X104] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X105,X107),s(fun(X105,bool),'const.util_prob.PREIMAGE_2'(s(fun(X105,X104),X106),s(fun(X104,bool),X108))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X104,happ(s(fun(X105,X104),X106),s(X105,X107))),s(fun(X104,bool),X108)))) ) | $spl0),
  inference(forward_demodulation,[],[f1249,f385])).
fof(f1249,plain,(
  ( ! [X107,X105,X108,X106,X104] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) != s(bool,'const.bool.IN_2'(s(X105,X107),s(fun(X105,bool),'const.util_prob.PREIMAGE_2'(s(fun(X105,X104),X106),s(fun(X104,bool),X108))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X104,happ(s(fun(X105,X104),X106),s(X105,X107))),s(fun(X104,bool),X108)))) ) | $spl0),
  inference(superposition,[],[f1132,f118])).
fof(f1179,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),'const.combin.K_1'(s(bool,X27))),s(X26,X28)))) ) | $spl0),
  inference(forward_demodulation,[],[f1153,f385])).
fof(f1153,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),'const.combin.K_1'(s(bool,X27))),s(X26,X28)))) ) | $spl0),
  inference(superposition,[],[f1132,f120])).
fof(f1177,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') != s(bool,X22) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X23)))) ) | $spl0),
  inference(forward_demodulation,[],[f1151,f385])).
fof(f1151,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X22),s(bool,X23)))) ) | $spl0),
  inference(superposition,[],[f1132,f107])).
fof(f1150,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,'const.combin.K_2'(s(bool,X19),s(X20,X21)))) ) | $spl0),
  inference(superposition,[],[f1132,f115])).
fof(f1132,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X7)) ) | $spl0),
  inference(subsumption_resolution,[],[f1106,f83])).
fof(f1106,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X7)) ) | $spl0),
  inference(superposition,[],[f154,f385])).
fof(f7618,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7535,f7449])).
fof(f7535,plain,(
  s(bool,'T_0') != s(bool,g(s(sK1,sK5('const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),bool,sK1)))) | ($spl0 | $spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f7439])).
fof(f7439,plain,(
  s(bool,f(s(sK1,sK5('const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),bool,sK1)))) != s(bool,g(s(sK1,sK5('const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),bool,sK1)))) | $spl0),
  inference(backward_demodulation,[],[f7437,f431])).
fof(f7614,plain,(
  $false | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7613,f83])).
fof(f7613,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f7612,f7449])).
fof(f7612,plain,(
  ( ! [X28,X26,X29,X27,X25] : (~p(s(bool,'const.bool.IN_2'(s(X25,X26),s(fun(X25,bool),'const.util_prob.PREIMAGE_2'(s(fun(X25,X27),X28),s(fun(X27,bool),X29))))))) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7486,f7449])).
fof(f7486,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X27,happ(s(fun(X25,X27),X28),s(X25,X26))),s(fun(X27,bool),X29))) | ~p(s(bool,'const.bool.IN_2'(s(X25,X26),s(fun(X25,bool),'const.util_prob.PREIMAGE_2'(s(fun(X25,X27),X28),s(fun(X27,bool),X29))))))) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f1267])).
fof(f7611,plain,(
  $false | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7610,f83])).
fof(f7610,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f7609,f7449])).
fof(f7609,plain,(
  ( ! [X30,X28,X31,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(X27,happ(s(fun(X28,X27),X29),s(X28,X30))),s(fun(X27,bool),X31))))) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7482,f7449])).
fof(f7482,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X28,X30),s(fun(X28,bool),'const.util_prob.PREIMAGE_2'(s(fun(X28,X27),X29),s(fun(X27,bool),X31))))) | ~p(s(bool,'const.bool.IN_2'(s(X27,happ(s(fun(X28,X27),X29),s(X28,X30))),s(fun(X27,bool),X31))))) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f1236])).
fof(f7608,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599,f7601,f7604,f7607])).
fof(f7605,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599,f7601,f7604])).
fof(f7602,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599,f7601])).
fof(f7600,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452,f7599])).
fof(f7596,plain,(
  $false | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7453,f83])).
fof(f7453,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f7449,f119])).
fof(f7595,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593,f7452])).
fof(f7594,plain,(
  $false | ($spl2 | $spl4)),
  inference(global_subsumption,[],[f83,f91,f94,f93,f103,f106,f107,f114,f111,f115,f120,f117,f121,f129,f133,f151,f154,f188,f256,f281,f383,f426,f451,f571,f572,f574,f590,f601,f602,f604,f642,f661,f684,f704,f724,f726,f1030,f1031,f1032,f1034,f767,f1045,f1056,f1058,f1060,f1072,f1074,f1076,f1188,f1190,f1193,f1195,f1231,f2420,f2421,f2423,f2424,f2484,f2485,f3761,f3762,f3763,f3764,f3765,f2522,f2521,f2515,f2513,f2510,f2509,f1063,f1065,f2414,f1089,f1086,f2536,f1087,f1263,f1085,f1082,f2537,f1083,f1262,f110,f1192,f2415,f1278,f1277,f1271,f1269,f1266,f1265,f1247,f1246,f1240,f1238,f1235,f1234,f118,f4161,f4142,f4141,f4132,f4131,f4129,f2538,f1189,f109,f656,f762,f4164,f4163,f4159,f4156,f4130,f3474,f3348,f3332,f3331,f3328,f3319,f3318,f2911,f2807,f2806,f2805,f2804,f2803,f2797,f2796,f2760,f2525,f2523,f2520,f2512,f2511,f2483,f2422,f2221,f2108,f2107,f2106,f2105,f2100,f2099,f2074,f1279,f1276,f1268,f1267,f1248,f1245,f1237,f1236,f1229,f1197,f1191,f1088,f1084,f1078,f1073,f1067,f1062,f1057,f727,f603,f573,f547,f546,f504,f503,f465,f464,f463,f462,f436,f435,f355,f297,f210,f153,f132,f119,f108,f98,f97,f7341,f3330,f3329,f7342,f7343,f1017,f7400,f7398,f433,f431,f82,f7396,f386,f7449,f7593])).
fof(f7447,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f7446,f129])).
fof(f7446,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),sK1)),s(fun(sK1,bool),'const.pred_set.UNIV_0'))) | $spl0),
  inference(forward_demodulation,[],[f7442,f107])).
fof(f7442,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),sK1)),s(fun(sK1,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))))) | $spl0),
  inference(backward_demodulation,[],[f7437,f7400])).
fof(f7445,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f7444,f129])).
fof(f7444,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),sK1)),s(fun(sK1,bool),'const.pred_set.UNIV_0'))) | $spl0),
  inference(forward_demodulation,[],[f7441,f107])).
fof(f7441,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),sK1)),s(fun(sK1,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))))) | $spl0),
  inference(backward_demodulation,[],[f7437,f7398])).
fof(f7428,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f7427,f129])).
fof(f7427,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),sK1)),s(fun(sK1,bool),'const.pred_set.UNIV_0'))) | $spl3),
  inference(forward_demodulation,[],[f7423,f107])).
fof(f7423,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),sK1)),s(fun(sK1,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))))) | $spl3),
  inference(backward_demodulation,[],[f7411,f7400])).
fof(f7411,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl3),
  inference(unit_resulting_resolution,[],[f119,f6593,f188])).
fof(f6593,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl3),
  inference(forward_demodulation,[],[f6592,f132])).
fof(f6592,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.EMPTY_0',sK1)),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) | $spl3),
  inference(forward_demodulation,[],[f6541,f1189])).
fof(f6541,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,sK4('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.EMPTY_0',sK1)),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.EMPTY_0',sK1)),s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))))) | $spl3),
  inference(unit_resulting_resolution,[],[f387,f110])).
fof(f7426,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f7425,f129])).
fof(f7425,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),sK1)),s(fun(sK1,bool),'const.pred_set.UNIV_0'))) | $spl3),
  inference(forward_demodulation,[],[f7422,f107])).
fof(f7422,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK1,sK4('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),sK1)),s(fun(sK1,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))))) | $spl3),
  inference(backward_demodulation,[],[f7411,f7398])).
fof(f7386,plain,(
  $false | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f7344,f6593])).
fof(f7344,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl1),
  inference(unit_resulting_resolution,[],[f384,f91])).
fof(f384,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl1),
  inference(cnf_transformation,[],[f384_D])).
fof(f384_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f7385,plain,(
  $false | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f7346,f6593])).
fof(f7346,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl1),
  inference(unit_resulting_resolution,[],[f121,f384,f151])).
fof(f7384,plain,(
  $false | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f7383,f6593])).
fof(f7383,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl1),
  inference(forward_demodulation,[],[f7347,f132])).
fof(f7347,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f133,f384,f151])).
fof(f7382,plain,(
  $false | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f7350,f6593])).
fof(f7350,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl1),
  inference(unit_resulting_resolution,[],[f121,f384,f151])).
fof(f7381,plain,(
  $false | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f7380,f6593])).
fof(f7380,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl1),
  inference(forward_demodulation,[],[f7351,f132])).
fof(f7351,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f133,f384,f151])).
fof(f7375,plain,(
  $false | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f7359,f6593])).
fof(f7359,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl1),
  inference(unit_resulting_resolution,[],[f119,f384,f188])).
fof(f4600,plain,(
  ( ! [X99] : (s(bool,'T_0') = s(bool,X99)) ) | $spl28),
  inference(cnf_transformation,[],[f4600_D])).
fof(f4600_D,plain,(
  ( ! [X99] : (s(bool,'T_0') = s(bool,X99)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f7307,plain,(
  ( ! [X78,X77] : (s(bool,X78) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X77),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X78)) ) | $spl32),
  inference(cnf_transformation,[],[f7307_D])).
fof(f7307_D,plain,(
  ( ! [X78,X77] : (s(bool,X78) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X77),s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X78)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f7229,plain,(
  ( ! [X78,X76,X77] : (s(bool,X78) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X77),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,X78)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f7100])).
fof(f7100,plain,(
  ( ! [X78,X76,X77] : (s(bool,X78) != s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X77),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,X78) | s(bool,'T_0') = s(bool,X76)) ) | $spl0),
  inference(superposition,[],[f1430,f4014])).
fof(f4958,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897,f4899,f4902,f4905,f4932,f4934,f4939,f4941,f4943,f4945,f4947,f4949,f4951,f4953,f4955,f4957])).
fof(f4957,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl28),
  inference(subsumption_resolution,[],[f4880,f93])).
fof(f4880,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(condensation,[],[f4879])).
fof(f4879,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl28),
  inference(condensation,[],[f4869])).
fof(f4869,plain,(
  ( ! [X94,X93,X91] : (s(bool,'T_0') = s(bool,X94) | p(s(bool,X93)) | p(s(bool,X94)) | s(bool,'T_0') = s(bool,X91)) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f2513])).
fof(f4955,plain,(
  ( ! [X17,X15,X13] : (s(fun(X13,bool),X15) = s(fun(X13,bool),X17)) ) | $spl28),
  inference(subsumption_resolution,[],[f4862,f4600])).
fof(f4862,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X13,sK4(happ(s(fun(X14,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X14,X16)),X17,X13)),s(fun(X13,bool),X15))) | s(fun(X13,bool),X15) = s(fun(X13,bool),X17)) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f1089])).
fof(f4953,plain,(
  ( ! [X4,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X4)) ) | $spl28),
  inference(subsumption_resolution,[],[f4861,f4600])).
fof(f4861,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK4('const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X4,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X4)) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f1086])).
fof(f4951,plain,(
  ( ! [X6,X8,X5] : (s(fun(X5,bool),X6) = s(fun(X5,bool),X8)) ) | $spl28),
  inference(subsumption_resolution,[],[f4859,f4600])).
fof(f4859,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X5,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X8,X5)),s(fun(X5,bool),X6))) | s(fun(X5,bool),X6) = s(fun(X5,bool),X8)) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f1087])).
fof(f4949,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(fun(X5,bool),X9) = s(fun(X5,bool),'const.util_prob.PREIMAGE_2'(s(fun(X5,X6),X7),s(fun(X6,bool),X8)))) ) | $spl28),
  inference(subsumption_resolution,[],[f4858,f4600])).
fof(f4858,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X6,happ(s(fun(X5,X6),X7),s(X5,sK4('const.util_prob.PREIMAGE_2'(s(fun(X5,X6),X7),s(fun(X6,bool),X8)),X9,X5)))),s(fun(X6,bool),X8))) | s(fun(X5,bool),X9) = s(fun(X5,bool),'const.util_prob.PREIMAGE_2'(s(fun(X5,X6),X7),s(fun(X6,bool),X8)))) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f1263])).
fof(f4947,plain,(
  ( ! [X17,X15,X13] : (s(fun(X13,bool),X15) = s(fun(X13,bool),X17)) ) | $spl28),
  inference(subsumption_resolution,[],[f4857,f4600])).
fof(f4857,plain,(
  ( ! [X14,X17,X15,X13,X16] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X13,sK4(X17,happ(s(fun(X14,fun(X13,bool)),'const.combin.K_1'(s(fun(X13,bool),X15))),s(X14,X16)),X13)),s(fun(X13,bool),X15))) | s(fun(X13,bool),X15) = s(fun(X13,bool),X17)) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f1085])).
fof(f4945,plain,(
  ( ! [X4,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X4)) ) | $spl28),
  inference(subsumption_resolution,[],[f4856,f4600])).
fof(f4856,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK4(X4,'const.combin.K_2'(s(fun(X0,bool),X1),s(X2,X3)),X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X4)) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f1082])).
fof(f4943,plain,(
  ( ! [X6,X8,X5] : (s(fun(X5,bool),X6) = s(fun(X5,bool),X8)) ) | $spl28),
  inference(subsumption_resolution,[],[f4854,f4600])).
fof(f4854,plain,(
  ( ! [X6,X8,X7,X5] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X5,sK4(X8,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X5,bool),X6),s(fun(X5,bool),X7)),X5)),s(fun(X5,bool),X6))) | s(fun(X5,bool),X6) = s(fun(X5,bool),X8)) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f1083])).
fof(f4941,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X2),X3),s(fun(X2,bool),X4)))) ) | $spl28),
  inference(subsumption_resolution,[],[f4853,f4600])).
fof(f4853,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X2,happ(s(fun(X0,X2),X3),s(X0,sK4(X1,'const.util_prob.PREIMAGE_2'(s(fun(X0,X2),X3),s(fun(X2,bool),X4)),X0)))),s(fun(X2,bool),X4))) | s(fun(X0,bool),X1) = s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X2),X3),s(fun(X2,bool),X4)))) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f1262])).
fof(f4939,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl28),
  inference(subsumption_resolution,[],[f4852,f4600])).
fof(f4852,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK4(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f110])).
fof(f4934,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(fun(X10,bool),X14) = s(fun(X10,bool),'const.util_prob.PREIMAGE_2'(s(fun(X10,X11),'const.combin.K_1'(s(X11,X12))),s(fun(X11,bool),X13)))) ) | $spl28),
  inference(subsumption_resolution,[],[f4832,f4600])).
fof(f4832,plain,(
  ( ! [X14,X12,X10,X13,X11] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X10,sK4('const.util_prob.PREIMAGE_2'(s(fun(X10,X11),'const.combin.K_1'(s(X11,X12))),s(fun(X11,bool),X13)),X14,X10)),s(fun(X10,bool),X14))) | s(fun(X10,bool),X14) = s(fun(X10,bool),'const.util_prob.PREIMAGE_2'(s(fun(X10,X11),'const.combin.K_1'(s(X11,X12))),s(fun(X11,bool),X13)))) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f4142])).
fof(f4932,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(fun(X5,bool),X6) = s(fun(X5,bool),'const.util_prob.PREIMAGE_2'(s(fun(X5,X7),'const.combin.K_1'(s(X7,X8))),s(fun(X7,bool),X9)))) ) | $spl28),
  inference(subsumption_resolution,[],[f4831,f4600])).
fof(f4831,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X5,sK4(X6,'const.util_prob.PREIMAGE_2'(s(fun(X5,X7),'const.combin.K_1'(s(X7,X8))),s(fun(X7,bool),X9)),X5)),s(fun(X5,bool),X6))) | s(fun(X5,bool),X6) = s(fun(X5,bool),'const.util_prob.PREIMAGE_2'(s(fun(X5,X7),'const.combin.K_1'(s(X7,X8))),s(fun(X7,bool),X9)))) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f4141])).
fof(f4905,plain,(
  ( ! [X12,X11,X9] : (s(fun(X9,bool),X11) = s(fun(X9,bool),X12)) ) | $spl28),
  inference(subsumption_resolution,[],[f4904,f4600])).
fof(f4904,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X9,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X12,X9)),s(fun(X9,bool),X11))) | s(fun(X9,bool),X11) = s(fun(X9,bool),X12)) ) | $spl28),
  inference(forward_demodulation,[],[f4662,f4600])).
fof(f4662,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'const.bool.IN_2'(s(X9,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X12,X9)),s(fun(X9,bool),X11))) != s(bool,'const.bool.IN_2'(s(X9,sK4('const.bool.COND_3'(s(bool,'T_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X12,X9)),s(fun(X9,bool),X12))) | s(fun(X9,bool),X11) = s(fun(X9,bool),X12)) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f1088])).
fof(f4902,plain,(
  ( ! [X12,X11,X9] : (s(fun(X9,bool),X11) = s(fun(X9,bool),X12)) ) | $spl28),
  inference(subsumption_resolution,[],[f4901,f4600])).
fof(f4901,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X9,sK4(X12,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X9)),s(fun(X9,bool),X11))) | s(fun(X9,bool),X11) = s(fun(X9,bool),X12)) ) | $spl28),
  inference(forward_demodulation,[],[f4661,f4600])).
fof(f4661,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'const.bool.IN_2'(s(X9,sK4(X12,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X9)),s(fun(X9,bool),X11))) != s(bool,'const.bool.IN_2'(s(X9,sK4(X12,'const.bool.COND_3'(s(bool,'T_0'),s(fun(X9,bool),X10),s(fun(X9,bool),X11)),X9)),s(fun(X9,bool),X12))) | s(fun(X9,bool),X11) = s(fun(X9,bool),X12)) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f1084])).
fof(f4899,plain,(
  ( ! [X2,X3] : (s(fun(X2,bool),'const.pred_set.EMPTY_0') = s(fun(X2,bool),X3)) ) | $spl28),
  inference(subsumption_resolution,[],[f4659,f4600])).
fof(f4659,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X2,sK4('const.pred_set.EMPTY_0',X3,X2)),s(fun(X2,bool),X3))) | s(fun(X2,bool),'const.pred_set.EMPTY_0') = s(fun(X2,bool),X3)) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f1073])).
fof(f4897,plain,(
  ( ! [X2,X3] : (s(fun(X2,bool),'const.pred_set.EMPTY_0') = s(fun(X2,bool),X3)) ) | $spl28),
  inference(subsumption_resolution,[],[f4656,f4600])).
fof(f4656,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X2,sK4(X3,'const.pred_set.EMPTY_0',X2)),s(fun(X2,bool),X3))) | s(fun(X2,bool),'const.pred_set.EMPTY_0') = s(fun(X2,bool),X3)) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f1057])).
fof(f4636,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f108])).
fof(f4891,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl28),
  inference(subsumption_resolution,[],[f4634,f4600])).
fof(f4634,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f97])).
fof(f4956,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897,f4899,f4902,f4905,f4932,f4934,f4939,f4941,f4943,f4945,f4947,f4949,f4951,f4953,f4955])).
fof(f4954,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897,f4899,f4902,f4905,f4932,f4934,f4939,f4941,f4943,f4945,f4947,f4949,f4951,f4953])).
fof(f4952,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897,f4899,f4902,f4905,f4932,f4934,f4939,f4941,f4943,f4945,f4947,f4949,f4951])).
fof(f4950,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897,f4899,f4902,f4905,f4932,f4934,f4939,f4941,f4943,f4945,f4947,f4949])).
fof(f4948,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897,f4899,f4902,f4905,f4932,f4934,f4939,f4941,f4943,f4945,f4947])).
fof(f4946,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897,f4899,f4902,f4905,f4932,f4934,f4939,f4941,f4943,f4945])).
fof(f4944,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897,f4899,f4902,f4905,f4932,f4934,f4939,f4941,f4943])).
fof(f4942,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897,f4899,f4902,f4905,f4932,f4934,f4939,f4941])).
fof(f4940,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897,f4899,f4902,f4905,f4932,f4934,f4939])).
fof(f4935,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897,f4899,f4902,f4905,f4932,f4934])).
fof(f4933,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897,f4899,f4902,f4905,f4932])).
fof(f4931,plain,(
  $false | ($spl0 | $spl28)),
  inference(subsumption_resolution,[],[f4823,f4600])).
fof(f4823,plain,(
  s(bool,'T_0') != s(bool,g(s(sK1,sK5('const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),bool,sK1)))) | ($spl0 | $spl28)),
  inference(backward_demodulation,[],[f4600,f1046])).
fof(f4930,plain,(
  $false | ($spl0 | $spl28)),
  inference(subsumption_resolution,[],[f4822,f4600])).
fof(f4822,plain,(
  s(bool,'T_0') != s(bool,g(s(sK1,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),bool,sK1)))) | ($spl0 | $spl28)),
  inference(backward_demodulation,[],[f4600,f1047])).
fof(f4929,plain,(
  $false | ($spl3 | $spl28)),
  inference(subsumption_resolution,[],[f4821,f4600])).
fof(f4821,plain,(
  s(bool,'T_0') != s(bool,g(s(sK1,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.EMPTY_0',bool,sK1)))) | ($spl3 | $spl28)),
  inference(backward_demodulation,[],[f4600,f432])).
fof(f432,plain,(
  s(bool,f(s(sK1,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.EMPTY_0',bool,sK1)))) != s(bool,g(s(sK1,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),'const.pred_set.EMPTY_0',bool,sK1)))) | $spl3),
  inference(unit_resulting_resolution,[],[f387,f117])).
fof(f4928,plain,(
  $false | ($spl3 | $spl28)),
  inference(subsumption_resolution,[],[f4820,f4600])).
fof(f4820,plain,(
  s(bool,'T_0') != s(bool,g(s(sK1,sK5('const.pred_set.EMPTY_0','const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),bool,sK1)))) | ($spl3 | $spl28)),
  inference(backward_demodulation,[],[f4600,f434])).
fof(f434,plain,(
  s(bool,f(s(sK1,sK5('const.pred_set.EMPTY_0','const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),bool,sK1)))) != s(bool,g(s(sK1,sK5('const.pred_set.EMPTY_0','const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3)),bool,sK1)))) | $spl3),
  inference(unit_resulting_resolution,[],[f387,f117])).
fof(f4927,plain,(
  $false | ($spl0 | $spl28)),
  inference(subsumption_resolution,[],[f4795,f4600])).
fof(f4795,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X0),s(bool,'T_0')))) ) | ($spl0 | $spl28)),
  inference(backward_demodulation,[],[f4600,f3904])).
fof(f4926,plain,(
  $false | ($spl0 | $spl28)),
  inference(subsumption_resolution,[],[f4781,f4600])).
fof(f4781,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X2)))) ) | ($spl0 | $spl28)),
  inference(backward_demodulation,[],[f4600,f3595])).
fof(f4912,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f4911,f83])).
fof(f4911,plain,(
  ~p(s(bool,'T_0')) | $spl28),
  inference(forward_demodulation,[],[f4910,f4600])).
fof(f4910,plain,(
  ( ! [X28,X26,X29,X27,X25] : (~p(s(bool,'const.bool.IN_2'(s(X25,X26),s(fun(X25,bool),'const.util_prob.PREIMAGE_2'(s(fun(X25,X27),X28),s(fun(X27,bool),X29))))))) ) | $spl28),
  inference(subsumption_resolution,[],[f4672,f4600])).
fof(f4672,plain,(
  ( ! [X28,X26,X29,X27,X25] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X27,happ(s(fun(X25,X27),X28),s(X25,X26))),s(fun(X27,bool),X29))) | ~p(s(bool,'const.bool.IN_2'(s(X25,X26),s(fun(X25,bool),'const.util_prob.PREIMAGE_2'(s(fun(X25,X27),X28),s(fun(X27,bool),X29))))))) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f1267])).
fof(f4909,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f4908,f83])).
fof(f4908,plain,(
  ~p(s(bool,'T_0')) | $spl28),
  inference(forward_demodulation,[],[f4907,f4600])).
fof(f4907,plain,(
  ( ! [X30,X28,X31,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(X27,happ(s(fun(X28,X27),X29),s(X28,X30))),s(fun(X27,bool),X31))))) ) | $spl28),
  inference(subsumption_resolution,[],[f4668,f4600])).
fof(f4668,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X28,X30),s(fun(X28,bool),'const.util_prob.PREIMAGE_2'(s(fun(X28,X27),X29),s(fun(X27,bool),X31))))) | ~p(s(bool,'const.bool.IN_2'(s(X27,happ(s(fun(X28,X27),X29),s(X28,X30))),s(fun(X27,bool),X31))))) ) | $spl28),
  inference(backward_demodulation,[],[f4600,f1236])).
fof(f4906,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897,f4899,f4902,f4905])).
fof(f4903,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897,f4899,f4902])).
fof(f4900,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897,f4899])).
fof(f4898,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636,f4897])).
fof(f4894,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f4637,f83])).
fof(f4637,plain,(
  ~p(s(bool,'T_0')) | $spl28),
  inference(backward_demodulation,[],[f4600,f119])).
fof(f4893,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891,f4636])).
fof(f4892,plain,(
  $false | $spl28),
  inference(global_subsumption,[],[f82,f83,f119,f91,f94,f93,f98,f97,f103,f106,f107,f108,f110,f109,f114,f111,f115,f120,f117,f118,f121,f129,f133,f132,f151,f153,f154,f188,f210,f256,f281,f297,f355,f383,f426,f433,f431,f435,f436,f451,f462,f463,f464,f465,f503,f504,f546,f547,f571,f572,f573,f574,f590,f601,f602,f603,f604,f642,f656,f661,f684,f704,f724,f726,f727,f1030,f1031,f1032,f1034,f762,f767,f1045,f1082,f1083,f1084,f1085,f1056,f1057,f1058,f1060,f1062,f1063,f1065,f1067,f1086,f1087,f1088,f1089,f1072,f1073,f1074,f1076,f1078,f1188,f1189,f1190,f1191,f1192,f1193,f1195,f1197,f1229,f1231,f1234,f1235,f1236,f1237,f1238,f1240,f1245,f1246,f1247,f1248,f1262,f1263,f1265,f1266,f1267,f1268,f1269,f1271,f1276,f1277,f1278,f1279,f2074,f2099,f2100,f2105,f2106,f2107,f2108,f2221,f2414,f2415,f2420,f2421,f2422,f2423,f2424,f2483,f2484,f2485,f2509,f2510,f2511,f2512,f2513,f2515,f2520,f2521,f2522,f2523,f2525,f2536,f2537,f2538,f2760,f2796,f2797,f2803,f2804,f2805,f2806,f2807,f2911,f3318,f3319,f3328,f3331,f3332,f3348,f3474,f3761,f3762,f3763,f3764,f3765,f4129,f4130,f4131,f4132,f4141,f4142,f4156,f4159,f4161,f4163,f4164,f4600,f4891])).
fof(f4601,plain,(
  ( ! [X101,X100] : (s(bool,X101) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X100))) | s(bool,'F_0') = s(bool,X101)) ) | $spl30),
  inference(cnf_transformation,[],[f4601_D])).
fof(f4601_D,plain,(
  ( ! [X101,X100] : (s(bool,X101) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X100))) | s(bool,'F_0') = s(bool,X101)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4515,plain,(
  ( ! [X101,X99,X100] : (s(bool,X101) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X100))) | s(bool,'T_0') = s(bool,X99) | s(bool,'F_0') = s(bool,X101)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f4380])).
fof(f4380,plain,(
  ( ! [X101,X99,X100] : (s(bool,X101) != s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X100))) | s(bool,'T_0') = s(bool,X99) | s(bool,'F_0') = s(bool,X101) | s(bool,'T_0') = s(bool,X99)) ) | $spl0),
  inference(superposition,[],[f1430,f3710])).
fof(f3781,plain,(
  ( ! [X16] : (s(bool,X16) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X16)) ) | $spl22),
  inference(cnf_transformation,[],[f3781_D])).
fof(f3781_D,plain,(
  ( ! [X16] : (s(bool,X16) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X16)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3774,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool))) ) | $spl12),
  inference(cnf_transformation,[],[f3774_D])).
fof(f3774_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3776,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl15),
  inference(cnf_transformation,[],[f3776_D])).
fof(f3776_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f3769,plain,(
  ( ! [X24,X23,X25] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X23) = s(fun(bool,bool),X24) | s(bool,'T_0') = s(bool,sK5(X23,X24,bool,bool)) | s(bool,X25) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X25)) ) | $spl0),
  inference(superposition,[],[f435,f1315])).
fof(f3783,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | $spl24),
  inference(cnf_transformation,[],[f3783_D])).
fof(f3783_D,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f3784,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X17)) ) | $spl26),
  inference(cnf_transformation,[],[f3784_D])).
fof(f3784_D,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X17)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3768,plain,(
  ( ! [X21,X22,X20] : (s(bool,X20) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X21) = s(fun(bool,bool),X22) | s(bool,'T_0') = s(bool,sK5(X21,X22,bool,bool)) | s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X20)) ) | $spl0),
  inference(superposition,[],[f435,f1315])).
fof(f3767,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19) | s(bool,'T_0') = s(bool,sK5(X18,X19,bool,bool)) | s(bool,'F_0') = s(bool,X17) | s(bool,'F_0') = s(bool,g(s(bool,'F_0')))) ) | $spl0),
  inference(superposition,[],[f435,f1315])).
fof(f3766,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,sK5(X14,X15,bool,bool)) | s(bool,X16) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X16)) ) | $spl0),
  inference(superposition,[],[f435,f1315])).
fof(f3777,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl16),
  inference(cnf_transformation,[],[f3777_D])).
fof(f3777_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3779,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl19),
  inference(cnf_transformation,[],[f3779_D])).
fof(f3779_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f3765,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK5(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f435,f151])).
fof(f3773,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl10),
  inference(cnf_transformation,[],[f3773_D])).
fof(f3773_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f3780,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl20),
  inference(cnf_transformation,[],[f3780_D])).
fof(f3780_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3764,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK5(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f435,f151])).
fof(f3763,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK5(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f435,f151])).
fof(f3762,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK5(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f435,f151])).
fof(f3761,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f435,f91])).
fof(f1035,plain,(
  ( ! [X57] : (s(bool,X57) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X57)) ) | $spl6),
  inference(cnf_transformation,[],[f1035_D])).
fof(f1035_D,plain,(
  ( ! [X57] : (s(bool,X57) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X57)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1037,plain,(
  s(fun(sK1,bool),'const.pred_set.UNIV_0') != s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) | $spl9),
  inference(cnf_transformation,[],[f1037_D])).
fof(f1037_D,plain,(
  s(fun(sK1,bool),'const.pred_set.UNIV_0') != s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f1045,plain,(
  ( ! [X57] : (s(fun(sK1,bool),'const.pred_set.UNIV_0') != s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) | s(bool,X57) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X57)) )),
  inference(forward_demodulation,[],[f813,f107])).
fof(f813,plain,(
  ( ! [X57] : (s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) != s(fun(sK1,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) | s(bool,X57) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f82,f151])).
fof(f1034,plain,(
  ( ! [X57] : (s(fun(sK1,bool),'const.pred_set.UNIV_0') != s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) | s(bool,X57) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X57)) )),
  inference(forward_demodulation,[],[f749,f107])).
fof(f749,plain,(
  ( ! [X57] : (s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) != s(fun(sK1,bool),'const.bool.COND_3'(s(bool,'T_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) | s(bool,X57) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f82,f151])).
fof(f385,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | $spl0),
  inference(cnf_transformation,[],[f385_D])).
fof(f385_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1017,plain,(
  ( ! [X60] : (s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) != s(fun(sK1,bool),'const.bool.COND_3'(s(bool,X60),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X60)) ) | $spl4),
  inference(cnf_transformation,[],[f1017_D])).
fof(f1017_D,plain,(
  ( ! [X60] : (s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) != s(fun(sK1,bool),'const.bool.COND_3'(s(bool,X60),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X60)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f684,plain,(
  ( ! [X60] : (s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) != s(fun(sK1,bool),'const.bool.COND_3'(s(bool,X60),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f82,f151])).
fof(f642,plain,(
  ( ! [X60] : (s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) != s(fun(sK1,bool),'const.bool.COND_3'(s(bool,X60),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3)))) )),
  inference(superposition,[],[f82,f151])).
fof(f387,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) | $spl3),
  inference(cnf_transformation,[],[f387_D])).
fof(f387_D,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f383,plain,(
  s(fun(sK1,bool),'const.pred_set.EMPTY_0') != s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3)))),
  inference(forward_demodulation,[],[f382,f108])).
fof(f382,plain,(
  s(fun(sK1,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK1,sK0),'const.combin.K_1'(s(sK0,sK2))),s(fun(sK0,bool),sK3))) != s(fun(sK1,bool),'const.bool.COND_3'(s(bool,'F_0'),s(fun(sK1,bool),'const.pred_set.UNIV_0'),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK3)))),
  inference(superposition,[],[f82,f91])).
% SZS output end Proof for thm.util_prob
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 72
% Passive clauses: 635
% Generated clauses: 7669
% Final active clauses: 17
% Final passive clauses: 39
% Input formulas: 25
% Initial clauses: 37
% 
% Duplicate literals: 453
% Trivial inequalities: 94
% Fw subsumption resolutions: 300
% Fw demodulations: 568
% Bw demodulations: 389
% Condensations: 827
% Global subsumptions: 37
% 
% Simple tautologies: 142
% Equational tautologies: 98
% Forward subsumptions: 4421
% Fw demodulations to eq. taut.: 18
% Bw demodulations to eq. taut.: 6
% 
% Unit resulting resolution: 434
% Factoring: 1
% Forward superposition: 2242
% Backward superposition: 1821
% Self superposition: 105
% Equality factoring: 60
% Equality resolution: 12
% 
% Split clauses: 16
% Split components: 41
% SAT solver clauses: 766
% SAT solver unit clauses: 67
% SAT solver binary clauses: 233
% 
% Sat splits: 16
% Sat splitting refutations: 60
% 
% Memory used [KB]: 5373
% Time elapsed: 0.300 s
% ------------------------------
% ------------------------------
% Success in time 1.616 s
