Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.relation
% (127525)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1142
% Passive clauses: 10663
% Generated clauses: 23764
% Final active clauses: 1126
% Final passive clauses: 9237
% Input formulas: 133
% Initial clauses: 493
% 
% Duplicate literals: 151
% Trivial inequalities: 198
% Fw subsumption resolutions: 1254
% Fw demodulations: 3192
% Bw demodulations: 14
% 
% Simple tautologies: 312
% Equational tautologies: 25
% Forward subsumptions: 7739
% 
% Binary resolution: 10252
% Factoring: 461
% Forward superposition: 2895
% Backward superposition: 4630
% Self superposition: 14
% Equality factoring: 2
% Equality resolution: 23
% 
% Split clauses: 370
% Split components: 742
% SAT solver clauses: 517
% SAT solver unit clauses: 9
% SAT solver binary clauses: 81
% 
% Sat splits: 370
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 517
% TWLsolver calls for satisfiability: 395
% 
% Memory used [KB]: 16247
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (127537)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 500
% Passive clauses: 3986
% Generated clauses: 13595
% Final active clauses: 451
% Final passive clauses: 3347
% Input formulas: 133
% Initial clauses: 493
% 
% Duplicate literals: 93
% Trivial inequalities: 2
% Fw subsumption resolutions: 413
% Fw demodulations: 1527
% Bw demodulations: 188
% Condensations: 43
% Global subsumptions: 23
% 
% Simple tautologies: 272
% Equational tautologies: 7
% Forward subsumptions: 2876
% 
% Binary resolution: 3184
% Factoring: 57
% Forward superposition: 2227
% Backward superposition: 1335
% Self superposition: 14
% Equality factoring: 2
% Equality resolution: 2
% 
% Split clauses: 49
% Split components: 98
% SAT solver clauses: 4799
% SAT solver unit clauses: 470
% SAT solver binary clauses: 129
% SAT solver learnt clauses: 46
% SAT solver learnt literals: 12
% 
% Sat splits: 49
% 
% TWLsolver clauses: 4707
% TWLsolver calls for satisfiability: 5362
% 
% Memory used [KB]: 18805
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% (127561)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 3320
% Passive clauses: 20142
% Generated clauses: 107971
% Final active clauses: 2767
% Final passive clauses: 6749
% Input formulas: 133
% Initial clauses: 493
% Discarded non-redundant clauses: 15307
% 
% Duplicate literals: 493
% Trivial inequalities: 65
% Fw subsumption resolutions: 3245
% Fw demodulations: 24937
% Bw demodulations: 4126
% Condensations: 1427
% Global subsumptions: 370
% 
% Simple tautologies: 182
% Equational tautologies: 116
% Forward subsumptions: 54668
% Bw demodulations to eq. taut.: 14
% 
% Unit resulting resolution: 35208
% Factoring: 1
% Forward superposition: 22624
% Backward superposition: 11034
% Self superposition: 47
% Equality factoring: 26
% Equality resolution: 10
% 
% Split clauses: 257
% Split components: 529
% SAT solver clauses: 22161
% SAT solver unit clauses: 3923
% SAT solver binary clauses: 7440
% 
% Sat splits: 257
% Sat splitting refutations: 1133
% 
% Memory used [KB]: 106181
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f52156,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f965,f966,f1189,f1026,f1028,f1029,f1030,f1033,f1123,f1291,f1293,f1295,f1297,f1298,f1391,f1492,f1493,f1517,f1519,f1520,f1521,f1524,f1614,f2699,f2701,f2702,f2703,f2706,f2796,f2853,f1187,f2914,f2916,f2917,f2918,f2921,f3011,f1190,f3114,f3116,f3117,f3118,f3121,f3211,f1679,f3442,f3444,f3446,f3448,f3449,f3542,f1467,f3809,f3811,f3813,f3815,f3816,f3909,f3972,f4280,f4365,f4732,f4817,f5232,f5228,f5005,f5343,f5344,f6144,f6330,f6145,f6332,f6157,f6328,f6158,f6169,f6171,f6172,f6173,f6176,f6266,f6866,f7365,f6334,f6336,f7376,f7378,f7379,f7380,f7383,f7473,f6337,f9073,f9259,f9074,f9261,f9086,f9257,f9087,f9098,f9100,f9101,f9102,f9105,f9195,f9796,f10307,f9263,f9265,f10318,f10320,f10321,f10322,f10325,f10415,f9266,f11041,f11043,f11044,f11045,f11048,f11138,f11201,f11315,f11317,f11319,f11321,f11322,f11415,f11494,f12058,f12060,f12062,f12064,f12065,f12158,f12442,f12511,f12513,f12514,f12515,f12518,f12608,f12672,f3090,f1678,f3089,f1676,f1185,f13800,f13802,f13803,f13804,f13807,f13897,f2847,f14075,f14077,f14079,f14081,f14082,f14175,f14239,f14912,f15087,f15085,f14923,f14925,f14926,f14927,f14930,f15020,f15573,f15759,f15574,f15761,f15586,f15757,f15587,f15598,f15600,f15601,f15602,f15605,f15695,f16296,f16756,f15763,f15765,f16767,f16769,f16770,f16771,f16774,f16864,f18091,f18093,f18094,f18095,f18098,f18188,f18400,f18402,f18403,f18404,f18407,f18497,f18558,f18566,f18568,f18570,f18572,f18573,f18666,f18724,f19189,f19191,f19192,f19193,f19196,f19286,f19462,f19464,f19466,f19468,f19469,f19562,f19620,f20498,f20500,f20502,f20504,f20505,f20598,f20882,f21219,f21221,f21223,f21225,f21226,f21319,f21377,f22047,f22049,f22051,f22053,f22054,f22147,f22216,f22267,f22269,f22271,f22273,f22274,f22367,f22436,f22932,f22934,f22936,f22938,f22939,f23032,f23086,f23799,f23801,f23802,f23803,f23806,f23896,f24782,f24783,f24784,f24785,f24786,f24787,f24816,f18251,f24817,f19349,f24818,f24795,f24819,f15088,f24820,f15766,f24821,f23953,f24848,f24875,f24860,f24876,f24861,f24877,f24862,f24878,f24863,f24879,f24864,f24880,f24865,f24881,f24866,f24882,f25240,f25239,f25235,f25012,f25011,f25010,f25352,f25353,f25496,f25498,f25500,f25502,f25503,f25596,f25881,f26601,f26603,f26605,f26607,f26608,f26701,f26770,f26847,f26849,f26851,f26853,f26854,f26947,f27016,f27165,f27170,f27611,f27657,f28040,f28086,f28250,f28252,f28254,f28256,f28257,f28350,f28404,f28421,f28423,f28425,f28427,f28428,f28521,f28575,f28766,f28796,f29052,f29082,f29192,f29194,f29195,f29196,f29199,f29289,f29358,f29360,f29361,f29362,f29365,f29455,f29735,f29767,f29347,f29877,f30026,f30194,f30261,f29513,f30445,f30540,f30594,f30888,f30942,f31000,f31054,f31449,f31451,f31452,f31453,f31456,f31546,f31627,f31696,f31698,f31699,f31700,f31703,f31793,f31949,f31951,f31952,f31953,f31956,f32046,f32127,f32206,f32208,f32209,f32210,f32213,f32303,f32368,f32459,f32461,f32462,f32463,f32466,f32556,f32888,f32925,f31858,f33069,f33071,f33072,f33073,f33076,f33166,f33231,f33404,f33406,f33408,f33410,f33411,f33504,f33580,f33771,f33808,f32621,f35112,f35300,f35113,f35302,f35125,f35298,f35126,f35304,f35306,f35137,f35139,f35140,f35141,f35144,f35234,f35307,f37007,f37195,f37008,f37197,f37020,f37193,f37021,f37199,f37201,f37032,f37034,f37035,f37036,f37039,f37129,f37202,f38323,f38372,f38554,f38647,f38696,f39011,f39060,f39113,f39163,f29918,f39343,f39531,f39344,f39533,f39356,f39529,f39357,f39535,f39537,f39368,f39370,f39371,f39372,f39375,f39465,f39538,f40588,f40590,f40591,f40592,f40595,f40685,f40766,f40835,f40837,f40838,f40839,f40842,f40932,f41142,f41144,f41145,f41146,f41149,f41239,f41676,f41713,f40997,f42152,f42189,f41304,f43152,f43340,f43153,f43342,f43165,f43338,f43166,f43344,f43346,f43177,f43179,f43180,f43181,f43184,f43274,f43347,f46426,f46614,f46427,f46616,f46439,f46612,f46440,f46618,f46620,f46451,f46453,f46454,f46455,f46458,f46548,f46621,f50855,f50857,f50858,f1191,f50859,f1193,f50862,f1192,f50952,f51215,f51243,f1194,f51005,f52144,f52145,f52148,f52150])).
fof(f52150,plain,(
  $false | ($spl0 | $spl74)),
  inference(subsumption_resolution,[],[f52149,f13736])).
fof(f13736,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK24,sK24,sK23,sK22))))) | $spl0),
  inference(resolution,[],[f12680,f598])).
fof(f598,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,sK40(X3,X2,X1,X0))))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f399])).
fof(f399,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,sK40(X3,X2,X1,X0))))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK40(X3,X2,X1,X0)),s(X0,X3)))))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK40])],[f398])).
fof(f398,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X5)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X5),s(X0,X3)))))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(rectify,[],[f397])).
fof(f397,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))) & (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))) <=> ? [X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X10)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X10),s(X6,X9))))))),
  file('thm.relation.TC_CASES|49|_E','thm.relation.EXTEND_RTC_TC_EQN')).
fof(f12680,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) | $spl0),
  inference(backward_demodulation,[],[f1185,f584])).
fof(f584,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))))),
  inference(cnf_transformation,[],[f385])).
fof(f385,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) & ~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25)))) & ! [X4] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X4),s(sK22,sK25)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK22,sK23,sK24,sK25])],[f275])).
fof(f275,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))))),
  inference(flattening,[],[f274])).
fof(f274,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ~! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))) => (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9)))) | ? [X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X10)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X10),s(X6,X9)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X9)))) => (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9)))) | ? [X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X10)))) & p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X10),s(X6,X9)))))))),
  file('thm.relation.TC_CASES|49|_E','thm.relation.TC_CASES|49|_E')).
fof(f52149,plain,(
  ~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK24,sK24,sK23,sK22))))) | ($spl0 | $spl74)),
  inference(forward_demodulation,[],[f52135,f948])).
fof(f948,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,sK157(X1,X0))) )),
  inference(cnf_transformation,[],[f583])).
fof(f583,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,sK157(X1,X0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK157])],[f267])).
fof(f267,plain,(
  ! [X0,X1] : ? [X2] : s(X0,X1) = s(X0,X2)),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X0] : ! [X1] : ? [X2] : s(X0,X1) = s(X0,X2)),
  inference(rectify,[],[f94])).
fof(f94,axiom,(
  ! [X6] : ! [X21] : ? [X8] : s(X6,X8) = s(X6,X21)),
  file('thm.relation.TC_CASES|49|_E','thm.bool.EXISTS_REFL')).
fof(f52135,plain,(
  ~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK157(sK40(sK24,sK24,sK23,sK22),sK22))))) | ($spl0 | $spl74)),
  inference(resolution,[],[f3798,f26833])).
fof(f26833,plain,(
  sP0(sK22,sK157(sK40(sK24,sK24,sK23,sK22),sK22),sK25,sK23) | ($spl0 | $spl74)),
  inference(resolution,[],[f22399,f15763])).
fof(f22399,plain,(
  ( ! [X66,X65] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X66),s(sK22,X65),s(sK22,sK24)))) | sP0(sK22,sK157(X65,sK22),sK25,X66)) ) | $spl0),
  inference(superposition,[],[f12858,f948])).
fof(f12858,plain,(
  ( ! [X237,X238] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X237),s(sK22,X238),s(sK22,sK24)))) | sP0(sK22,X238,sK25,X237)) ) | $spl0),
  inference(superposition,[],[f610,f1185])).
fof(f610,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,X2,X3,X1) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f403])).
fof(f403,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | sP0(X0,X2,X3,X1)) & (~sP0(X0,X2,X3,X1) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f350])).
fof(f350,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> sP0(X0,X2,X3,X1))),
  inference(definition_folding,[],[f284,f349])).
fof(f349,plain,(
  ! [X0,X2,X3,X1] : (sP0(X0,X2,X3,X1) <=> ! [X4] : (? [X8,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X9))))) | ? [X5,X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f284,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : (? [X8,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X9))))) | ? [X5,X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(flattening,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : ((? [X8,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X9))))) | ? [X5,X6,X7] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7)))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(ennf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : ((! [X8,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X9))))) & ! [X5,X6,X7] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7)))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : ((! [X8,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X8))),s(X0,X9)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X9))))) & ! [X5,X6,X7] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7)))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X7,X21,X22] : (p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X21),s(X6,X22)))) <=> ! [X11] : ((! [X8,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X11),s(X6,X8))),s(X6,X10))))) & ! [X8,X10,X9] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X11),s(X6,X8))),s(X6,X10)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X11),s(X6,X10))),s(X6,X9))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X11),s(X6,X8))),s(X6,X9)))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X11),s(X6,X21))),s(X6,X22))))))),
  file('thm.relation.TC_CASES|49|_E','thm.relation.TC_DEF')).
fof(f3798,plain,(
  ( ! [X4] : (~sP0(sK22,X4,sK25,sK23) | ~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X4))))) )),
  inference(resolution,[],[f586,f611])).
fof(f611,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~sP0(X0,X2,X3,X1)) )),
  inference(cnf_transformation,[],[f403])).
fof(f586,plain,(
  ( ! [X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X4),s(sK22,sK25)))) | ~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X4))))) )),
  inference(cnf_transformation,[],[f385])).
fof(f52148,plain,(
  $false | ($spl0 | $spl34)),
  inference(subsumption_resolution,[],[f52147,f962])).
fof(f962,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK25,sK24,sK23,sK22)))))),
  inference(resolution,[],[f584,f598])).
fof(f52147,plain,(
  ~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK25,sK24,sK23,sK22))))) | ($spl0 | $spl34)),
  inference(forward_demodulation,[],[f52134,f948])).
fof(f52134,plain,(
  ~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK157(sK40(sK25,sK24,sK23,sK22),sK22))))) | ($spl0 | $spl34)),
  inference(resolution,[],[f3798,f26832])).
fof(f26832,plain,(
  sP0(sK22,sK157(sK40(sK25,sK24,sK23,sK22),sK22),sK25,sK23) | ($spl0 | $spl34)),
  inference(resolution,[],[f22399,f12772])).
fof(f12772,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) | ($spl0 | $spl34)),
  inference(backward_demodulation,[],[f1185,f6334])).
fof(f52145,plain,(
  $false | ($spl0 | $spl74)),
  inference(subsumption_resolution,[],[f52132,f13736])).
fof(f52132,plain,(
  ~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK24,sK24,sK23,sK22))))) | ($spl0 | $spl74)),
  inference(resolution,[],[f3798,f22253])).
fof(f22253,plain,(
  sP0(sK22,sK40(sK24,sK24,sK23,sK22),sK25,sK23) | ($spl0 | $spl74)),
  inference(resolution,[],[f12858,f15763])).
fof(f52144,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f52131,f962])).
fof(f52131,plain,(
  ~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK25,sK24,sK23,sK22))))) | $spl34),
  inference(resolution,[],[f3798,f10979])).
fof(f10979,plain,(
  sP0(sK22,sK40(sK25,sK24,sK23,sK22),sK25,sK23) | $spl34),
  inference(resolution,[],[f6334,f610])).
fof(f51005,plain,(
  ( ! [X25] : (p(s(bool,X25)) | s(bool,X25) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl226),
  inference(cnf_transformation,[],[f51005_D])).
fof(f51005_D,plain,(
  ( ! [X25] : (p(s(bool,X25)) | s(bool,X25) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f1194,plain,(
  ( ! [X89] : (p(s(bool,X89))) ) | $spl14),
  inference(cnf_transformation,[],[f1194_D])).
fof(f1194_D,plain,(
  ( ! [X89] : (p(s(bool,X89))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f51243,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X1))) ) | ($spl0 | $spl24 | $spl172 | $spl180)),
  inference(resolution,[],[f50919,f24747])).
fof(f24747,plain,(
  ( ! [X6,X8,X7] : (~sP15(X7,X8,X6) | p(s(bool,X6)) | s(bool,X7) != s(bool,X8) | p(s(bool,X8))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24744,f845])).
fof(f845,plain,(
  ( ! [X2,X0,X1] : (sP16(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f541])).
fof(f541,plain,(
  ! [X0,X1,X2] : ((~sP16(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP16(X0,X1,X2)))),
  inference(rectify,[],[f540])).
fof(f540,plain,(
  ! [X1,X0,X2] : ((~sP16(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP16(X1,X0,X2)))),
  inference(flattening,[],[f539])).
fof(f539,plain,(
  ! [X1,X0,X2] : ((~sP16(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP16(X1,X0,X2)))),
  inference(nnf_transformation,[],[f375])).
fof(f375,plain,(
  ! [X1,X0,X2] : (sP16(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP16])])).
fof(f24744,plain,(
  ( ! [X4,X5,X3] : (~sP16(X4,X5,X3) | p(s(bool,X3)) | ~sP15(X5,X4,X3) | s(bool,X4) != s(bool,X5)) ) | ($spl0 | $spl24)),
  inference(duplicate_literal_removal,[],[f24738])).
fof(f24738,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~sP16(X4,X5,X3) | ~sP15(X5,X4,X3) | p(s(bool,X3)) | s(bool,X4) != s(bool,X5)) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24736,f853])).
fof(f853,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP18(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f544])).
fof(f544,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP18(X2,X1,X0)) & (~sP18(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f378])).
fof(f378,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP18(X2,X1,X0))),
  inference(definition_folding,[],[f225,f377,f376,f375,f374])).
fof(f374,plain,(
  ! [X0,X1,X2] : (sP15(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP15])])).
fof(f376,plain,(
  ! [X2,X0,X1] : (sP17(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP17])])).
fof(f377,plain,(
  ! [X2,X1,X0] : (sP18(X2,X1,X0) <=> (sP15(X0,X1,X2) & sP16(X1,X0,X2) & sP17(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP18])])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X20,X18,X19] : ((p(s(bool,X19)) <=> s(bool,X18) = s(bool,X20)) <=> ((p(s(bool,X19)) | p(s(bool,X18)) | p(s(bool,X20))) & (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X18))) & (p(s(bool,X18)) | ~p(s(bool,X20)) | ~p(s(bool,X19))) & (p(s(bool,X20)) | ~p(s(bool,X18)) | ~p(s(bool,X19)))))),
  file('thm.relation.TC_CASES|49|_E','thm.sat.sat_2|49|_prove_ltb')).
fof(f24736,plain,(
  ( ! [X10,X11,X9] : (sP18(X9,X10,X11) | p(s(bool,X9)) | ~sP16(X10,X11,X9) | ~sP15(X11,X10,X9)) ) | ($spl0 | $spl24)),
  inference(subsumption_resolution,[],[f24733,f2847])).
fof(f24733,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | sP18(X9,X10,X11) | p(s(bool,X9)) | ~sP16(X10,X11,X9) | ~sP15(X11,X10,X9)) ) | $spl0),
  inference(resolution,[],[f20556,f837])).
fof(f837,plain,(
  ( ! [X2,X0,X1] : (sP18(X0,X1,X2) | p(s(bool,X0)) | ~sP17(X0,X2,X1) | ~sP16(X1,X2,X0) | ~sP15(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f535])).
fof(f535,plain,(
  ! [X0,X1,X2] : ((~sP18(X0,X1,X2) | (sP15(X2,X1,X0) & sP16(X1,X2,X0) & sP17(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP15(X2,X1,X0) | ~sP16(X1,X2,X0) | ~sP17(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP18(X0,X1,X2)))),
  inference(rectify,[],[f534])).
fof(f534,plain,(
  ! [X2,X1,X0] : ((~sP18(X2,X1,X0) | (sP15(X0,X1,X2) & sP16(X1,X0,X2) & sP17(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP15(X0,X1,X2) | ~sP16(X1,X0,X2) | ~sP17(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP18(X2,X1,X0)))),
  inference(flattening,[],[f533])).
fof(f533,plain,(
  ! [X2,X1,X0] : ((~sP18(X2,X1,X0) | (sP15(X0,X1,X2) & sP16(X1,X0,X2) & sP17(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP15(X0,X1,X2) | ~sP16(X1,X0,X2) | ~sP17(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP18(X2,X1,X0)))),
  inference(nnf_transformation,[],[f377])).
fof(f20556,plain,(
  ( ! [X208,X206,X207] : (sP17(X206,X207,X208) | s(bool,X206) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12719,f841])).
fof(f841,plain,(
  ( ! [X2,X0,X1] : (sP17(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f538])).
fof(f538,plain,(
  ! [X0,X1,X2] : ((~sP17(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP17(X0,X1,X2)))),
  inference(rectify,[],[f537])).
fof(f537,plain,(
  ! [X2,X0,X1] : ((~sP17(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP17(X2,X0,X1)))),
  inference(flattening,[],[f536])).
fof(f536,plain,(
  ! [X2,X0,X1] : ((~sP17(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP17(X2,X0,X1)))),
  inference(nnf_transformation,[],[f376])).
fof(f12719,plain,(
  ( ! [X261] : (~p(s(bool,X261)) | s(bool,X261) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(backward_demodulation,[],[f1185,f1119])).
fof(f1119,plain,(
  ( ! [X261] : (~p(s(bool,X261)) | s(bool,X261) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f584,f919])).
fof(f919,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f572])).
fof(f572,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f571])).
fof(f571,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & (((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0)))))),
  inference(ennf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) => p(s(bool,X1))) & (p(s(bool,X1)) => p(s(bool,X0)))))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X15,X16] : (s(bool,X15) = s(bool,X16) <=> ((p(s(bool,X15)) => p(s(bool,X16))) & (p(s(bool,X16)) => p(s(bool,X15)))))),
  file('thm.relation.TC_CASES|49|_E','thm.bool.EQ_IMP_THM')).
fof(f50919,plain,(
  ( ! [X161,X160] : (sP15('const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X160,X161)) ) | ($spl172 | $spl180)),
  inference(resolution,[],[f50846,f849])).
fof(f849,plain,(
  ( ! [X2,X0,X1] : (sP15(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f543])).
fof(f543,plain,(
  ! [X0,X1,X2] : ((~sP15(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP15(X0,X1,X2)))),
  inference(flattening,[],[f542])).
fof(f542,plain,(
  ! [X0,X1,X2] : ((~sP15(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP15(X0,X1,X2)))),
  inference(nnf_transformation,[],[f374])).
fof(f50846,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) | ($spl172 | $spl180)),
  inference(forward_demodulation,[],[f37201,f37193])).
fof(f51215,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X3))) ) | ($spl0 | $spl24 | $spl172 | $spl180)),
  inference(resolution,[],[f50918,f24746])).
fof(f24746,plain,(
  ( ! [X4,X5,X3] : (~sP15(X4,X5,X3) | p(s(bool,X3)) | s(bool,X4) != s(bool,X5) | p(s(bool,X4))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24744,f844])).
fof(f844,plain,(
  ( ! [X2,X0,X1] : (sP16(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f541])).
fof(f50918,plain,(
  ( ! [X158,X159] : (sP15(X158,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X159)) ) | ($spl172 | $spl180)),
  inference(resolution,[],[f50846,f848])).
fof(f848,plain,(
  ( ! [X2,X0,X1] : (sP15(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f543])).
fof(f50952,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | ($spl172 | $spl180)),
  inference(resolution,[],[f50846,f952])).
fof(f952,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f776])).
fof(f776,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f509])).
fof(f509,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f508])).
fof(f508,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f323])).
fof(f323,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f107])).
fof(f107,axiom,(
  ! [X8,X33,X10,X34] : ((s(bool,X8) = s(bool,X33) & (p(s(bool,X33)) => s(bool,X10) = s(bool,X34))) => ((p(s(bool,X8)) => p(s(bool,X10))) <=> (p(s(bool,X33)) => p(s(bool,X34)))))),
  file('thm.relation.TC_CASES|49|_E','thm.bool.IMP_CONG')).
fof(f1192,plain,(
  ( ! [X85,X84] : (s(bool,X84) != s(bool,X85) | ~p(s(bool,X84))) ) | $spl10),
  inference(cnf_transformation,[],[f1192_D])).
fof(f1192_D,plain,(
  ( ! [X85,X84] : (s(bool,X84) != s(bool,X85) | ~p(s(bool,X84))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f50862,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | s(bool,X46) != s(bool,X47)) ) | ($spl172 | $spl180)),
  inference(resolution,[],[f50846,f779])).
fof(f779,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f509])).
fof(f1193,plain,(
  ( ! [X90] : (~p(s(bool,X90))) ) | $spl12),
  inference(cnf_transformation,[],[f1193_D])).
fof(f1193_D,plain,(
  ( ! [X90] : (~p(s(bool,X90))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f50859,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | ($spl172 | $spl180)),
  inference(resolution,[],[f50846,f778])).
fof(f778,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f509])).
fof(f1191,plain,(
  ( ! [X79,X77] : (s(bool,X77) != s(bool,X79) | p(s(bool,X77))) ) | $spl8),
  inference(cnf_transformation,[],[f1191_D])).
fof(f1191_D,plain,(
  ( ! [X79,X77] : (s(bool,X77) != s(bool,X79) | p(s(bool,X77))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f50858,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | ($spl172 | $spl180)),
  inference(resolution,[],[f50846,f777])).
fof(f777,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f509])).
fof(f50857,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | s(bool,X31) != s(bool,X32)) ) | ($spl172 | $spl180)),
  inference(resolution,[],[f50846,f775])).
fof(f775,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f509])).
fof(f50855,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | ($spl172 | $spl180)),
  inference(resolution,[],[f50846,f773])).
fof(f773,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f509])).
fof(f46621,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | $spl224),
  inference(cnf_transformation,[],[f46621_D])).
fof(f46621_D,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f46548,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | $spl74),
  inference(resolution,[],[f19126,f952])).
fof(f19126,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl74),
  inference(resolution,[],[f15763,f599])).
fof(f599,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK40(X3,X2,X1,X0)),s(X0,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f399])).
fof(f46458,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24))) | s(bool,X66) != s(bool,X67)) ) | $spl74),
  inference(resolution,[],[f19126,f779])).
fof(f46455,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | $spl74),
  inference(resolution,[],[f19126,f778])).
fof(f46454,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | $spl74),
  inference(resolution,[],[f19126,f777])).
fof(f46453,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24))) | s(bool,X51) != s(bool,X52)) ) | $spl74),
  inference(resolution,[],[f19126,f775])).
fof(f46451,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | $spl74),
  inference(resolution,[],[f19126,f773])).
fof(f46620,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl222),
  inference(cnf_transformation,[],[f46620_D])).
fof(f46620_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f46618,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl220),
  inference(cnf_transformation,[],[f46618_D])).
fof(f46618_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f46440,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl74),
  inference(resolution,[],[f19126,f944])).
fof(f944,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f345])).
fof(f345,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(ennf_transformation,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f260])).
fof(f260,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) => (p(s(bool,'const.relation.RC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X6] : ! [X7,X8,X10] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X10)))) => (p(s(bool,'const.relation.RC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X10)))) | p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X10))))))),
  file('thm.relation.TC_CASES|49|_E','thm.relation.RTC_TC_RC')).
fof(f46612,plain,(
  s(sK22,sK24) = s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)) | $spl214),
  inference(cnf_transformation,[],[f46612_D])).
fof(f46612_D,plain,(
  s(sK22,sK24) = s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f46439,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | s(sK22,sK24) = s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)) | $spl74),
  inference(resolution,[],[f19126,f941])).
fof(f941,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f581])).
fof(f581,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) & ((s(X0,X2) != s(X0,X3) & ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f580])).
fof(f580,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (s(X0,X2) = s(X0,X3) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))) & ((s(X0,X2) != s(X0,X3) & ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f258])).
fof(f258,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X6] : ! [X7,X8,X10] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X10)))) <=> (s(X6,X8) = s(X6,X10) | p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X10))))))),
  file('thm.relation.TC_CASES|49|_E','thm.relation.RTC_CASES_TC')).
fof(f46616,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl218),
  inference(cnf_transformation,[],[f46616_D])).
fof(f46616_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f46427,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | s(sK22,sK24) = s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)) | $spl74),
  inference(resolution,[],[f19126,f619])).
fof(f619,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK52(X3,X2,X1,X0)),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f410])).
fof(f410,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,sK52(X3,X2,X1,X0))))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,sK52(X3,X2,X1,X0)),s(X0,X3)))))) & ((s(X0,X2) != s(X0,X3) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK52])],[f409])).
fof(f409,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X5)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X5),s(X0,X3)))))) & ((s(X0,X2) != s(X0,X3) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(rectify,[],[f408])).
fof(f408,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))) & ((s(X0,X2) != s(X0,X3) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(flattening,[],[f407])).
fof(f407,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | (s(X0,X2) = s(X0,X3) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3))))))) & ((s(X0,X2) != s(X0,X3) & ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> (s(X0,X2) = s(X0,X3) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X4),s(X0,X3)))))))),
  inference(rectify,[],[f114])).
fof(f114,axiom,(
  ! [X6] : ! [X7,X8,X10] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X10)))) <=> (s(X6,X8) = s(X6,X10) | ? [X12] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X12)))) & p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X12),s(X6,X10)))))))),
  file('thm.relation.TC_CASES|49|_E','thm.relation.RTC_CASES|49|')).
fof(f46614,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)))),s(sK22,sK52(sK24,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22),sK23,sK22))))) | $spl216),
  inference(cnf_transformation,[],[f46614_D])).
fof(f46614_D,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)))),s(sK22,sK52(sK24,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22),sK23,sK22))))) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f46426,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)))),s(sK22,sK52(sK24,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22),sK23,sK22))))) | s(sK22,sK24) = s(sK22,sK40(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)) | $spl74),
  inference(resolution,[],[f19126,f618])).
fof(f618,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,sK52(X3,X2,X1,X0))))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f410])).
fof(f43347,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | $spl212),
  inference(cnf_transformation,[],[f43347_D])).
fof(f43347_D,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f43274,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl34)),
  inference(resolution,[],[f18028,f952])).
fof(f18028,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | ($spl0 | $spl34)),
  inference(resolution,[],[f12772,f599])).
fof(f43184,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24))) | s(bool,X66) != s(bool,X67)) ) | ($spl0 | $spl34)),
  inference(resolution,[],[f18028,f779])).
fof(f43181,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl34)),
  inference(resolution,[],[f18028,f778])).
fof(f43180,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl34)),
  inference(resolution,[],[f18028,f777])).
fof(f43179,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24))) | s(bool,X51) != s(bool,X52)) ) | ($spl0 | $spl34)),
  inference(resolution,[],[f18028,f775])).
fof(f43177,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl34)),
  inference(resolution,[],[f18028,f773])).
fof(f43346,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl210),
  inference(cnf_transformation,[],[f43346_D])).
fof(f43346_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f43344,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl208),
  inference(cnf_transformation,[],[f43344_D])).
fof(f43344_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f43166,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | ($spl0 | $spl34)),
  inference(resolution,[],[f18028,f944])).
fof(f43338,plain,(
  s(sK22,sK24) = s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)) | $spl202),
  inference(cnf_transformation,[],[f43338_D])).
fof(f43338_D,plain,(
  s(sK22,sK24) = s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f43165,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | s(sK22,sK24) = s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)) | ($spl0 | $spl34)),
  inference(resolution,[],[f18028,f941])).
fof(f43342,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl206),
  inference(cnf_transformation,[],[f43342_D])).
fof(f43342_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f43153,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | s(sK22,sK24) = s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)) | ($spl0 | $spl34)),
  inference(resolution,[],[f18028,f619])).
fof(f43340,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)))),s(sK22,sK52(sK24,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22))))) | $spl204),
  inference(cnf_transformation,[],[f43340_D])).
fof(f43340_D,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)))),s(sK22,sK52(sK24,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22))))) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f43152,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)))),s(sK22,sK52(sK24,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22))))) | s(sK22,sK24) = s(sK22,sK40(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)) | ($spl0 | $spl34)),
  inference(resolution,[],[f18028,f618])).
fof(f41304,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) | $spl200),
  inference(cnf_transformation,[],[f41304_D])).
fof(f41304_D,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f42189,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22)))) | p(s(bool,X1))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f41206,f24747])).
fof(f41206,plain,(
  ( ! [X180,X181] : (sP15('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))),X180,X181)) ) | $spl0),
  inference(resolution,[],[f40825,f849])).
fof(f40825,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))))) | $spl0),
  inference(resolution,[],[f40554,f943])).
fof(f943,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f581])).
fof(f40554,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))))) | $spl0),
  inference(resolution,[],[f13736,f624])).
fof(f624,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(ennf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) => p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7,X8,X10] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X10)))) => p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X10)))))),
  file('thm.relation.TC_CASES|49|_E','thm.relation.TC_RULES|split|0')).
fof(f42152,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22)))) | p(s(bool,X3))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f41205,f24746])).
fof(f41205,plain,(
  ( ! [X178,X179] : (sP15(X178,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))),X179)) ) | $spl0),
  inference(resolution,[],[f40825,f848])).
fof(f40997,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) | $spl198),
  inference(cnf_transformation,[],[f40997_D])).
fof(f40997_D,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f41713,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22)))) | p(s(bool,X1))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f40899,f24747])).
fof(f40899,plain,(
  ( ! [X213,X214] : (sP15('const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))),X213,X214)) ) | $spl0),
  inference(resolution,[],[f40554,f849])).
fof(f41676,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22)))) | p(s(bool,X3))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f40898,f24746])).
fof(f40898,plain,(
  ( ! [X212,X211] : (sP15(X211,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))),X212)) ) | $spl0),
  inference(resolution,[],[f40554,f848])).
fof(f41239,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) | $spl0),
  inference(resolution,[],[f40825,f952])).
fof(f41149,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22)))) | s(bool,X66) != s(bool,X67)) ) | $spl0),
  inference(resolution,[],[f40825,f779])).
fof(f41146,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) | $spl0),
  inference(resolution,[],[f40825,f778])).
fof(f41145,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) | $spl0),
  inference(resolution,[],[f40825,f777])).
fof(f41144,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22)))) | s(bool,X51) != s(bool,X52)) ) | $spl0),
  inference(resolution,[],[f40825,f775])).
fof(f41142,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) | $spl0),
  inference(resolution,[],[f40825,f773])).
fof(f40932,plain,(
  ( ! [X265,X266] : (p(s(bool,X265)) | p(s(bool,X266)) | s(bool,X266) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) | $spl0),
  inference(resolution,[],[f40554,f952])).
fof(f40842,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22)))) | s(bool,X99) != s(bool,X100)) ) | $spl0),
  inference(resolution,[],[f40554,f779])).
fof(f40839,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | ~p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X91) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) | $spl0),
  inference(resolution,[],[f40554,f778])).
fof(f40838,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | s(bool,X86) != s(bool,X88) | s(bool,X87) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) | $spl0),
  inference(resolution,[],[f40554,f777])).
fof(f40837,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22)))) | s(bool,X84) != s(bool,X85)) ) | $spl0),
  inference(resolution,[],[f40554,f775])).
fof(f40835,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) | $spl0),
  inference(resolution,[],[f40554,f773])).
fof(f40766,plain,(
  ( ! [X69] : (p(s(bool,X69)) | s(bool,X69) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) | $spl196),
  inference(cnf_transformation,[],[f40766_D])).
fof(f40766_D,plain,(
  ( ! [X69] : (p(s(bool,X69)) | s(bool,X69) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f40685,plain,(
  ( ! [X257,X256] : (p(s(bool,X256)) | p(s(bool,X257)) | s(bool,X257) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) | $spl0),
  inference(resolution,[],[f13736,f952])).
fof(f40595,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | ~p(s(bool,X90)) | s(bool,X89) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK24,sK24,sK23,sK22)))) | s(bool,X90) != s(bool,X91)) ) | $spl0),
  inference(resolution,[],[f13736,f779])).
fof(f40592,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,X80)) | ~p(s(bool,X81)) | p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) | $spl0),
  inference(resolution,[],[f13736,f778])).
fof(f40591,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) | $spl0),
  inference(resolution,[],[f13736,f777])).
fof(f40590,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,X74)) | ~p(s(bool,X75)) | s(bool,X74) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK24,sK24,sK23,sK22)))) | s(bool,X75) != s(bool,X76)) ) | $spl0),
  inference(resolution,[],[f13736,f775])).
fof(f40588,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,X68)) | p(s(bool,X69)) | s(bool,X68) != s(bool,X70) | s(bool,X69) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK24,sK24,sK23,sK22))))) ) | $spl0),
  inference(resolution,[],[f13736,f773])).
fof(f39538,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | $spl194),
  inference(cnf_transformation,[],[f39538_D])).
fof(f39538_D,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f39465,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl32)),
  inference(resolution,[],[f34391,f952])).
fof(f34391,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | ($spl0 | $spl32)),
  inference(forward_demodulation,[],[f6332,f1185])).
fof(f39375,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24))) | s(bool,X66) != s(bool,X67)) ) | ($spl0 | $spl32)),
  inference(resolution,[],[f34391,f779])).
fof(f39372,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl32)),
  inference(resolution,[],[f34391,f778])).
fof(f39371,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl32)),
  inference(resolution,[],[f34391,f777])).
fof(f39370,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24))) | s(bool,X51) != s(bool,X52)) ) | ($spl0 | $spl32)),
  inference(resolution,[],[f34391,f775])).
fof(f39368,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl32)),
  inference(resolution,[],[f34391,f773])).
fof(f39537,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl192),
  inference(cnf_transformation,[],[f39537_D])).
fof(f39537_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f39535,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl190),
  inference(cnf_transformation,[],[f39535_D])).
fof(f39535_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f39357,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | ($spl0 | $spl32)),
  inference(resolution,[],[f34391,f944])).
fof(f39529,plain,(
  s(sK22,sK24) = s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)) | $spl184),
  inference(cnf_transformation,[],[f39529_D])).
fof(f39529_D,plain,(
  s(sK22,sK24) = s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f39356,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | s(sK22,sK24) = s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)) | ($spl0 | $spl32)),
  inference(resolution,[],[f34391,f941])).
fof(f39533,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl188),
  inference(cnf_transformation,[],[f39533_D])).
fof(f39533_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f39344,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | s(sK22,sK24) = s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)) | ($spl0 | $spl32)),
  inference(resolution,[],[f34391,f619])).
fof(f39531,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)))),s(sK22,sK52(sK24,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22))))) | $spl186),
  inference(cnf_transformation,[],[f39531_D])).
fof(f39531_D,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)))),s(sK22,sK52(sK24,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22))))) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f39343,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)))),s(sK22,sK52(sK24,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22))))) | s(sK22,sK24) = s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)) | ($spl0 | $spl32)),
  inference(resolution,[],[f34391,f618])).
fof(f29918,plain,(
  s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | $spl144),
  inference(cnf_transformation,[],[f29918_D])).
fof(f29918_D,plain,(
  s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f39163,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,X9)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f34320,f20558])).
fof(f20558,plain,(
  ( ! [X212,X213,X214] : (sP16(X213,X212,X214) | s(bool,X212) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12719,f844])).
fof(f34320,plain,(
  ( ! [X173,X172] : (~sP16(X172,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X173) | ~p(s(bool,X172)) | p(s(bool,X173))) ) | $spl0),
  inference(forward_demodulation,[],[f1577,f1185])).
fof(f1577,plain,(
  ( ! [X173,X172] : (~p(s(bool,X172)) | p(s(bool,X173)) | ~sP16(X172,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)),X173)) )),
  inference(resolution,[],[f1016,f842])).
fof(f842,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP16(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f541])).
fof(f1016,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))))),
  inference(resolution,[],[f584,f943])).
fof(f39113,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f34319,f20559])).
fof(f20559,plain,(
  ( ! [X216,X215,X217] : (sP16(X215,X216,X217) | s(bool,X215) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12719,f845])).
fof(f34319,plain,(
  ( ! [X171,X170] : (~sP16('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X170,X171) | ~p(s(bool,X170)) | p(s(bool,X171))) ) | $spl0),
  inference(forward_demodulation,[],[f1576,f1185])).
fof(f1576,plain,(
  ( ! [X171,X170] : (~p(s(bool,X170)) | p(s(bool,X171)) | ~sP16('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)),X170,X171)) )),
  inference(resolution,[],[f1016,f842])).
fof(f39060,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,X9)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f34317,f20555])).
fof(f20555,plain,(
  ( ! [X204,X205,X203] : (sP17(X204,X203,X205) | s(bool,X203) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12719,f840])).
fof(f840,plain,(
  ( ! [X2,X0,X1] : (sP17(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f538])).
fof(f34317,plain,(
  ( ! [X167,X166] : (~sP17(X166,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X167) | ~p(s(bool,X166)) | p(s(bool,X167))) ) | $spl0),
  inference(forward_demodulation,[],[f1574,f1185])).
fof(f1574,plain,(
  ( ! [X167,X166] : (~p(s(bool,X166)) | p(s(bool,X167)) | ~sP17(X166,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)),X167)) )),
  inference(resolution,[],[f1016,f838])).
fof(f838,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP17(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f538])).
fof(f39011,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f34316,f20556])).
fof(f34316,plain,(
  ( ! [X165,X164] : (~sP17('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X164,X165) | ~p(s(bool,X164)) | p(s(bool,X165))) ) | $spl0),
  inference(forward_demodulation,[],[f1573,f1185])).
fof(f1573,plain,(
  ( ! [X165,X164] : (~p(s(bool,X164)) | p(s(bool,X165)) | ~sP17('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)),X164,X165)) )),
  inference(resolution,[],[f1016,f838])).
fof(f38696,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,X9)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f34307,f20542])).
fof(f20542,plain,(
  ( ! [X165,X166,X164] : (sP13(X165,X164,X166) | s(bool,X164) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12719,f823])).
fof(f823,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f530])).
fof(f530,plain,(
  ! [X0,X1,X2] : ((~sP13(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP13(X0,X1,X2)))),
  inference(flattening,[],[f529])).
fof(f529,plain,(
  ! [X0,X1,X2] : ((~sP13(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP13(X0,X1,X2)))),
  inference(nnf_transformation,[],[f371])).
fof(f371,plain,(
  ! [X0,X1,X2] : (sP13(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP13])])).
fof(f34307,plain,(
  ( ! [X140,X141] : (~sP13(X140,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X141) | ~p(s(bool,X140)) | p(s(bool,X141))) ) | $spl0),
  inference(forward_demodulation,[],[f1561,f1185])).
fof(f1561,plain,(
  ( ! [X140,X141] : (~p(s(bool,X140)) | p(s(bool,X141)) | ~sP13(X140,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)),X141)) )),
  inference(resolution,[],[f1016,f821])).
fof(f821,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP13(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f530])).
fof(f38647,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f34305,f20543])).
fof(f20543,plain,(
  ( ! [X167,X169,X168] : (sP13(X167,X168,X169) | s(bool,X167) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12719,f824])).
fof(f824,plain,(
  ( ! [X2,X0,X1] : (sP13(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f530])).
fof(f34305,plain,(
  ( ! [X138,X139] : (~sP13('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X138,X139) | ~p(s(bool,X138)) | p(s(bool,X139))) ) | $spl0),
  inference(forward_demodulation,[],[f1560,f1185])).
fof(f1560,plain,(
  ( ! [X138,X139] : (~p(s(bool,X138)) | p(s(bool,X139)) | ~sP13('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)),X138,X139)) )),
  inference(resolution,[],[f1016,f821])).
fof(f38554,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f34303,f20529])).
fof(f20529,plain,(
  ( ! [X127,X125,X126] : (sP11(X125,X126,X127) | s(bool,X125) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12719,f807])).
fof(f807,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f523])).
fof(f523,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f522])).
fof(f522,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f521])).
fof(f521,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f368])).
fof(f368,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f34303,plain,(
  ( ! [X107,X106] : (~sP11('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X106,X107) | p(s(bool,X106)) | p(s(bool,X107))) ) | $spl0),
  inference(forward_demodulation,[],[f1544,f1185])).
fof(f1544,plain,(
  ( ! [X107,X106] : (p(s(bool,X106)) | p(s(bool,X107)) | ~sP11('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)),X106,X107)) )),
  inference(resolution,[],[f1016,f804])).
fof(f804,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f523])).
fof(f38372,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,X9)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f31402,f20512])).
fof(f20512,plain,(
  ( ! [X76,X74,X75] : (sP9(X75,X76,X74) | s(bool,X74) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12719,f788])).
fof(f788,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f515])).
fof(f515,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f514])).
fof(f514,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f513])).
fof(f513,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f31402,plain,(
  ( ! [X80,X81] : (~sP9(X80,X81,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | ~p(s(bool,X80)) | p(s(bool,X81))) ) | $spl0),
  inference(forward_demodulation,[],[f1531,f1185])).
fof(f1531,plain,(
  ( ! [X80,X81] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP9(X80,X81,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f1016,f787])).
fof(f787,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f515])).
fof(f38323,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f31401,f20513])).
fof(f20513,plain,(
  ( ! [X78,X79,X77] : (sP9(X77,X78,X79) | s(bool,X77) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12719,f790])).
fof(f790,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f515])).
fof(f31401,plain,(
  ( ! [X78,X79] : (~sP9('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X78,X79) | p(s(bool,X78)) | ~p(s(bool,X79))) ) | $spl0),
  inference(forward_demodulation,[],[f1530,f1185])).
fof(f1530,plain,(
  ( ! [X78,X79] : (p(s(bool,X78)) | ~p(s(bool,X79)) | ~sP9('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)),X78,X79)) )),
  inference(resolution,[],[f1016,f787])).
fof(f37202,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | $spl182),
  inference(cnf_transformation,[],[f37202_D])).
fof(f37202_D,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f37129,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | $spl72),
  inference(resolution,[],[f15761,f952])).
fof(f37039,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24))) | s(bool,X66) != s(bool,X67)) ) | $spl72),
  inference(resolution,[],[f15761,f779])).
fof(f37036,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | $spl72),
  inference(resolution,[],[f15761,f778])).
fof(f37035,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | $spl72),
  inference(resolution,[],[f15761,f777])).
fof(f37034,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24))) | s(bool,X51) != s(bool,X52)) ) | $spl72),
  inference(resolution,[],[f15761,f775])).
fof(f37032,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | $spl72),
  inference(resolution,[],[f15761,f773])).
fof(f37201,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl180),
  inference(cnf_transformation,[],[f37201_D])).
fof(f37201_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f37199,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl178),
  inference(cnf_transformation,[],[f37199_D])).
fof(f37199_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f37021,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl72),
  inference(resolution,[],[f15761,f944])).
fof(f37193,plain,(
  s(sK22,sK24) = s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)) | $spl172),
  inference(cnf_transformation,[],[f37193_D])).
fof(f37193_D,plain,(
  s(sK22,sK24) = s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f37020,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | s(sK22,sK24) = s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)) | $spl72),
  inference(resolution,[],[f15761,f941])).
fof(f37197,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl176),
  inference(cnf_transformation,[],[f37197_D])).
fof(f37197_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f37008,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | s(sK22,sK24) = s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)) | $spl72),
  inference(resolution,[],[f15761,f619])).
fof(f37195,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)))),s(sK22,sK52(sK24,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22),sK23,sK22))))) | $spl174),
  inference(cnf_transformation,[],[f37195_D])).
fof(f37195_D,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)))),s(sK22,sK52(sK24,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22),sK23,sK22))))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f37007,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)))),s(sK22,sK52(sK24,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22),sK23,sK22))))) | s(sK22,sK24) = s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)) | $spl72),
  inference(resolution,[],[f15761,f618])).
fof(f35307,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | $spl170),
  inference(cnf_transformation,[],[f35307_D])).
fof(f35307_D,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f35234,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl29)),
  inference(resolution,[],[f15082,f952])).
fof(f15082,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | ($spl0 | $spl29)),
  inference(subsumption_resolution,[],[f14899,f12771])).
fof(f12771,plain,(
  s(sK22,sK24) != s(sK22,sK40(sK25,sK24,sK23,sK22)) | ($spl0 | $spl29)),
  inference(backward_demodulation,[],[f1185,f6327])).
fof(f6327,plain,(
  s(sK22,sK25) != s(sK22,sK40(sK25,sK24,sK23,sK22)) | $spl29),
  inference(cnf_transformation,[],[f6327_D])).
fof(f6327_D,plain,(
  s(sK22,sK25) != s(sK22,sK40(sK25,sK24,sK23,sK22)) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f14899,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | s(sK22,sK24) = s(sK22,sK40(sK25,sK24,sK23,sK22)) | $spl0),
  inference(resolution,[],[f12683,f619])).
fof(f12683,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) | $spl0),
  inference(backward_demodulation,[],[f1185,f963])).
fof(f963,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25))))),
  inference(resolution,[],[f584,f599])).
fof(f35144,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24))) | s(bool,X66) != s(bool,X67)) ) | ($spl0 | $spl29)),
  inference(resolution,[],[f15082,f779])).
fof(f35141,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl29)),
  inference(resolution,[],[f15082,f778])).
fof(f35140,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl29)),
  inference(resolution,[],[f15082,f777])).
fof(f35139,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24))) | s(bool,X51) != s(bool,X52)) ) | ($spl0 | $spl29)),
  inference(resolution,[],[f15082,f775])).
fof(f35137,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl29)),
  inference(resolution,[],[f15082,f773])).
fof(f35306,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl168),
  inference(cnf_transformation,[],[f35306_D])).
fof(f35306_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f35304,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl166),
  inference(cnf_transformation,[],[f35304_D])).
fof(f35304_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f35126,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | ($spl0 | $spl29)),
  inference(resolution,[],[f15082,f944])).
fof(f35298,plain,(
  s(sK22,sK24) = s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)) | $spl160),
  inference(cnf_transformation,[],[f35298_D])).
fof(f35298_D,plain,(
  s(sK22,sK24) = s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f35125,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | s(sK22,sK24) = s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)) | ($spl0 | $spl29)),
  inference(resolution,[],[f15082,f941])).
fof(f35302,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl164),
  inference(cnf_transformation,[],[f35302_D])).
fof(f35302_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f35113,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | s(sK22,sK24) = s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)) | ($spl0 | $spl29)),
  inference(resolution,[],[f15082,f619])).
fof(f35300,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)))),s(sK22,sK52(sK24,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22))))) | $spl162),
  inference(cnf_transformation,[],[f35300_D])).
fof(f35300_D,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)))),s(sK22,sK52(sK24,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22))))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f35112,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)))),s(sK22,sK52(sK24,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22),sK23,sK22))))) | s(sK22,sK24) = s(sK22,sK52(sK24,sK40(sK25,sK24,sK23,sK22),sK23,sK22)) | ($spl0 | $spl29)),
  inference(resolution,[],[f15082,f618])).
fof(f32621,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))))) ) | $spl154),
  inference(cnf_transformation,[],[f32621_D])).
fof(f32621_D,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f33808,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22)))) | p(s(bool,X1))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f32523,f24747])).
fof(f32523,plain,(
  ( ! [X180,X181] : (sP15('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))),X180,X181)) )),
  inference(resolution,[],[f31686,f849])).
fof(f31686,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22)))))),
  inference(resolution,[],[f31415,f943])).
fof(f31415,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22)))))),
  inference(resolution,[],[f962,f624])).
fof(f33771,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22)))) | p(s(bool,X3))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f32522,f24746])).
fof(f32522,plain,(
  ( ! [X178,X179] : (sP15(X178,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))),X179)) )),
  inference(resolution,[],[f31686,f848])).
fof(f33580,plain,(
  ( ! [X108] : (~p(s(bool,X108)) | s(bool,X108) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl158),
  inference(cnf_transformation,[],[f33580_D])).
fof(f33580_D,plain,(
  ( ! [X108] : (~p(s(bool,X108)) | s(bool,X108) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f33504,plain,(
  ( ! [X302,X301] : (p(s(bool,X301)) | ~p(s(bool,X302)) | s(bool,X302) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl3),
  inference(resolution,[],[f1186,f952])).
fof(f1186,plain,(
  ~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) | $spl3),
  inference(cnf_transformation,[],[f1186_D])).
fof(f1186_D,plain,(
  ~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f33411,plain,(
  ( ! [X130,X128,X129] : (~p(s(bool,X128)) | ~p(s(bool,X129)) | s(bool,X129) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22)))) | s(bool,X128) != s(bool,X130)) ) | $spl3),
  inference(resolution,[],[f1186,f779])).
fof(f33410,plain,(
  ( ! [X127,X125,X126] : (p(s(bool,X125)) | ~p(s(bool,X126)) | ~p(s(bool,X127)) | s(bool,X126) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl3),
  inference(resolution,[],[f1186,f778])).
fof(f33408,plain,(
  ( ! [X121,X120,X119] : (p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl3),
  inference(resolution,[],[f1186,f777])).
fof(f33406,plain,(
  ( ! [X114,X115,X113] : (~p(s(bool,X113)) | ~p(s(bool,X114)) | s(bool,X114) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22)))) | s(bool,X113) != s(bool,X115)) ) | $spl3),
  inference(resolution,[],[f1186,f775])).
fof(f33404,plain,(
  ( ! [X109,X107,X108] : (p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X107) != s(bool,X109) | s(bool,X108) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl3),
  inference(resolution,[],[f1186,f773])).
fof(f33231,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl156),
  inference(cnf_transformation,[],[f33231_D])).
fof(f33231_D,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f33166,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl2),
  inference(resolution,[],[f32196,f952])).
fof(f32196,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22))))) | $spl2),
  inference(resolution,[],[f31915,f943])).
fof(f31915,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22))))) | $spl2),
  inference(resolution,[],[f1187,f624])).
fof(f33076,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22)))) | s(bool,X66) != s(bool,X67)) ) | $spl2),
  inference(resolution,[],[f32196,f779])).
fof(f33073,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl2),
  inference(resolution,[],[f32196,f778])).
fof(f33072,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl2),
  inference(resolution,[],[f32196,f777])).
fof(f33071,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22)))) | s(bool,X51) != s(bool,X52)) ) | $spl2),
  inference(resolution,[],[f32196,f775])).
fof(f33069,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl2),
  inference(resolution,[],[f32196,f773])).
fof(f31858,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))))) ) | $spl148),
  inference(cnf_transformation,[],[f31858_D])).
fof(f31858_D,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f32925,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X1) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22)))) | p(s(bool,X1))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f31760,f24747])).
fof(f31760,plain,(
  ( ! [X213,X214] : (sP15('const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))),X213,X214)) )),
  inference(resolution,[],[f31415,f849])).
fof(f32888,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,X3) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22)))) | p(s(bool,X3))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f31759,f24746])).
fof(f31759,plain,(
  ( ! [X212,X211] : (sP15(X211,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))),X212)) )),
  inference(resolution,[],[f31415,f848])).
fof(f32556,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))))) )),
  inference(resolution,[],[f31686,f952])).
fof(f32466,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22)))) | s(bool,X66) != s(bool,X67)) )),
  inference(resolution,[],[f31686,f779])).
fof(f32463,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))))) )),
  inference(resolution,[],[f31686,f778])).
fof(f32462,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))))) )),
  inference(resolution,[],[f31686,f777])).
fof(f32461,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22)))) | s(bool,X51) != s(bool,X52)) )),
  inference(resolution,[],[f31686,f775])).
fof(f32459,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))))) )),
  inference(resolution,[],[f31686,f773])).
fof(f32368,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl152),
  inference(cnf_transformation,[],[f32368_D])).
fof(f32368_D,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f32303,plain,(
  ( ! [X265,X266] : (p(s(bool,X265)) | p(s(bool,X266)) | s(bool,X266) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl2),
  inference(resolution,[],[f31915,f952])).
fof(f32213,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22)))) | s(bool,X99) != s(bool,X100)) ) | $spl2),
  inference(resolution,[],[f31915,f779])).
fof(f32210,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | ~p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X91) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl2),
  inference(resolution,[],[f31915,f778])).
fof(f32209,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | s(bool,X86) != s(bool,X88) | s(bool,X87) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl2),
  inference(resolution,[],[f31915,f777])).
fof(f32208,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22)))) | s(bool,X84) != s(bool,X85)) ) | $spl2),
  inference(resolution,[],[f31915,f775])).
fof(f32206,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl2),
  inference(resolution,[],[f31915,f773])).
fof(f32127,plain,(
  ( ! [X69] : (p(s(bool,X69)) | s(bool,X69) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl150),
  inference(cnf_transformation,[],[f32127_D])).
fof(f32127_D,plain,(
  ( ! [X69] : (p(s(bool,X69)) | s(bool,X69) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f32046,plain,(
  ( ! [X257,X256] : (p(s(bool,X256)) | p(s(bool,X257)) | s(bool,X257) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl2),
  inference(resolution,[],[f1187,f952])).
fof(f31956,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | ~p(s(bool,X90)) | s(bool,X89) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22)))) | s(bool,X90) != s(bool,X91)) ) | $spl2),
  inference(resolution,[],[f1187,f779])).
fof(f31953,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,X80)) | ~p(s(bool,X81)) | p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl2),
  inference(resolution,[],[f1187,f778])).
fof(f31952,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl2),
  inference(resolution,[],[f1187,f777])).
fof(f31951,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,X74)) | ~p(s(bool,X75)) | s(bool,X74) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22)))) | s(bool,X75) != s(bool,X76)) ) | $spl2),
  inference(resolution,[],[f1187,f775])).
fof(f31949,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,X68)) | p(s(bool,X69)) | s(bool,X68) != s(bool,X70) | s(bool,X69) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) ) | $spl2),
  inference(resolution,[],[f1187,f773])).
fof(f31793,plain,(
  ( ! [X265,X266] : (p(s(bool,X265)) | p(s(bool,X266)) | s(bool,X266) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))))) )),
  inference(resolution,[],[f31415,f952])).
fof(f31703,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22)))) | s(bool,X99) != s(bool,X100)) )),
  inference(resolution,[],[f31415,f779])).
fof(f31700,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | ~p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X91) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))))) )),
  inference(resolution,[],[f31415,f778])).
fof(f31699,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | s(bool,X86) != s(bool,X88) | s(bool,X87) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))))) )),
  inference(resolution,[],[f31415,f777])).
fof(f31698,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22)))) | s(bool,X84) != s(bool,X85)) )),
  inference(resolution,[],[f31415,f775])).
fof(f31696,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK40(sK25,sK24,sK23,sK22))))) )),
  inference(resolution,[],[f31415,f773])).
fof(f31627,plain,(
  ( ! [X69] : (p(s(bool,X69)) | s(bool,X69) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK25,sK24,sK23,sK22))))) ) | $spl146),
  inference(cnf_transformation,[],[f31627_D])).
fof(f31627_D,plain,(
  ( ! [X69] : (p(s(bool,X69)) | s(bool,X69) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK25,sK24,sK23,sK22))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f31546,plain,(
  ( ! [X257,X256] : (p(s(bool,X256)) | p(s(bool,X257)) | s(bool,X257) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK25,sK24,sK23,sK22))))) )),
  inference(resolution,[],[f962,f952])).
fof(f31456,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | ~p(s(bool,X90)) | s(bool,X89) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK25,sK24,sK23,sK22)))) | s(bool,X90) != s(bool,X91)) )),
  inference(resolution,[],[f962,f779])).
fof(f31453,plain,(
  ( ! [X80,X81,X82] : (p(s(bool,X80)) | ~p(s(bool,X81)) | p(s(bool,X82)) | s(bool,X82) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK25,sK24,sK23,sK22))))) )),
  inference(resolution,[],[f962,f778])).
fof(f31452,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK25,sK24,sK23,sK22))))) )),
  inference(resolution,[],[f962,f777])).
fof(f31451,plain,(
  ( ! [X76,X74,X75] : (p(s(bool,X74)) | ~p(s(bool,X75)) | s(bool,X74) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK25,sK24,sK23,sK22)))) | s(bool,X75) != s(bool,X76)) )),
  inference(resolution,[],[f962,f775])).
fof(f31449,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,X68)) | p(s(bool,X69)) | s(bool,X68) != s(bool,X70) | s(bool,X69) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK40(sK25,sK24,sK23,sK22))))) )),
  inference(resolution,[],[f962,f773])).
fof(f31054,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13860,f25556])).
fof(f25556,plain,(
  ( ! [X212,X213,X214] : (sP16(X213,X212,X214) | s(bool,X212) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f24303,f844])).
fof(f24303,plain,(
  ( ! [X228] : (~p(s(bool,X228)) | s(bool,X228) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(forward_demodulation,[],[f1610,f1185])).
fof(f1610,plain,(
  ( ! [X228] : (s(bool,X228) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | ~p(s(bool,X228))) )),
  inference(resolution,[],[f1016,f919])).
fof(f13860,plain,(
  ( ! [X206,X205] : (~sP16(X205,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X206) | p(s(bool,X206)) | ~p(s(bool,X205))) ) | $spl0),
  inference(resolution,[],[f12680,f842])).
fof(f31000,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13859,f25557])).
fof(f25557,plain,(
  ( ! [X216,X215,X217] : (sP16(X215,X216,X217) | s(bool,X215) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f24303,f845])).
fof(f13859,plain,(
  ( ! [X204,X203] : (~sP16('const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X203,X204) | p(s(bool,X204)) | ~p(s(bool,X203))) ) | $spl0),
  inference(resolution,[],[f12680,f842])).
fof(f30942,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13857,f25553])).
fof(f25553,plain,(
  ( ! [X204,X205,X203] : (sP17(X204,X203,X205) | s(bool,X203) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f24303,f840])).
fof(f13857,plain,(
  ( ! [X200,X199] : (~sP17(X199,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X200) | p(s(bool,X200)) | ~p(s(bool,X199))) ) | $spl0),
  inference(resolution,[],[f12680,f838])).
fof(f30888,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13856,f25554])).
fof(f25554,plain,(
  ( ! [X208,X206,X207] : (sP17(X206,X207,X208) | s(bool,X206) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f24303,f841])).
fof(f13856,plain,(
  ( ! [X198,X197] : (~sP17('const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X197,X198) | p(s(bool,X198)) | ~p(s(bool,X197))) ) | $spl0),
  inference(resolution,[],[f12680,f838])).
fof(f30594,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13844,f25540])).
fof(f25540,plain,(
  ( ! [X165,X166,X164] : (sP13(X165,X164,X166) | s(bool,X164) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f24303,f823])).
fof(f13844,plain,(
  ( ! [X173,X174] : (~sP13(X173,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X174) | p(s(bool,X174)) | ~p(s(bool,X173))) ) | $spl0),
  inference(resolution,[],[f12680,f821])).
fof(f30540,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13843,f25541])).
fof(f25541,plain,(
  ( ! [X167,X169,X168] : (sP13(X167,X168,X169) | s(bool,X167) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f24303,f824])).
fof(f13843,plain,(
  ( ! [X171,X172] : (~sP13('const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X171,X172) | p(s(bool,X172)) | ~p(s(bool,X171))) ) | $spl0),
  inference(resolution,[],[f12680,f821])).
fof(f30445,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13827,f25527])).
fof(f25527,plain,(
  ( ! [X127,X125,X126] : (sP11(X125,X126,X127) | s(bool,X125) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f24303,f807])).
fof(f13827,plain,(
  ( ! [X140,X139] : (~sP11('const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X139,X140) | p(s(bool,X140)) | p(s(bool,X139))) ) | $spl0),
  inference(resolution,[],[f12680,f804])).
fof(f29513,plain,(
  ( ! [X34,X32] : (sP0(sK22,sK157(sK25,sK22),sK25,X32) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X32),s(sK22,sK24),s(sK22,sK24)))) ) | $spl142),
  inference(cnf_transformation,[],[f29513_D])).
fof(f29513_D,plain,(
  ( ! [X34,X32] : (sP0(sK22,sK157(sK25,sK22),sK25,X32) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X32),s(sK22,sK24),s(sK22,sK24)))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f30261,plain,(
  ( ! [X2,X0,X1] : (sP0(sK22,sK157(sK25,sK22),sK25,X0) | p(s(bool,X1)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X0),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X2))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f29422,f24747])).
fof(f29422,plain,(
  ( ! [X233,X234,X232] : (sP15('const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X232),s(sK22,sK24),s(sK22,sK24)),X233,X234) | sP0(sK22,sK157(sK25,sK22),sK25,X232)) ) | $spl0),
  inference(resolution,[],[f28574,f849])).
fof(f28574,plain,(
  ( ! [X9] : (p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X9),s(sK22,sK24),s(sK22,sK24)))) | sP0(sK22,sK157(sK25,sK22),sK25,X9)) ) | $spl0),
  inference(subsumption_resolution,[],[f28414,f951])).
fof(f951,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0] : ! [X1,X2] : p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X2))))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X6] : ! [X7,X8] : p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X8))))),
  file('thm.relation.TC_CASES|49|_E','thm.relation.RTC_RULES|split|0')).
fof(f28414,plain,(
  ( ! [X9] : (sP0(sK22,sK157(sK25,sK22),sK25,X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X9),s(sK22,sK24),s(sK22,sK24)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),X9),s(sK22,sK24),s(sK22,sK24))))) ) | $spl0),
  inference(resolution,[],[f26964,f944])).
fof(f26964,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X0),s(sK22,sK24),s(sK22,sK24)))) | sP0(sK22,sK157(sK25,sK22),sK25,X0)) ) | $spl0),
  inference(superposition,[],[f22399,f1185])).
fof(f30194,plain,(
  ( ! [X4,X5,X3] : (sP0(sK22,sK157(sK25,sK22),sK25,X3) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X3),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X5))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f29421,f24746])).
fof(f29421,plain,(
  ( ! [X231,X229,X230] : (sP15(X230,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X229),s(sK22,sK24),s(sK22,sK24)),X231) | sP0(sK22,sK157(sK25,sK22),sK25,X229)) ) | $spl0),
  inference(resolution,[],[f28574,f848])).
fof(f30026,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13814,f25510])).
fof(f25510,plain,(
  ( ! [X76,X74,X75] : (sP9(X75,X76,X74) | s(bool,X74) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f24303,f788])).
fof(f13814,plain,(
  ( ! [X114,X113] : (~sP9(X113,X114,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X114)) | ~p(s(bool,X113))) ) | $spl0),
  inference(resolution,[],[f12680,f787])).
fof(f29877,plain,(
  ( ! [X28,X29] : (~p(s(bool,X28)) | p(s(bool,X29)) | s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13813,f25511])).
fof(f25511,plain,(
  ( ! [X78,X79,X77] : (sP9(X77,X78,X79) | s(bool,X77) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f24303,f790])).
fof(f13813,plain,(
  ( ! [X111,X112] : (~sP9('const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X111,X112) | ~p(s(bool,X112)) | p(s(bool,X111))) ) | $spl0),
  inference(resolution,[],[f12680,f787])).
fof(f29347,plain,(
  ( ! [X34,X32] : (sP0(sK22,sK25,sK157(sK25,sK22),X32) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X32),s(sK22,sK24),s(sK22,sK24)))) ) | $spl140),
  inference(cnf_transformation,[],[f29347_D])).
fof(f29347_D,plain,(
  ( ! [X34,X32] : (sP0(sK22,sK25,sK157(sK25,sK22),X32) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X32),s(sK22,sK24),s(sK22,sK24)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f29767,plain,(
  ( ! [X2,X0,X1] : (sP0(sK22,sK25,sK157(sK25,sK22),X0) | p(s(bool,X1)) | s(bool,X2) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X0),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X2))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f29256,f24747])).
fof(f29256,plain,(
  ( ! [X233,X234,X232] : (sP15('const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X232),s(sK22,sK24),s(sK22,sK24)),X233,X234) | sP0(sK22,sK25,sK157(sK25,sK22),X232)) ) | $spl0),
  inference(resolution,[],[f28403,f849])).
fof(f28403,plain,(
  ( ! [X9] : (p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X9),s(sK22,sK24),s(sK22,sK24)))) | sP0(sK22,sK25,sK157(sK25,sK22),X9)) ) | $spl0),
  inference(subsumption_resolution,[],[f28243,f951])).
fof(f28243,plain,(
  ( ! [X9] : (sP0(sK22,sK25,sK157(sK25,sK22),X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X9),s(sK22,sK24),s(sK22,sK24)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),X9),s(sK22,sK24),s(sK22,sK24))))) ) | $spl0),
  inference(resolution,[],[f26728,f944])).
fof(f26728,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X0),s(sK22,sK24),s(sK22,sK24)))) | sP0(sK22,sK25,sK157(sK25,sK22),X0)) ) | $spl0),
  inference(superposition,[],[f22189,f1185])).
fof(f22189,plain,(
  ( ! [X66,X65] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X66),s(sK22,sK24),s(sK22,X65)))) | sP0(sK22,sK25,sK157(X65,sK22),X66)) ) | $spl0),
  inference(superposition,[],[f12857,f948])).
fof(f12857,plain,(
  ( ! [X235,X236] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X235),s(sK22,sK24),s(sK22,X236)))) | sP0(sK22,sK25,X236,X235)) ) | $spl0),
  inference(superposition,[],[f610,f1185])).
fof(f29735,plain,(
  ( ! [X4,X5,X3] : (sP0(sK22,sK25,sK157(sK25,sK22),X3) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X3),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X5))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f29255,f24746])).
fof(f29255,plain,(
  ( ! [X231,X229,X230] : (sP15(X230,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X229),s(sK22,sK24),s(sK22,sK24)),X231) | sP0(sK22,sK25,sK157(sK25,sK22),X229)) ) | $spl0),
  inference(resolution,[],[f28403,f848])).
fof(f29455,plain,(
  ( ! [X318,X317,X319] : (sP0(sK22,sK157(sK25,sK22),sK25,X317) | p(s(bool,X318)) | p(s(bool,X319)) | s(bool,X319) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X317),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f28574,f952])).
fof(f29365,plain,(
  ( ! [X61,X62,X60,X63] : (sP0(sK22,sK157(sK25,sK22),sK25,X60) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X60),s(sK22,sK24),s(sK22,sK24))) | s(bool,X62) != s(bool,X63)) ) | $spl0),
  inference(resolution,[],[f28574,f779])).
fof(f29362,plain,(
  ( ! [X50,X48,X51,X49] : (sP0(sK22,sK157(sK25,sK22),sK25,X48) | p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X48),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f28574,f778])).
fof(f29361,plain,(
  ( ! [X47,X45,X46,X44] : (sP0(sK22,sK157(sK25,sK22),sK25,X44) | p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X44),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f28574,f777])).
fof(f29360,plain,(
  ( ! [X43,X41,X42,X40] : (sP0(sK22,sK157(sK25,sK22),sK25,X40) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X40),s(sK22,sK24),s(sK22,sK24))) | s(bool,X42) != s(bool,X43)) ) | $spl0),
  inference(resolution,[],[f28574,f775])).
fof(f29358,plain,(
  ( ! [X35,X33,X34,X32] : (sP0(sK22,sK157(sK25,sK22),sK25,X32) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X32),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f28574,f773])).
fof(f29289,plain,(
  ( ! [X318,X317,X319] : (sP0(sK22,sK25,sK157(sK25,sK22),X317) | p(s(bool,X318)) | p(s(bool,X319)) | s(bool,X319) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X317),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f28403,f952])).
fof(f29199,plain,(
  ( ! [X61,X62,X60,X63] : (sP0(sK22,sK25,sK157(sK25,sK22),X60) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X60),s(sK22,sK24),s(sK22,sK24))) | s(bool,X62) != s(bool,X63)) ) | $spl0),
  inference(resolution,[],[f28403,f779])).
fof(f29196,plain,(
  ( ! [X50,X48,X51,X49] : (sP0(sK22,sK25,sK157(sK25,sK22),X48) | p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X48),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f28403,f778])).
fof(f29195,plain,(
  ( ! [X47,X45,X46,X44] : (sP0(sK22,sK25,sK157(sK25,sK22),X44) | p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X44),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f28403,f777])).
fof(f29194,plain,(
  ( ! [X43,X41,X42,X40] : (sP0(sK22,sK25,sK157(sK25,sK22),X40) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X40),s(sK22,sK24),s(sK22,sK24))) | s(bool,X42) != s(bool,X43)) ) | $spl0),
  inference(resolution,[],[f28403,f775])).
fof(f29192,plain,(
  ( ! [X35,X33,X34,X32] : (sP0(sK22,sK25,sK157(sK25,sK22),X32) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X32),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f28403,f773])).
fof(f29082,plain,(
  ( ! [X0,X1] : (sP0(sK22,sK157(sK25,sK22),sK25,X0) | p(s(bool,X1)) | s(bool,X1) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X0),s(sK22,sK24),s(sK22,sK24)))) ) | $spl138),
  inference(cnf_transformation,[],[f29082_D])).
fof(f29082_D,plain,(
  ( ! [X0,X1] : (sP0(sK22,sK157(sK25,sK22),sK25,X0) | p(s(bool,X1)) | s(bool,X1) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X0),s(sK22,sK24),s(sK22,sK24)))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f29052,plain,(
  ( ! [X2,X0,X1] : (sP0(sK22,sK157(sK25,sK22),sK25,X0) | p(s(bool,X1)) | s(bool,X1) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X0),s(sK22,sK24),s(sK22,sK24))) | ~p(s(bool,X2))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f28060,f24844])).
fof(f24844,plain,(
  ( ! [X2,X0,X1] : (~sP17(X1,X2,X0) | p(s(bool,X0)) | s(bool,X0) = s(bool,X2) | ~p(s(bool,X1))) ) | ($spl0 | $spl24)),
  inference(subsumption_resolution,[],[f24837,f847])).
fof(f847,plain,(
  ( ! [X2,X0,X1] : (sP15(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f543])).
fof(f24837,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~sP17(X1,X2,X0) | ~sP15(X2,X0,X1) | s(bool,X0) = s(bool,X2) | ~p(s(bool,X1))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24835,f852])).
fof(f852,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP18(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f544])).
fof(f24835,plain,(
  ( ! [X10,X11,X9] : (sP18(X10,X9,X11) | p(s(bool,X9)) | ~sP17(X10,X11,X9) | ~sP15(X11,X9,X10)) ) | ($spl0 | $spl24)),
  inference(subsumption_resolution,[],[f24832,f2847])).
fof(f24832,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | sP18(X10,X9,X11) | p(s(bool,X9)) | ~sP17(X10,X11,X9) | ~sP15(X11,X9,X10)) ) | $spl0),
  inference(resolution,[],[f20559,f836])).
fof(f836,plain,(
  ( ! [X2,X0,X1] : (sP18(X0,X1,X2) | p(s(bool,X1)) | ~sP17(X0,X2,X1) | ~sP16(X1,X2,X0) | ~sP15(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f535])).
fof(f28060,plain,(
  ( ! [X2,X0,X1] : (sP17(X0,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X1),s(sK22,sK24),s(sK22,sK24)),X2) | sP0(sK22,sK157(sK25,sK22),sK25,X1)) ) | $spl0),
  inference(superposition,[],[f23355,f1185])).
fof(f23355,plain,(
  ( ! [X97,X95,X98,X96] : (sP17(X96,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X97),s(sK22,X95),s(sK22,sK24)),X98) | sP0(sK22,sK157(X95,sK22),sK25,X97)) ) | $spl0),
  inference(superposition,[],[f22324,f948])).
fof(f22324,plain,(
  ( ! [X285,X283,X284,X286] : (sP17(X285,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X284),s(sK22,X283),s(sK22,sK24)),X286) | sP0(sK22,X283,sK25,X284)) ) | $spl0),
  inference(resolution,[],[f12858,f840])).
fof(f28796,plain,(
  ( ! [X0,X1] : (sP0(sK22,sK25,sK157(sK25,sK22),X0) | p(s(bool,X1)) | s(bool,X1) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X0),s(sK22,sK24),s(sK22,sK24)))) ) | $spl136),
  inference(cnf_transformation,[],[f28796_D])).
fof(f28796_D,plain,(
  ( ! [X0,X1] : (sP0(sK22,sK25,sK157(sK25,sK22),X0) | p(s(bool,X1)) | s(bool,X1) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X0),s(sK22,sK24),s(sK22,sK24)))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f28766,plain,(
  ( ! [X2,X0,X1] : (sP0(sK22,sK25,sK157(sK25,sK22),X0) | p(s(bool,X1)) | s(bool,X1) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X0),s(sK22,sK24),s(sK22,sK24))) | ~p(s(bool,X2))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f27641,f24844])).
fof(f27641,plain,(
  ( ! [X2,X0,X1] : (sP17(X0,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X1),s(sK22,sK24),s(sK22,sK24)),X2) | sP0(sK22,sK25,sK157(sK25,sK22),X1)) ) | $spl0),
  inference(superposition,[],[f22772,f1185])).
fof(f22772,plain,(
  ( ! [X97,X95,X98,X96] : (sP17(X96,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X97),s(sK22,sK24),s(sK22,X95)),X98) | sP0(sK22,sK25,sK157(X95,sK22),X97)) ) | $spl0),
  inference(superposition,[],[f22104,f948])).
fof(f22104,plain,(
  ( ! [X285,X283,X284,X286] : (sP17(X285,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X284),s(sK22,sK24),s(sK22,X283)),X286) | sP0(sK22,sK25,X283,X284)) ) | $spl0),
  inference(resolution,[],[f12857,f840])).
fof(f28575,plain,(
  ( ! [X36,X34] : (sP0(sK22,sK157(sK25,sK22),sK25,X34) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X34),s(sK22,sK24),s(sK22,sK24)))) ) | $spl134),
  inference(cnf_transformation,[],[f28575_D])).
fof(f28575_D,plain,(
  ( ! [X36,X34] : (sP0(sK22,sK157(sK25,sK22),sK25,X34) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X34),s(sK22,sK24),s(sK22,sK24)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f28521,plain,(
  ( ! [X329,X328,X330] : (sP0(sK22,sK157(sK25,sK22),sK25,X328) | p(s(bool,X329)) | ~p(s(bool,X330)) | s(bool,X330) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X328),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f26964,f952])).
fof(f28428,plain,(
  ( ! [X64,X62,X65,X63] : (sP0(sK22,sK157(sK25,sK22),sK25,X62) | ~p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X62),s(sK22,sK24),s(sK22,sK24))) | s(bool,X63) != s(bool,X65)) ) | $spl0),
  inference(resolution,[],[f26964,f779])).
fof(f28427,plain,(
  ( ! [X61,X59,X60,X58] : (sP0(sK22,sK157(sK25,sK22),sK25,X58) | p(s(bool,X59)) | ~p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X60) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X58),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f26964,f778])).
fof(f28425,plain,(
  ( ! [X52,X50,X53,X51] : (sP0(sK22,sK157(sK25,sK22),sK25,X50) | p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,X53) | s(bool,X52) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X50),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f26964,f777])).
fof(f28423,plain,(
  ( ! [X45,X43,X44,X42] : (sP0(sK22,sK157(sK25,sK22),sK25,X42) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X42),s(sK22,sK24),s(sK22,sK24))) | s(bool,X43) != s(bool,X45)) ) | $spl0),
  inference(resolution,[],[f26964,f775])).
fof(f28421,plain,(
  ( ! [X37,X35,X36,X34] : (sP0(sK22,sK157(sK25,sK22),sK25,X34) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X34),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f26964,f773])).
fof(f28404,plain,(
  ( ! [X36,X34] : (sP0(sK22,sK25,sK157(sK25,sK22),X34) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X34),s(sK22,sK24),s(sK22,sK24)))) ) | $spl132),
  inference(cnf_transformation,[],[f28404_D])).
fof(f28404_D,plain,(
  ( ! [X36,X34] : (sP0(sK22,sK25,sK157(sK25,sK22),X34) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X34),s(sK22,sK24),s(sK22,sK24)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f28350,plain,(
  ( ! [X329,X328,X330] : (sP0(sK22,sK25,sK157(sK25,sK22),X328) | p(s(bool,X329)) | ~p(s(bool,X330)) | s(bool,X330) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X328),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f26728,f952])).
fof(f28257,plain,(
  ( ! [X64,X62,X65,X63] : (sP0(sK22,sK25,sK157(sK25,sK22),X62) | ~p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X62),s(sK22,sK24),s(sK22,sK24))) | s(bool,X63) != s(bool,X65)) ) | $spl0),
  inference(resolution,[],[f26728,f779])).
fof(f28256,plain,(
  ( ! [X61,X59,X60,X58] : (sP0(sK22,sK25,sK157(sK25,sK22),X58) | p(s(bool,X59)) | ~p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X60) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X58),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f26728,f778])).
fof(f28254,plain,(
  ( ! [X52,X50,X53,X51] : (sP0(sK22,sK25,sK157(sK25,sK22),X50) | p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,X53) | s(bool,X52) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X50),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f26728,f777])).
fof(f28252,plain,(
  ( ! [X45,X43,X44,X42] : (sP0(sK22,sK25,sK157(sK25,sK22),X42) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X42),s(sK22,sK24),s(sK22,sK24))) | s(bool,X43) != s(bool,X45)) ) | $spl0),
  inference(resolution,[],[f26728,f775])).
fof(f28250,plain,(
  ( ! [X37,X35,X36,X34] : (sP0(sK22,sK25,sK157(sK25,sK22),X34) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X34),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f26728,f773])).
fof(f28086,plain,(
  ( ! [X2,X0,X1] : (sP0(sK22,sK157(X0,sK22),sK25,X1) | s(bool,X2) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X1),s(sK22,X0),s(sK22,sK24))) | p(s(bool,X2))) ) | $spl130),
  inference(cnf_transformation,[],[f28086_D])).
fof(f28086_D,plain,(
  ( ! [X2,X0,X1] : (sP0(sK22,sK157(X0,sK22),sK25,X1) | s(bool,X2) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X1),s(sK22,X0),s(sK22,sK24))) | p(s(bool,X2))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f28040,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(sK22,sK157(X0,sK22),sK25,X1) | p(s(bool,X2)) | s(bool,X2) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X1),s(sK22,X0),s(sK22,sK24))) | ~p(s(bool,X3))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f23355,f24844])).
fof(f27657,plain,(
  ( ! [X2,X0,X1] : (sP0(sK22,sK25,sK157(X0,sK22),X1) | s(bool,X2) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X1),s(sK22,sK24),s(sK22,X0))) | p(s(bool,X2))) ) | $spl128),
  inference(cnf_transformation,[],[f27657_D])).
fof(f27657_D,plain,(
  ( ! [X2,X0,X1] : (sP0(sK22,sK25,sK157(X0,sK22),X1) | s(bool,X2) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X1),s(sK22,sK24),s(sK22,X0))) | p(s(bool,X2))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f27611,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(sK22,sK25,sK157(X0,sK22),X1) | p(s(bool,X2)) | s(bool,X2) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X1),s(sK22,sK24),s(sK22,X0))) | ~p(s(bool,X3))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f22772,f24844])).
fof(f27170,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | s(bool,X0) = s(bool,X1) | p(s(bool,X1))) ) | $spl126),
  inference(cnf_transformation,[],[f27170_D])).
fof(f27170_D,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | s(bool,X0) = s(bool,X1) | p(s(bool,X1))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f27165,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | ~p(s(bool,X2))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f25553,f24844])).
fof(f27016,plain,(
  ( ! [X47,X50,X48] : (sP0(sK22,sK157(X47,sK22),sK25,X48) | s(bool,X50) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X48),s(sK22,X47),s(sK22,sK24))) | ~p(s(bool,X50))) ) | $spl124),
  inference(cnf_transformation,[],[f27016_D])).
fof(f27016_D,plain,(
  ( ! [X47,X50,X48] : (sP0(sK22,sK157(X47,sK22),sK25,X48) | s(bool,X50) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X48),s(sK22,X47),s(sK22,sK24))) | ~p(s(bool,X50))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f26947,plain,(
  ( ! [X441,X443,X444,X442] : (sP0(sK22,sK157(X441,sK22),sK25,X442) | p(s(bool,X443)) | ~p(s(bool,X444)) | s(bool,X444) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X442),s(sK22,X441),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f22399,f952])).
fof(f26854,plain,(
  ( ! [X85,X83,X86,X84,X82] : (sP0(sK22,sK157(X82,sK22),sK25,X83) | ~p(s(bool,X84)) | ~p(s(bool,X85)) | s(bool,X85) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X83),s(sK22,X82),s(sK22,sK24))) | s(bool,X84) != s(bool,X86)) ) | $spl0),
  inference(resolution,[],[f22399,f779])).
fof(f26853,plain,(
  ( ! [X80,X78,X81,X79,X77] : (sP0(sK22,sK157(X77,sK22),sK25,X78) | p(s(bool,X79)) | ~p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X80) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X78),s(sK22,X77),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f22399,f778])).
fof(f26851,plain,(
  ( ! [X70,X68,X71,X69,X67] : (sP0(sK22,sK157(X67,sK22),sK25,X68) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,X71) | s(bool,X70) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X68),s(sK22,X67),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f22399,f777])).
fof(f26849,plain,(
  ( ! [X61,X59,X57,X60,X58] : (sP0(sK22,sK157(X57,sK22),sK25,X58) | ~p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X58),s(sK22,X57),s(sK22,sK24))) | s(bool,X59) != s(bool,X61)) ) | $spl0),
  inference(resolution,[],[f22399,f775])).
fof(f26847,plain,(
  ( ! [X47,X50,X48,X51,X49] : (sP0(sK22,sK157(X47,sK22),sK25,X48) | p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,X51) | s(bool,X50) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X48),s(sK22,X47),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f22399,f773])).
fof(f26770,plain,(
  ( ! [X47,X50,X48] : (sP0(sK22,sK25,sK157(X47,sK22),X48) | s(bool,X50) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X48),s(sK22,sK24),s(sK22,X47))) | ~p(s(bool,X50))) ) | $spl122),
  inference(cnf_transformation,[],[f26770_D])).
fof(f26770_D,plain,(
  ( ! [X47,X50,X48] : (sP0(sK22,sK25,sK157(X47,sK22),X48) | s(bool,X50) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X48),s(sK22,sK24),s(sK22,X47))) | ~p(s(bool,X50))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f26701,plain,(
  ( ! [X441,X443,X444,X442] : (sP0(sK22,sK25,sK157(X441,sK22),X442) | p(s(bool,X443)) | ~p(s(bool,X444)) | s(bool,X444) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X442),s(sK22,sK24),s(sK22,X441)))) ) | $spl0),
  inference(resolution,[],[f22189,f952])).
fof(f26608,plain,(
  ( ! [X85,X83,X86,X84,X82] : (sP0(sK22,sK25,sK157(X82,sK22),X83) | ~p(s(bool,X84)) | ~p(s(bool,X85)) | s(bool,X85) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X83),s(sK22,sK24),s(sK22,X82))) | s(bool,X84) != s(bool,X86)) ) | $spl0),
  inference(resolution,[],[f22189,f779])).
fof(f26607,plain,(
  ( ! [X80,X78,X81,X79,X77] : (sP0(sK22,sK25,sK157(X77,sK22),X78) | p(s(bool,X79)) | ~p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X80) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X78),s(sK22,sK24),s(sK22,X77)))) ) | $spl0),
  inference(resolution,[],[f22189,f778])).
fof(f26605,plain,(
  ( ! [X70,X68,X71,X69,X67] : (sP0(sK22,sK25,sK157(X67,sK22),X68) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,X71) | s(bool,X70) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X68),s(sK22,sK24),s(sK22,X67)))) ) | $spl0),
  inference(resolution,[],[f22189,f777])).
fof(f26603,plain,(
  ( ! [X61,X59,X57,X60,X58] : (sP0(sK22,sK25,sK157(X57,sK22),X58) | ~p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X58),s(sK22,sK24),s(sK22,X57))) | s(bool,X59) != s(bool,X61)) ) | $spl0),
  inference(resolution,[],[f22189,f775])).
fof(f26601,plain,(
  ( ! [X47,X50,X48,X51,X49] : (sP0(sK22,sK25,sK157(X47,sK22),X48) | p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,X51) | s(bool,X50) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X48),s(sK22,sK24),s(sK22,X47)))) ) | $spl0),
  inference(resolution,[],[f22189,f773])).
fof(f25881,plain,(
  ( ! [X26,X24] : (s(bool,X24) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | s(bool,X24) != s(bool,X26) | ~p(s(bool,X26))) ) | $spl120),
  inference(cnf_transformation,[],[f25881_D])).
fof(f25881_D,plain,(
  ( ! [X26,X24] : (s(bool,X24) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | s(bool,X24) != s(bool,X26) | ~p(s(bool,X26))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f25596,plain,(
  ( ! [X318,X320,X319] : (s(bool,X318) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X319)) | ~p(s(bool,X320)) | s(bool,X318) != s(bool,X320)) ) | $spl0),
  inference(resolution,[],[f24303,f952])).
fof(f25503,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,X52) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X54) | s(bool,X53) != s(bool,X55)) ) | $spl0),
  inference(resolution,[],[f24303,f779])).
fof(f25502,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X48) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X49)) | ~p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X48) != s(bool,X50)) ) | $spl0),
  inference(resolution,[],[f24303,f778])).
fof(f25500,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X40) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X40) != s(bool,X42)) ) | $spl0),
  inference(resolution,[],[f24303,f777])).
fof(f25498,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X32) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,X35)) ) | $spl0),
  inference(resolution,[],[f24303,f775])).
fof(f25496,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26)) ) | $spl0),
  inference(resolution,[],[f24303,f773])).
fof(f25353,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f25226,f1193])).
fof(f25226,plain,(
  ( ! [X1402,X1404,X1406,X1403,X1405] : (p(s(bool,happ(s(fun(X1402,bool),happ(s(fun(X1402,fun(X1402,bool)),sK51(X1403,X1404,X1405,X1402)),s(X1402,X1406))),s(X1402,X1406))))) ) | $spl12),
  inference(resolution,[],[f1193,f615])).
fof(f615,plain,(
  ( ! [X2,X0,X8,X3,X1] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK51(X3,X2,X1,X0)),s(X0,X8))),s(X0,X8))))) )),
  inference(cnf_transformation,[],[f406])).
fof(f406,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ! [X9] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,sK47(X9,X3,X2,X1,X0)))),s(X0,sK47(X9,X3,X2,X1,X0))))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK48(X9,X3,X2,X1,X0)))),s(X0,sK49(X9,X3,X2,X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,sK49(X9,X3,X2,X1,X0)))),s(X0,sK50(X9,X3,X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,sK48(X9,X3,X2,X1,X0)))),s(X0,sK50(X9,X3,X2,X1,X0)))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,X2))),s(X0,X3)))))) & ((! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK51(X3,X2,X1,X0)),s(X0,X8))),s(X0,X8)))) & ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK51(X3,X2,X1,X0)),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK51(X3,X2,X1,X0)),s(X0,X5))),s(X0,X7))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),sK51(X3,X2,X1,X0)),s(X0,X2))),s(X0,X3))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK47,sK48,sK49,sK50,sK51])],[f405])).
fof(f405,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ! [X9] : (? [X13] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,X13))),s(X0,X13)))) | ? [X10,X11,X12] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X10))),s(X0,X11)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,X11))),s(X0,X12)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,X10))),s(X0,X12))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X9),s(X0,X2))),s(X0,X3)))))) & (? [X4] : (! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X8)))) & ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(rectify,[],[f404])).
fof(f404,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | ! [X4] : (? [X8] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X8)))) | ? [X5,X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3)))))) & (? [X4] : (! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X8)))) & ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))) | p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : (? [X8] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X8)))) | ? [X5,X6,X7] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(flattening,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : ((? [X8] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X8)))) | ? [X5,X6,X7] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7)))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(ennf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : ((! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X8)))) & ! [X5,X6,X7] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7)))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.relation.RTC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) <=> ! [X4] : ((! [X8] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X8))),s(X0,X8)))) & ! [X5,X6,X7] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X6))),s(X0,X7))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X5))),s(X0,X7)))))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X4),s(X0,X2))),s(X0,X3))))))),
  inference(rectify,[],[f126])).
fof(f126,axiom,(
  ! [X6] : ! [X7,X21,X22] : (p(s(bool,'const.relation.RTC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X21),s(X6,X22)))) <=> ! [X11] : ((! [X8] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X11),s(X6,X8))),s(X6,X8)))) & ! [X8,X10,X9] : ((p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X10)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X11),s(X6,X10))),s(X6,X9))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X11),s(X6,X8))),s(X6,X9)))))) => p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X11),s(X6,X21))),s(X6,X22))))))),
  file('thm.relation.TC_CASES|49|_E','thm.relation.RTC_DEF')).
fof(f25352,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f25179,f1193])).
fof(f25179,plain,(
  ( ! [X1217,X1214,X1216,X1215] : (p(s(bool,'const.relation.RTC_3'(s(fun(X1214,fun(X1214,bool)),X1215),s(X1214,X1216),s(X1214,X1217))))) ) | $spl12),
  inference(resolution,[],[f1193,f615])).
fof(f25010,plain,(
  $false | ($spl0 | $spl12)),
  inference(resolution,[],[f1193,f12680])).
fof(f25011,plain,(
  $false | ($spl0 | $spl12 | $spl34)),
  inference(resolution,[],[f1193,f12772])).
fof(f25012,plain,(
  $false | ($spl12 | $spl74)),
  inference(resolution,[],[f1193,f15763])).
fof(f25235,plain,(
  $false | $spl12),
  inference(resolution,[],[f1193,f951])).
fof(f25239,plain,(
  $false | ($spl0 | $spl12)),
  inference(resolution,[],[f1193,f12683])).
fof(f25240,plain,(
  $false | ($spl0 | $spl12)),
  inference(resolution,[],[f1193,f13737])).
fof(f13737,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) | $spl0),
  inference(resolution,[],[f12680,f599])).
fof(f24882,plain,(
  ( ! [X50] : (p(s(bool,X50)) | s(bool,X50) = s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl118),
  inference(cnf_transformation,[],[f24882_D])).
fof(f24882_D,plain,(
  ( ! [X50] : (p(s(bool,X50)) | s(bool,X50) = s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f24866,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | s(bool,X50) = s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | ~p(s(bool,X51))) ) | ($spl0 | $spl24 | $spl63)),
  inference(resolution,[],[f24844,f18623])).
fof(f18623,plain,(
  ( ! [X140,X141] : (sP17(X140,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)),X141)) ) | $spl63),
  inference(resolution,[],[f15084,f840])).
fof(f15084,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) | $spl63),
  inference(cnf_transformation,[],[f15084_D])).
fof(f15084_D,plain,(
  ~p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f24881,plain,(
  ( ! [X48] : (p(s(bool,X48)) | s(bool,X48) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl116),
  inference(cnf_transformation,[],[f24881_D])).
fof(f24881_D,plain,(
  ( ! [X48] : (p(s(bool,X48)) | s(bool,X48) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f24865,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | s(bool,X48) = s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | ~p(s(bool,X49))) ) | ($spl0 | $spl21 | $spl24)),
  inference(resolution,[],[f24844,f19519])).
fof(f19519,plain,(
  ( ! [X149,X148] : (sP17(X148,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)),X149)) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f16928,f840])).
fof(f16928,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) | ($spl0 | $spl21)),
  inference(forward_demodulation,[],[f1677,f1185])).
fof(f1677,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | $spl21),
  inference(cnf_transformation,[],[f1677_D])).
fof(f1677_D,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f24880,plain,(
  ( ! [X46] : (p(s(bool,X46)) | s(bool,X46) = s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK24)))) ) | $spl114),
  inference(cnf_transformation,[],[f24880_D])).
fof(f24880_D,plain,(
  ( ! [X46] : (p(s(bool,X46)) | s(bool,X46) = s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK24)))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f24864,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | s(bool,X46) = s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK24))) | ~p(s(bool,X47))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24844,f12725])).
fof(f12725,plain,(
  ( ! [X229,X230] : (sP17(X229,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK24)),X230)) ) | $spl0),
  inference(backward_demodulation,[],[f1185,f1348])).
fof(f1348,plain,(
  ( ! [X229,X230] : (sP17(X229,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25)),X230)) )),
  inference(resolution,[],[f585,f840])).
fof(f585,plain,(
  ~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25))))),
  inference(cnf_transformation,[],[f385])).
fof(f24879,plain,(
  ( ! [X43,X44] : (p(s(bool,X43)) | sP0(sK22,sK25,sK25,X44) | s(bool,X43) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X44),s(sK22,sK24),s(sK22,sK24)))) ) | $spl112),
  inference(cnf_transformation,[],[f24879_D])).
fof(f24879_D,plain,(
  ( ! [X43,X44] : (p(s(bool,X43)) | sP0(sK22,sK25,sK25,X44) | s(bool,X43) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X44),s(sK22,sK24),s(sK22,sK24)))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f24863,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | s(bool,X43) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X44),s(sK22,sK24),s(sK22,sK24))) | ~p(s(bool,X45)) | sP0(sK22,sK25,sK25,X44)) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24844,f22757])).
fof(f22757,plain,(
  ( ! [X2,X0,X1] : (sP17(X0,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X1),s(sK22,sK24),s(sK22,sK24)),X2) | sP0(sK22,sK25,sK25,X1)) ) | $spl0),
  inference(superposition,[],[f22104,f1185])).
fof(f24878,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | sP0(sK22,X41,sK25,X40) | s(bool,X39) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X40),s(sK22,X41),s(sK22,sK24)))) ) | $spl110),
  inference(cnf_transformation,[],[f24878_D])).
fof(f24878_D,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | sP0(sK22,X41,sK25,X40) | s(bool,X39) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X40),s(sK22,X41),s(sK22,sK24)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f24862,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,X39)) | s(bool,X39) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X40),s(sK22,X41),s(sK22,sK24))) | ~p(s(bool,X42)) | sP0(sK22,X41,sK25,X40)) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24844,f22324])).
fof(f24877,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | sP0(sK22,sK25,X37,X36) | s(bool,X35) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X36),s(sK22,sK24),s(sK22,X37)))) ) | $spl108),
  inference(cnf_transformation,[],[f24877_D])).
fof(f24877_D,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X35)) | sP0(sK22,sK25,X37,X36) | s(bool,X35) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X36),s(sK22,sK24),s(sK22,X37)))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f24861,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,X35)) | s(bool,X35) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X36),s(sK22,sK24),s(sK22,X37))) | ~p(s(bool,X38)) | sP0(sK22,sK25,X37,X36)) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24844,f22104])).
fof(f24876,plain,(
  ( ! [X33] : (p(s(bool,X33)) | s(bool,X33) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl106),
  inference(cnf_transformation,[],[f24876_D])).
fof(f24876_D,plain,(
  ( ! [X33] : (p(s(bool,X33)) | s(bool,X33) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f24860,plain,(
  ( ! [X33,X34] : (p(s(bool,X33)) | s(bool,X33) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | ~p(s(bool,X34))) ) | ($spl0 | $spl21 | $spl24)),
  inference(resolution,[],[f24844,f21276])).
fof(f21276,plain,(
  ( ! [X144,X143] : (sP17(X143,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)),X144)) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f19453,f840])).
fof(f19453,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) | ($spl0 | $spl21)),
  inference(resolution,[],[f16928,f943])).
fof(f24875,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X4) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | s(bool,X3) = s(bool,X4)) ) | $spl104),
  inference(cnf_transformation,[],[f24875_D])).
fof(f24875_D,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X4) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | s(bool,X3) = s(bool,X4)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f24848,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,X4) | ~p(s(bool,X5)) | s(bool,X4) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24844,f20555])).
fof(f23953,plain,(
  ( ! [X34,X32] : (sP0(sK22,sK25,sK25,X32) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X32),s(sK22,sK24),s(sK22,sK24)))) ) | $spl100),
  inference(cnf_transformation,[],[f23953_D])).
fof(f23953_D,plain,(
  ( ! [X34,X32] : (sP0(sK22,sK25,sK25,X32) | p(s(bool,X34)) | s(bool,X34) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X32),s(sK22,sK24),s(sK22,sK24)))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f24821,plain,(
  ( ! [X54,X56,X55] : (p(s(bool,X54)) | s(bool,X56) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X55),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X56)) | sP0(sK22,sK25,sK25,X55)) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24747,f23863])).
fof(f23863,plain,(
  ( ! [X233,X234,X232] : (sP15('const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X232),s(sK22,sK24),s(sK22,sK24)),X233,X234) | sP0(sK22,sK25,sK25,X232)) ) | $spl0),
  inference(resolution,[],[f23085,f849])).
fof(f23085,plain,(
  ( ! [X9] : (p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X9),s(sK22,sK24),s(sK22,sK24)))) | sP0(sK22,sK25,sK25,X9)) ) | $spl0),
  inference(subsumption_resolution,[],[f22925,f951])).
fof(f22925,plain,(
  ( ! [X9] : (sP0(sK22,sK25,sK25,X9) | p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X9),s(sK22,sK24),s(sK22,sK24)))) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),X9),s(sK22,sK24),s(sK22,sK24))))) ) | $spl0),
  inference(resolution,[],[f22174,f944])).
fof(f22174,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X0),s(sK22,sK24),s(sK22,sK24)))) | sP0(sK22,sK25,sK25,X0)) ) | $spl0),
  inference(superposition,[],[f12857,f1185])).
fof(f15766,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl78),
  inference(cnf_transformation,[],[f15766_D])).
fof(f15766_D,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f24820,plain,(
  ( ! [X52,X53] : (p(s(bool,X52)) | s(bool,X53) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24))) | p(s(bool,X53))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24747,f15662])).
fof(f15662,plain,(
  ( ! [X180,X181] : (sP15('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)),X180,X181)) ) | $spl0),
  inference(resolution,[],[f13737,f849])).
fof(f15088,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl66),
  inference(cnf_transformation,[],[f15088_D])).
fof(f15088_D,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f24819,plain,(
  ( ! [X50,X51] : (p(s(bool,X50)) | s(bool,X51) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | p(s(bool,X51))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24747,f12768])).
fof(f12768,plain,(
  ( ! [X180,X181] : (sP15('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)),X180,X181)) ) | $spl0),
  inference(backward_demodulation,[],[f1185,f6233])).
fof(f6233,plain,(
  ( ! [X180,X181] : (sP15('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)),X180,X181)) )),
  inference(resolution,[],[f963,f849])).
fof(f24795,plain,(
  ( ! [X34] : (s(bool,X34) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X34))) ) | $spl102),
  inference(cnf_transformation,[],[f24795_D])).
fof(f24795_D,plain,(
  ( ! [X34] : (s(bool,X34) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X34))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f24818,plain,(
  ( ! [X48,X49] : (p(s(bool,X48)) | s(bool,X49) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X49))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24747,f12749])).
fof(f12749,plain,(
  ( ! [X180,X181] : (sP15('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X180,X181)) ) | $spl0),
  inference(backward_demodulation,[],[f1185,f1581])).
fof(f1581,plain,(
  ( ! [X180,X181] : (sP15('const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)),X180,X181)) )),
  inference(resolution,[],[f1016,f849])).
fof(f19349,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl86),
  inference(cnf_transformation,[],[f19349_D])).
fof(f19349_D,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f24817,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | s(bool,X47) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24))) | p(s(bool,X47))) ) | ($spl0 | $spl24 | $spl74)),
  inference(resolution,[],[f24747,f19253])).
fof(f19253,plain,(
  ( ! [X213,X214] : (sP15('const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)),X213,X214)) ) | $spl74),
  inference(resolution,[],[f15763,f849])).
fof(f18251,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl80),
  inference(cnf_transformation,[],[f18251_D])).
fof(f18251_D,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f24816,plain,(
  ( ! [X45,X44] : (p(s(bool,X44)) | s(bool,X45) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | p(s(bool,X45))) ) | ($spl0 | $spl24 | $spl34)),
  inference(resolution,[],[f24747,f18155])).
fof(f18155,plain,(
  ( ! [X213,X214] : (sP15('const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)),X213,X214)) ) | ($spl0 | $spl34)),
  inference(resolution,[],[f12772,f849])).
fof(f24787,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | s(bool,X40) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X41),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X40)) | sP0(sK22,sK25,sK25,X41)) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24746,f23862])).
fof(f23862,plain,(
  ( ! [X231,X229,X230] : (sP15(X230,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X229),s(sK22,sK24),s(sK22,sK24)),X231) | sP0(sK22,sK25,sK25,X229)) ) | $spl0),
  inference(resolution,[],[f23085,f848])).
fof(f24786,plain,(
  ( ! [X37,X38] : (p(s(bool,X37)) | s(bool,X38) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24))) | p(s(bool,X38))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24746,f15661])).
fof(f15661,plain,(
  ( ! [X178,X179] : (sP15(X178,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)),X179)) ) | $spl0),
  inference(resolution,[],[f13737,f848])).
fof(f24785,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | s(bool,X36) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | p(s(bool,X36))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24746,f12767])).
fof(f12767,plain,(
  ( ! [X178,X179] : (sP15(X178,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)),X179)) ) | $spl0),
  inference(backward_demodulation,[],[f1185,f6232])).
fof(f6232,plain,(
  ( ! [X178,X179] : (sP15(X178,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)),X179)) )),
  inference(resolution,[],[f963,f848])).
fof(f24784,plain,(
  ( ! [X33,X34] : (p(s(bool,X33)) | s(bool,X34) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X34))) ) | ($spl0 | $spl24)),
  inference(resolution,[],[f24746,f12748])).
fof(f12748,plain,(
  ( ! [X178,X179] : (sP15(X178,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)),X179)) ) | $spl0),
  inference(backward_demodulation,[],[f1185,f1580])).
fof(f1580,plain,(
  ( ! [X178,X179] : (sP15(X178,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)),X179)) )),
  inference(resolution,[],[f1016,f848])).
fof(f24783,plain,(
  ( ! [X31,X32] : (p(s(bool,X31)) | s(bool,X32) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24))) | p(s(bool,X32))) ) | ($spl0 | $spl24 | $spl74)),
  inference(resolution,[],[f24746,f19252])).
fof(f19252,plain,(
  ( ! [X212,X211] : (sP15(X211,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)),X212)) ) | $spl74),
  inference(resolution,[],[f15763,f848])).
fof(f24782,plain,(
  ( ! [X30,X29] : (p(s(bool,X29)) | s(bool,X30) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | p(s(bool,X30))) ) | ($spl0 | $spl24 | $spl34)),
  inference(resolution,[],[f24746,f18154])).
fof(f18154,plain,(
  ( ! [X212,X211] : (sP15(X211,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)),X212)) ) | ($spl0 | $spl34)),
  inference(resolution,[],[f12772,f848])).
fof(f23896,plain,(
  ( ! [X318,X317,X319] : (sP0(sK22,sK25,sK25,X317) | p(s(bool,X318)) | p(s(bool,X319)) | s(bool,X319) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X317),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f23085,f952])).
fof(f23806,plain,(
  ( ! [X61,X62,X60,X63] : (sP0(sK22,sK25,sK25,X60) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X60),s(sK22,sK24),s(sK22,sK24))) | s(bool,X62) != s(bool,X63)) ) | $spl0),
  inference(resolution,[],[f23085,f779])).
fof(f23803,plain,(
  ( ! [X50,X48,X51,X49] : (sP0(sK22,sK25,sK25,X48) | p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | s(bool,X51) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X48),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f23085,f778])).
fof(f23802,plain,(
  ( ! [X47,X45,X46,X44] : (sP0(sK22,sK25,sK25,X44) | p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X45) != s(bool,X47) | s(bool,X46) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X44),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f23085,f777])).
fof(f23801,plain,(
  ( ! [X43,X41,X42,X40] : (sP0(sK22,sK25,sK25,X40) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X40),s(sK22,sK24),s(sK22,sK24))) | s(bool,X42) != s(bool,X43)) ) | $spl0),
  inference(resolution,[],[f23085,f775])).
fof(f23799,plain,(
  ( ! [X35,X33,X34,X32] : (sP0(sK22,sK25,sK25,X32) | p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),X32),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f23085,f773])).
fof(f23086,plain,(
  ( ! [X36,X34] : (sP0(sK22,sK25,sK25,X34) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X34),s(sK22,sK24),s(sK22,sK24)))) ) | $spl98),
  inference(cnf_transformation,[],[f23086_D])).
fof(f23086_D,plain,(
  ( ! [X36,X34] : (sP0(sK22,sK25,sK25,X34) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X34),s(sK22,sK24),s(sK22,sK24)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f23032,plain,(
  ( ! [X329,X328,X330] : (sP0(sK22,sK25,sK25,X328) | p(s(bool,X329)) | ~p(s(bool,X330)) | s(bool,X330) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X328),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f22174,f952])).
fof(f22939,plain,(
  ( ! [X64,X62,X65,X63] : (sP0(sK22,sK25,sK25,X62) | ~p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X62),s(sK22,sK24),s(sK22,sK24))) | s(bool,X63) != s(bool,X65)) ) | $spl0),
  inference(resolution,[],[f22174,f779])).
fof(f22938,plain,(
  ( ! [X61,X59,X60,X58] : (sP0(sK22,sK25,sK25,X58) | p(s(bool,X59)) | ~p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X60) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X58),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f22174,f778])).
fof(f22936,plain,(
  ( ! [X52,X50,X53,X51] : (sP0(sK22,sK25,sK25,X50) | p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,X53) | s(bool,X52) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X50),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f22174,f777])).
fof(f22934,plain,(
  ( ! [X45,X43,X44,X42] : (sP0(sK22,sK25,sK25,X42) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X42),s(sK22,sK24),s(sK22,sK24))) | s(bool,X43) != s(bool,X45)) ) | $spl0),
  inference(resolution,[],[f22174,f775])).
fof(f22932,plain,(
  ( ! [X37,X35,X36,X34] : (sP0(sK22,sK25,sK25,X34) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X34),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f22174,f773])).
fof(f22436,plain,(
  ( ! [X47,X50,X48] : (sP0(sK22,X47,sK25,X48) | s(bool,X50) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X48),s(sK22,X47),s(sK22,sK24))) | ~p(s(bool,X50))) ) | $spl96),
  inference(cnf_transformation,[],[f22436_D])).
fof(f22436_D,plain,(
  ( ! [X47,X50,X48] : (sP0(sK22,X47,sK25,X48) | s(bool,X50) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X48),s(sK22,X47),s(sK22,sK24))) | ~p(s(bool,X50))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f22367,plain,(
  ( ! [X441,X443,X444,X442] : (sP0(sK22,X441,sK25,X442) | p(s(bool,X443)) | ~p(s(bool,X444)) | s(bool,X444) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X442),s(sK22,X441),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12858,f952])).
fof(f22274,plain,(
  ( ! [X85,X83,X86,X84,X82] : (sP0(sK22,X82,sK25,X83) | ~p(s(bool,X84)) | ~p(s(bool,X85)) | s(bool,X85) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X83),s(sK22,X82),s(sK22,sK24))) | s(bool,X84) != s(bool,X86)) ) | $spl0),
  inference(resolution,[],[f12858,f779])).
fof(f22273,plain,(
  ( ! [X80,X78,X81,X79,X77] : (sP0(sK22,X77,sK25,X78) | p(s(bool,X79)) | ~p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X80) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X78),s(sK22,X77),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12858,f778])).
fof(f22271,plain,(
  ( ! [X70,X68,X71,X69,X67] : (sP0(sK22,X67,sK25,X68) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,X71) | s(bool,X70) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X68),s(sK22,X67),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12858,f777])).
fof(f22269,plain,(
  ( ! [X61,X59,X57,X60,X58] : (sP0(sK22,X57,sK25,X58) | ~p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X58),s(sK22,X57),s(sK22,sK24))) | s(bool,X59) != s(bool,X61)) ) | $spl0),
  inference(resolution,[],[f12858,f775])).
fof(f22267,plain,(
  ( ! [X47,X50,X48,X51,X49] : (sP0(sK22,X47,sK25,X48) | p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,X51) | s(bool,X50) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X48),s(sK22,X47),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12858,f773])).
fof(f22216,plain,(
  ( ! [X47,X50,X48] : (sP0(sK22,sK25,X47,X48) | s(bool,X50) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X48),s(sK22,sK24),s(sK22,X47))) | ~p(s(bool,X50))) ) | $spl94),
  inference(cnf_transformation,[],[f22216_D])).
fof(f22216_D,plain,(
  ( ! [X47,X50,X48] : (sP0(sK22,sK25,X47,X48) | s(bool,X50) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X48),s(sK22,sK24),s(sK22,X47))) | ~p(s(bool,X50))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f22147,plain,(
  ( ! [X441,X443,X444,X442] : (sP0(sK22,sK25,X441,X442) | p(s(bool,X443)) | ~p(s(bool,X444)) | s(bool,X444) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X442),s(sK22,sK24),s(sK22,X441)))) ) | $spl0),
  inference(resolution,[],[f12857,f952])).
fof(f22054,plain,(
  ( ! [X85,X83,X86,X84,X82] : (sP0(sK22,sK25,X82,X83) | ~p(s(bool,X84)) | ~p(s(bool,X85)) | s(bool,X85) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X83),s(sK22,sK24),s(sK22,X82))) | s(bool,X84) != s(bool,X86)) ) | $spl0),
  inference(resolution,[],[f12857,f779])).
fof(f22053,plain,(
  ( ! [X80,X78,X81,X79,X77] : (sP0(sK22,sK25,X77,X78) | p(s(bool,X79)) | ~p(s(bool,X80)) | ~p(s(bool,X81)) | s(bool,X80) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X78),s(sK22,sK24),s(sK22,X77)))) ) | $spl0),
  inference(resolution,[],[f12857,f778])).
fof(f22051,plain,(
  ( ! [X70,X68,X71,X69,X67] : (sP0(sK22,sK25,X67,X68) | p(s(bool,X69)) | ~p(s(bool,X70)) | s(bool,X69) != s(bool,X71) | s(bool,X70) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X68),s(sK22,sK24),s(sK22,X67)))) ) | $spl0),
  inference(resolution,[],[f12857,f777])).
fof(f22049,plain,(
  ( ! [X61,X59,X57,X60,X58] : (sP0(sK22,sK25,X57,X58) | ~p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X58),s(sK22,sK24),s(sK22,X57))) | s(bool,X59) != s(bool,X61)) ) | $spl0),
  inference(resolution,[],[f12857,f775])).
fof(f22047,plain,(
  ( ! [X47,X50,X48,X51,X49] : (sP0(sK22,sK25,X47,X48) | p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X49) != s(bool,X51) | s(bool,X50) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),X48),s(sK22,sK24),s(sK22,X47)))) ) | $spl0),
  inference(resolution,[],[f12857,f773])).
fof(f21377,plain,(
  ( ! [X22] : (~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl92),
  inference(cnf_transformation,[],[f21377_D])).
fof(f21377_D,plain,(
  ( ! [X22] : (~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f21319,plain,(
  ( ! [X216,X215] : (p(s(bool,X215)) | ~p(s(bool,X216)) | s(bool,X216) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f19453,f952])).
fof(f21226,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X43) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X42) != s(bool,X44)) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f19453,f779])).
fof(f21225,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f19453,f778])).
fof(f21223,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f19453,f777])).
fof(f21221,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X27) != s(bool,X29)) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f19453,f775])).
fof(f21219,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f19453,f773])).
fof(f20882,plain,(
  ( ! [X26,X24] : (s(bool,X24) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | s(bool,X24) != s(bool,X26) | ~p(s(bool,X26))) ) | $spl90),
  inference(cnf_transformation,[],[f20882_D])).
fof(f20882_D,plain,(
  ( ! [X26,X24] : (s(bool,X24) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | s(bool,X24) != s(bool,X26) | ~p(s(bool,X26))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f20598,plain,(
  ( ! [X318,X320,X319] : (s(bool,X318) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X319)) | ~p(s(bool,X320)) | s(bool,X318) != s(bool,X320)) ) | $spl0),
  inference(resolution,[],[f12719,f952])).
fof(f20505,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,X52) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X54) | s(bool,X53) != s(bool,X55)) ) | $spl0),
  inference(resolution,[],[f12719,f779])).
fof(f20504,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X48) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X49)) | ~p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X48) != s(bool,X50)) ) | $spl0),
  inference(resolution,[],[f12719,f778])).
fof(f20502,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X40) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X40) != s(bool,X42)) ) | $spl0),
  inference(resolution,[],[f12719,f777])).
fof(f20500,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X32) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,X35)) ) | $spl0),
  inference(resolution,[],[f12719,f775])).
fof(f20498,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26)) ) | $spl0),
  inference(resolution,[],[f12719,f773])).
fof(f19620,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl88),
  inference(cnf_transformation,[],[f19620_D])).
fof(f19620_D,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f19562,plain,(
  ( ! [X220,X221] : (p(s(bool,X220)) | ~p(s(bool,X221)) | s(bool,X221) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f16928,f952])).
fof(f19469,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X48) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X47) != s(bool,X49)) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f16928,f779])).
fof(f19468,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f16928,f778])).
fof(f19466,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X38)) | ~p(s(bool,X39)) | s(bool,X38) != s(bool,X40) | s(bool,X39) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f16928,f777])).
fof(f19464,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X32) != s(bool,X34)) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f16928,f775])).
fof(f19462,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl21)),
  inference(resolution,[],[f16928,f773])).
fof(f19286,plain,(
  ( ! [X265,X266] : (p(s(bool,X265)) | p(s(bool,X266)) | s(bool,X266) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl74),
  inference(resolution,[],[f15763,f952])).
fof(f19196,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X99) != s(bool,X100)) ) | $spl74),
  inference(resolution,[],[f15763,f779])).
fof(f19193,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | ~p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X91) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl74),
  inference(resolution,[],[f15763,f778])).
fof(f19192,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | s(bool,X86) != s(bool,X88) | s(bool,X87) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl74),
  inference(resolution,[],[f15763,f777])).
fof(f19191,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X84) != s(bool,X85)) ) | $spl74),
  inference(resolution,[],[f15763,f775])).
fof(f19189,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl74),
  inference(resolution,[],[f15763,f773])).
fof(f18724,plain,(
  ( ! [X19] : (~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl84),
  inference(cnf_transformation,[],[f18724_D])).
fof(f18724_D,plain,(
  ( ! [X19] : (~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f18666,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | ~p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl63),
  inference(resolution,[],[f15084,f952])).
fof(f18573,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X39) != s(bool,X41)) ) | $spl63),
  inference(resolution,[],[f15084,f779])).
fof(f18572,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl63),
  inference(resolution,[],[f15084,f778])).
fof(f18570,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl63),
  inference(resolution,[],[f15084,f777])).
fof(f18568,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X25) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X24) != s(bool,X26)) ) | $spl63),
  inference(resolution,[],[f15084,f775])).
fof(f18566,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl63),
  inference(resolution,[],[f15084,f773])).
fof(f18558,plain,(
  ( ! [X25] : (p(s(bool,X25)) | s(bool,X25) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl82),
  inference(cnf_transformation,[],[f18558_D])).
fof(f18558_D,plain,(
  ( ! [X25] : (p(s(bool,X25)) | s(bool,X25) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f18497,plain,(
  ( ! [X212,X213] : (p(s(bool,X212)) | p(s(bool,X213)) | s(bool,X213) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl62),
  inference(resolution,[],[f15085,f952])).
fof(f18407,plain,(
  ( ! [X47,X45,X46] : (p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X45) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X46) != s(bool,X47)) ) | $spl62),
  inference(resolution,[],[f15085,f779])).
fof(f18404,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | s(bool,X38) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl62),
  inference(resolution,[],[f15085,f778])).
fof(f18403,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X33) != s(bool,X35) | s(bool,X34) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl62),
  inference(resolution,[],[f15085,f777])).
fof(f18402,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X30) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X31) != s(bool,X32)) ) | $spl62),
  inference(resolution,[],[f15085,f775])).
fof(f18400,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl62),
  inference(resolution,[],[f15085,f773])).
fof(f18188,plain,(
  ( ! [X265,X266] : (p(s(bool,X265)) | p(s(bool,X266)) | s(bool,X266) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl34)),
  inference(resolution,[],[f12772,f952])).
fof(f18098,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X99) != s(bool,X100)) ) | ($spl0 | $spl34)),
  inference(resolution,[],[f12772,f779])).
fof(f18095,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | ~p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X91) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl34)),
  inference(resolution,[],[f12772,f778])).
fof(f18094,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | s(bool,X86) != s(bool,X88) | s(bool,X87) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl34)),
  inference(resolution,[],[f12772,f777])).
fof(f18093,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X84) != s(bool,X85)) ) | ($spl0 | $spl34)),
  inference(resolution,[],[f12772,f775])).
fof(f18091,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | ($spl0 | $spl34)),
  inference(resolution,[],[f12772,f773])).
fof(f16864,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13737,f952])).
fof(f16774,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X66) != s(bool,X67)) ) | $spl0),
  inference(resolution,[],[f13737,f779])).
fof(f16771,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13737,f778])).
fof(f16770,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13737,f777])).
fof(f16769,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X51) != s(bool,X52)) ) | $spl0),
  inference(resolution,[],[f13737,f775])).
fof(f16767,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13737,f773])).
fof(f15765,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) | $spl76),
  inference(cnf_transformation,[],[f15765_D])).
fof(f15765_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f15763,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) | $spl74),
  inference(cnf_transformation,[],[f15763_D])).
fof(f15763_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f16756,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) | $spl0),
  inference(resolution,[],[f13737,f944])).
fof(f16296,plain,(
  $false | ($spl0 | $spl68)),
  inference(subsumption_resolution,[],[f16295,f12680])).
fof(f16295,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) | $spl68),
  inference(subsumption_resolution,[],[f15779,f3795])).
fof(f3795,plain,(
  ~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK24))))),
  inference(resolution,[],[f586,f584])).
fof(f15779,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK24)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) | $spl68),
  inference(superposition,[],[f598,f15757])).
fof(f15695,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13737,f952])).
fof(f15605,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X66) != s(bool,X67)) ) | $spl0),
  inference(resolution,[],[f13737,f779])).
fof(f15602,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13737,f778])).
fof(f15601,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13737,f777])).
fof(f15600,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X51) != s(bool,X52)) ) | $spl0),
  inference(resolution,[],[f13737,f775])).
fof(f15598,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f13737,f773])).
fof(f15587,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) | $spl0),
  inference(resolution,[],[f13737,f944])).
fof(f15757,plain,(
  s(sK22,sK24) = s(sK22,sK40(sK24,sK24,sK23,sK22)) | $spl68),
  inference(cnf_transformation,[],[f15757_D])).
fof(f15757_D,plain,(
  s(sK22,sK24) = s(sK22,sK40(sK24,sK24,sK23,sK22)) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f15586,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)),s(sK22,sK24)))) | s(sK22,sK24) = s(sK22,sK40(sK24,sK24,sK23,sK22)) | $spl0),
  inference(resolution,[],[f13737,f941])).
fof(f15761,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | $spl72),
  inference(cnf_transformation,[],[f15761_D])).
fof(f15761_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f15574,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK24)))) | s(sK22,sK24) = s(sK22,sK40(sK24,sK24,sK23,sK22)) | $spl0),
  inference(resolution,[],[f13737,f619])).
fof(f15759,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)))),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22))))) | $spl70),
  inference(cnf_transformation,[],[f15759_D])).
fof(f15759_D,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)))),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22))))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f15573,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK24,sK24,sK23,sK22)))),s(sK22,sK52(sK24,sK40(sK24,sK24,sK23,sK22),sK23,sK22))))) | s(sK22,sK24) = s(sK22,sK40(sK24,sK24,sK23,sK22)) | $spl0),
  inference(resolution,[],[f13737,f618])).
fof(f15020,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12683,f952])).
fof(f14930,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X66) != s(bool,X67)) ) | $spl0),
  inference(resolution,[],[f12683,f779])).
fof(f14927,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12683,f778])).
fof(f14926,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12683,f777])).
fof(f14925,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24))) | s(bool,X51) != s(bool,X52)) ) | $spl0),
  inference(resolution,[],[f12683,f775])).
fof(f14923,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12683,f773])).
fof(f15085,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) | $spl62),
  inference(cnf_transformation,[],[f15085_D])).
fof(f15085_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f15087,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) | $spl64),
  inference(cnf_transformation,[],[f15087_D])).
fof(f15087_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f14912,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK24)))) | $spl0),
  inference(resolution,[],[f12683,f944])).
fof(f14239,plain,(
  ( ! [X36,X34] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X34)))) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X34),s(sK22,sK24)))) ) | $spl60),
  inference(cnf_transformation,[],[f14239_D])).
fof(f14239_D,plain,(
  ( ! [X36,X34] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X34)))) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X34),s(sK22,sK24)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f14175,plain,(
  ( ! [X329,X328,X330] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X328)))) | p(s(bool,X329)) | ~p(s(bool,X330)) | s(bool,X330) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X328),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12682,f952])).
fof(f12682,plain,(
  ( ! [X4] : (~p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X4),s(sK22,sK24)))) | ~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X4))))) ) | $spl0),
  inference(backward_demodulation,[],[f1185,f586])).
fof(f14082,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X62)))) | ~p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X62),s(sK22,sK24))) | s(bool,X63) != s(bool,X65)) ) | $spl0),
  inference(resolution,[],[f12682,f779])).
fof(f14081,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X60) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X58),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12682,f778])).
fof(f14079,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X50)))) | p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,X53) | s(bool,X52) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X50),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12682,f777])).
fof(f14077,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X42)))) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X42),s(sK22,sK24))) | s(bool,X43) != s(bool,X45)) ) | $spl0),
  inference(resolution,[],[f12682,f775])).
fof(f14075,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X34)))) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X34),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12682,f773])).
fof(f2847,plain,(
  ( ! [X78] : (s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X78))) ) | $spl24),
  inference(cnf_transformation,[],[f2847_D])).
fof(f2847_D,plain,(
  ( ! [X78] : (s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | p(s(bool,X78))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f13897,plain,(
  ( ! [X265,X266] : (p(s(bool,X265)) | p(s(bool,X266)) | s(bool,X266) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12680,f952])).
fof(f13807,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | s(bool,X99) != s(bool,X100)) ) | $spl0),
  inference(resolution,[],[f12680,f779])).
fof(f13804,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | ~p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X91) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12680,f778])).
fof(f13803,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | s(bool,X86) != s(bool,X88) | s(bool,X87) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12680,f777])).
fof(f13802,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | s(bool,X84) != s(bool,X85)) ) | $spl0),
  inference(resolution,[],[f12680,f775])).
fof(f13800,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f12680,f773])).
fof(f1185,plain,(
  s(sK22,sK24) = s(sK22,sK25) | $spl0),
  inference(cnf_transformation,[],[f1185_D])).
fof(f1185_D,plain,(
  s(sK22,sK24) = s(sK22,sK25) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1676,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK52(sK25,sK24,sK23,sK22))))) | $spl18),
  inference(cnf_transformation,[],[f1676_D])).
fof(f1676_D,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK52(sK25,sK24,sK23,sK22))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3089,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK52(sK25,sK24,sK23,sK22))))) | s(sK22,sK24) = s(sK22,sK25)),
  inference(resolution,[],[f1016,f618])).
fof(f1678,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | $spl20),
  inference(cnf_transformation,[],[f1678_D])).
fof(f1678_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f3090,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | s(sK22,sK24) = s(sK22,sK25)),
  inference(resolution,[],[f1016,f619])).
fof(f12672,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl58),
  inference(cnf_transformation,[],[f12672_D])).
fof(f12672_D,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f12608,plain,(
  ( ! [X265,X266] : (p(s(bool,X265)) | p(s(bool,X266)) | s(bool,X266) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl46),
  inference(resolution,[],[f9263,f952])).
fof(f12518,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25))) | s(bool,X99) != s(bool,X100)) ) | $spl46),
  inference(resolution,[],[f9263,f779])).
fof(f12515,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | ~p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X91) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl46),
  inference(resolution,[],[f9263,f778])).
fof(f12514,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | s(bool,X86) != s(bool,X88) | s(bool,X87) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl46),
  inference(resolution,[],[f9263,f777])).
fof(f12513,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25))) | s(bool,X84) != s(bool,X85)) ) | $spl46),
  inference(resolution,[],[f9263,f775])).
fof(f12511,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl46),
  inference(resolution,[],[f9263,f773])).
fof(f12442,plain,(
  ( ! [X26,X24] : (s(bool,X24) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | s(bool,X24) != s(bool,X26) | ~p(s(bool,X26))) ) | $spl56),
  inference(cnf_transformation,[],[f12442_D])).
fof(f12442_D,plain,(
  ( ! [X26,X24] : (s(bool,X24) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | s(bool,X24) != s(bool,X26) | ~p(s(bool,X26))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f12158,plain,(
  ( ! [X318,X320,X319] : (s(bool,X318) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | p(s(bool,X319)) | ~p(s(bool,X320)) | s(bool,X318) != s(bool,X320)) )),
  inference(resolution,[],[f1119,f952])).
fof(f12065,plain,(
  ( ! [X54,X52,X55,X53] : (s(bool,X52) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | ~p(s(bool,X53)) | ~p(s(bool,X54)) | s(bool,X52) != s(bool,X54) | s(bool,X53) != s(bool,X55)) )),
  inference(resolution,[],[f1119,f779])).
fof(f12064,plain,(
  ( ! [X50,X48,X51,X49] : (s(bool,X48) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | p(s(bool,X49)) | ~p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X48) != s(bool,X50)) )),
  inference(resolution,[],[f1119,f778])).
fof(f12062,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X40) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X43) | s(bool,X40) != s(bool,X42)) )),
  inference(resolution,[],[f1119,f777])).
fof(f12060,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X32) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,X35)) )),
  inference(resolution,[],[f1119,f775])).
fof(f12058,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) = s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X24) != s(bool,X26)) )),
  inference(resolution,[],[f1119,f773])).
fof(f11494,plain,(
  ( ! [X108] : (~p(s(bool,X108)) | s(bool,X108) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK24)))) ) | $spl54),
  inference(cnf_transformation,[],[f11494_D])).
fof(f11494_D,plain,(
  ( ! [X108] : (~p(s(bool,X108)) | s(bool,X108) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK24)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f11415,plain,(
  ( ! [X302,X301] : (p(s(bool,X301)) | ~p(s(bool,X302)) | s(bool,X302) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK24)))) )),
  inference(resolution,[],[f3795,f952])).
fof(f11322,plain,(
  ( ! [X130,X128,X129] : (~p(s(bool,X128)) | ~p(s(bool,X129)) | s(bool,X129) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK24))) | s(bool,X128) != s(bool,X130)) )),
  inference(resolution,[],[f3795,f779])).
fof(f11321,plain,(
  ( ! [X127,X125,X126] : (p(s(bool,X125)) | ~p(s(bool,X126)) | ~p(s(bool,X127)) | s(bool,X126) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK24)))) )),
  inference(resolution,[],[f3795,f778])).
fof(f11319,plain,(
  ( ! [X121,X120,X119] : (p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK24)))) )),
  inference(resolution,[],[f3795,f777])).
fof(f11317,plain,(
  ( ! [X114,X115,X113] : (~p(s(bool,X113)) | ~p(s(bool,X114)) | s(bool,X114) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK24))) | s(bool,X113) != s(bool,X115)) )),
  inference(resolution,[],[f3795,f775])).
fof(f11315,plain,(
  ( ! [X109,X107,X108] : (p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X107) != s(bool,X109) | s(bool,X108) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK24)))) )),
  inference(resolution,[],[f3795,f773])).
fof(f11201,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl52),
  inference(cnf_transformation,[],[f11201_D])).
fof(f11201_D,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f11138,plain,(
  ( ! [X265,X266] : (p(s(bool,X265)) | p(s(bool,X266)) | s(bool,X266) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl34),
  inference(resolution,[],[f6334,f952])).
fof(f11048,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25))) | s(bool,X99) != s(bool,X100)) ) | $spl34),
  inference(resolution,[],[f6334,f779])).
fof(f11045,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | ~p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X91) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl34),
  inference(resolution,[],[f6334,f778])).
fof(f11044,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | s(bool,X86) != s(bool,X88) | s(bool,X87) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl34),
  inference(resolution,[],[f6334,f777])).
fof(f11043,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25))) | s(bool,X84) != s(bool,X85)) ) | $spl34),
  inference(resolution,[],[f6334,f775])).
fof(f11041,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl34),
  inference(resolution,[],[f6334,f773])).
fof(f9266,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl50),
  inference(cnf_transformation,[],[f9266_D])).
fof(f9266_D,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f10415,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl20),
  inference(resolution,[],[f1678,f952])).
fof(f10325,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25))) | s(bool,X66) != s(bool,X67)) ) | $spl20),
  inference(resolution,[],[f1678,f779])).
fof(f10322,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl20),
  inference(resolution,[],[f1678,f778])).
fof(f10321,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl20),
  inference(resolution,[],[f1678,f777])).
fof(f10320,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25))) | s(bool,X51) != s(bool,X52)) ) | $spl20),
  inference(resolution,[],[f1678,f775])).
fof(f10318,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl20),
  inference(resolution,[],[f1678,f773])).
fof(f9265,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | $spl48),
  inference(cnf_transformation,[],[f9265_D])).
fof(f9265_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f9263,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | $spl46),
  inference(cnf_transformation,[],[f9263_D])).
fof(f9263_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f10307,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | $spl20),
  inference(resolution,[],[f1678,f944])).
fof(f9796,plain,(
  $false | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f9795,f1016])).
fof(f9795,plain,(
  ~p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) | ($spl1 | $spl40)),
  inference(subsumption_resolution,[],[f9794,f1184])).
fof(f1184,plain,(
  s(sK22,sK24) != s(sK22,sK25) | $spl1),
  inference(cnf_transformation,[],[f1184_D])).
fof(f1184_D,plain,(
  s(sK22,sK24) != s(sK22,sK25) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f9794,plain,(
  s(sK22,sK24) = s(sK22,sK25) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) | $spl40),
  inference(subsumption_resolution,[],[f9279,f585])).
fof(f9279,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25)))) | s(sK22,sK24) = s(sK22,sK25) | ~p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) | $spl40),
  inference(superposition,[],[f618,f9257])).
fof(f9195,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl20),
  inference(resolution,[],[f1678,f952])).
fof(f9105,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25))) | s(bool,X66) != s(bool,X67)) ) | $spl20),
  inference(resolution,[],[f1678,f779])).
fof(f9102,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl20),
  inference(resolution,[],[f1678,f778])).
fof(f9101,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl20),
  inference(resolution,[],[f1678,f777])).
fof(f9100,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25))) | s(bool,X51) != s(bool,X52)) ) | $spl20),
  inference(resolution,[],[f1678,f775])).
fof(f9098,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl20),
  inference(resolution,[],[f1678,f773])).
fof(f9087,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | $spl20),
  inference(resolution,[],[f1678,f944])).
fof(f9257,plain,(
  s(sK22,sK25) = s(sK22,sK52(sK25,sK24,sK23,sK22)) | $spl40),
  inference(cnf_transformation,[],[f9257_D])).
fof(f9257_D,plain,(
  s(sK22,sK25) = s(sK22,sK52(sK25,sK24,sK23,sK22)) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f9086,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | s(sK22,sK25) = s(sK22,sK52(sK25,sK24,sK23,sK22)) | $spl20),
  inference(resolution,[],[f1678,f941])).
fof(f9261,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK52(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK25)))) | $spl44),
  inference(cnf_transformation,[],[f9261_D])).
fof(f9261_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK52(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK25)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f9074,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK52(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK25)))) | s(sK22,sK25) = s(sK22,sK52(sK25,sK24,sK23,sK22)) | $spl20),
  inference(resolution,[],[f1678,f619])).
fof(f9259,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)))),s(sK22,sK52(sK25,sK52(sK25,sK24,sK23,sK22),sK23,sK22))))) | $spl42),
  inference(cnf_transformation,[],[f9259_D])).
fof(f9259_D,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)))),s(sK22,sK52(sK25,sK52(sK25,sK24,sK23,sK22),sK23,sK22))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f9073,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)))),s(sK22,sK52(sK25,sK52(sK25,sK24,sK23,sK22),sK23,sK22))))) | s(sK22,sK25) = s(sK22,sK52(sK25,sK24,sK23,sK22)) | $spl20),
  inference(resolution,[],[f1678,f618])).
fof(f6337,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) | $spl38),
  inference(cnf_transformation,[],[f6337_D])).
fof(f6337_D,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f7473,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) )),
  inference(resolution,[],[f963,f952])).
fof(f7383,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25))) | s(bool,X66) != s(bool,X67)) )),
  inference(resolution,[],[f963,f779])).
fof(f7380,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) )),
  inference(resolution,[],[f963,f778])).
fof(f7379,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) )),
  inference(resolution,[],[f963,f777])).
fof(f7378,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25))) | s(bool,X51) != s(bool,X52)) )),
  inference(resolution,[],[f963,f775])).
fof(f7376,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) )),
  inference(resolution,[],[f963,f773])).
fof(f6336,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | $spl36),
  inference(cnf_transformation,[],[f6336_D])).
fof(f6336_D,plain,(
  p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f6334,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | $spl34),
  inference(cnf_transformation,[],[f6334_D])).
fof(f6334_D,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f7365,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25))))),
  inference(resolution,[],[f963,f944])).
fof(f6866,plain,(
  $false | $spl28),
  inference(subsumption_resolution,[],[f6865,f584])).
fof(f6865,plain,(
  ~p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) | $spl28),
  inference(subsumption_resolution,[],[f6350,f585])).
fof(f6350,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25)))) | ~p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) | $spl28),
  inference(superposition,[],[f598,f6328])).
fof(f6266,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) )),
  inference(resolution,[],[f963,f952])).
fof(f6176,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25))) | s(bool,X66) != s(bool,X67)) )),
  inference(resolution,[],[f963,f779])).
fof(f6173,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) )),
  inference(resolution,[],[f963,f778])).
fof(f6172,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) )),
  inference(resolution,[],[f963,f777])).
fof(f6171,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25))) | s(bool,X51) != s(bool,X52)) )),
  inference(resolution,[],[f963,f775])).
fof(f6169,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) )),
  inference(resolution,[],[f963,f773])).
fof(f6158,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | p(s(bool,'const.relation.RC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25))))),
  inference(resolution,[],[f963,f944])).
fof(f6328,plain,(
  s(sK22,sK25) = s(sK22,sK40(sK25,sK24,sK23,sK22)) | $spl28),
  inference(cnf_transformation,[],[f6328_D])).
fof(f6328_D,plain,(
  s(sK22,sK25) = s(sK22,sK40(sK25,sK24,sK23,sK22)) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f6157,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | s(sK22,sK25) = s(sK22,sK40(sK25,sK24,sK23,sK22))),
  inference(resolution,[],[f963,f941])).
fof(f6332,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK25)))) | $spl32),
  inference(cnf_transformation,[],[f6332_D])).
fof(f6332_D,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK25)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f6145,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22)),s(sK22,sK25)))) | s(sK22,sK25) = s(sK22,sK40(sK25,sK24,sK23,sK22))),
  inference(resolution,[],[f963,f619])).
fof(f6330,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)))),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22))))) | $spl30),
  inference(cnf_transformation,[],[f6330_D])).
fof(f6330_D,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)))),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f6144,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK40(sK25,sK24,sK23,sK22)))),s(sK22,sK52(sK25,sK40(sK25,sK24,sK23,sK22),sK23,sK22))))) | s(sK22,sK25) = s(sK22,sK40(sK25,sK24,sK23,sK22))),
  inference(resolution,[],[f963,f618])).
fof(f5344,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5219,f4872])).
fof(f4872,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl10),
  inference(equality_resolution,[],[f1192])).
fof(f5219,plain,(
  ( ! [X1402,X1404,X1406,X1403,X1405] : (p(s(bool,happ(s(fun(X1402,bool),happ(s(fun(X1402,fun(X1402,bool)),sK51(X1403,X1404,X1405,X1402)),s(X1402,X1406))),s(X1402,X1406))))) ) | $spl10),
  inference(resolution,[],[f4872,f615])).
fof(f5343,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f5172,f4872])).
fof(f5172,plain,(
  ( ! [X1217,X1214,X1216,X1215] : (p(s(bool,'const.relation.RTC_3'(s(fun(X1214,fun(X1214,bool)),X1215),s(X1214,X1216),s(X1214,X1217))))) ) | $spl10),
  inference(resolution,[],[f4872,f615])).
fof(f5005,plain,(
  $false | $spl10),
  inference(resolution,[],[f4872,f584])).
fof(f5228,plain,(
  $false | $spl10),
  inference(resolution,[],[f4872,f951])).
fof(f5232,plain,(
  $false | $spl10),
  inference(resolution,[],[f4872,f1016])).
fof(f4817,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f4594,f4420])).
fof(f4420,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl8),
  inference(equality_resolution,[],[f1191])).
fof(f4594,plain,(
  ( ! [X616] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X616))))) ) | $spl8),
  inference(resolution,[],[f4420,f586])).
fof(f4732,plain,(
  $false | $spl8),
  inference(resolution,[],[f4420,f585])).
fof(f4365,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f4142,f1194])).
fof(f4142,plain,(
  ( ! [X616] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X616))))) ) | $spl14),
  inference(resolution,[],[f1194,f586])).
fof(f4280,plain,(
  $false | $spl14),
  inference(resolution,[],[f1194,f585])).
fof(f3972,plain,(
  ( ! [X36,X34] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X34)))) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X34),s(sK22,sK25)))) ) | $spl26),
  inference(cnf_transformation,[],[f3972_D])).
fof(f3972_D,plain,(
  ( ! [X36,X34] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X34)))) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X34),s(sK22,sK25)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3909,plain,(
  ( ! [X329,X328,X330] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X328)))) | p(s(bool,X329)) | ~p(s(bool,X330)) | s(bool,X330) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X328),s(sK22,sK25)))) )),
  inference(resolution,[],[f586,f952])).
fof(f3816,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X62)))) | ~p(s(bool,X63)) | ~p(s(bool,X64)) | s(bool,X64) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X62),s(sK22,sK25))) | s(bool,X63) != s(bool,X65)) )),
  inference(resolution,[],[f586,f779])).
fof(f3815,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X58)))) | p(s(bool,X59)) | ~p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X60) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X58),s(sK22,sK25)))) )),
  inference(resolution,[],[f586,f778])).
fof(f3813,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X50)))) | p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,X53) | s(bool,X52) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X50),s(sK22,sK25)))) )),
  inference(resolution,[],[f586,f777])).
fof(f3811,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X42)))) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X44) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X42),s(sK22,sK25))) | s(bool,X43) != s(bool,X45)) )),
  inference(resolution,[],[f586,f775])).
fof(f3809,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,X34)))) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,X37) | s(bool,X36) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,X34),s(sK22,sK25)))) )),
  inference(resolution,[],[f586,f773])).
fof(f1467,plain,(
  ( ! [X108] : (~p(s(bool,X108)) | s(bool,X108) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25)))) ) | $spl16),
  inference(cnf_transformation,[],[f1467_D])).
fof(f1467_D,plain,(
  ( ! [X108] : (~p(s(bool,X108)) | s(bool,X108) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3542,plain,(
  ( ! [X302,X301] : (p(s(bool,X301)) | ~p(s(bool,X302)) | s(bool,X302) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25)))) )),
  inference(resolution,[],[f585,f952])).
fof(f3449,plain,(
  ( ! [X130,X128,X129] : (~p(s(bool,X128)) | ~p(s(bool,X129)) | s(bool,X129) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25))) | s(bool,X128) != s(bool,X130)) )),
  inference(resolution,[],[f585,f779])).
fof(f3448,plain,(
  ( ! [X127,X125,X126] : (p(s(bool,X125)) | ~p(s(bool,X126)) | ~p(s(bool,X127)) | s(bool,X126) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25)))) )),
  inference(resolution,[],[f585,f778])).
fof(f3446,plain,(
  ( ! [X121,X120,X119] : (p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25)))) )),
  inference(resolution,[],[f585,f777])).
fof(f3444,plain,(
  ( ! [X114,X115,X113] : (~p(s(bool,X113)) | ~p(s(bool,X114)) | s(bool,X114) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25))) | s(bool,X113) != s(bool,X115)) )),
  inference(resolution,[],[f585,f775])).
fof(f3442,plain,(
  ( ! [X109,X107,X108] : (p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X107) != s(bool,X109) | s(bool,X108) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25)))) )),
  inference(resolution,[],[f585,f773])).
fof(f1679,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) ) | $spl22),
  inference(cnf_transformation,[],[f1679_D])).
fof(f1679_D,plain,(
  ( ! [X45] : (p(s(bool,X45)) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3211,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f1016,f952])).
fof(f3121,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | s(bool,X66) != s(bool,X67)) )),
  inference(resolution,[],[f1016,f779])).
fof(f3118,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f1016,f778])).
fof(f3117,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f1016,f777])).
fof(f3116,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | s(bool,X51) != s(bool,X52)) )),
  inference(resolution,[],[f1016,f775])).
fof(f3114,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f1016,f773])).
fof(f1190,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) ) | $spl6),
  inference(cnf_transformation,[],[f1190_D])).
fof(f1190_D,plain,(
  ( ! [X78] : (p(s(bool,X78)) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f3011,plain,(
  ( ! [X265,X266] : (p(s(bool,X265)) | p(s(bool,X266)) | s(bool,X266) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f584,f952])).
fof(f2921,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | s(bool,X99) != s(bool,X100)) )),
  inference(resolution,[],[f584,f779])).
fof(f2918,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | ~p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X91) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f584,f778])).
fof(f2917,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | s(bool,X86) != s(bool,X88) | s(bool,X87) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f584,f777])).
fof(f2916,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | s(bool,X84) != s(bool,X85)) )),
  inference(resolution,[],[f584,f775])).
fof(f2914,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f584,f773])).
fof(f1187,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) | $spl2),
  inference(cnf_transformation,[],[f1187_D])).
fof(f1187_D,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f2853,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) | s(sK22,sK24) = s(sK22,sK25)),
  inference(resolution,[],[f584,f622])).
fof(f622,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,sK53(X3,X2,X1,X0))))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f411])).
fof(f411,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,sK53(X3,X2,X1,X0))))) & s(X0,X2) != s(X0,sK53(X3,X2,X1,X0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK53])],[f288])).
fof(f288,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & s(X0,X2) != s(X0,X4)))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3)) | ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & s(X0,X2) != s(X0,X4)))),
  inference(ennf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) & s(X0,X2) != s(X0,X3)) => ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & s(X0,X2) != s(X0,X4)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0] : ! [X1,X2,X3] : ((p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3)))) & ~s(X0,X2) = s(X0,X3)) => ? [X4] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X4)))) & ~s(X0,X2) = s(X0,X4)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6] : ! [X7,X8,X10] : ((p(s(bool,'const.relation.TC_3'(s(fun(X6,fun(X6,bool)),X7),s(X6,X8),s(X6,X10)))) & ~s(X6,X8) = s(X6,X10)) => ? [X9] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X8))),s(X6,X9)))) & ~s(X6,X8) = s(X6,X9)))),
  file('thm.relation.TC_CASES|49|_E','thm.relation.TC_implies_one_step')).
fof(f2796,plain,(
  ( ! [X265,X266] : (p(s(bool,X265)) | p(s(bool,X266)) | s(bool,X266) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f1682,f952])).
fof(f1682,plain,(
  p(s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) | $spl0),
  inference(backward_demodulation,[],[f1185,f584])).
fof(f2706,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | s(bool,X99) != s(bool,X100)) ) | $spl0),
  inference(resolution,[],[f1682,f779])).
fof(f2703,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | ~p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X91) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f1682,f778])).
fof(f2702,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | s(bool,X86) != s(bool,X88) | s(bool,X87) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f1682,f777])).
fof(f2701,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24))) | s(bool,X84) != s(bool,X85)) ) | $spl0),
  inference(resolution,[],[f1682,f775])).
fof(f2699,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK24)))) ) | $spl0),
  inference(resolution,[],[f1682,f773])).
fof(f1614,plain,(
  ( ! [X233,X232] : (p(s(bool,X232)) | p(s(bool,X233)) | s(bool,X233) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f1016,f952])).
fof(f1524,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X65)) | ~p(s(bool,X66)) | s(bool,X65) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | s(bool,X66) != s(bool,X67)) )),
  inference(resolution,[],[f1016,f779])).
fof(f1521,plain,(
  ( ! [X57,X58,X56] : (p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | s(bool,X58) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f1016,f778])).
fof(f1520,plain,(
  ( ! [X54,X55,X53] : (p(s(bool,X53)) | p(s(bool,X54)) | s(bool,X53) != s(bool,X55) | s(bool,X54) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f1016,f777])).
fof(f1519,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X50)) | ~p(s(bool,X51)) | s(bool,X50) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | s(bool,X51) != s(bool,X52)) )),
  inference(resolution,[],[f1016,f775])).
fof(f1517,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X44) != s(bool,X46) | s(bool,X45) != s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f1016,f773])).
fof(f1493,plain,(
  p(s(bool,'const.relation.RTC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK52(sK25,sK24,sK23,sK22)),s(sK22,sK25)))) | s(sK22,sK24) = s(sK22,sK25)),
  inference(resolution,[],[f1016,f619])).
fof(f1492,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK52(sK25,sK24,sK23,sK22))))) | s(sK22,sK24) = s(sK22,sK25)),
  inference(resolution,[],[f1016,f618])).
fof(f1391,plain,(
  ( ! [X302,X301] : (p(s(bool,X301)) | ~p(s(bool,X302)) | s(bool,X302) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25)))) )),
  inference(resolution,[],[f585,f952])).
fof(f1298,plain,(
  ( ! [X130,X128,X129] : (~p(s(bool,X128)) | ~p(s(bool,X129)) | s(bool,X129) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25))) | s(bool,X128) != s(bool,X130)) )),
  inference(resolution,[],[f585,f779])).
fof(f1297,plain,(
  ( ! [X127,X125,X126] : (p(s(bool,X125)) | ~p(s(bool,X126)) | ~p(s(bool,X127)) | s(bool,X126) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25)))) )),
  inference(resolution,[],[f585,f778])).
fof(f1295,plain,(
  ( ! [X121,X120,X119] : (p(s(bool,X119)) | ~p(s(bool,X120)) | s(bool,X119) != s(bool,X121) | s(bool,X120) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25)))) )),
  inference(resolution,[],[f585,f777])).
fof(f1293,plain,(
  ( ! [X114,X115,X113] : (~p(s(bool,X113)) | ~p(s(bool,X114)) | s(bool,X114) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25))) | s(bool,X113) != s(bool,X115)) )),
  inference(resolution,[],[f585,f775])).
fof(f1291,plain,(
  ( ! [X109,X107,X108] : (p(s(bool,X107)) | ~p(s(bool,X108)) | s(bool,X107) != s(bool,X109) | s(bool,X108) != s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK25)))) )),
  inference(resolution,[],[f585,f773])).
fof(f1123,plain,(
  ( ! [X265,X266] : (p(s(bool,X265)) | p(s(bool,X266)) | s(bool,X266) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f584,f952])).
fof(f1033,plain,(
  ( ! [X99,X100,X98] : (p(s(bool,X98)) | ~p(s(bool,X99)) | s(bool,X98) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | s(bool,X99) != s(bool,X100)) )),
  inference(resolution,[],[f584,f779])).
fof(f1030,plain,(
  ( ! [X90,X91,X89] : (p(s(bool,X89)) | ~p(s(bool,X90)) | p(s(bool,X91)) | s(bool,X91) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f584,f778])).
fof(f1029,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | s(bool,X86) != s(bool,X88) | s(bool,X87) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f584,f777])).
fof(f1028,plain,(
  ( ! [X85,X83,X84] : (p(s(bool,X83)) | ~p(s(bool,X84)) | s(bool,X83) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25))) | s(bool,X84) != s(bool,X85)) )),
  inference(resolution,[],[f584,f775])).
fof(f1026,plain,(
  ( ! [X78,X79,X77] : (p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X77) != s(bool,X79) | s(bool,X78) != s(bool,'const.relation.TC_3'(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24),s(sK22,sK25)))) )),
  inference(resolution,[],[f584,f773])).
fof(f1189,plain,(
  s(sK22,sK24) != s(sK22,sK53(sK25,sK24,sK23,sK22)) | $spl5),
  inference(cnf_transformation,[],[f1189_D])).
fof(f1189_D,plain,(
  s(sK22,sK24) != s(sK22,sK53(sK25,sK24,sK23,sK22)) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f966,plain,(
  s(sK22,sK24) != s(sK22,sK53(sK25,sK24,sK23,sK22)) | s(sK22,sK24) = s(sK22,sK25)),
  inference(resolution,[],[f584,f623])).
fof(f623,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) != s(X0,sK53(X3,X2,X1,X0)) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.relation.TC_3'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f411])).
fof(f965,plain,(
  p(s(bool,happ(s(fun(sK22,bool),happ(s(fun(sK22,fun(sK22,bool)),sK23),s(sK22,sK24))),s(sK22,sK53(sK25,sK24,sK23,sK22))))) | s(sK22,sK24) = s(sK22,sK25)),
  inference(resolution,[],[f584,f622])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1090
% Passive clauses: 31312
% Generated clauses: 51148
% Final active clauses: 952
% Final passive clauses: 27331
% Input formulas: 133
% Initial clauses: 368
% 
% Selected by SInE selection: 76
% SInE iterations: 1
% 
% Duplicate literals: 884
% Fw subsumption resolutions: 749
% Fw demodulations: 5075
% Bw demodulations: 105
% 
% Simple tautologies: 1871
% Equational tautologies: 32
% Forward subsumptions: 11556
% 
% Binary resolution: 16765
% Forward superposition: 22291
% Backward superposition: 5184
% Equality resolution: 5
% 
% Split clauses: 430
% Split components: 933
% SAT solver clauses: 457
% SAT solver unit clauses: 14
% SAT solver binary clauses: 91
% 
% Sat splits: 430
% Sat splitting refutations: 24
% 
% Memory used [KB]: 25330
% Time elapsed: 1.649 s
% ------------------------------
% ------------------------------
% Success in time 13.391 s
