Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.extreal
% (79244)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 500
% Passive clauses: 3844
% Generated clauses: 20344
% Final active clauses: 470
% Final passive clauses: 3099
% Input formulas: 69
% Initial clauses: 168
% 
% Duplicate literals: 104
% Trivial inequalities: 14
% Fw subsumption resolutions: 701
% Fw demodulations: 2394
% Bw demodulations: 28
% 
% Simple tautologies: 151
% Equational tautologies: 18
% Forward subsumptions: 13087
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 4078
% Factoring: 42
% Forward superposition: 6136
% Backward superposition: 6844
% Self superposition: 119
% Equality factoring: 8
% Equality resolution: 4
% 
% Split clauses: 136
% Split components: 272
% SAT solver clauses: 241
% SAT solver unit clauses: 10
% SAT solver binary clauses: 77
% 
% Sat splits: 136
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 240
% TWLsolver calls for satisfiability: 124
% 
% Memory used [KB]: 21364
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f8004,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f8000,f7697,f3849,f7698,f295,f467,f331,f481,f4435,f271,f447,f452,f454,f7224,f7405,f270,f445,f336,f486,f7702,f2212,f7700,f299,f469,f2323,f355,f503])).
fof(f503,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) ) | $spl88),
  inference(cnf_transformation,[],[f503_D])).
fof(f503_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f355,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | ~p(s(bool,X0)) | p(s(bool,X1))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(flattening,[],[f233])).
fof(f233,plain,(
  ! [X0,X1] : ((~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) | (~p(s(bool,X0)) | p(s(bool,X1)))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1))))))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (~p(s(bool,X0)) | p(s(bool,X1))))),
  inference(ennf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X0))),s(bool,X1)))) <=> (p(s(bool,X0)) => p(s(bool,X1))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X36,X37] : (p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'==>_0'),s(bool,X36))),s(bool,X37)))) <=> (p(s(bool,X36)) => p(s(bool,X37))))),
  file('thm.extreal.extreal_distinct|split|2','==>_HL_0')).
fof(f2323,plain,(
  $false | ($spl9 | $spl88 | $spl760)),
  inference(subsumption_resolution,[],[f2312,f454])).
fof(f2312,plain,(
  p(s(bool,'F_0')) | ($spl88 | $spl760)),
  inference(condensation,[],[f2274])).
fof(f2274,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0))) ) | ($spl88 | $spl760)),
  inference(backward_demodulation,[],[f2212,f503])).
fof(f469,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl28),
  inference(cnf_transformation,[],[f469_D])).
fof(f469_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f299,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X27] : (s(bool,'F_0') = s(bool,X27) <=> ~p(s(bool,X27)))),
  file('thm.extreal.extreal_distinct|split|2','thm.bool.EQ_CLAUSES|split|2')).
fof(f7700,plain,(
  ( ! [X126] : (p(s(bool,happ(s(fun('type.realax.real',bool),X126),s('type.realax.real',sK9))))) ) | $spl1670),
  inference(cnf_transformation,[],[f7700_D])).
fof(f7700_D,plain,(
  ( ! [X126] : (p(s(bool,happ(s(fun('type.realax.real',bool),X126),s('type.realax.real',sK9))))) ) <=> ~$spl1670),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1670])])).
fof(f2212,plain,(
  ( ! [X31] : (s(bool,'F_0') = s(bool,X31)) ) | $spl760),
  inference(cnf_transformation,[],[f2212_D])).
fof(f2212_D,plain,(
  ( ! [X31] : (s(bool,'F_0') = s(bool,X31)) ) <=> ~$spl760),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl760])])).
fof(f7702,plain,(
  ( ! [X128,X129] : (s(bool,'F_0') = s(bool,X128) | p(s(bool,happ(s(fun('type.realax.real',bool),X129),s('type.realax.real',sK9))))) ) | ($spl28 | $spl44 | $spl1594)),
  inference(forward_demodulation,[],[f7595,f481])).
fof(f7595,plain,(
  ( ! [X127,X128,X129] : (p(s(bool,happ(s(fun('type.realax.real',bool),X129),s('type.realax.real',sK9)))) | s(bool,'F_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X127),s(bool,X128),s(fun('type.realax.real',bool),X129)))) ) | ($spl28 | $spl1594)),
  inference(superposition,[],[f469,f7405])).
fof(f486,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))) ) | $spl54),
  inference(cnf_transformation,[],[f486_D])).
fof(f486_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f336,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,happ(s(fun('type.realax.real',X0),X4),s('type.realax.real',X1))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X1))),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X4)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X16] : ! [X6,X29,X30,X17] : s(X16,happ(s(fun('type.realax.real',X16),X17),s('type.realax.real',X6))) = s(X16,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))),s(X16,X29),s(X16,X30),s(fun('type.realax.real',X16),X17)))),
  file('thm.extreal.extreal_distinct|split|2','thm.extreal.extreal_case_def|split|2')).
fof(f445,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK9))) | $spl0),
  inference(cnf_transformation,[],[f445_D])).
fof(f445_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK9))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f270,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK9)))),
  inference(cnf_transformation,[],[f200])).
fof(f200,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',sK9)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9])],[f171])).
fof(f171,plain,(
  ? [X0] : s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0] : s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6)))),
  file('thm.extreal.extreal_distinct|split|2','thm.extreal.extreal_distinct|split|2')).
fof(f7405,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X1),s('type.realax.real',sK9))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X1)))) ) | $spl1594),
  inference(cnf_transformation,[],[f7405_D])).
fof(f7405_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X1),s('type.realax.real',sK9))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X1)))) ) <=> ~$spl1594),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1594])])).
fof(f7224,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.realax.real',X0),X1),s('type.realax.real',sK9))) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X2),s(X0,X3),s(fun('type.realax.real',X0),X1)))) ) | ($spl0 | $spl54)),
  inference(superposition,[],[f486,f445])).
fof(f454,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f454_D])).
fof(f454_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f452,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f276,f271])).
fof(f276,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f30])).
fof(f30,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.extreal.extreal_distinct|split|2','thm.bool.NOT_CLAUSES|split|1')).
fof(f447,plain,(
  p(s(bool,'T_0')) | $spl2),
  inference(cnf_transformation,[],[f447_D])).
fof(f447_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f271,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.extreal_distinct|split|2',aHL_TRUTH)).
fof(f4435,plain,(
  $false | ($spl2 | $spl9 | $spl1258)),
  inference(subsumption_resolution,[],[f4180,f447])).
fof(f4180,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl1258)),
  inference(backward_demodulation,[],[f3849,f454])).
fof(f481,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) ) | $spl44),
  inference(cnf_transformation,[],[f481_D])).
fof(f481_D,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f331,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2,X3] : s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X16] : ! [X29,X30,X17] : s(X16,X30) = s(X16,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X16,X29),s(X16,X30),s(fun('type.realax.real',X16),X17)))),
  file('thm.extreal.extreal_distinct|split|2','thm.extreal.extreal_case_def|split|1')).
fof(f467,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl24),
  inference(cnf_transformation,[],[f467_D])).
fof(f467_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f295,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X27] : (s(bool,'T_0') = s(bool,X27) <=> p(s(bool,X27)))),
  file('thm.extreal.extreal_distinct|split|2','thm.bool.EQ_CLAUSES|split|0')).
fof(f7698,plain,(
  ( ! [X123] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X123),s('type.realax.real',sK9))))) ) | $spl1668),
  inference(cnf_transformation,[],[f7698_D])).
fof(f7698_D,plain,(
  ( ! [X123] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X123),s('type.realax.real',sK9))))) ) <=> ~$spl1668),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1668])])).
fof(f3849,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6)) ) | $spl1258),
  inference(cnf_transformation,[],[f3849_D])).
fof(f3849_D,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl1258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1258])])).
fof(f7697,plain,(
  ( ! [X123,X122] : (s(bool,'T_0') = s(bool,X122) | ~p(s(bool,happ(s(fun('type.realax.real',bool),X123),s('type.realax.real',sK9))))) ) | ($spl24 | $spl44 | $spl1594)),
  inference(forward_demodulation,[],[f7593,f481])).
fof(f7593,plain,(
  ( ! [X123,X121,X122] : (~p(s(bool,happ(s(fun('type.realax.real',bool),X123),s('type.realax.real',sK9)))) | s(bool,'T_0') = s(bool,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(bool,X121),s(bool,X122),s(fun('type.realax.real',bool),X123)))) ) | ($spl24 | $spl1594)),
  inference(superposition,[],[f467,f7405])).
fof(f8000,plain,(
  $false | ($spl1668 | $spl1670)),
  inference(subsumption_resolution,[],[f7698,f7700])).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 264
% Passive clauses: 1272
% Generated clauses: 8081
% Final active clauses: 208
% Final passive clauses: 497
% Input formulas: 69
% Initial clauses: 168
% 
% Duplicate literals: 62
% Trivial inequalities: 27
% Fw subsumption resolutions: 359
% Fw demodulations: 1251
% Bw demodulations: 551
% Condensations: 221
% Global subsumptions: 14
% 
% Simple tautologies: 148
% Equational tautologies: 12
% Forward subsumptions: 3445
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 1000
% Factoring: 4
% Forward superposition: 1690
% Backward superposition: 1339
% Self superposition: 59
% Equality factoring: 8
% Equality resolution: 10
% 
% Split clauses: 63
% Split components: 126
% SAT solver clauses: 1385
% SAT solver unit clauses: 219
% SAT solver binary clauses: 71
% SAT solver learnt clauses: 29
% SAT solver learnt literals: 10
% 
% Sat splits: 63
% Sat splitting refutations: 19
% 
% TWLsolver clauses: 1328
% TWLsolver calls for satisfiability: 1380
% 
% Memory used [KB]: 7675
% Time elapsed: 0.300 s
% ------------------------------
% ------------------------------
% Success in time 0.93 s
