Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.extreal
% (52301)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1797
% Passive clauses: 8661
% Generated clauses: 42417
% Final active clauses: 1629
% Final passive clauses: 6436
% Input formulas: 261
% Initial clauses: 489
% 
% Duplicate literals: 261
% Trivial inequalities: 138
% Fw subsumption resolutions: 2775
% Fw demodulations: 4255
% Bw demodulations: 117
% 
% Simple tautologies: 267
% Equational tautologies: 27
% Forward subsumptions: 25596
% 
% Binary resolution: 9436
% Factoring: 1229
% Forward superposition: 12550
% Backward superposition: 11007
% Self superposition: 39
% Equality factoring: 2
% Equality resolution: 29
% 
% Split clauses: 477
% Split components: 956
% SAT solver clauses: 825
% SAT solver unit clauses: 57
% SAT solver binary clauses: 130
% 
% Sat splits: 477
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 824
% TWLsolver calls for satisfiability: 709
% 
% Memory used [KB]: 19829
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.extreal
% (52315)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 876
% Passive clauses: 3833
% Generated clauses: 30813
% Final active clauses: 715
% Final passive clauses: 2415
% Input formulas: 261
% Initial clauses: 489
% 
% Duplicate literals: 115
% Trivial inequalities: 125
% Fw subsumption resolutions: 1323
% Fw demodulations: 4217
% Bw demodulations: 630
% Condensations: 76
% Global subsumptions: 36
% 
% Simple tautologies: 251
% Equational tautologies: 23
% Forward subsumptions: 16173
% Fw demodulations to eq. taut.: 5
% 
% Binary resolution: 3502
% Factoring: 111
% Forward superposition: 9014
% Backward superposition: 6755
% Self superposition: 72
% Equality factoring: 2
% Equality resolution: 33
% 
% Split clauses: 134
% Split components: 268
% SAT solver clauses: 4803
% SAT solver unit clauses: 679
% SAT solver binary clauses: 191
% SAT solver learnt clauses: 72
% SAT solver learnt literals: 12
% 
% Sat splits: 134
% Sat splitting refutations: 54
% 
% TWLsolver clauses: 4659
% TWLsolver calls for satisfiability: 3856
% 
% Memory used [KB]: 24050
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f45330,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5583,f5599,f5590,f5601,f5606,f16104,f16114,f16113,f20966,f20970,f20968,f24932,f24935,f24933,f24936,f24938,f24940,f24944,f24946,f24948,f24950,f24954,f24956,f24958,f24960,f25010,f25012,f25062,f25064,f25116,f25168,f25196,f25307,f25309,f25382,f25432,f25445,f25447,f25449,f25462,f26147,f26220,f26281,f26283,f26815,f26936,f26949,f26951,f27002,f27004,f27006,f27011,f27013,f27015,f27017,f27030,f27032,f27034,f27101,f27103,f27149,f27151,f27208,f27223,f27849,f27851,f27924,f27974,f27987,f27989,f27991,f28004,f28689,f28762,f28823,f28825,f29357,f29478,f29491,f29493,f29544,f29546,f29548,f29553,f29555,f29557,f29559,f29572,f29574,f29576,f29643,f29645,f29691,f29693,f29750,f29765,f30313,f30315,f30317,f30345,f30347,f30349,f30431,f30433,f30455,f30457,f30539,f30541,f30563,f30565,f30647,f30649,f30671,f30673,f30707,f30709,f30711,f30713,f32417,f32418,f32419,f32420,f32421,f32422,f32423,f32424,f32425,f32426,f30762,f30767,f30802,f30807,f30812,f30847,f30879,f30908,f30936,f30938,f30940,f30959,f30961,f30980,f30982,f31001,f31003,f31022,f31024,f31054,f31081,f31108,f31285,f31287,f31289,f31291,f31293,f31295,f31297,f31314,f31316,f31318,f31379,f31381,f31383,f31385,f31387,f31389,f31391,f31393,f31395,f31397,f31399,f31401,f31403,f31405,f31407,f31409,f31411,f31413,f31415,f31426,f31433,f31435,f31437,f31439,f31482,f31496,f31537,f31551,f31553,f31582,f31584,f31592,f31594,f31623,f31625,f31634,f31636,f31664,f31666,f31675,f31677,f31714,f31730,f31781,f31783,f31806,f31808,f31831,f31833,f31856,f31906,f31908,f31910,f31938,f31940,f31942,f31965,f31967,f31986,f31996,f32009,f32033,f32036,f32038,f32040,f32042,f32053,f32058,f32075,f32078,f32080,f32082,f32084,f32086,f32088,f32090,f32150,f32152,f32161,f32163,f32165,f32167,f32169,f32171,f32173,f32175,f32229,f32231,f32285,f32287,f32291,f32293,f32295,f32298,f32300,f32302,f32304,f32307,f32323,f32325,f32331,f32333,f32335,f32337,f32340,f32342,f32344,f32346,f32348,f32350,f32352,f32354,f32356,f32358,f32360,f32362,f32364,f32366,f32368,f32370,f32372,f32427,f32428,f32429,f32430,f32431,f32432,f32433,f32434,f32435,f32436,f32446,f32447,f32448,f32449,f32450,f32451,f32452,f32453,f32454,f32455,f32456,f32457,f32458,f32459,f32460,f32461,f32462,f32463,f32464,f32465,f32466,f32467,f32468,f32469,f32470,f32471,f32472,f32473,f32474,f32475,f32476,f32477,f32478,f32479,f32480,f32481,f32482,f32483,f32484,f32485,f32486,f32487,f32488,f32489,f32490,f32491,f32492,f32493,f32494,f32495,f32496,f32497,f32498,f32499,f32500,f32501,f32502,f32503,f32504,f32505,f32506,f32507,f32508,f32509,f32510,f32511,f32512,f32513,f32514,f32515,f32516,f32517,f32518,f32519,f32520,f32521,f32522,f32523,f32524,f32525,f32526,f32527,f32528,f32529,f32530,f32531,f32532,f32533,f32534,f32535,f32536,f32537,f32538,f32539,f32540,f32541,f32542,f32543,f32544,f32545,f32546,f32547,f32548,f32549,f32550,f32551,f32552,f32553,f32554,f32555,f32556,f32557,f32558,f32559,f32560,f32561,f32562,f32563,f32564,f32565,f32566,f32567,f32568,f32569,f32570,f32571,f32572,f32573,f32574,f32575,f32576,f32577,f32578,f32579,f32580,f32581,f32582,f32583,f32584,f32585,f32586,f32587,f32588,f32589,f32590,f32591,f32592,f32593,f32594,f32595,f32596,f32597,f32598,f32599,f32600,f32601,f32602,f32603,f32604,f32605,f32606,f32607,f32608,f32609,f32610,f32611,f32612,f32613,f32614,f32615,f32616,f32617,f32618,f32619,f32620,f32621,f32622,f32623,f32624,f32625,f32626,f32627,f32628,f32629,f32630,f32631,f32632,f32633,f32634,f32635,f32636,f32637,f32638,f32639,f32640,f32641,f32642,f32643,f32644,f32645,f32646,f32647,f32648,f32649,f32650,f32651,f32652,f32653,f32654,f32655,f32656,f32657,f32658,f32659,f32660,f32661,f32662,f32663,f32664,f32665,f32666,f32667,f32668,f32669,f32670,f32671,f32672,f32673,f32674,f32675,f32676,f32677,f32678,f32679,f32680,f32681,f32682,f32683,f32684,f32685,f32686,f32687,f32688,f32689,f32690,f32691,f32692,f32693,f32694,f32695,f32696,f32697,f32698,f32699,f32700,f32701,f32702,f32703,f32704,f32705,f32706,f32707,f32708,f32709,f32710,f32711,f32712,f32713,f32714,f32715,f32716,f32717,f32718,f32719,f32720,f32721,f32722,f32723,f32724,f32725,f32726,f32727,f32728,f32729,f32730,f32731,f32732,f32733,f32734,f32735,f32736,f32737,f32738,f32739,f32740,f32741,f32742,f32743,f32744,f32745,f32746,f32747,f32748,f32749,f32750,f32751,f32752,f32753,f32754,f32755,f32756,f32757,f32758,f32759,f32760,f32761,f32762,f32763,f32764,f32765,f32766,f32767,f32768,f32769,f32770,f32771,f32772,f32773,f32774,f32775,f32776,f32777,f32778,f32779,f32780,f32781,f32782,f32783,f32784,f32785,f32786,f32787,f32788,f32789,f32790,f32791,f32792,f32793,f32794,f32795,f32796,f32797,f32798,f32799,f32800,f32801,f32802,f32803,f32804,f32805,f32806,f32807,f32808,f32809,f32810,f32811,f32812,f32813,f32814,f32815,f32816,f32817,f32818,f32819,f32820,f32821,f32822,f32823,f32824,f32825,f32826,f32827,f32828,f32829,f32830,f32831,f32832,f32833,f32834,f32835,f32836,f32837,f32838,f32839,f32840,f32841,f32842,f32843,f32844,f32845,f32846,f32847,f32848,f32849,f32850,f32851,f32852,f32853,f32854,f32855,f32856,f32857,f32858,f32859,f32860,f32861,f32862,f32863,f32864,f32865,f32866,f32867,f32868,f32869,f32870,f32871,f32872,f32873,f32874,f32875,f32876,f32877,f32878,f32879,f32880,f32881,f32882,f32883,f32884,f32885,f32886,f32887,f32888,f32889,f32890,f32891,f32892,f32893,f32894,f32895,f32896,f32897,f32898,f32899,f32900,f32901,f32902,f32903,f32904,f32905,f32906,f32907,f32908,f32909,f32910,f32911,f32912,f32913,f32914,f32915,f32916,f32917,f32918,f32919,f32920,f32921,f32922,f32923,f32924,f32925,f32926,f32927,f32928,f32929,f32930,f32931,f32932,f32933,f32934,f32935,f32936,f32937,f32938,f32939,f32940,f32941,f32942,f32943,f32944,f32945,f32946,f32947,f32948,f32949,f32950,f32951,f32952,f32953,f32954,f32955,f32956,f32957,f32958,f32959,f32960,f32961,f32962,f32963,f32964,f32965,f32966,f32967,f32968,f32969,f32970,f32971,f32972,f32973,f32974,f32975,f32976,f32977,f32978,f32979,f32980,f32981,f32982,f32983,f32984,f32985,f32986,f32987,f32988,f32989,f32990,f32991,f32992,f32993,f32994,f32995,f32996,f32997,f32998,f32999,f33000,f33001,f33002,f33003,f33004,f33005,f33006,f33007,f33008,f33009,f33010,f33011,f33012,f33013,f33014,f33015,f33016,f33017,f33018,f33019,f33020,f33021,f33022,f33023,f33024,f33025,f33026,f33027,f33028,f33029,f33030,f33031,f33032,f33033,f33034,f33035,f33036,f33037,f33038,f33039,f33040,f33041,f33042,f33043,f33044,f33045,f33046,f33047,f33048,f33049,f33050,f33051,f33052,f33053,f33054,f33055,f33056,f33057,f33058,f33059,f33060,f33061,f33062,f33063,f33064,f33065,f33066,f33067,f33068,f33069,f33070,f33071,f33072,f33073,f33074,f33075,f33076,f33077,f33078,f33079,f33080,f33081,f33082,f33083,f33084,f33085,f33086,f33087,f33088,f33089,f33090,f33091,f33092,f33093,f33094,f33095,f33096,f33097,f33098,f33099,f33100,f33101,f33102,f33103,f33104,f33105,f33106,f33107,f33108,f33109,f33110,f33111,f33112,f33113,f33114,f33115,f33116,f33117,f33118,f33119,f33120,f33121,f33122,f33123,f33124,f33125,f33126,f33127,f33128,f33129,f33130,f33131,f33132,f33133,f33134,f33135,f33136,f33137,f33138,f33139,f33140,f33141,f33142,f33143,f33144,f33145,f33146,f33147,f33148,f33149,f33150,f33151,f33152,f33153,f33154,f33155,f33156,f33157,f33158,f33159,f33160,f33161,f33162,f33163,f33164,f33165,f33166,f33167,f33168,f33169,f33170,f33171,f33172,f33173,f33174,f33175,f33176,f33177,f33178,f33179,f33180,f33181,f33182,f33183,f33184,f33185,f33186,f33187,f33188,f33189,f33190,f33191,f33192,f33193,f33194,f33195,f33196,f33197,f33198,f33199,f33200,f33201,f33202,f33203,f33204,f33205,f33206,f33207,f33208,f33209,f33210,f33211,f33212,f33213,f33214,f33215,f33216,f33217,f33218,f33219,f33220,f33221,f33222,f33223,f33224,f33225,f33226,f33227,f33228,f33229,f33230,f33231,f33232,f33233,f33234,f33235,f33236,f33237,f33238,f33239,f33240,f33241,f33242,f33243,f33244,f33245,f33246,f33247,f33248,f33249,f33250,f33251,f33252,f33253,f33254,f33255,f33256,f33257,f33258,f33259,f33260,f33261,f33262,f33263,f33264,f33265,f33266,f33267,f33268,f33269,f33270,f33271,f33272,f33273,f33274,f33275,f33276,f33277,f33278,f33279,f33280,f33281,f33282,f33283,f33284,f33285,f33286,f33287,f33288,f33289,f33290,f33291,f33292,f33293,f33294,f33295,f33296,f33297,f33298,f33299,f33300,f33301,f33302,f33303,f33304,f33305,f33306,f33307,f33308,f33309,f33310,f33311,f33312,f33313,f33314,f33315,f33316,f33317,f33318,f33319,f33320,f33321,f33322,f33323,f33324,f33325,f33326,f33327,f33328,f33329,f33330,f33331,f33332,f33333,f33334,f33335,f33336,f33337,f33338,f33339,f33340,f33341,f33342,f33343,f33344,f33345,f33346,f33347,f33348,f33349,f33350,f33351,f33352,f33353,f33354,f33355,f33356,f33357,f33358,f33359,f33360,f33361,f33362,f33363,f33364,f33365,f33366,f33367,f33368,f33369,f33370,f33371,f33372,f33373,f33374,f33375,f33376,f33377,f33378,f33379,f33380,f33381,f33382,f33383,f33384,f33385,f33386,f33387,f33388,f33389,f33390,f33391,f33392,f33393,f33394,f33395,f33396,f33397,f33398,f33399,f33400,f33401,f33402,f33403,f33404,f33405,f33406,f33407,f33408,f33409,f33410,f33411,f33412,f33413,f33414,f33415,f33416,f33417,f33418,f33419,f33420,f33421,f33422,f33423,f33424,f33425,f33426,f33427,f33428,f33429,f33430,f33431,f33432,f33433,f33434,f33435,f33436,f33437,f33438,f33439,f33440,f33441,f33442,f33443,f33444,f33445,f33446,f33447,f44484,f44490,f44491,f44488,f44486,f44747,f44749,f44750,f44751,f44752,f44753,f44754,f44755,f44756,f44757,f44758,f44759,f44760,f44761,f44762,f44763,f44764,f44765,f44766,f44767,f44768,f44769,f44770,f44771,f44772,f44773,f44774,f44775,f44776,f44777,f44778,f44779,f45045,f45048,f45065,f45067,f45071,f45073,f45084,f45086,f45091,f45093,f45098,f45100,f45105,f45107,f45114,f45116,f45122,f45124,f45127,f45131,f45133,f45139,f45160,f45180,f45182,f45184,f45188,f45190,f45194,f45196,f45217,f45219,f45248,f45250,f45279,f45281,f45293,f45295,f45318,f45320])).
fof(f45320,plain,(
  $false | ($spl13 | $spl20)),
  inference(subsumption_resolution,[],[f45319,f24068])).
fof(f24068,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl13),
  inference(unit_resulting_resolution,[],[f853,f20974,f1258])).
fof(f1258,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f814])).
fof(f814,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f660])).
fof(f660,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f535,f659,f658,f657,f656])).
fof(f656,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f657,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f658,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f659,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f535,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f534])).
fof(f534,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X39,X40,X41] : ((p(s(bool,X41)) <=> s(bool,X39) = s(bool,X40)) <=> ((p(s(bool,X41)) | p(s(bool,X40)) | p(s(bool,X39))) & (p(s(bool,X41)) | ~p(s(bool,X39)) | ~p(s(bool,X40))) & (p(s(bool,X40)) | ~p(s(bool,X39)) | ~p(s(bool,X41))) & (p(s(bool,X39)) | ~p(s(bool,X40)) | ~p(s(bool,X41)))))),
  file('thm.extreal.le_addr_imp','thm.sat.sat_2|49|_prove_ltb')).
fof(f20974,plain,(
  ~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl13),
  inference(unit_resulting_resolution,[],[f20970,f1115])).
fof(f1115,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f582])).
fof(f582,plain,(
  ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))),
  inference(ennf_transformation,[],[f462])).
fof(f462,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))),
  inference(rectify,[],[f105])).
fof(f105,axiom,(
  ! [X6,X7] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))))),
  file('thm.extreal.le_addr_imp','thm.extreal.lt_imp_le')).
fof(f853,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(cnf_transformation,[],[f115])).
fof(f115,axiom,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_le_def|split|2')).
fof(f45319,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ($spl13 | $spl20)),
  inference(forward_demodulation,[],[f45025,f886])).
fof(f886,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f81])).
fof(f81,axiom,(
  ! [X17] : s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X17)))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_add_def|split|1')).
fof(f45025,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | ($spl13 | $spl20)),
  inference(backward_demodulation,[],[f44488,f24023])).
fof(f24023,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))))) | $spl13),
  inference(unit_resulting_resolution,[],[f1592,f20974,f1255])).
fof(f1255,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f814])).
fof(f1592,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f841,f951])).
fof(f951,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f677])).
fof(f677,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X21] : (s(bool,'T_0') = s(bool,X21) <=> p(s(bool,X21)))),
  file('thm.extreal.le_addr_imp','thm.bool.EQ_CLAUSES|split|0')).
fof(f841,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))))))),
  inference(cnf_transformation,[],[f670])).
fof(f670,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f560])).
fof(f560,plain,(
  ? [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) & ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(ennf_transformation,[],[f263])).
fof(f263,plain,(
  ~! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X7)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X7)))) => p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))))))),
  file('thm.extreal.le_addr_imp','thm.extreal.le_addr_imp')).
fof(f45318,plain,(
  $false | ($spl13 | $spl20)),
  inference(subsumption_resolution,[],[f45317,f24026])).
fof(f24026,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f853,f20974,f1258])).
fof(f45317,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | ($spl13 | $spl20)),
  inference(forward_demodulation,[],[f45024,f886])).
fof(f45024,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | ($spl13 | $spl20)),
  inference(backward_demodulation,[],[f44488,f24021])).
fof(f24021,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f1592,f20974,f1255])).
fof(f45295,plain,(
  $false | ($spl13 | $spl20)),
  inference(subsumption_resolution,[],[f45294,f22497])).
fof(f22497,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f853,f20973,f1256])).
fof(f1256,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f814])).
fof(f20973,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl13),
  inference(unit_resulting_resolution,[],[f20970,f1147])).
fof(f1147,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f761])).
fof(f761,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) & (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(nnf_transformation,[],[f489])).
fof(f489,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f488])).
fof(f488,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X6,X7] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) <=> ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_lt_def')).
fof(f45294,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | ($spl13 | $spl20)),
  inference(forward_demodulation,[],[f45002,f886])).
fof(f45002,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | ($spl13 | $spl20)),
  inference(backward_demodulation,[],[f44488,f22582])).
fof(f22582,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f1592,f20973,f1257])).
fof(f1257,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f814])).
fof(f45293,plain,(
  $false | ($spl13 | $spl20)),
  inference(subsumption_resolution,[],[f45292,f22539])).
fof(f22539,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl13),
  inference(unit_resulting_resolution,[],[f853,f20973,f1256])).
fof(f45292,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ($spl13 | $spl20)),
  inference(forward_demodulation,[],[f45001,f886])).
fof(f45001,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | ($spl13 | $spl20)),
  inference(backward_demodulation,[],[f44488,f22580])).
fof(f22580,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))))) | $spl13),
  inference(unit_resulting_resolution,[],[f1592,f20973,f1257])).
fof(f45281,plain,(
  $false | ($spl13 | $spl20)),
  inference(subsumption_resolution,[],[f45280,f21251])).
fof(f21251,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl13),
  inference(unit_resulting_resolution,[],[f853,f20970,f1258])).
fof(f45280,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | ($spl13 | $spl20)),
  inference(forward_demodulation,[],[f44990,f886])).
fof(f44990,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | ($spl13 | $spl20)),
  inference(backward_demodulation,[],[f44488,f21207])).
fof(f21207,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))))) | $spl13),
  inference(unit_resulting_resolution,[],[f1592,f20970,f1255])).
fof(f45279,plain,(
  $false | ($spl13 | $spl20)),
  inference(subsumption_resolution,[],[f45278,f21210])).
fof(f21210,plain,(
  ~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f853,f20970,f1258])).
fof(f45278,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | ($spl13 | $spl20)),
  inference(forward_demodulation,[],[f44989,f886])).
fof(f44989,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | ($spl13 | $spl20)),
  inference(backward_demodulation,[],[f44488,f21205])).
fof(f21205,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f1592,f20970,f1255])).
fof(f45250,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45249,f19999])).
fof(f19999,plain,(
  ~sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f853,f13947,f1258])).
fof(f13947,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0'))))),
  inference(forward_demodulation,[],[f13788,f1639])).
fof(f1639,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f841,f956])).
fof(f956,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f679])).
fof(f679,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f344])).
fof(f344,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X21] : (s(bool,'F_0') = s(bool,X21) <=> ~p(s(bool,X21)))),
  file('thm.extreal.le_addr_imp','thm.bool.EQ_CLAUSES|split|2')).
fof(f13788,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))))))))),
  inference(unit_resulting_resolution,[],[f1592,f1228])).
fof(f1228,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f800])).
fof(f800,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f532])).
fof(f532,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f257])).
fof(f257,axiom,(
  ! [X48,X49,X50] : (p(s(bool,happ(s(fun(X48,bool),'=_1'(s(X48,X49))),s(X48,X50)))) <=> s(X48,X49) = s(X48,X50))),
  file('thm.extreal.le_addr_imp',arity_1_HL_EQ)).
fof(f45249,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl20),
  inference(forward_demodulation,[],[f44961,f886])).
fof(f44961,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | $spl20),
  inference(backward_demodulation,[],[f44488,f19955])).
fof(f19955,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f1592,f13947,f1255])).
fof(f45248,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45247,f19958])).
fof(f19958,plain,(
  ~sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f853,f13947,f1258])).
fof(f45247,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | $spl20),
  inference(forward_demodulation,[],[f44960,f886])).
fof(f44960,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl20),
  inference(backward_demodulation,[],[f44488,f19953])).
fof(f19953,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0')),
  inference(unit_resulting_resolution,[],[f1592,f13947,f1255])).
fof(f45219,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45218,f18978])).
fof(f18978,plain,(
  ~sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f853,f13946,f1258])).
fof(f13946,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0'))))),
  inference(forward_demodulation,[],[f13798,f1639])).
fof(f13798,plain,(
  ~p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))))))),s(bool,'T_0'))))),
  inference(unit_resulting_resolution,[],[f1592,f1228])).
fof(f45218,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl20),
  inference(forward_demodulation,[],[f44932,f886])).
fof(f44932,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | $spl20),
  inference(backward_demodulation,[],[f44488,f18934])).
fof(f18934,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f1592,f13946,f1255])).
fof(f45217,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45216,f18937])).
fof(f18937,plain,(
  ~sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f853,f13946,f1258])).
fof(f45216,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | $spl20),
  inference(forward_demodulation,[],[f44931,f886])).
fof(f44931,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl20),
  inference(backward_demodulation,[],[f44488,f18932])).
fof(f18932,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0')),
  inference(unit_resulting_resolution,[],[f1592,f13946,f1255])).
fof(f45196,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45195,f14427])).
fof(f14427,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f14167,f1048])).
fof(f1048,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f404])).
fof(f404,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f403])).
fof(f403,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f154])).
fof(f154,axiom,(
  ! [X8] : ! [X23,X24] : s(X8,X24) = s(X8,'const.bool.COND_3'(s(bool,'F_0'),s(X8,X23),s(X8,X24)))),
  file('thm.extreal.le_addr_imp','thm.bool.COND_CLAUSES|split|1')).
fof(f14167,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2))))),s(X0,X2))))) )),
  inference(unit_resulting_resolution,[],[f1048,f1229])).
fof(f1229,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'=_1'(s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f800])).
fof(f45195,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl20),
  inference(forward_demodulation,[],[f44911,f886])).
fof(f44911,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl20),
  inference(backward_demodulation,[],[f44488,f16554])).
fof(f16554,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))))),s('type.extreal.extreal',sK12))))),
  inference(unit_resulting_resolution,[],[f8160,f1228])).
fof(f8160,plain,(
  s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),
  inference(subsumption_resolution,[],[f8067,f1359])).
fof(f1359,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f866,f842])).
fof(f842,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f20])).
fof(f20,axiom,(
  p(s(bool,'T_0'))),
  file('thm.extreal.le_addr_imp','thm.bool.TRUTH')).
fof(f866,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f671])).
fof(f671,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f264])).
fof(f264,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f40])).
fof(f40,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.extreal.le_addr_imp','thm.bool.NOT_CLAUSES|split|2')).
fof(f8067,plain,(
  p(s(bool,'F_0')) | s('type.extreal.extreal',sK12) != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),
  inference(superposition,[],[f1142,f1639])).
fof(f1142,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | s('type.extreal.extreal',X0) != s('type.extreal.extreal',X1)) )),
  inference(cnf_transformation,[],[f758])).
fof(f758,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal',X1)) & ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & s('type.extreal.extreal',X0) != s('type.extreal.extreal',X1)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f757])).
fof(f757,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal',X1))) & ((~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & s('type.extreal.extreal',X0) != s('type.extreal.extreal',X1)) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(nnf_transformation,[],[f486])).
fof(f486,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) <=> (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal',X1)))),
  inference(rectify,[],[f106])).
fof(f106,axiom,(
  ! [X6,X7] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) <=> (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) | s('type.extreal.extreal',X6) = s('type.extreal.extreal',X7)))),
  file('thm.extreal.le_addr_imp','thm.extreal.le_lt')).
fof(f45194,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45193,f14427])).
fof(f45193,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl20),
  inference(forward_demodulation,[],[f44910,f886])).
fof(f44910,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))))) | $spl20),
  inference(backward_demodulation,[],[f44488,f16553])).
fof(f16553,plain,(
  ~p(s(bool,happ(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',sK12))),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))))))),
  inference(unit_resulting_resolution,[],[f8160,f1228])).
fof(f45190,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45189,f14848])).
fof(f14848,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f853,f6090,f1256])).
fof(f6090,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0'))))))))))),
  inference(backward_demodulation,[],[f6088,f861])).
fof(f861,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(cnf_transformation,[],[f72])).
fof(f72,axiom,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  file('thm.extreal.le_addr_imp','thm.extreal.lt_|48|2')).
fof(f6088,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0')),
  inference(subsumption_resolution,[],[f6076,f842])).
fof(f6076,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.ZERO_0')),
  inference(superposition,[],[f977,f894])).
fof(f894,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f176])).
fof(f176,axiom,(
  ! [X22] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X22)))),
  file('thm.extreal.le_addr_imp','thm.numeral.numeral_lte|split|0')).
fof(f977,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f690])).
fof(f690,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))))),
  inference(nnf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f228])).
fof(f228,axiom,(
  ! [X22] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X22),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X22))),
  file('thm.extreal.le_addr_imp','thm.arithmetic.LESS_EQ_|48|')).
fof(f45189,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | $spl20),
  inference(forward_demodulation,[],[f44907,f886])).
fof(f44907,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl20),
  inference(backward_demodulation,[],[f44488,f14899])).
fof(f14899,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0')),
  inference(unit_resulting_resolution,[],[f1592,f6090,f1257])).
fof(f45188,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45187,f14873])).
fof(f14873,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f853,f6090,f1256])).
fof(f45187,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl20),
  inference(forward_demodulation,[],[f44906,f886])).
fof(f44906,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | $spl20),
  inference(backward_demodulation,[],[f44488,f14897])).
fof(f14897,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f1592,f6090,f1257])).
fof(f45184,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45183,f1472])).
fof(f1472,plain,(
  ( ! [X2,X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f874,f1254])).
fof(f1254,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f813])).
fof(f813,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f812])).
fof(f812,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f656])).
fof(f874,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  file('thm.extreal.le_addr_imp','thm.extreal.le_infty|split|1')).
fof(f45183,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'const.combin.I_1'(s(bool,'F_0')),'F_0') | $spl20),
  inference(forward_demodulation,[],[f44903,f886])).
fof(f44903,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'const.combin.I_1'(s(bool,'F_0')),'F_0') | $spl20),
  inference(backward_demodulation,[],[f44488,f12674])).
fof(f12674,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'const.combin.I_1'(s(bool,'F_0')),'F_0')),
  inference(forward_demodulation,[],[f12588,f1639])).
fof(f12588,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'const.combin.I_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),'F_0')),
  inference(unit_resulting_resolution,[],[f3470,f1406,f1415,f1359,f1242])).
fof(f1242,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f805])).
fof(f805,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f804])).
fof(f804,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f803])).
fof(f803,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f659])).
fof(f1415,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),X1)) )),
  inference(unit_resulting_resolution,[],[f841,f1249])).
fof(f1249,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f811])).
fof(f811,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f810])).
fof(f810,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f809])).
fof(f809,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f657])).
fof(f1406,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),X1)) )),
  inference(unit_resulting_resolution,[],[f841,f1245])).
fof(f1245,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f808])).
fof(f808,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f807])).
fof(f807,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f806])).
fof(f806,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f658])).
fof(f3470,plain,(
  ( ! [X0] : (~sP5('F_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f1359,f1041,f1258])).
fof(f1041,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f395])).
fof(f395,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f394])).
fof(f394,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : ! [X6] : s(X8,X6) = s(X8,'const.combin.I_1'(s(X8,X6)))),
  file('thm.extreal.le_addr_imp','thm.combin.I_THM')).
fof(f45182,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45181,f1471])).
fof(f1471,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.PosInf_0')),X2)) )),
  inference(unit_resulting_resolution,[],[f874,f1253])).
fof(f1253,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f813])).
fof(f45181,plain,(
  ~sP2('const.combin.I_1'(s(bool,'F_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0') | $spl20),
  inference(forward_demodulation,[],[f44902,f886])).
fof(f44902,plain,(
  ~sP2('const.combin.I_1'(s(bool,'F_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'F_0') | $spl20),
  inference(backward_demodulation,[],[f44488,f12671])).
fof(f12671,plain,(
  ~sP2('const.combin.I_1'(s(bool,'F_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'F_0')),
  inference(forward_demodulation,[],[f12657,f1639])).
fof(f12657,plain,(
  ~sP2('const.combin.I_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'F_0')),
  inference(unit_resulting_resolution,[],[f3482,f1407,f1418,f1359,f1242])).
fof(f1418,plain,(
  ( ! [X0,X1] : (sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f841,f1250])).
fof(f1250,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f811])).
fof(f1407,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1359,f1246])).
fof(f1246,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f808])).
fof(f3482,plain,(
  ( ! [X0] : (~sP5('F_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f1359,f1041,f1258])).
fof(f45180,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45179,f1471])).
fof(f45179,plain,(
  ~sP2('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0') | $spl20),
  inference(forward_demodulation,[],[f44901,f886])).
fof(f44901,plain,(
  ~sP2('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'F_0') | $spl20),
  inference(backward_demodulation,[],[f44488,f12658])).
fof(f12658,plain,(
  ~sP2('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'F_0')),
  inference(unit_resulting_resolution,[],[f3463,f1404,f1418,f1359,f1242])).
fof(f1404,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1359,f1245])).
fof(f3463,plain,(
  ~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1639,f1258])).
fof(f45160,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45159,f1472])).
fof(f45159,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0','F_0') | $spl20),
  inference(forward_demodulation,[],[f44882,f886])).
fof(f44882,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'F_0','F_0') | $spl20),
  inference(backward_demodulation,[],[f44488,f12430])).
fof(f12430,plain,(
  ~sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f3475,f1407,f1415,f1359,f1241])).
fof(f1241,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f805])).
fof(f3475,plain,(
  ~sP5('F_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f1359,f1639,f1258])).
fof(f45139,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45138,f1468])).
fof(f1468,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X2),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f874,f1244])).
fof(f1244,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f808])).
fof(f45138,plain,(
  ~sP4('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl20),
  inference(forward_demodulation,[],[f44862,f886])).
fof(f44862,plain,(
  ~sP4('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | $spl20),
  inference(backward_demodulation,[],[f44488,f12259])).
fof(f12259,plain,(
  ~sP4('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f1425,f3449,f1418,f842,f1240])).
fof(f1240,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f805])).
fof(f3449,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0')),
  inference(unit_resulting_resolution,[],[f842,f1592,f1257])).
fof(f1425,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f842,f1254])).
fof(f45133,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45132,f9466])).
fof(f9466,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f853,f1355,f1256])).
fof(f1355,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(forward_demodulation,[],[f859,f858])).
fof(f858,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f121])).
fof(f121,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.extreal.le_addr_imp','thm.numeral.numeral_distrib|split|13')).
fof(f859,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(cnf_transformation,[],[f73])).
fof(f73,axiom,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  file('thm.extreal.le_addr_imp','thm.extreal.lt_|48||49|')).
fof(f45132,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | $spl20),
  inference(forward_demodulation,[],[f44857,f886])).
fof(f44857,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl20),
  inference(backward_demodulation,[],[f44488,f9509])).
fof(f9509,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0')),
  inference(unit_resulting_resolution,[],[f1592,f1355,f1257])).
fof(f45131,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45130,f9487])).
fof(f9487,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f853,f1355,f1256])).
fof(f45130,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl20),
  inference(forward_demodulation,[],[f44856,f886])).
fof(f44856,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | $spl20),
  inference(backward_demodulation,[],[f44488,f9507])).
fof(f9507,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f1592,f1355,f1257])).
fof(f45127,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f44853,f886])).
fof(f44853,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))) | $spl20),
  inference(backward_demodulation,[],[f44488,f8160])).
fof(f45124,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45123,f6989])).
fof(f6989,plain,(
  sP5(sK27('const.arithmetic.ZERO_0'),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f853,f6174,f1256])).
fof(f6174,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0')))),
  inference(unit_resulting_resolution,[],[f6088,f1066])).
fof(f1066,plain,(
  ( ! [X0] : (p(s(bool,sK27(X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f734])).
fof(f734,plain,(
  ! [X0] : ((~p(s(bool,sK27(X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,sK27(X0)))) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK27(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27])],[f733])).
fof(f733,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(flattening,[],[f732])).
fof(f732,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(nnf_transformation,[],[f427])).
fof(f427,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(rectify,[],[f235])).
fof(f235,axiom,(
  ! [X22] : ? [X47] : ((p(s(bool,X47)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X22)) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.num.num',X22))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X47),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_pow_def|split|1')).
fof(f45123,plain,(
  ~sP5(sK27('const.arithmetic.ZERO_0'),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | $spl20),
  inference(forward_demodulation,[],[f44850,f886])).
fof(f44850,plain,(
  ~sP5(sK27('const.arithmetic.ZERO_0'),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl20),
  inference(backward_demodulation,[],[f44488,f7028])).
fof(f7028,plain,(
  ~sP5(sK27('const.arithmetic.ZERO_0'),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0')),
  inference(unit_resulting_resolution,[],[f1592,f6174,f1257])).
fof(f45122,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45121,f7008])).
fof(f7008,plain,(
  sP5(sK27('const.arithmetic.ZERO_0'),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f853,f6174,f1256])).
fof(f45121,plain,(
  ~sP5(sK27('const.arithmetic.ZERO_0'),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl20),
  inference(forward_demodulation,[],[f44849,f886])).
fof(f44849,plain,(
  ~sP5(sK27('const.arithmetic.ZERO_0'),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | $spl20),
  inference(backward_demodulation,[],[f44488,f7026])).
fof(f7026,plain,(
  ~sP5(sK27('const.arithmetic.ZERO_0'),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f1592,f6174,f1257])).
fof(f45116,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45115,f6307])).
fof(f6307,plain,(
  sP5(sK28('const.arithmetic.ZERO_0'),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f853,f6173,f1256])).
fof(f6173,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0')))),
  inference(unit_resulting_resolution,[],[f6088,f1069])).
fof(f1069,plain,(
  ( ! [X0] : (p(s(bool,sK28(X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f737])).
fof(f737,plain,(
  ! [X0] : ((~p(s(bool,sK28(X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,sK28(X0)))) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK28(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK28])],[f736])).
fof(f736,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(flattening,[],[f735])).
fof(f735,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f428])).
fof(f428,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.num.num',X0))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(rectify,[],[f244])).
fof(f244,axiom,(
  ! [X22] : ? [X47] : ((p(s(bool,X47)) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X22)) & s('type.extreal.extreal','const.extreal.extreal_pow_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.num.num',X22))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X47),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X22))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_pow_def|split|2')).
fof(f45115,plain,(
  ~sP5(sK28('const.arithmetic.ZERO_0'),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | $spl20),
  inference(forward_demodulation,[],[f44844,f886])).
fof(f44844,plain,(
  ~sP5(sK28('const.arithmetic.ZERO_0'),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl20),
  inference(backward_demodulation,[],[f44488,f6344])).
fof(f6344,plain,(
  ~sP5(sK28('const.arithmetic.ZERO_0'),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0')),
  inference(unit_resulting_resolution,[],[f1592,f6173,f1257])).
fof(f45114,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45113,f6325])).
fof(f6325,plain,(
  sP5(sK28('const.arithmetic.ZERO_0'),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f853,f6173,f1256])).
fof(f45113,plain,(
  ~sP5(sK28('const.arithmetic.ZERO_0'),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl20),
  inference(forward_demodulation,[],[f44843,f886])).
fof(f44843,plain,(
  ~sP5(sK28('const.arithmetic.ZERO_0'),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | $spl20),
  inference(backward_demodulation,[],[f44488,f6342])).
fof(f6342,plain,(
  ~sP5(sK28('const.arithmetic.ZERO_0'),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f1592,f6173,f1257])).
fof(f45107,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45106,f4511])).
fof(f4511,plain,(
  ( ! [X0] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f853,f908,f1256])).
fof(f908,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))))) )),
  inference(cnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X22] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X22))))))),
  file('thm.extreal.le_addr_imp','thm.extreal.le_num')).
fof(f45106,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0')) ) | $spl20),
  inference(forward_demodulation,[],[f44837,f886])).
fof(f44837,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'T_0')) ) | $spl20),
  inference(backward_demodulation,[],[f44488,f4550])).
fof(f4550,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1592,f908,f1257])).
fof(f45105,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45104,f4530])).
fof(f4530,plain,(
  ( ! [X0] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f853,f908,f1256])).
fof(f45104,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) ) | $spl20),
  inference(forward_demodulation,[],[f44836,f886])).
fof(f44836,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))))) ) | $spl20),
  inference(backward_demodulation,[],[f44488,f4548])).
fof(f4548,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))) )),
  inference(unit_resulting_resolution,[],[f1592,f908,f1257])).
fof(f45100,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45099,f4325])).
fof(f4325,plain,(
  sP5(sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f853,f3245,f1256])).
fof(f3245,plain,(
  p(s(bool,sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(equality_resolution,[],[f1081])).
fof(f1081,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK32(X0)))) )),
  inference(cnf_transformation,[],[f749])).
fof(f749,plain,(
  ! [X0] : ((~p(s(bool,sK32(X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK32(X0)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK32(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK32])],[f748])).
fof(f748,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(flattening,[],[f747])).
fof(f747,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(nnf_transformation,[],[f432])).
fof(f432,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(rectify,[],[f237])).
fof(f237,axiom,(
  ! [X7] : ? [X47] : ((p(s(bool,X47)) <=> s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X47),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_mul_def|split|5')).
fof(f45099,plain,(
  ~sP5(sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | $spl20),
  inference(forward_demodulation,[],[f44832,f886])).
fof(f44832,plain,(
  ~sP5(sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl20),
  inference(backward_demodulation,[],[f44488,f4362])).
fof(f4362,plain,(
  ~sP5(sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0')),
  inference(unit_resulting_resolution,[],[f1592,f3245,f1257])).
fof(f45098,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45097,f4343])).
fof(f4343,plain,(
  sP5(sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f853,f3245,f1256])).
fof(f45097,plain,(
  ~sP5(sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl20),
  inference(forward_demodulation,[],[f44831,f886])).
fof(f44831,plain,(
  ~sP5(sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | $spl20),
  inference(backward_demodulation,[],[f44488,f4360])).
fof(f4360,plain,(
  ~sP5(sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f1592,f3245,f1257])).
fof(f45093,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45092,f4061])).
fof(f4061,plain,(
  sP5(sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f853,f3233,f1256])).
fof(f3233,plain,(
  p(s(bool,sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(equality_resolution,[],[f1078])).
fof(f1078,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK31(X0)))) )),
  inference(cnf_transformation,[],[f746])).
fof(f746,plain,(
  ! [X0] : ((~p(s(bool,sK31(X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK31(X0)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK31(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK31])],[f745])).
fof(f745,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(flattening,[],[f744])).
fof(f744,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f431])).
fof(f431,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(rectify,[],[f242])).
fof(f242,axiom,(
  ! [X7] : ? [X47] : ((p(s(bool,X47)) <=> s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X7))))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X47),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_mul_def|split|7')).
fof(f45092,plain,(
  ~sP5(sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | $spl20),
  inference(forward_demodulation,[],[f44827,f886])).
fof(f44827,plain,(
  ~sP5(sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl20),
  inference(backward_demodulation,[],[f44488,f4092])).
fof(f4092,plain,(
  ~sP5(sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0')),
  inference(unit_resulting_resolution,[],[f1592,f3233,f1257])).
fof(f45091,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45090,f4076])).
fof(f4076,plain,(
  sP5(sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f853,f3233,f1256])).
fof(f45090,plain,(
  ~sP5(sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl20),
  inference(forward_demodulation,[],[f44826,f886])).
fof(f44826,plain,(
  ~sP5(sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | $spl20),
  inference(backward_demodulation,[],[f44488,f4090])).
fof(f4090,plain,(
  ~sP5(sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f1592,f3233,f1257])).
fof(f45086,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45085,f3873])).
fof(f3873,plain,(
  sP5(sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f853,f3157,f1256])).
fof(f3157,plain,(
  p(s(bool,sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(equality_resolution,[],[f1075])).
fof(f1075,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK30(X0)))) )),
  inference(cnf_transformation,[],[f743])).
fof(f743,plain,(
  ! [X0] : ((~p(s(bool,sK30(X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK30(X0)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK30(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK30])],[f742])).
fof(f742,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(flattening,[],[f741])).
fof(f741,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f430])).
fof(f430,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(rectify,[],[f239])).
fof(f239,axiom,(
  ! [X6] : ? [X47] : ((p(s(bool,X47)) <=> s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))),s('type.extreal.extreal','const.extreal.PosInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X47),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6))),s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_mul_def|split|6')).
fof(f45085,plain,(
  ~sP5(sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | $spl20),
  inference(forward_demodulation,[],[f44822,f886])).
fof(f44822,plain,(
  ~sP5(sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl20),
  inference(backward_demodulation,[],[f44488,f3903])).
fof(f3903,plain,(
  ~sP5(sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0')),
  inference(unit_resulting_resolution,[],[f1592,f3157,f1257])).
fof(f45084,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45083,f3887])).
fof(f3887,plain,(
  sP5(sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f853,f3157,f1256])).
fof(f45083,plain,(
  ~sP5(sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl20),
  inference(forward_demodulation,[],[f44821,f886])).
fof(f44821,plain,(
  ~sP5(sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | $spl20),
  inference(backward_demodulation,[],[f44488,f3901])).
fof(f3901,plain,(
  ~sP5(sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f1592,f3157,f1257])).
fof(f45073,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45072,f3429])).
fof(f3429,plain,(
  sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f842,f853,f1256])).
fof(f45072,plain,(
  ~sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl20),
  inference(forward_demodulation,[],[f44811,f886])).
fof(f44811,plain,(
  ~sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | $spl20),
  inference(backward_demodulation,[],[f44488,f3452])).
fof(f3452,plain,(
  ~sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f842,f1592,f1257])).
fof(f45071,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45070,f3418])).
fof(f3418,plain,(
  sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f842,f853,f1256])).
fof(f45070,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | $spl20),
  inference(forward_demodulation,[],[f44810,f886])).
fof(f44810,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl20),
  inference(backward_demodulation,[],[f44488,f3449])).
fof(f45067,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45066,f3473])).
fof(f3473,plain,(
  ~sP5('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(unit_resulting_resolution,[],[f1359,f853,f1258])).
fof(f45066,plain,(
  sP5('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl20),
  inference(forward_demodulation,[],[f44807,f886])).
fof(f44807,plain,(
  sP5('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | $spl20),
  inference(backward_demodulation,[],[f44488,f3410])).
fof(f3410,plain,(
  sP5('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f1359,f1592,f1255])).
fof(f45065,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45064,f3461])).
fof(f3461,plain,(
  ~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f1359,f853,f1258])).
fof(f45064,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0') | $spl20),
  inference(forward_demodulation,[],[f44806,f886])).
fof(f44806,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl20),
  inference(backward_demodulation,[],[f44488,f3408])).
fof(f3408,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0')),
  inference(unit_resulting_resolution,[],[f1359,f1592,f1255])).
fof(f45048,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45047,f1591])).
fof(f1591,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f951])).
fof(f45047,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl20),
  inference(forward_demodulation,[],[f45046,f853])).
fof(f45046,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl20),
  inference(forward_demodulation,[],[f44790,f886])).
fof(f44790,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | $spl20),
  inference(backward_demodulation,[],[f44488,f1639])).
fof(f45045,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f45044,f853])).
fof(f45044,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))) | $spl20),
  inference(forward_demodulation,[],[f44789,f886])).
fof(f44789,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',sK13))))) | $spl20),
  inference(backward_demodulation,[],[f44488,f1592])).
fof(f44779,plain,(
  $false | ($spl13 | $spl18)),
  inference(subsumption_resolution,[],[f44737,f24067])).
fof(f24067,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) ) | $spl13),
  inference(unit_resulting_resolution,[],[f888,f20974,f1258])).
fof(f888,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(cnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))),
  inference(rectify,[],[f111])).
fof(f111,axiom,(
  ! [X17] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X17)))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_le_def|split|1')).
fof(f44737,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | ($spl13 | $spl18)),
  inference(backward_demodulation,[],[f44486,f24023])).
fof(f44778,plain,(
  $false | ($spl13 | $spl18)),
  inference(subsumption_resolution,[],[f44736,f24025])).
fof(f24025,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)),'T_0')) ) | $spl13),
  inference(unit_resulting_resolution,[],[f888,f20974,f1258])).
fof(f44736,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | ($spl13 | $spl18)),
  inference(backward_demodulation,[],[f44486,f24021])).
fof(f44777,plain,(
  $false | ($spl13 | $spl18)),
  inference(subsumption_resolution,[],[f44714,f22496])).
fof(f22496,plain,(
  ( ! [X0] : (sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)),'T_0')) ) | $spl13),
  inference(unit_resulting_resolution,[],[f888,f20973,f1256])).
fof(f44714,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | ($spl13 | $spl18)),
  inference(backward_demodulation,[],[f44486,f22582])).
fof(f44776,plain,(
  $false | ($spl13 | $spl18)),
  inference(subsumption_resolution,[],[f44713,f22538])).
fof(f22538,plain,(
  ( ! [X0] : (sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) ) | $spl13),
  inference(unit_resulting_resolution,[],[f888,f20973,f1256])).
fof(f44713,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | ($spl13 | $spl18)),
  inference(backward_demodulation,[],[f44486,f22580])).
fof(f44775,plain,(
  $false | ($spl13 | $spl18)),
  inference(subsumption_resolution,[],[f44702,f21250])).
fof(f21250,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) ) | $spl13),
  inference(unit_resulting_resolution,[],[f888,f20970,f1258])).
fof(f44702,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | ($spl13 | $spl18)),
  inference(backward_demodulation,[],[f44486,f21207])).
fof(f44774,plain,(
  $false | ($spl13 | $spl18)),
  inference(subsumption_resolution,[],[f44701,f21209])).
fof(f21209,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)),'T_0')) ) | $spl13),
  inference(unit_resulting_resolution,[],[f888,f20970,f1258])).
fof(f44701,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | ($spl13 | $spl18)),
  inference(backward_demodulation,[],[f44486,f21205])).
fof(f44773,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44673,f19998])).
fof(f19998,plain,(
  ( ! [X0] : (~sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f888,f13947,f1258])).
fof(f44673,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | $spl18),
  inference(backward_demodulation,[],[f44486,f19955])).
fof(f44772,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44672,f19957])).
fof(f19957,plain,(
  ( ! [X0] : (~sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f888,f13947,f1258])).
fof(f44672,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl18),
  inference(backward_demodulation,[],[f44486,f19953])).
fof(f44771,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44644,f18977])).
fof(f18977,plain,(
  ( ! [X0] : (~sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f888,f13946,f1258])).
fof(f44644,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | $spl18),
  inference(backward_demodulation,[],[f44486,f18934])).
fof(f44770,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44643,f18936])).
fof(f18936,plain,(
  ( ! [X0] : (~sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f888,f13946,f1258])).
fof(f44643,plain,(
  sP5(happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl18),
  inference(backward_demodulation,[],[f44486,f18932])).
fof(f44769,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44619,f14847])).
fof(f14847,plain,(
  ( ! [X0] : (sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f888,f6090,f1256])).
fof(f44619,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl18),
  inference(backward_demodulation,[],[f44486,f14899])).
fof(f44768,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44618,f14872])).
fof(f14872,plain,(
  ( ! [X0] : (sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f888,f6090,f1256])).
fof(f44618,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.num.|48|_0')))))))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | $spl18),
  inference(backward_demodulation,[],[f44486,f14897])).
fof(f44767,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44569,f9465])).
fof(f9465,plain,(
  ( ! [X0] : (sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f888,f1355,f1256])).
fof(f44569,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl18),
  inference(backward_demodulation,[],[f44486,f9509])).
fof(f44766,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44568,f9486])).
fof(f9486,plain,(
  ( ! [X0] : (sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f888,f1355,f1256])).
fof(f44568,plain,(
  ~sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | $spl18),
  inference(backward_demodulation,[],[f44486,f9507])).
fof(f44765,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44562,f6988])).
fof(f6988,plain,(
  ( ! [X0] : (sP5(sK27('const.arithmetic.ZERO_0'),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f888,f6174,f1256])).
fof(f44562,plain,(
  ~sP5(sK27('const.arithmetic.ZERO_0'),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl18),
  inference(backward_demodulation,[],[f44486,f7028])).
fof(f44764,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44561,f7007])).
fof(f7007,plain,(
  ( ! [X0] : (sP5(sK27('const.arithmetic.ZERO_0'),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f888,f6174,f1256])).
fof(f44561,plain,(
  ~sP5(sK27('const.arithmetic.ZERO_0'),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | $spl18),
  inference(backward_demodulation,[],[f44486,f7026])).
fof(f44763,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44556,f6306])).
fof(f6306,plain,(
  ( ! [X0] : (sP5(sK28('const.arithmetic.ZERO_0'),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f888,f6173,f1256])).
fof(f44556,plain,(
  ~sP5(sK28('const.arithmetic.ZERO_0'),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl18),
  inference(backward_demodulation,[],[f44486,f6344])).
fof(f44762,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44555,f6324])).
fof(f6324,plain,(
  ( ! [X0] : (sP5(sK28('const.arithmetic.ZERO_0'),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f888,f6173,f1256])).
fof(f44555,plain,(
  ~sP5(sK28('const.arithmetic.ZERO_0'),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | $spl18),
  inference(backward_demodulation,[],[f44486,f6342])).
fof(f44761,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44549,f4510])).
fof(f4510,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f888,f908,f1256])).
fof(f44549,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'T_0')) ) | $spl18),
  inference(backward_demodulation,[],[f44486,f4550])).
fof(f44760,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44548,f4529])).
fof(f4529,plain,(
  ( ! [X0,X1] : (sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f888,f908,f1256])).
fof(f44548,plain,(
  ( ! [X0] : (~sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))))) ) | $spl18),
  inference(backward_demodulation,[],[f44486,f4548])).
fof(f44759,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44544,f4324])).
fof(f4324,plain,(
  ( ! [X0] : (sP5(sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f888,f3245,f1256])).
fof(f44544,plain,(
  ~sP5(sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl18),
  inference(backward_demodulation,[],[f44486,f4362])).
fof(f44758,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44543,f4342])).
fof(f4342,plain,(
  ( ! [X0] : (sP5(sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f888,f3245,f1256])).
fof(f44543,plain,(
  ~sP5(sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | $spl18),
  inference(backward_demodulation,[],[f44486,f4360])).
fof(f44757,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44539,f4060])).
fof(f4060,plain,(
  ( ! [X0] : (sP5(sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f888,f3233,f1256])).
fof(f44539,plain,(
  ~sP5(sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl18),
  inference(backward_demodulation,[],[f44486,f4092])).
fof(f44756,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44538,f4075])).
fof(f4075,plain,(
  ( ! [X0] : (sP5(sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f888,f3233,f1256])).
fof(f44538,plain,(
  ~sP5(sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | $spl18),
  inference(backward_demodulation,[],[f44486,f4090])).
fof(f44755,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44534,f3874])).
fof(f3874,plain,(
  ( ! [X0] : (sP5(sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f888,f3157,f1256])).
fof(f44534,plain,(
  ~sP5(sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl18),
  inference(backward_demodulation,[],[f44486,f3903])).
fof(f44754,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44533,f3888])).
fof(f3888,plain,(
  ( ! [X0] : (sP5(sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f888,f3157,f1256])).
fof(f44533,plain,(
  ~sP5(sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | $spl18),
  inference(backward_demodulation,[],[f44486,f3901])).
fof(f44753,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44523,f3430])).
fof(f3430,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f842,f888,f1256])).
fof(f44523,plain,(
  ~sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | $spl18),
  inference(backward_demodulation,[],[f44486,f3452])).
fof(f44752,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44522,f3419])).
fof(f3419,plain,(
  ( ! [X0] : (sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f842,f888,f1256])).
fof(f44522,plain,(
  ~sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl18),
  inference(backward_demodulation,[],[f44486,f3449])).
fof(f44751,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44519,f3474])).
fof(f3474,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) )),
  inference(unit_resulting_resolution,[],[f1359,f888,f1258])).
fof(f44519,plain,(
  sP5('F_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | $spl18),
  inference(backward_demodulation,[],[f44486,f3410])).
fof(f44750,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44518,f3462])).
fof(f3462,plain,(
  ( ! [X0] : (~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f1359,f888,f1258])).
fof(f44518,plain,(
  sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13)))),'T_0') | $spl18),
  inference(backward_demodulation,[],[f44486,f3408])).
fof(f44749,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44748,f1591])).
fof(f44748,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl18),
  inference(forward_demodulation,[],[f44502,f888])).
fof(f44502,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | $spl18),
  inference(backward_demodulation,[],[f44486,f1639])).
fof(f44747,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f44501,f888])).
fof(f44501,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',sK13))))) | $spl18),
  inference(backward_demodulation,[],[f44486,f1592])).
fof(f44486,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) | $spl18),
  inference(cnf_transformation,[],[f44486_D])).
fof(f44486_D,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f44488,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | $spl20),
  inference(cnf_transformation,[],[f44488_D])).
fof(f44488_D,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f44491,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)),
  inference(subsumption_resolution,[],[f44420,f1603])).
fof(f1603,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))),
  inference(unit_resulting_resolution,[],[f840,f952])).
fof(f952,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f677])).
fof(f840,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))))),
  inference(cnf_transformation,[],[f670])).
fof(f44420,plain,(
  s(bool,'T_0') != s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)),
  inference(superposition,[],[f1592,f1125])).
fof(f1125,plain,(
  ( ! [X0,X1] : (s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f601])).
fof(f601,plain,(
  ! [X0,X1] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f600])).
fof(f600,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) | s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(ennf_transformation,[],[f473])).
fof(f473,plain,(
  ! [X0,X1] : ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(flattening,[],[f472])).
fof(f472,plain,(
  ! [X0,X1] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X1))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6,X7] : ((~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) & ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6)) => s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',X7))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7))))))),
  file('thm.extreal.le_addr_imp','thm.extreal.le_addr')).
fof(f44490,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)),
  inference(subsumption_resolution,[],[f44489,f1591])).
fof(f44489,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)),
  inference(forward_demodulation,[],[f44419,f1603])).
fof(f44419,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)),
  inference(superposition,[],[f1639,f1125])).
fof(f44484,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)),
  inference(subsumption_resolution,[],[f44483,f1591])).
fof(f44483,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)),
  inference(forward_demodulation,[],[f44410,f1603])).
fof(f44410,plain,(
  s(bool,'F_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',sK12) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',sK12)),
  inference(superposition,[],[f1125,f1639])).
fof(f33447,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24939,f32303])).
fof(f32303,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) != s(X0,X2)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1229])).
fof(f24939,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f956])).
fof(f33446,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24941,f32303])).
fof(f24941,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f24933,f1126])).
fof(f1126,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f603])).
fof(f603,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f602])).
fof(f602,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f474])).
fof(f474,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X23,X24] : ((p(s(bool,X23)) => p(s(bool,X24))) => ((p(s(bool,X24)) => p(s(bool,X23))) => s(bool,X23) = s(bool,X24)))),
  file('thm.extreal.le_addr_imp','thm.bool.IMP_ANTISYM_AX')).
fof(f33445,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24942,f32303])).
fof(f24942,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1359,f24933,f1126])).
fof(f33444,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24947,f1359])).
fof(f24947,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1591,f24933,f1126])).
fof(f33443,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24949,f24933])).
fof(f24949,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1592,f24933,f1126])).
fof(f33442,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24951,f32303])).
fof(f24951,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f24933,f1126])).
fof(f33441,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24952,f32303])).
fof(f24952,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1359,f24933,f1126])).
fof(f33440,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24957,f1359])).
fof(f24957,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1591,f24933,f1126])).
fof(f33439,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24959,f24933])).
fof(f24959,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1592,f24933,f1126])).
fof(f33438,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24961,f1359])).
fof(f24961,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1667,f24933,f1205])).
fof(f1205,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f795])).
fof(f795,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f794])).
fof(f794,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f793])).
fof(f793,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f654])).
fof(f654,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f1667,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1359,f842,f1207])).
fof(f1207,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f795])).
fof(f33437,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24962,f24933])).
fof(f24962,plain,(
  p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f18715,f24933,f1205])).
fof(f18715,plain,(
  sP1('T_0',happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')))),
  inference(unit_resulting_resolution,[],[f842,f13946,f1207])).
fof(f33436,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24963,f24933])).
fof(f24963,plain,(
  p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f19736,f24933,f1205])).
fof(f19736,plain,(
  sP1('T_0',happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')))),
  inference(unit_resulting_resolution,[],[f842,f13947,f1207])).
fof(f33435,plain,(
  $false | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f24964,f24933])).
fof(f24964,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f20988,f24933,f1205])).
fof(f20988,plain,(
  sP1('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f842,f20970,f1207])).
fof(f33434,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24966,f24933])).
fof(f24966,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1715,f24933,f1205])).
fof(f1715,plain,(
  sP1('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))),
  inference(unit_resulting_resolution,[],[f1359,f1681])).
fof(f1681,plain,(
  ( ! [X0] : (sP1(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1678,f842])).
fof(f1678,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) )),
  inference(superposition,[],[f1210,f1603])).
fof(f1210,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f795])).
fof(f33433,plain,(
  $false | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f24967,f24933])).
fof(f24967,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f22308,f24933,f1205])).
fof(f22308,plain,(
  sP1('F_0','const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | $spl13),
  inference(unit_resulting_resolution,[],[f1359,f20973,f1210])).
fof(f33432,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24968,f24933])).
fof(f24968,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1890,f24933,f1205])).
fof(f1890,plain,(
  sP1('F_0',sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f1724,f1210])).
fof(f1724,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f1041,f1066])).
fof(f33431,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24969,f24933])).
fof(f24969,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6903,f24933,f1205])).
fof(f6903,plain,(
  sP1('F_0',sK27('const.arithmetic.ZERO_0'))),
  inference(unit_resulting_resolution,[],[f1359,f6174,f1210])).
fof(f33430,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24970,f24933])).
fof(f24970,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1735,f24933,f1205])).
fof(f1735,plain,(
  sP1('F_0',sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f1721,f1210])).
fof(f1721,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f845,f1066])).
fof(f845,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f122])).
fof(f122,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.extreal.le_addr_imp','thm.prim_rec.PRE|split|0')).
fof(f33429,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24971,f24933])).
fof(f24971,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1775,f24933,f1205])).
fof(f1775,plain,(
  sP1('F_0',sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f1722,f1210])).
fof(f1722,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(unit_resulting_resolution,[],[f846,f1066])).
fof(f846,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(cnf_transformation,[],[f248])).
fof(f248,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_size_def|split|1')).
fof(f33428,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24972,f24933])).
fof(f24972,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1826,f24933,f1205])).
fof(f1826,plain,(
  sP1('F_0',sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f1723,f1210])).
fof(f1723,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))),
  inference(unit_resulting_resolution,[],[f849,f1066])).
fof(f849,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f246])).
fof(f246,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_size_def|split|0')).
fof(f33427,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24973,f24933])).
fof(f24973,plain,(
  p(s(bool,sK28('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2526,f24933,f1205])).
fof(f2526,plain,(
  sP1('F_0',sK28('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f1756,f1210])).
fof(f1756,plain,(
  p(s(bool,sK28('const.combin.I_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f1041,f1069])).
fof(f33426,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24974,f24933])).
fof(f24974,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6221,f24933,f1205])).
fof(f6221,plain,(
  sP1('F_0',sK28('const.arithmetic.ZERO_0'))),
  inference(unit_resulting_resolution,[],[f1359,f6173,f1210])).
fof(f33425,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24975,f24933])).
fof(f24975,plain,(
  p(s(bool,sK28('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2064,f24933,f1205])).
fof(f2064,plain,(
  sP1('F_0',sK28('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f1753,f1210])).
fof(f1753,plain,(
  p(s(bool,sK28('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f845,f1069])).
fof(f33424,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24976,f24933])).
fof(f24976,plain,(
  p(s(bool,sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2216,f24933,f1205])).
fof(f2216,plain,(
  sP1('F_0',sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f1754,f1210])).
fof(f1754,plain,(
  p(s(bool,sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(unit_resulting_resolution,[],[f846,f1069])).
fof(f33423,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24977,f24933])).
fof(f24977,plain,(
  p(s(bool,sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2354,f24933,f1205])).
fof(f2354,plain,(
  sP1('F_0',sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f1755,f1210])).
fof(f1755,plain,(
  p(s(bool,sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))),
  inference(unit_resulting_resolution,[],[f849,f1069])).
fof(f33422,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24978,f24933])).
fof(f24978,plain,(
  p(s(bool,sK29('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3258,f24933,f1205])).
fof(f3258,plain,(
  sP1('F_0',sK29('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f3145,f1210])).
fof(f3145,plain,(
  p(s(bool,sK29('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(equality_resolution,[],[f1072])).
fof(f1072,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK29(X0)))) )),
  inference(cnf_transformation,[],[f740])).
fof(f740,plain,(
  ! [X0] : ((~p(s(bool,sK29(X0))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK29(X0)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,sK29(X0)),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK29])],[f739])).
fof(f739,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(flattening,[],[f738])).
fof(f738,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X1)))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(nnf_transformation,[],[f429])).
fof(f429,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X1),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  inference(rectify,[],[f240])).
fof(f240,axiom,(
  ! [X6] : ? [X47] : ((p(s(bool,X47)) <=> s('type.realax.real',X6) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & s('type.extreal.extreal','const.extreal.extreal_mul_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X6))),s('type.extreal.extreal','const.extreal.NegInf_0'))) = s('type.extreal.extreal','const.bool.COND_3'(s(bool,X47),s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.extreal.extreal','const.bool.COND_3'(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6))),s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_mul_def|split|4')).
fof(f33421,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24979,f24933])).
fof(f24979,plain,(
  p(s(bool,sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3866,f24933,f1205])).
fof(f3866,plain,(
  sP1('F_0',sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f3157,f1210])).
fof(f33420,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24980,f24933])).
fof(f24980,plain,(
  p(s(bool,sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f4053,f24933,f1205])).
fof(f4053,plain,(
  sP1('F_0',sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f3233,f1210])).
fof(f33419,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24981,f24933])).
fof(f24981,plain,(
  p(s(bool,sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f4317,f24933,f1205])).
fof(f4317,plain,(
  sP1('F_0',sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f3245,f1210])).
fof(f33418,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24982,f24933])).
fof(f24982,plain,(
  p(s(bool,sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2766,f24933,f1205])).
fof(f2766,plain,(
  sP1('F_0',sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f1761,f1210])).
fof(f1761,plain,(
  p(s(bool,sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(unit_resulting_resolution,[],[f848,f1084])).
fof(f1084,plain,(
  ( ! [X0] : (p(s(bool,sK33(X0))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f752])).
fof(f752,plain,(
  ! [X0] : ((~p(s(bool,sK33(X0))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | p(s(bool,sK33(X0)))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,sK33(X0)),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK33])],[f751])).
fof(f751,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | p(s(bool,X1))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(flattening,[],[f750])).
fof(f750,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))) & ((s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) & s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) | p(s(bool,X1)))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(nnf_transformation,[],[f433])).
fof(f433,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X0))) = s('type.realax.real','const.bool.COND_3'(s(bool,X1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X0))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  inference(rectify,[],[f241])).
fof(f241,axiom,(
  ! [X6] : ? [X47] : ((p(s(bool,X47)) <=> (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6))) & s('type.realax.real','const.extreal.real_1'(s('type.extreal.extreal',X6))) = s('type.realax.real','const.bool.COND_3'(s(bool,X47),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.min.@_1'(s(fun('type.realax.real',bool),'_dst_x0x1_2'(s(fun('type.extreal.extreal',bool),'=_1'(s('type.extreal.extreal',X6))),s(fun('type.realax.real','type.extreal.extreal'),'const.extreal.Normal_0'))))))))),
  file('thm.extreal.le_addr_imp','thm.extreal.real_def')).
fof(f848,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(cnf_transformation,[],[f230])).
fof(f230,axiom,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_ainv_def|split|1')).
fof(f33417,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24983,f24933])).
fof(f24983,plain,(
  p(s(bool,sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3086,f24933,f1205])).
fof(f3086,plain,(
  sP1('F_0',sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f1787,f1210])).
fof(f1787,plain,(
  p(s(bool,sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))),
  inference(unit_resulting_resolution,[],[f847,f1085])).
fof(f1085,plain,(
  ( ! [X0] : (p(s(bool,sK33(X0))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f752])).
fof(f847,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f231])).
fof(f231,axiom,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_ainv_def|split|0')).
fof(f33416,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24984,f24933])).
fof(f24984,plain,(
  p(s(bool,sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2932,f24933,f1205])).
fof(f2932,plain,(
  sP1('F_0',sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f1762,f1210])).
fof(f1762,plain,(
  p(s(bool,sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))))),
  inference(unit_resulting_resolution,[],[f1041,f1084])).
fof(f33415,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24985,f24933])).
fof(f24985,plain,(
  p(s(bool,sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3172,f24933,f1205])).
fof(f3172,plain,(
  sP1('F_0',sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(unit_resulting_resolution,[],[f1359,f1788,f1210])).
fof(f1788,plain,(
  p(s(bool,sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(unit_resulting_resolution,[],[f1041,f1085])).
fof(f33414,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24986,f1359])).
fof(f24986,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1674,f24933,f1205])).
fof(f1674,plain,(
  sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1671])).
fof(f1671,plain,(
  ( ! [X0] : (sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),X0) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1668,f842])).
fof(f1668,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),X0)) )),
  inference(superposition,[],[f1207,f1603])).
fof(f33413,plain,(
  $false | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f24990,f1359])).
fof(f24990,plain,(
  p(s(bool,'F_0')) | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f22307,f24933,f1205])).
fof(f22307,plain,(
  sP1('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'F_0') | $spl13),
  inference(unit_resulting_resolution,[],[f1359,f20973,f1207])).
fof(f33412,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24991,f1359])).
fof(f24991,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1889,f24933,f1205])).
fof(f1889,plain,(
  sP1(sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1724,f1207])).
fof(f33411,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24992,f1359])).
fof(f24992,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6902,f24933,f1205])).
fof(f6902,plain,(
  sP1(sK27('const.arithmetic.ZERO_0'),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f6174,f1207])).
fof(f33410,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24993,f1359])).
fof(f24993,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1734,f24933,f1205])).
fof(f1734,plain,(
  sP1(sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1721,f1207])).
fof(f33409,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24994,f1359])).
fof(f24994,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1774,f24933,f1205])).
fof(f1774,plain,(
  sP1(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1722,f1207])).
fof(f33408,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24995,f1359])).
fof(f24995,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1825,f24933,f1205])).
fof(f1825,plain,(
  sP1(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1723,f1207])).
fof(f33407,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24996,f1359])).
fof(f24996,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f2525,f24933,f1205])).
fof(f2525,plain,(
  sP1(sK28('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1756,f1207])).
fof(f33406,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24997,f1359])).
fof(f24997,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6220,f24933,f1205])).
fof(f6220,plain,(
  sP1(sK28('const.arithmetic.ZERO_0'),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f6173,f1207])).
fof(f33405,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24998,f1359])).
fof(f24998,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f2063,f24933,f1205])).
fof(f2063,plain,(
  sP1(sK28('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1753,f1207])).
fof(f33404,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f24999,f1359])).
fof(f24999,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f2215,f24933,f1205])).
fof(f2215,plain,(
  sP1(sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1754,f1207])).
fof(f33403,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25000,f1359])).
fof(f25000,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f2353,f24933,f1205])).
fof(f2353,plain,(
  sP1(sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1755,f1207])).
fof(f33402,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25001,f1359])).
fof(f25001,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f3257,f24933,f1205])).
fof(f3257,plain,(
  sP1(sK29('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f3145,f1207])).
fof(f33401,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25002,f1359])).
fof(f25002,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f3865,f24933,f1205])).
fof(f3865,plain,(
  sP1(sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f3157,f1207])).
fof(f33400,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25003,f1359])).
fof(f25003,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f4052,f24933,f1205])).
fof(f4052,plain,(
  sP1(sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f3233,f1207])).
fof(f33399,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25004,f1359])).
fof(f25004,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f4316,f24933,f1205])).
fof(f4316,plain,(
  sP1(sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f3245,f1207])).
fof(f33398,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25005,f1359])).
fof(f25005,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f2765,f24933,f1205])).
fof(f2765,plain,(
  sP1(sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1761,f1207])).
fof(f33397,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25006,f1359])).
fof(f25006,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f3085,f24933,f1205])).
fof(f3085,plain,(
  sP1(sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1787,f1207])).
fof(f33396,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25007,f1359])).
fof(f25007,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f2931,f24933,f1205])).
fof(f2931,plain,(
  sP1(sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1762,f1207])).
fof(f33395,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25008,f1359])).
fof(f25008,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f3171,f24933,f1205])).
fof(f3171,plain,(
  sP1(sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),'F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1788,f1207])).
fof(f33394,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25017,f1359])).
fof(f25017,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1677,f24933,f1205])).
fof(f1677,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1359,f842,f1210])).
fof(f33393,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25018,f1359])).
fof(f25018,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1715,f24933,f1205])).
fof(f33392,plain,(
  $false | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f25019,f1359])).
fof(f25019,plain,(
  p(s(bool,'F_0')) | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f22308,f24933,f1205])).
fof(f33391,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25020,f1359])).
fof(f25020,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1890,f24933,f1205])).
fof(f33390,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25021,f1359])).
fof(f25021,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6903,f24933,f1205])).
fof(f33389,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25022,f1359])).
fof(f25022,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1735,f24933,f1205])).
fof(f33388,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25023,f1359])).
fof(f25023,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1775,f24933,f1205])).
fof(f33387,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25024,f1359])).
fof(f25024,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1826,f24933,f1205])).
fof(f33386,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25025,f1359])).
fof(f25025,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f2526,f24933,f1205])).
fof(f33385,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25026,f1359])).
fof(f25026,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6221,f24933,f1205])).
fof(f33384,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25027,f1359])).
fof(f25027,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f2064,f24933,f1205])).
fof(f33383,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25028,f1359])).
fof(f25028,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f2216,f24933,f1205])).
fof(f33382,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25029,f1359])).
fof(f25029,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f2354,f24933,f1205])).
fof(f33381,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25030,f1359])).
fof(f25030,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f3258,f24933,f1205])).
fof(f33380,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25031,f1359])).
fof(f25031,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f3866,f24933,f1205])).
fof(f33379,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25032,f1359])).
fof(f25032,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f4053,f24933,f1205])).
fof(f33378,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25033,f1359])).
fof(f25033,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f4317,f24933,f1205])).
fof(f33377,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25034,f1359])).
fof(f25034,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f2766,f24933,f1205])).
fof(f33376,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25035,f1359])).
fof(f25035,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f3086,f24933,f1205])).
fof(f33375,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25036,f1359])).
fof(f25036,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f2932,f24933,f1205])).
fof(f33374,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25037,f1359])).
fof(f25037,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f3172,f24933,f1205])).
fof(f33373,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25038,f24933])).
fof(f25038,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1674,f24933,f1205])).
fof(f33372,plain,(
  $false | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f25039,f24933])).
fof(f25039,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f20989,f24933,f1205])).
fof(f20989,plain,(
  sP1('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),'T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f842,f20970,f1210])).
fof(f33371,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25040,f24933])).
fof(f25040,plain,(
  p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f18716,f24933,f1205])).
fof(f18716,plain,(
  sP1(happ(s(fun(bool,bool),'=_1'(s(bool,'F_0'))),s(bool,'T_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f842,f13946,f1210])).
fof(f33370,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25041,f24933])).
fof(f25041,plain,(
  p(s(bool,happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f19737,f24933,f1205])).
fof(f19737,plain,(
  sP1(happ(s(fun(bool,bool),'=_1'(s(bool,'T_0'))),s(bool,'F_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f842,f13947,f1210])).
fof(f33369,plain,(
  $false | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f25042,f24933])).
fof(f25042,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f22307,f24933,f1205])).
fof(f33368,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25043,f24933])).
fof(f25043,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1889,f24933,f1205])).
fof(f33367,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25044,f24933])).
fof(f25044,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6902,f24933,f1205])).
fof(f33366,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25045,f24933])).
fof(f25045,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1734,f24933,f1205])).
fof(f33365,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25046,f24933])).
fof(f25046,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1774,f24933,f1205])).
fof(f33364,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25047,f24933])).
fof(f25047,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1825,f24933,f1205])).
fof(f33363,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25048,f24933])).
fof(f25048,plain,(
  p(s(bool,sK28('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2525,f24933,f1205])).
fof(f33362,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25049,f24933])).
fof(f25049,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6220,f24933,f1205])).
fof(f33361,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25050,f24933])).
fof(f25050,plain,(
  p(s(bool,sK28('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2063,f24933,f1205])).
fof(f33360,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25051,f24933])).
fof(f25051,plain,(
  p(s(bool,sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2215,f24933,f1205])).
fof(f33359,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25052,f24933])).
fof(f25052,plain,(
  p(s(bool,sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2353,f24933,f1205])).
fof(f33358,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25053,f24933])).
fof(f25053,plain,(
  p(s(bool,sK29('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3257,f24933,f1205])).
fof(f33357,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25054,f24933])).
fof(f25054,plain,(
  p(s(bool,sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3865,f24933,f1205])).
fof(f33356,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25055,f24933])).
fof(f25055,plain,(
  p(s(bool,sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f4052,f24933,f1205])).
fof(f33355,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25056,f24933])).
fof(f25056,plain,(
  p(s(bool,sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f4316,f24933,f1205])).
fof(f33354,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25057,f24933])).
fof(f25057,plain,(
  p(s(bool,sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2765,f24933,f1205])).
fof(f33353,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25058,f24933])).
fof(f25058,plain,(
  p(s(bool,sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3085,f24933,f1205])).
fof(f33352,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25059,f24933])).
fof(f25059,plain,(
  p(s(bool,sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2931,f24933,f1205])).
fof(f33351,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25060,f24933])).
fof(f25060,plain,(
  p(s(bool,sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3171,f24933,f1205])).
fof(f33350,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25195,f27837])).
fof(f27837,plain,(
  ( ! [X0] : (sP5('T_0','T_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f1401,f1410,f24933,f1242])).
fof(f1410,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f842,f1248])).
fof(f1248,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f811])).
fof(f1401,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f842,f1244])).
fof(f1419,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f842,f1252])).
fof(f1252,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f813])).
fof(f25195,plain,(
  ( ! [X0] : (~sP5('T_0','T_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f842,f842,f24933,f1239])).
fof(f1239,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f805])).
fof(f33349,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25197,f30348])).
fof(f30348,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1250])).
fof(f25197,plain,(
  ~sP3('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1401,f24933,f1241])).
fof(f3451,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f842,f1591,f1257])).
fof(f1422,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f842,f1253])).
fof(f33348,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25198,f30348])).
fof(f25198,plain,(
  ~sP3('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1401,f24933,f1241])).
fof(f33347,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25204,f30348])).
fof(f25204,plain,(
  ~sP3('T_0','F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f1401,f24933,f1241])).
fof(f7025,plain,(
  ~sP5(sK27('const.arithmetic.ZERO_0'),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1591,f6174,f1257])).
fof(f33346,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25205,f30348])).
fof(f25205,plain,(
  ~sP3('T_0','F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f1401,f24933,f1241])).
fof(f6341,plain,(
  ~sP5(sK28('const.arithmetic.ZERO_0'),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f1591,f6173,f1257])).
fof(f33345,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25208,f30348])).
fof(f25208,plain,(
  ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'const.combin.I_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1423,f3482,f1402,f24933,f1241])).
fof(f1402,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) )),
  inference(unit_resulting_resolution,[],[f840,f1244])).
fof(f1423,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f840,f1253])).
fof(f33344,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25209,f30348])).
fof(f25209,plain,(
  ~sP3('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1460,f3482,f1457,f24933,f1241])).
fof(f1457,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f850,f1244])).
fof(f850,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  inference(cnf_transformation,[],[f210])).
fof(f210,axiom,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0'))))),
  file('thm.extreal.le_addr_imp','thm.extreal.lt_infty|split|2')).
fof(f1460,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f850,f1253])).
fof(f33343,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25210,f30348])).
fof(f25210,plain,(
  ~sP3('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1424,f3482,f1403,f24933,f1241])).
fof(f1403,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f844,f1244])).
fof(f844,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f149])).
fof(f149,axiom,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.extreal.le_addr_imp','thm.numeral.numeral_distrib|split|35')).
fof(f1424,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f844,f1253])).
fof(f33342,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25211,f30348])).
fof(f25211,plain,(
  ~sP3(sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),'const.combin.I_1'(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1894,f3482,f1891,f24933,f1241])).
fof(f1891,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f1724,f1244])).
fof(f1894,plain,(
  ( ! [X0,X1] : (sP2(X0,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),X1)) )),
  inference(unit_resulting_resolution,[],[f1724,f1253])).
fof(f33341,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25212,f30348])).
fof(f25212,plain,(
  ~sP3(sK27('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f6948,f24933,f1241])).
fof(f6948,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK27('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f6174,f1244])).
fof(f6906,plain,(
  ~sP5('T_0',sK27('const.arithmetic.ZERO_0'),'F_0')),
  inference(unit_resulting_resolution,[],[f842,f1359,f6174,f1239])).
fof(f6985,plain,(
  ( ! [X0,X1] : (sP2(X0,sK27('const.arithmetic.ZERO_0'),X1)) )),
  inference(unit_resulting_resolution,[],[f6174,f1253])).
fof(f33340,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25215,f30348])).
fof(f25215,plain,(
  ~sP3(sK27('const.arithmetic.ZERO_0'),'const.combin.I_1'(s(bool,sK27('const.arithmetic.ZERO_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f3482,f6948,f24933,f1241])).
fof(f33339,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25216,f30348])).
fof(f25216,plain,(
  ~sP3(sK27('const.arithmetic.ZERO_0'),'F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f6948,f24933,f1241])).
fof(f6907,plain,(
  ~sP5(sK27('const.arithmetic.ZERO_0'),sK27('const.arithmetic.ZERO_0'),'F_0')),
  inference(unit_resulting_resolution,[],[f6174,f1359,f6174,f1239])).
fof(f33338,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25217,f30348])).
fof(f25217,plain,(
  ~sP3(sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),'const.combin.I_1'(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1739,f3482,f1736,f24933,f1241])).
fof(f1736,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f1721,f1244])).
fof(f1739,plain,(
  ( ! [X0,X1] : (sP2(X0,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),X1)) )),
  inference(unit_resulting_resolution,[],[f1721,f1253])).
fof(f33337,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25218,f30348])).
fof(f25218,plain,(
  ~sP3(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),'const.combin.I_1'(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1779,f3482,f1776,f24933,f1241])).
fof(f1776,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1722,f1244])).
fof(f1779,plain,(
  ( ! [X0,X1] : (sP2(X0,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),X1)) )),
  inference(unit_resulting_resolution,[],[f1722,f1253])).
fof(f33336,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25219,f30348])).
fof(f25219,plain,(
  ~sP3(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'const.combin.I_1'(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1830,f3482,f1827,f24933,f1241])).
fof(f1827,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1723,f1244])).
fof(f1830,plain,(
  ( ! [X0,X1] : (sP2(X0,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),X1)) )),
  inference(unit_resulting_resolution,[],[f1723,f1253])).
fof(f33335,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25220,f30348])).
fof(f25220,plain,(
  ~sP3(sK28('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f6266,f24933,f1241])).
fof(f6266,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,sK28('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f6173,f1244])).
fof(f6224,plain,(
  ~sP5('T_0',sK28('const.arithmetic.ZERO_0'),'F_0')),
  inference(unit_resulting_resolution,[],[f842,f1359,f6173,f1239])).
fof(f6303,plain,(
  ( ! [X0,X1] : (sP2(X0,sK28('const.arithmetic.ZERO_0'),X1)) )),
  inference(unit_resulting_resolution,[],[f6173,f1253])).
fof(f33334,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25223,f30348])).
fof(f25223,plain,(
  ~sP3(sK28('const.arithmetic.ZERO_0'),'const.combin.I_1'(s(bool,sK28('const.arithmetic.ZERO_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f3482,f6266,f24933,f1241])).
fof(f33333,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25224,f30348])).
fof(f25224,plain,(
  ~sP3(sK28('const.arithmetic.ZERO_0'),'F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f6266,f24933,f1241])).
fof(f6225,plain,(
  ~sP5(sK28('const.arithmetic.ZERO_0'),sK28('const.arithmetic.ZERO_0'),'F_0')),
  inference(unit_resulting_resolution,[],[f6173,f1359,f6173,f1239])).
fof(f33332,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25225,f30348])).
fof(f25225,plain,(
  ~sP3('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1404,f24933,f1241])).
fof(f33331,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25226,f30348])).
fof(f25226,plain,(
  ~sP3(sK27('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f1404,f24933,f1241])).
fof(f33330,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25227,f30348])).
fof(f25227,plain,(
  ~sP3(sK28('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f1404,f24933,f1241])).
fof(f33329,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25228,f30348])).
fof(f25228,plain,(
  ~sP3('T_0','F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f1404,f24933,f1241])).
fof(f33328,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25229,f30348])).
fof(f25229,plain,(
  ~sP3(sK27('const.arithmetic.ZERO_0'),'F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f1404,f24933,f1241])).
fof(f33327,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25230,f30348])).
fof(f25230,plain,(
  ~sP3('T_0','F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f1404,f24933,f1241])).
fof(f33326,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25231,f30348])).
fof(f25231,plain,(
  ~sP3(sK28('const.arithmetic.ZERO_0'),'F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f1404,f24933,f1241])).
fof(f33325,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25232,f30348])).
fof(f25232,plain,(
  ~sP3('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1406,f24933,f1241])).
fof(f33324,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25234,f30348])).
fof(f25234,plain,(
  ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'const.combin.I_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1423,f3482,f1407,f24933,f1241])).
fof(f33323,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25235,f30348])).
fof(f25235,plain,(
  ~sP3('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1460,f3482,f1407,f24933,f1241])).
fof(f33322,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25236,f30348])).
fof(f25236,plain,(
  ~sP3('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1424,f3482,f1407,f24933,f1241])).
fof(f33321,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25237,f30348])).
fof(f25237,plain,(
  ~sP3(sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),'const.combin.I_1'(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1894,f3482,f1407,f24933,f1241])).
fof(f33320,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25238,f30348])).
fof(f25238,plain,(
  ~sP3(sK27('const.arithmetic.ZERO_0'),'const.combin.I_1'(s(bool,sK27('const.arithmetic.ZERO_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f3482,f1407,f24933,f1241])).
fof(f33319,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25239,f30348])).
fof(f25239,plain,(
  ~sP3(sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),'const.combin.I_1'(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1739,f3482,f1407,f24933,f1241])).
fof(f33318,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25240,f30348])).
fof(f25240,plain,(
  ~sP3(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),'const.combin.I_1'(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1779,f3482,f1407,f24933,f1241])).
fof(f33317,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25241,f30348])).
fof(f25241,plain,(
  ~sP3(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'const.combin.I_1'(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1830,f3482,f1407,f24933,f1241])).
fof(f33316,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25242,f30348])).
fof(f25242,plain,(
  ~sP3(sK28('const.arithmetic.ZERO_0'),'const.combin.I_1'(s(bool,sK28('const.arithmetic.ZERO_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f3482,f1407,f24933,f1241])).
fof(f33315,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25250,f30348])).
fof(f25250,plain,(
  ~sP3('const.combin.I_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f3470,f1407,f24933,f1241])).
fof(f1426,plain,(
  ( ! [X0,X1] : (sP2('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f840,f1254])).
fof(f33314,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25251,f30348])).
fof(f25251,plain,(
  ~sP3('const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1461,f3470,f1407,f24933,f1241])).
fof(f1461,plain,(
  ( ! [X0,X1] : (sP2('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f850,f1254])).
fof(f33313,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25252,f30348])).
fof(f25252,plain,(
  ~sP3('const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1427,f3470,f1407,f24933,f1241])).
fof(f1427,plain,(
  ( ! [X0,X1] : (sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f844,f1254])).
fof(f33312,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25253,f30348])).
fof(f25253,plain,(
  ~sP3('const.combin.I_1'(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))),sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1895,f3470,f1407,f24933,f1241])).
fof(f1895,plain,(
  ( ! [X0,X1] : (sP2(sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1724,f1254])).
fof(f33311,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25254,f30348])).
fof(f25254,plain,(
  ~sP3('const.combin.I_1'(s(bool,sK27('const.arithmetic.ZERO_0'))),sK27('const.arithmetic.ZERO_0'),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6986,f3470,f1407,f24933,f1241])).
fof(f6986,plain,(
  ( ! [X0,X1] : (sP2(sK27('const.arithmetic.ZERO_0'),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f6174,f1254])).
fof(f33310,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25255,f30348])).
fof(f25255,plain,(
  ~sP3('const.combin.I_1'(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))),sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1740,f3470,f1407,f24933,f1241])).
fof(f1740,plain,(
  ( ! [X0,X1] : (sP2(sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1721,f1254])).
fof(f33309,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25256,f30348])).
fof(f25256,plain,(
  ~sP3('const.combin.I_1'(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))),sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1780,f3470,f1407,f24933,f1241])).
fof(f1780,plain,(
  ( ! [X0,X1] : (sP2(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1722,f1254])).
fof(f33308,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25257,f30348])).
fof(f25257,plain,(
  ~sP3('const.combin.I_1'(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1831,f3470,f1407,f24933,f1241])).
fof(f1831,plain,(
  ( ! [X0,X1] : (sP2(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1723,f1254])).
fof(f33307,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25258,f30348])).
fof(f25258,plain,(
  ~sP3('const.combin.I_1'(s(bool,sK28('const.arithmetic.ZERO_0'))),sK28('const.arithmetic.ZERO_0'),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6304,f3470,f1407,f24933,f1241])).
fof(f6304,plain,(
  ( ! [X0,X1] : (sP2(sK28('const.arithmetic.ZERO_0'),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f6173,f1254])).
fof(f33306,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25264,f30316])).
fof(f30316,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1246])).
fof(f25264,plain,(
  ~sP4('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1410,f24933,f1241])).
fof(f33305,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25265,f30316])).
fof(f25265,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1410,f24933,f1241])).
fof(f33304,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25270,f30316])).
fof(f25270,plain,(
  ~sP4('T_0','F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f1410,f24933,f1241])).
fof(f33303,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f25271,f30316])).
fof(f25271,plain,(
  ~sP4('T_0','F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f1410,f24933,f1241])).
fof(f33302,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f26130,f30316])).
fof(f26130,plain,(
  ~sP4(sK27('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f6983,f24933,f1241])).
fof(f6983,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK27('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f6174,f1248])).
fof(f33301,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f26133,f30316])).
fof(f26133,plain,(
  ~sP4(sK27('const.arithmetic.ZERO_0'),'F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f6983,f24933,f1241])).
fof(f33300,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f26798,f30316])).
fof(f26798,plain,(
  ~sP4(sK28('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f6301,f24933,f1241])).
fof(f6301,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK28('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f6173,f1248])).
fof(f33299,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f26801,f30316])).
fof(f26801,plain,(
  ~sP4(sK28('const.arithmetic.ZERO_0'),'F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f6301,f24933,f1241])).
fof(f33298,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f26974,f30316])).
fof(f26974,plain,(
  ~sP4('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1413,f24933,f1241])).
fof(f1413,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1359,f1249])).
fof(f33297,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f26975,f30316])).
fof(f26975,plain,(
  ~sP4('T_0','F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f1413,f24933,f1241])).
fof(f33296,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f26976,f30316])).
fof(f26976,plain,(
  ~sP4('T_0','F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f1413,f24933,f1241])).
fof(f33295,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f26977,f30316])).
fof(f26977,plain,(
  ~sP4(sK27('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f1413,f24933,f1241])).
fof(f33294,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f26978,f30316])).
fof(f26978,plain,(
  ~sP4(sK27('const.arithmetic.ZERO_0'),'F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f1413,f24933,f1241])).
fof(f33293,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f26979,f30316])).
fof(f26979,plain,(
  ~sP4(sK28('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f1413,f24933,f1241])).
fof(f33292,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f26980,f30316])).
fof(f26980,plain,(
  ~sP4(sK28('const.arithmetic.ZERO_0'),'F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f1413,f24933,f1241])).
fof(f33291,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27186,f30316])).
fof(f27186,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1415,f24933,f1241])).
fof(f33290,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27739,f30348])).
fof(f27739,plain,(
  ~sP3('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1401,f24933,f1242])).
fof(f33289,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27740,f30348])).
fof(f27740,plain,(
  ~sP3('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1401,f24933,f1242])).
fof(f33288,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27746,f30348])).
fof(f27746,plain,(
  ~sP3('T_0','F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f1401,f24933,f1242])).
fof(f33287,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27747,f30348])).
fof(f27747,plain,(
  ~sP3('T_0','F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f1401,f24933,f1242])).
fof(f33286,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27750,f30348])).
fof(f27750,plain,(
  ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'const.combin.I_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1423,f3482,f1402,f24933,f1242])).
fof(f33285,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27751,f30348])).
fof(f27751,plain,(
  ~sP3('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1460,f3482,f1457,f24933,f1242])).
fof(f33284,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27752,f30348])).
fof(f27752,plain,(
  ~sP3('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1424,f3482,f1403,f24933,f1242])).
fof(f33283,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27753,f30348])).
fof(f27753,plain,(
  ~sP3(sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),'const.combin.I_1'(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1894,f3482,f1891,f24933,f1242])).
fof(f33282,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27754,f30348])).
fof(f27754,plain,(
  ~sP3(sK27('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f6948,f24933,f1242])).
fof(f33281,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27757,f30348])).
fof(f27757,plain,(
  ~sP3(sK27('const.arithmetic.ZERO_0'),'const.combin.I_1'(s(bool,sK27('const.arithmetic.ZERO_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f3482,f6948,f24933,f1242])).
fof(f33280,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27758,f30348])).
fof(f27758,plain,(
  ~sP3(sK27('const.arithmetic.ZERO_0'),'F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f6948,f24933,f1242])).
fof(f33279,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27759,f30348])).
fof(f27759,plain,(
  ~sP3(sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),'const.combin.I_1'(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1739,f3482,f1736,f24933,f1242])).
fof(f33278,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27760,f30348])).
fof(f27760,plain,(
  ~sP3(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),'const.combin.I_1'(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1779,f3482,f1776,f24933,f1242])).
fof(f33277,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27761,f30348])).
fof(f27761,plain,(
  ~sP3(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'const.combin.I_1'(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1830,f3482,f1827,f24933,f1242])).
fof(f33276,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27762,f30348])).
fof(f27762,plain,(
  ~sP3(sK28('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f6266,f24933,f1242])).
fof(f33275,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27765,f30348])).
fof(f27765,plain,(
  ~sP3(sK28('const.arithmetic.ZERO_0'),'const.combin.I_1'(s(bool,sK28('const.arithmetic.ZERO_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f3482,f6266,f24933,f1242])).
fof(f33274,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27766,f30348])).
fof(f27766,plain,(
  ~sP3(sK28('const.arithmetic.ZERO_0'),'F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f6266,f24933,f1242])).
fof(f33273,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27767,f30348])).
fof(f27767,plain,(
  ~sP3('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1404,f24933,f1242])).
fof(f33272,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27768,f30348])).
fof(f27768,plain,(
  ~sP3(sK27('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f1404,f24933,f1242])).
fof(f33271,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27769,f30348])).
fof(f27769,plain,(
  ~sP3(sK28('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f1404,f24933,f1242])).
fof(f33270,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27770,f30348])).
fof(f27770,plain,(
  ~sP3('T_0','F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f1404,f24933,f1242])).
fof(f33269,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27771,f30348])).
fof(f27771,plain,(
  ~sP3(sK27('const.arithmetic.ZERO_0'),'F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f1404,f24933,f1242])).
fof(f33268,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27772,f30348])).
fof(f27772,plain,(
  ~sP3('T_0','F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f1404,f24933,f1242])).
fof(f33267,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27773,f30348])).
fof(f27773,plain,(
  ~sP3(sK28('const.arithmetic.ZERO_0'),'F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f1404,f24933,f1242])).
fof(f33266,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27774,f30348])).
fof(f27774,plain,(
  ~sP3('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1406,f24933,f1242])).
fof(f33265,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27776,f30348])).
fof(f27776,plain,(
  ~sP3('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'const.combin.I_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1423,f3482,f1407,f24933,f1242])).
fof(f33264,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27777,f30348])).
fof(f27777,plain,(
  ~sP3('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1460,f3482,f1407,f24933,f1242])).
fof(f33263,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27778,f30348])).
fof(f27778,plain,(
  ~sP3('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1424,f3482,f1407,f24933,f1242])).
fof(f33262,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27779,f30348])).
fof(f27779,plain,(
  ~sP3(sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),'const.combin.I_1'(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1894,f3482,f1407,f24933,f1242])).
fof(f33261,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27780,f30348])).
fof(f27780,plain,(
  ~sP3(sK27('const.arithmetic.ZERO_0'),'const.combin.I_1'(s(bool,sK27('const.arithmetic.ZERO_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f3482,f1407,f24933,f1242])).
fof(f33260,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27781,f30348])).
fof(f27781,plain,(
  ~sP3(sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),'const.combin.I_1'(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1739,f3482,f1407,f24933,f1242])).
fof(f33259,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27782,f30348])).
fof(f27782,plain,(
  ~sP3(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),'const.combin.I_1'(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1779,f3482,f1407,f24933,f1242])).
fof(f33258,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27783,f30348])).
fof(f27783,plain,(
  ~sP3(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'const.combin.I_1'(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1830,f3482,f1407,f24933,f1242])).
fof(f33257,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27784,f30348])).
fof(f27784,plain,(
  ~sP3(sK28('const.arithmetic.ZERO_0'),'const.combin.I_1'(s(bool,sK28('const.arithmetic.ZERO_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f3482,f1407,f24933,f1242])).
fof(f33256,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27792,f30348])).
fof(f27792,plain,(
  ~sP3('const.combin.I_1'(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f3470,f1407,f24933,f1242])).
fof(f33255,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27793,f30348])).
fof(f27793,plain,(
  ~sP3('const.combin.I_1'(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1461,f3470,f1407,f24933,f1242])).
fof(f33254,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27794,f30348])).
fof(f27794,plain,(
  ~sP3('const.combin.I_1'(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1427,f3470,f1407,f24933,f1242])).
fof(f33253,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27795,f30348])).
fof(f27795,plain,(
  ~sP3('const.combin.I_1'(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))),sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1895,f3470,f1407,f24933,f1242])).
fof(f33252,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27796,f30348])).
fof(f27796,plain,(
  ~sP3('const.combin.I_1'(s(bool,sK27('const.arithmetic.ZERO_0'))),sK27('const.arithmetic.ZERO_0'),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6986,f3470,f1407,f24933,f1242])).
fof(f33251,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27797,f30348])).
fof(f27797,plain,(
  ~sP3('const.combin.I_1'(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))),sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1740,f3470,f1407,f24933,f1242])).
fof(f33250,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27798,f30348])).
fof(f27798,plain,(
  ~sP3('const.combin.I_1'(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))),sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1780,f3470,f1407,f24933,f1242])).
fof(f33249,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27799,f30348])).
fof(f27799,plain,(
  ~sP3('const.combin.I_1'(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))),sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1831,f3470,f1407,f24933,f1242])).
fof(f33248,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27800,f30348])).
fof(f27800,plain,(
  ~sP3('const.combin.I_1'(s(bool,sK28('const.arithmetic.ZERO_0'))),sK28('const.arithmetic.ZERO_0'),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6304,f3470,f1407,f24933,f1242])).
fof(f33247,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27806,f30316])).
fof(f27806,plain,(
  ~sP4('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1410,f24933,f1242])).
fof(f33246,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27807,f30316])).
fof(f27807,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1410,f24933,f1242])).
fof(f33245,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27812,f30316])).
fof(f27812,plain,(
  ~sP4('T_0','F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f1410,f24933,f1242])).
fof(f33244,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f27813,f30316])).
fof(f27813,plain,(
  ~sP4('T_0','F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f1410,f24933,f1242])).
fof(f33243,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f28019,f30806])).
fof(f30806,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1603,f24933,f1258])).
fof(f28019,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f1401,f1411,f24933,f1242])).
fof(f1411,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) )),
  inference(unit_resulting_resolution,[],[f840,f1248])).
fof(f33242,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f28183,f30806])).
fof(f28183,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f1401,f1458,f24933,f1242])).
fof(f1458,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f850,f1248])).
fof(f33241,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f28347,f30806])).
fof(f28347,plain,(
  sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f1401,f1412,f24933,f1242])).
fof(f1412,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f844,f1248])).
fof(f33240,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f28511,f30806])).
fof(f28511,plain,(
  sP5(sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f1401,f1892,f24933,f1242])).
fof(f1892,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f1724,f1248])).
fof(f33239,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f28672,f30316])).
fof(f28672,plain,(
  ~sP4(sK27('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f6983,f24933,f1242])).
fof(f33238,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f28675,f30316])).
fof(f28675,plain,(
  ~sP4(sK27('const.arithmetic.ZERO_0'),'F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f6983,f24933,f1242])).
fof(f33237,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f28851,f30806])).
fof(f28851,plain,(
  sP5(sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f1401,f1737,f24933,f1242])).
fof(f1737,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f1721,f1248])).
fof(f33236,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29015,f30806])).
fof(f29015,plain,(
  sP5(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f1401,f1777,f24933,f1242])).
fof(f1777,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1722,f1248])).
fof(f33235,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29179,f30806])).
fof(f29179,plain,(
  sP5(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f1401,f1828,f24933,f1242])).
fof(f1828,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1723,f1248])).
fof(f33234,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29340,f30316])).
fof(f29340,plain,(
  ~sP4(sK28('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f6301,f24933,f1242])).
fof(f33233,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29343,f30316])).
fof(f29343,plain,(
  ~sP4(sK28('const.arithmetic.ZERO_0'),'F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f6301,f24933,f1242])).
fof(f33232,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29516,f30316])).
fof(f29516,plain,(
  ~sP4('T_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1413,f24933,f1242])).
fof(f33231,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29517,f30316])).
fof(f29517,plain,(
  ~sP4('T_0','F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f1413,f24933,f1242])).
fof(f33230,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29518,f30316])).
fof(f29518,plain,(
  ~sP4('T_0','F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f1413,f24933,f1242])).
fof(f33229,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29519,f30316])).
fof(f29519,plain,(
  ~sP4(sK27('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f1413,f24933,f1242])).
fof(f33228,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29520,f30316])).
fof(f29520,plain,(
  ~sP4(sK27('const.arithmetic.ZERO_0'),'F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f1413,f24933,f1242])).
fof(f33227,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29521,f30316])).
fof(f29521,plain,(
  ~sP4(sK28('const.arithmetic.ZERO_0'),'F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f1413,f24933,f1242])).
fof(f33226,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29522,f30316])).
fof(f29522,plain,(
  ~sP4(sK28('const.arithmetic.ZERO_0'),'F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f1413,f24933,f1242])).
fof(f33225,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29632,f6906])).
fof(f29632,plain,(
  sP5('T_0',sK27('const.arithmetic.ZERO_0'),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f6948,f1413,f24933,f1242])).
fof(f33224,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29637,f6907])).
fof(f29637,plain,(
  sP5(sK27('const.arithmetic.ZERO_0'),sK27('const.arithmetic.ZERO_0'),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6984,f6948,f1413,f24933,f1242])).
fof(f6984,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,sK27('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f6174,f1252])).
fof(f33223,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29680,f6224])).
fof(f29680,plain,(
  sP5('T_0',sK28('const.arithmetic.ZERO_0'),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f6266,f1413,f24933,f1242])).
fof(f33222,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29689,f6225])).
fof(f29689,plain,(
  sP5(sK28('const.arithmetic.ZERO_0'),sK28('const.arithmetic.ZERO_0'),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6302,f6266,f1413,f24933,f1242])).
fof(f6302,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,sK28('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f6173,f1252])).
fof(f33221,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f29728,f30316])).
fof(f29728,plain,(
  ~sP4('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1415,f24933,f1242])).
fof(f33220,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30156,f30766])).
fof(f30766,plain,(
  ( ! [X0] : (~sP5(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'F_0')) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1639,f24933,f1258])).
fof(f30156,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1420,f1404,f1418,f24933,f1242])).
fof(f1420,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) )),
  inference(unit_resulting_resolution,[],[f840,f1252])).
fof(f33219,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30157,f30766])).
fof(f30157,plain,(
  sP5('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1459,f1404,f1418,f24933,f1242])).
fof(f1459,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f850,f1252])).
fof(f33218,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30158,f30766])).
fof(f30158,plain,(
  sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1421,f1404,f1418,f24933,f1242])).
fof(f1421,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f844,f1252])).
fof(f33217,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30159,f30766])).
fof(f30159,plain,(
  sP5(sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1893,f1404,f1418,f24933,f1242])).
fof(f1893,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f1724,f1252])).
fof(f33216,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30161,f30766])).
fof(f30161,plain,(
  sP5(sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1738,f1404,f1418,f24933,f1242])).
fof(f1738,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f1721,f1252])).
fof(f33215,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30162,f30766])).
fof(f30162,plain,(
  sP5(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1778,f1404,f1418,f24933,f1242])).
fof(f1778,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1722,f1252])).
fof(f33214,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30163,f30766])).
fof(f30163,plain,(
  sP5(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1829,f1404,f1418,f24933,f1242])).
fof(f1829,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1723,f1252])).
fof(f33213,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30312,f30316])).
fof(f30312,plain,(
  ( ! [X0] : (~sP4('T_0','T_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f842,f842,f24933,f1243])).
fof(f1243,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f808])).
fof(f33212,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30344,f30348])).
fof(f30344,plain,(
  ( ! [X0] : (~sP3('T_0','T_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f842,f842,f24933,f1247])).
fof(f1247,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f811])).
fof(f33211,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30350,f24933])).
fof(f30350,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f24933,f24933,f1251])).
fof(f1251,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f813])).
fof(f33210,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30351,f24933])).
fof(f30351,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1420,f24933,f24933,f1251])).
fof(f33209,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30352,f24933])).
fof(f30352,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1459,f24933,f24933,f1251])).
fof(f33208,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30353,f24933])).
fof(f30353,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1421,f24933,f24933,f1251])).
fof(f33207,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30354,f24933])).
fof(f30354,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1893,f24933,f24933,f1251])).
fof(f33206,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30355,f24933])).
fof(f30355,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6984,f24933,f24933,f1251])).
fof(f33205,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30356,f24933])).
fof(f30356,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1738,f24933,f24933,f1251])).
fof(f33204,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30357,f24933])).
fof(f30357,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1778,f24933,f24933,f1251])).
fof(f33203,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30358,f24933])).
fof(f30358,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1829,f24933,f24933,f1251])).
fof(f33202,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30359,f24933])).
fof(f30359,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6302,f24933,f24933,f1251])).
fof(f33201,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30361,f24933])).
fof(f30361,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1423,f24933,f24933,f1251])).
fof(f33200,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30362,f24933])).
fof(f30362,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1460,f24933,f24933,f1251])).
fof(f33199,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30363,f24933])).
fof(f30363,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1424,f24933,f24933,f1251])).
fof(f33198,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30364,f24933])).
fof(f30364,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1894,f24933,f24933,f1251])).
fof(f33197,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30365,f24933])).
fof(f30365,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f24933,f24933,f1251])).
fof(f33196,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30366,f24933])).
fof(f30366,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1739,f24933,f24933,f1251])).
fof(f33195,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30367,f24933])).
fof(f30367,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1779,f24933,f24933,f1251])).
fof(f33194,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30368,f24933])).
fof(f30368,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1830,f24933,f24933,f1251])).
fof(f33193,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30369,f24933])).
fof(f30369,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f24933,f24933,f1251])).
fof(f33192,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30370,f24933])).
fof(f30370,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1425,f24933,f24933,f1251])).
fof(f33191,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30371,f24933])).
fof(f30371,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f24933,f24933,f1251])).
fof(f33190,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30372,f24933])).
fof(f30372,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1461,f24933,f24933,f1251])).
fof(f33189,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30373,f24933])).
fof(f30373,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1427,f24933,f24933,f1251])).
fof(f33188,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30374,f24933])).
fof(f30374,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1895,f24933,f24933,f1251])).
fof(f33187,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30375,f24933])).
fof(f30375,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6986,f24933,f24933,f1251])).
fof(f33186,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30376,f24933])).
fof(f30376,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1740,f24933,f24933,f1251])).
fof(f33185,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30377,f24933])).
fof(f30377,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1780,f24933,f24933,f1251])).
fof(f33184,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30378,f24933])).
fof(f30378,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1831,f24933,f24933,f1251])).
fof(f33183,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30379,f24933])).
fof(f30379,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6304,f24933,f24933,f1251])).
fof(f33182,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30381,f24933])).
fof(f30381,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1423,f1359,f24933,f1251])).
fof(f33181,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30382,f24933])).
fof(f30382,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1460,f1359,f24933,f1251])).
fof(f33180,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30383,f24933])).
fof(f30383,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1424,f1359,f24933,f1251])).
fof(f33179,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30384,f24933])).
fof(f30384,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1894,f1359,f24933,f1251])).
fof(f33178,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30385,f24933])).
fof(f30385,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f1359,f24933,f1251])).
fof(f33177,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30386,f24933])).
fof(f30386,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1739,f1359,f24933,f1251])).
fof(f33176,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30387,f24933])).
fof(f30387,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1779,f1359,f24933,f1251])).
fof(f33175,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30388,f24933])).
fof(f30388,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1830,f1359,f24933,f1251])).
fof(f33174,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30389,f24933])).
fof(f30389,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f1359,f24933,f1251])).
fof(f33173,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30390,f24933])).
fof(f30390,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1425,f1359,f24933,f1251])).
fof(f33172,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30391,f24933])).
fof(f30391,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f1359,f24933,f1251])).
fof(f33171,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30392,f24933])).
fof(f30392,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1461,f1359,f24933,f1251])).
fof(f33170,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30393,f24933])).
fof(f30393,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1427,f1359,f24933,f1251])).
fof(f33169,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30394,f24933])).
fof(f30394,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1895,f1359,f24933,f1251])).
fof(f33168,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30395,f24933])).
fof(f30395,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6986,f1359,f24933,f1251])).
fof(f33167,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30396,f24933])).
fof(f30396,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1740,f1359,f24933,f1251])).
fof(f33166,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30397,f24933])).
fof(f30397,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1780,f1359,f24933,f1251])).
fof(f33165,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30398,f24933])).
fof(f30398,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1831,f1359,f24933,f1251])).
fof(f33164,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30399,f24933])).
fof(f30399,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6304,f1359,f24933,f1251])).
fof(f33163,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30401,f24933])).
fof(f30401,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1420,f24933,f24933,f1251])).
fof(f33162,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30402,f24933])).
fof(f30402,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1459,f24933,f24933,f1251])).
fof(f33161,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30403,f24933])).
fof(f30403,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1421,f24933,f24933,f1251])).
fof(f33160,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30404,f24933])).
fof(f30404,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1893,f24933,f24933,f1251])).
fof(f33159,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30405,f24933])).
fof(f30405,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6984,f24933,f24933,f1251])).
fof(f33158,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30406,f24933])).
fof(f30406,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1738,f24933,f24933,f1251])).
fof(f33157,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30407,f24933])).
fof(f30407,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1778,f24933,f24933,f1251])).
fof(f33156,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30408,f24933])).
fof(f30408,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1829,f24933,f24933,f1251])).
fof(f33155,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30409,f24933])).
fof(f30409,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6302,f24933,f24933,f1251])).
fof(f33154,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30410,f24933])).
fof(f30410,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f24933,f24933,f1251])).
fof(f33153,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30411,f24933])).
fof(f30411,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1423,f24933,f24933,f1251])).
fof(f33152,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30412,f24933])).
fof(f30412,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1460,f24933,f24933,f1251])).
fof(f33151,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30413,f24933])).
fof(f30413,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1424,f24933,f24933,f1251])).
fof(f33150,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30414,f24933])).
fof(f30414,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1894,f24933,f24933,f1251])).
fof(f33149,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30415,f24933])).
fof(f30415,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f24933,f24933,f1251])).
fof(f33148,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30416,f24933])).
fof(f30416,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1739,f24933,f24933,f1251])).
fof(f33147,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30417,f24933])).
fof(f30417,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1779,f24933,f24933,f1251])).
fof(f33146,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30418,f24933])).
fof(f30418,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1830,f24933,f24933,f1251])).
fof(f33145,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30419,f24933])).
fof(f30419,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f24933,f24933,f1251])).
fof(f33144,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30420,f24933])).
fof(f30420,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1425,f24933,f24933,f1251])).
fof(f33143,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30421,f24933])).
fof(f30421,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f24933,f24933,f1251])).
fof(f33142,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30422,f24933])).
fof(f30422,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1461,f24933,f24933,f1251])).
fof(f33141,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30423,f24933])).
fof(f30423,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1427,f24933,f24933,f1251])).
fof(f33140,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30424,f24933])).
fof(f30424,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1895,f24933,f24933,f1251])).
fof(f33139,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30425,f24933])).
fof(f30425,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6986,f24933,f24933,f1251])).
fof(f33138,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30426,f24933])).
fof(f30426,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1740,f24933,f24933,f1251])).
fof(f33137,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30427,f24933])).
fof(f30427,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1780,f24933,f24933,f1251])).
fof(f33136,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30428,f24933])).
fof(f30428,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1831,f24933,f24933,f1251])).
fof(f33135,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30429,f24933])).
fof(f30429,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6304,f24933,f24933,f1251])).
fof(f33134,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30435,f24933])).
fof(f30435,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1420,f1359,f24933,f1251])).
fof(f33133,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30436,f24933])).
fof(f30436,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1459,f1359,f24933,f1251])).
fof(f33132,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30437,f24933])).
fof(f30437,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1421,f1359,f24933,f1251])).
fof(f33131,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30438,f24933])).
fof(f30438,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1893,f1359,f24933,f1251])).
fof(f33130,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30439,f24933])).
fof(f30439,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6984,f1359,f24933,f1251])).
fof(f33129,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30440,f24933])).
fof(f30440,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1738,f1359,f24933,f1251])).
fof(f33128,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30441,f24933])).
fof(f30441,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1778,f1359,f24933,f1251])).
fof(f33127,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30442,f24933])).
fof(f30442,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1829,f1359,f24933,f1251])).
fof(f33126,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30443,f24933])).
fof(f30443,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6302,f1359,f24933,f1251])).
fof(f33125,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30444,f24933])).
fof(f30444,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1425,f1359,f24933,f1251])).
fof(f33124,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30445,f24933])).
fof(f30445,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f1359,f24933,f1251])).
fof(f33123,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30446,f24933])).
fof(f30446,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1461,f1359,f24933,f1251])).
fof(f33122,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30447,f24933])).
fof(f30447,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1427,f1359,f24933,f1251])).
fof(f33121,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30448,f24933])).
fof(f30448,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1895,f1359,f24933,f1251])).
fof(f33120,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30449,f24933])).
fof(f30449,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6986,f1359,f24933,f1251])).
fof(f33119,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30450,f24933])).
fof(f30450,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1740,f1359,f24933,f1251])).
fof(f33118,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30451,f24933])).
fof(f30451,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1780,f1359,f24933,f1251])).
fof(f33117,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30452,f24933])).
fof(f30452,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1831,f1359,f24933,f1251])).
fof(f33116,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30453,f24933])).
fof(f30453,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6304,f1359,f24933,f1251])).
fof(f33115,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30459,f24933])).
fof(f30459,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1420,f24933,f24933,f1251])).
fof(f33114,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30460,f24933])).
fof(f30460,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1459,f24933,f24933,f1251])).
fof(f33113,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30461,f24933])).
fof(f30461,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1421,f24933,f24933,f1251])).
fof(f33112,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30462,f24933])).
fof(f30462,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1893,f24933,f24933,f1251])).
fof(f33111,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30463,f24933])).
fof(f30463,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6984,f24933,f24933,f1251])).
fof(f33110,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30464,f24933])).
fof(f30464,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1738,f24933,f24933,f1251])).
fof(f33109,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30465,f24933])).
fof(f30465,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1778,f24933,f24933,f1251])).
fof(f33108,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30466,f24933])).
fof(f30466,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1829,f24933,f24933,f1251])).
fof(f33107,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30467,f24933])).
fof(f30467,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6302,f24933,f24933,f1251])).
fof(f33106,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30468,f24933])).
fof(f30468,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f24933,f24933,f1251])).
fof(f33105,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30469,f24933])).
fof(f30469,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1423,f24933,f24933,f1251])).
fof(f33104,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30470,f24933])).
fof(f30470,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1460,f24933,f24933,f1251])).
fof(f33103,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30471,f24933])).
fof(f30471,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1424,f24933,f24933,f1251])).
fof(f33102,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30472,f24933])).
fof(f30472,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1894,f24933,f24933,f1251])).
fof(f33101,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30473,f24933])).
fof(f30473,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f24933,f24933,f1251])).
fof(f33100,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30474,f24933])).
fof(f30474,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1739,f24933,f24933,f1251])).
fof(f33099,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30475,f24933])).
fof(f30475,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1779,f24933,f24933,f1251])).
fof(f33098,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30476,f24933])).
fof(f30476,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1830,f24933,f24933,f1251])).
fof(f33097,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30477,f24933])).
fof(f30477,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f24933,f24933,f1251])).
fof(f33096,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30478,f24933])).
fof(f30478,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1425,f24933,f24933,f1251])).
fof(f33095,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30479,f24933])).
fof(f30479,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f24933,f24933,f1251])).
fof(f33094,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30480,f24933])).
fof(f30480,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1461,f24933,f24933,f1251])).
fof(f33093,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30481,f24933])).
fof(f30481,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1427,f24933,f24933,f1251])).
fof(f33092,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30482,f24933])).
fof(f30482,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1895,f24933,f24933,f1251])).
fof(f33091,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30483,f24933])).
fof(f30483,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6986,f24933,f24933,f1251])).
fof(f33090,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30484,f24933])).
fof(f30484,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1740,f24933,f24933,f1251])).
fof(f33089,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30485,f24933])).
fof(f30485,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1780,f24933,f24933,f1251])).
fof(f33088,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30486,f24933])).
fof(f30486,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1831,f24933,f24933,f1251])).
fof(f33087,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30487,f24933])).
fof(f30487,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6304,f24933,f24933,f1251])).
fof(f33086,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30489,f24933])).
fof(f30489,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1420,f1359,f24933,f1251])).
fof(f33085,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30490,f24933])).
fof(f30490,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1459,f1359,f24933,f1251])).
fof(f33084,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30491,f24933])).
fof(f30491,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1421,f1359,f24933,f1251])).
fof(f33083,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30492,f24933])).
fof(f30492,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1893,f1359,f24933,f1251])).
fof(f33082,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30493,f24933])).
fof(f30493,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6984,f1359,f24933,f1251])).
fof(f33081,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30494,f24933])).
fof(f30494,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1738,f1359,f24933,f1251])).
fof(f33080,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30495,f24933])).
fof(f30495,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1778,f1359,f24933,f1251])).
fof(f33079,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30496,f24933])).
fof(f30496,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1829,f1359,f24933,f1251])).
fof(f33078,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30497,f24933])).
fof(f30497,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6302,f1359,f24933,f1251])).
fof(f33077,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30498,f24933])).
fof(f30498,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f1359,f24933,f1251])).
fof(f33076,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30499,f24933])).
fof(f30499,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1423,f1359,f24933,f1251])).
fof(f33075,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30500,f24933])).
fof(f30500,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1460,f1359,f24933,f1251])).
fof(f33074,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30501,f24933])).
fof(f30501,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1424,f1359,f24933,f1251])).
fof(f33073,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30502,f24933])).
fof(f30502,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1894,f1359,f24933,f1251])).
fof(f33072,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30503,f24933])).
fof(f30503,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f1359,f24933,f1251])).
fof(f33071,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30504,f24933])).
fof(f30504,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1739,f1359,f24933,f1251])).
fof(f33070,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30505,f24933])).
fof(f30505,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1779,f1359,f24933,f1251])).
fof(f33069,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30506,f24933])).
fof(f30506,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1830,f1359,f24933,f1251])).
fof(f33068,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30507,f24933])).
fof(f30507,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f1359,f24933,f1251])).
fof(f33067,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30508,f24933])).
fof(f30508,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f24933,f24933,f1251])).
fof(f33066,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30509,f24933])).
fof(f30509,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1420,f24933,f24933,f1251])).
fof(f33065,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30510,f24933])).
fof(f30510,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1459,f24933,f24933,f1251])).
fof(f33064,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30511,f24933])).
fof(f30511,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1421,f24933,f24933,f1251])).
fof(f33063,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30512,f24933])).
fof(f30512,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1893,f24933,f24933,f1251])).
fof(f33062,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30513,f24933])).
fof(f30513,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6984,f24933,f24933,f1251])).
fof(f33061,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30514,f24933])).
fof(f30514,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1738,f24933,f24933,f1251])).
fof(f33060,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30515,f24933])).
fof(f30515,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1778,f24933,f24933,f1251])).
fof(f33059,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30516,f24933])).
fof(f30516,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1829,f24933,f24933,f1251])).
fof(f33058,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30517,f24933])).
fof(f30517,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6302,f24933,f24933,f1251])).
fof(f33057,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30518,f24933])).
fof(f30518,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f24933,f24933,f1251])).
fof(f33056,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30519,f24933])).
fof(f30519,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1423,f24933,f24933,f1251])).
fof(f33055,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30520,f24933])).
fof(f30520,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1460,f24933,f24933,f1251])).
fof(f33054,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30521,f24933])).
fof(f30521,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1424,f24933,f24933,f1251])).
fof(f33053,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30522,f24933])).
fof(f30522,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1894,f24933,f24933,f1251])).
fof(f33052,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30523,f24933])).
fof(f30523,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f24933,f24933,f1251])).
fof(f33051,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30524,f24933])).
fof(f30524,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1739,f24933,f24933,f1251])).
fof(f33050,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30525,f24933])).
fof(f30525,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1779,f24933,f24933,f1251])).
fof(f33049,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30526,f24933])).
fof(f30526,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1830,f24933,f24933,f1251])).
fof(f33048,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30527,f24933])).
fof(f30527,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f24933,f24933,f1251])).
fof(f33047,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30529,f24933])).
fof(f30529,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f24933,f24933,f1251])).
fof(f33046,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30530,f24933])).
fof(f30530,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1461,f24933,f24933,f1251])).
fof(f33045,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30531,f24933])).
fof(f30531,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1427,f24933,f24933,f1251])).
fof(f33044,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30532,f24933])).
fof(f30532,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1895,f24933,f24933,f1251])).
fof(f33043,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30533,f24933])).
fof(f30533,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6986,f24933,f24933,f1251])).
fof(f33042,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30534,f24933])).
fof(f30534,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1740,f24933,f24933,f1251])).
fof(f33041,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30535,f24933])).
fof(f30535,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1780,f24933,f24933,f1251])).
fof(f33040,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30536,f24933])).
fof(f30536,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1831,f24933,f24933,f1251])).
fof(f33039,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30537,f24933])).
fof(f30537,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6304,f24933,f24933,f1251])).
fof(f33038,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30542,f24933])).
fof(f30542,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f1359,f24933,f1251])).
fof(f33037,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30543,f24933])).
fof(f30543,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1423,f1359,f24933,f1251])).
fof(f33036,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30544,f24933])).
fof(f30544,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1460,f1359,f24933,f1251])).
fof(f33035,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30545,f24933])).
fof(f30545,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1424,f1359,f24933,f1251])).
fof(f33034,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30546,f24933])).
fof(f30546,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1894,f1359,f24933,f1251])).
fof(f33033,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30547,f24933])).
fof(f30547,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f1359,f24933,f1251])).
fof(f33032,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30548,f24933])).
fof(f30548,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1739,f1359,f24933,f1251])).
fof(f33031,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30549,f24933])).
fof(f30549,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1779,f1359,f24933,f1251])).
fof(f33030,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30550,f24933])).
fof(f30550,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1830,f1359,f24933,f1251])).
fof(f33029,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30551,f24933])).
fof(f30551,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f1359,f24933,f1251])).
fof(f33028,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30553,f24933])).
fof(f30553,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f1359,f24933,f1251])).
fof(f33027,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30554,f24933])).
fof(f30554,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1461,f1359,f24933,f1251])).
fof(f33026,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30555,f24933])).
fof(f30555,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1427,f1359,f24933,f1251])).
fof(f33025,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30556,f24933])).
fof(f30556,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1895,f1359,f24933,f1251])).
fof(f33024,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30557,f24933])).
fof(f30557,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6986,f1359,f24933,f1251])).
fof(f33023,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30558,f24933])).
fof(f30558,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1740,f1359,f24933,f1251])).
fof(f33022,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30559,f24933])).
fof(f30559,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1780,f1359,f24933,f1251])).
fof(f33021,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30560,f24933])).
fof(f30560,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1831,f1359,f24933,f1251])).
fof(f33020,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30561,f24933])).
fof(f30561,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6304,f1359,f24933,f1251])).
fof(f33019,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30566,f24933])).
fof(f30566,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f24933,f24933,f1251])).
fof(f33018,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30567,f24933])).
fof(f30567,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1420,f24933,f24933,f1251])).
fof(f33017,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30568,f24933])).
fof(f30568,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1459,f24933,f24933,f1251])).
fof(f33016,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30569,f24933])).
fof(f30569,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1421,f24933,f24933,f1251])).
fof(f33015,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30570,f24933])).
fof(f30570,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1893,f24933,f24933,f1251])).
fof(f33014,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30571,f24933])).
fof(f30571,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6984,f24933,f24933,f1251])).
fof(f33013,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30572,f24933])).
fof(f30572,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1738,f24933,f24933,f1251])).
fof(f33012,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30573,f24933])).
fof(f30573,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1778,f24933,f24933,f1251])).
fof(f33011,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30574,f24933])).
fof(f30574,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1829,f24933,f24933,f1251])).
fof(f33010,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30575,f24933])).
fof(f30575,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6302,f24933,f24933,f1251])).
fof(f33009,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30577,f24933])).
fof(f30577,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1423,f24933,f24933,f1251])).
fof(f33008,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30578,f24933])).
fof(f30578,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1460,f24933,f24933,f1251])).
fof(f33007,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30579,f24933])).
fof(f30579,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1424,f24933,f24933,f1251])).
fof(f33006,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30580,f24933])).
fof(f30580,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1894,f24933,f24933,f1251])).
fof(f33005,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30581,f24933])).
fof(f30581,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f24933,f24933,f1251])).
fof(f33004,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30582,f24933])).
fof(f30582,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1739,f24933,f24933,f1251])).
fof(f33003,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30583,f24933])).
fof(f30583,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1779,f24933,f24933,f1251])).
fof(f33002,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30584,f24933])).
fof(f30584,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1830,f24933,f24933,f1251])).
fof(f33001,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30585,f24933])).
fof(f30585,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f24933,f24933,f1251])).
fof(f33000,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30586,f24933])).
fof(f30586,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1425,f24933,f24933,f1251])).
fof(f32999,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30587,f24933])).
fof(f30587,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f24933,f24933,f1251])).
fof(f32998,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30588,f24933])).
fof(f30588,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1461,f24933,f24933,f1251])).
fof(f32997,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30589,f24933])).
fof(f30589,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1427,f24933,f24933,f1251])).
fof(f32996,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30590,f24933])).
fof(f30590,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1895,f24933,f24933,f1251])).
fof(f32995,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30591,f24933])).
fof(f30591,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6986,f24933,f24933,f1251])).
fof(f32994,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30592,f24933])).
fof(f30592,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1740,f24933,f24933,f1251])).
fof(f32993,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30593,f24933])).
fof(f30593,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1780,f24933,f24933,f1251])).
fof(f32992,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30594,f24933])).
fof(f30594,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1831,f24933,f24933,f1251])).
fof(f32991,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30595,f24933])).
fof(f30595,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6304,f24933,f24933,f1251])).
fof(f32990,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30596,f24933])).
fof(f30596,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f1359,f24933,f1251])).
fof(f32989,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30597,f24933])).
fof(f30597,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1420,f1359,f24933,f1251])).
fof(f32988,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30598,f24933])).
fof(f30598,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1459,f1359,f24933,f1251])).
fof(f32987,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30599,f24933])).
fof(f30599,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1421,f1359,f24933,f1251])).
fof(f32986,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30600,f24933])).
fof(f30600,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1893,f1359,f24933,f1251])).
fof(f32985,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30601,f24933])).
fof(f30601,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6984,f1359,f24933,f1251])).
fof(f32984,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30602,f24933])).
fof(f30602,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1738,f1359,f24933,f1251])).
fof(f32983,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30603,f24933])).
fof(f30603,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1778,f1359,f24933,f1251])).
fof(f32982,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30604,f24933])).
fof(f30604,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1829,f1359,f24933,f1251])).
fof(f32981,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30605,f24933])).
fof(f30605,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6302,f1359,f24933,f1251])).
fof(f32980,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30607,f24933])).
fof(f30607,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1423,f1359,f24933,f1251])).
fof(f32979,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30608,f24933])).
fof(f30608,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1460,f1359,f24933,f1251])).
fof(f32978,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30609,f24933])).
fof(f30609,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1424,f1359,f24933,f1251])).
fof(f32977,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30610,f24933])).
fof(f30610,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1894,f1359,f24933,f1251])).
fof(f32976,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30611,f24933])).
fof(f30611,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f1359,f24933,f1251])).
fof(f32975,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30612,f24933])).
fof(f30612,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1739,f1359,f24933,f1251])).
fof(f32974,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30613,f24933])).
fof(f30613,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1779,f1359,f24933,f1251])).
fof(f32973,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30614,f24933])).
fof(f30614,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1830,f1359,f24933,f1251])).
fof(f32972,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30615,f24933])).
fof(f30615,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f1359,f24933,f1251])).
fof(f32971,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30616,f24933])).
fof(f30616,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f24933,f24933,f1251])).
fof(f32970,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30617,f24933])).
fof(f30617,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1420,f24933,f24933,f1251])).
fof(f32969,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30618,f24933])).
fof(f30618,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1459,f24933,f24933,f1251])).
fof(f32968,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30619,f24933])).
fof(f30619,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1421,f24933,f24933,f1251])).
fof(f32967,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30620,f24933])).
fof(f30620,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1893,f24933,f24933,f1251])).
fof(f32966,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30621,f24933])).
fof(f30621,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6984,f24933,f24933,f1251])).
fof(f32965,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30622,f24933])).
fof(f30622,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1738,f24933,f24933,f1251])).
fof(f32964,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30623,f24933])).
fof(f30623,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1778,f24933,f24933,f1251])).
fof(f32963,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30624,f24933])).
fof(f30624,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1829,f24933,f24933,f1251])).
fof(f32962,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30625,f24933])).
fof(f30625,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6302,f24933,f24933,f1251])).
fof(f32961,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30626,f24933])).
fof(f30626,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f24933,f24933,f1251])).
fof(f32960,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30627,f24933])).
fof(f30627,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1423,f24933,f24933,f1251])).
fof(f32959,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30628,f24933])).
fof(f30628,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1460,f24933,f24933,f1251])).
fof(f32958,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30629,f24933])).
fof(f30629,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1424,f24933,f24933,f1251])).
fof(f32957,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30630,f24933])).
fof(f30630,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1894,f24933,f24933,f1251])).
fof(f32956,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30631,f24933])).
fof(f30631,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f24933,f24933,f1251])).
fof(f32955,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30632,f24933])).
fof(f30632,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1739,f24933,f24933,f1251])).
fof(f32954,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30633,f24933])).
fof(f30633,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1779,f24933,f24933,f1251])).
fof(f32953,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30634,f24933])).
fof(f30634,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1830,f24933,f24933,f1251])).
fof(f32952,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30635,f24933])).
fof(f30635,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f24933,f24933,f1251])).
fof(f32951,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30637,f24933])).
fof(f30637,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f24933,f24933,f1251])).
fof(f32950,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30638,f24933])).
fof(f30638,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1461,f24933,f24933,f1251])).
fof(f32949,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30639,f24933])).
fof(f30639,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1427,f24933,f24933,f1251])).
fof(f32948,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30640,f24933])).
fof(f30640,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1895,f24933,f24933,f1251])).
fof(f32947,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30641,f24933])).
fof(f30641,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6986,f24933,f24933,f1251])).
fof(f32946,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30642,f24933])).
fof(f30642,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1740,f24933,f24933,f1251])).
fof(f32945,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30643,f24933])).
fof(f30643,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1780,f24933,f24933,f1251])).
fof(f32944,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30644,f24933])).
fof(f30644,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1831,f24933,f24933,f1251])).
fof(f32943,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30645,f24933])).
fof(f30645,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6304,f24933,f24933,f1251])).
fof(f32942,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30650,f24933])).
fof(f30650,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f1359,f24933,f1251])).
fof(f32941,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30651,f24933])).
fof(f30651,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1420,f1359,f24933,f1251])).
fof(f32940,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30652,f24933])).
fof(f30652,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1459,f1359,f24933,f1251])).
fof(f32939,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30653,f24933])).
fof(f30653,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1421,f1359,f24933,f1251])).
fof(f32938,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30654,f24933])).
fof(f30654,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1893,f1359,f24933,f1251])).
fof(f32937,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30655,f24933])).
fof(f30655,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6984,f1359,f24933,f1251])).
fof(f32936,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30656,f24933])).
fof(f30656,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1738,f1359,f24933,f1251])).
fof(f32935,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30657,f24933])).
fof(f30657,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1778,f1359,f24933,f1251])).
fof(f32934,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30658,f24933])).
fof(f30658,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1829,f1359,f24933,f1251])).
fof(f32933,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30659,f24933])).
fof(f30659,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6302,f1359,f24933,f1251])).
fof(f32932,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30661,f24933])).
fof(f30661,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1426,f1359,f24933,f1251])).
fof(f32931,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30662,f24933])).
fof(f30662,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1461,f1359,f24933,f1251])).
fof(f32930,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30663,f24933])).
fof(f30663,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1427,f1359,f24933,f1251])).
fof(f32929,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30664,f24933])).
fof(f30664,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1895,f1359,f24933,f1251])).
fof(f32928,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30665,f24933])).
fof(f30665,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6986,f1359,f24933,f1251])).
fof(f32927,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30666,f24933])).
fof(f30666,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1740,f1359,f24933,f1251])).
fof(f32926,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30667,f24933])).
fof(f30667,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1780,f1359,f24933,f1251])).
fof(f32925,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30668,f24933])).
fof(f30668,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1831,f1359,f24933,f1251])).
fof(f32924,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30669,f24933])).
fof(f30669,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6304,f1359,f24933,f1251])).
fof(f32923,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30674,f32303])).
fof(f30674,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f3451,f24933,f1255])).
fof(f32922,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30675,f32303])).
fof(f30675,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3452,f24933,f1255])).
fof(f32921,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30676,f32303])).
fof(f30676,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f3448,f24933,f1255])).
fof(f3448,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f842,f1591,f1257])).
fof(f32920,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30677,f32303])).
fof(f30677,plain,(
  s(bool,'F_0') = s(bool,sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7723,f24933,f1255])).
fof(f7723,plain,(
  ~sP5('T_0','F_0',sK28('const.arithmetic.ZERO_0'))),
  inference(unit_resulting_resolution,[],[f6250,f1238])).
fof(f1238,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f805])).
fof(f6250,plain,(
  ~sP4('T_0',sK28('const.arithmetic.ZERO_0'),'F_0')),
  inference(unit_resulting_resolution,[],[f842,f1359,f6173,f1243])).
fof(f32919,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30678,f32303])).
fof(f30678,plain,(
  s(bool,'F_0') = s(bool,sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f9925,f24933,f1255])).
fof(f9925,plain,(
  ~sP5('T_0','F_0',sK27('const.arithmetic.ZERO_0'))),
  inference(unit_resulting_resolution,[],[f6932,f1238])).
fof(f6932,plain,(
  ~sP4('T_0',sK27('const.arithmetic.ZERO_0'),'F_0')),
  inference(unit_resulting_resolution,[],[f842,f1359,f6174,f1243])).
fof(f32918,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30679,f32303])).
fof(f30679,plain,(
  s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3449,f24933,f1255])).
fof(f32917,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30680,f32303])).
fof(f30680,plain,(
  s(bool,'F_0') = s(bool,sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6906,f24933,f1255])).
fof(f32916,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30681,f32303])).
fof(f30681,plain,(
  s(bool,'F_0') = s(bool,sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6224,f24933,f1255])).
fof(f32915,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30691,f32303])).
fof(f30691,plain,(
  s(bool,'const.combin.I_1'(s(bool,'T_0'))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | $spl14),
  inference(unit_resulting_resolution,[],[f3701,f24933,f1255])).
fof(f3701,plain,(
  ~sP5('F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'const.combin.I_1'(s(bool,'T_0')))),
  inference(superposition,[],[f3482,f1603])).
fof(f32914,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30693,f32303])).
fof(f30693,plain,(
  s(bool,'const.combin.I_1'(s(bool,'T_0'))) = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13))) | $spl14),
  inference(unit_resulting_resolution,[],[f3685,f24933,f1255])).
fof(f3685,plain,(
  ~sP5('F_0','const.combin.I_1'(s(bool,'T_0')),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))),
  inference(superposition,[],[f3470,f1603])).
fof(f32913,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30698,f32303])).
fof(f30698,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f7025,f24933,f1255])).
fof(f32912,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30699,f32303])).
fof(f30699,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f7027,f24933,f1255])).
fof(f7027,plain,(
  ~sP5(sK27('const.arithmetic.ZERO_0'),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1591,f6174,f1257])).
fof(f32911,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30700,f32303])).
fof(f30700,plain,(
  s(bool,'F_0') = s(bool,sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f15665,f24933,f1255])).
fof(f15665,plain,(
  ~sP5(sK27('const.arithmetic.ZERO_0'),'F_0',sK27('const.arithmetic.ZERO_0'))),
  inference(unit_resulting_resolution,[],[f6933,f1238])).
fof(f6933,plain,(
  ~sP4(sK27('const.arithmetic.ZERO_0'),sK27('const.arithmetic.ZERO_0'),'F_0')),
  inference(unit_resulting_resolution,[],[f6174,f1359,f6174,f1243])).
fof(f32910,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30701,f32303])).
fof(f30701,plain,(
  s(bool,'F_0') = s(bool,sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6907,f24933,f1255])).
fof(f32909,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30702,f32303])).
fof(f30702,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6341,f24933,f1255])).
fof(f32908,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30703,f32303])).
fof(f30703,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6343,f24933,f1255])).
fof(f6343,plain,(
  ~sP5(sK28('const.arithmetic.ZERO_0'),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1591,f6173,f1257])).
fof(f32907,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30704,f32303])).
fof(f30704,plain,(
  s(bool,'F_0') = s(bool,sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f14085,f24933,f1255])).
fof(f14085,plain,(
  ~sP5(sK28('const.arithmetic.ZERO_0'),'F_0',sK28('const.arithmetic.ZERO_0'))),
  inference(unit_resulting_resolution,[],[f6251,f1238])).
fof(f6251,plain,(
  ~sP4(sK28('const.arithmetic.ZERO_0'),sK28('const.arithmetic.ZERO_0'),'F_0')),
  inference(unit_resulting_resolution,[],[f6173,f1359,f6173,f1243])).
fof(f32906,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30705,f32303])).
fof(f30705,plain,(
  s(bool,'F_0') = s(bool,sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6225,f24933,f1255])).
fof(f32905,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30878,f32133])).
fof(f32133,plain,(
  ( ! [X0,X1] : (sP7('T_0',X0,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1428,f24933,f1343])).
fof(f1343,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f1265])).
fof(f1265,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f817])).
fof(f817,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f816])).
fof(f816,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f815])).
fof(f815,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f662])).
fof(f662,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f1428,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f842,f1267])).
fof(f1267,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f819])).
fof(f819,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f818])).
fof(f818,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f661])).
fof(f661,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f30878,plain,(
  ( ! [X0,X1] : (~sP7('T_0',X0,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f842,f24933,f1260])).
fof(f1260,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f817])).
fof(f32904,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30907,f32133])).
fof(f30907,plain,(
  ( ! [X0,X1] : (~sP7('T_0',X0,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f842,f24933,f1261])).
fof(f1261,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f817])).
fof(f32903,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30935,f30939])).
fof(f30939,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1269])).
fof(f1269,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f819])).
fof(f30935,plain,(
  ( ! [X0] : (~sP6('T_0','T_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f842,f842,f24933,f1266])).
fof(f1266,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f819])).
fof(f32902,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30943,f1359])).
fof(f30943,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f5149,f24933,f1270])).
fof(f1270,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f821])).
fof(f821,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f820])).
fof(f820,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f663])).
fof(f663,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f537,f662,f661])).
fof(f537,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f536])).
fof(f536,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X39,X40,X41] : ((p(s(bool,X41)) <=> (p(s(bool,X40)) & p(s(bool,X39)))) <=> ((p(s(bool,X41)) | ~p(s(bool,X40)) | ~p(s(bool,X39))) & (p(s(bool,X40)) | ~p(s(bool,X41))) & (p(s(bool,X39)) | ~p(s(bool,X41)))))),
  file('thm.extreal.le_addr_imp','thm.sat.sat_23_prove_ltb')).
fof(f5149,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f5135,f1259])).
fof(f1259,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f817])).
fof(f5135,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f842,f1359,f842,f1266])).
fof(f32901,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30944,f1359])).
fof(f30944,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6434,f24933,f1270])).
fof(f6434,plain,(
  ~sP7('F_0',sK28('const.arithmetic.ZERO_0'),'T_0')),
  inference(unit_resulting_resolution,[],[f1359,f842,f6173,f1275])).
fof(f1275,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f821])).
fof(f32900,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30945,f1359])).
fof(f30945,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7118,f24933,f1270])).
fof(f7118,plain,(
  ~sP7('F_0',sK27('const.arithmetic.ZERO_0'),'T_0')),
  inference(unit_resulting_resolution,[],[f1359,f842,f6174,f1275])).
fof(f32899,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30946,f1359])).
fof(f30946,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6442,f24933,f1270])).
fof(f6442,plain,(
  ~sP7('F_0','T_0',sK28('const.arithmetic.ZERO_0'))),
  inference(unit_resulting_resolution,[],[f1359,f842,f6173,f1275])).
fof(f32898,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30947,f1359])).
fof(f30947,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6443,f24933,f1270])).
fof(f6443,plain,(
  ~sP7('F_0',sK28('const.arithmetic.ZERO_0'),sK28('const.arithmetic.ZERO_0'))),
  inference(unit_resulting_resolution,[],[f1359,f6173,f6173,f1275])).
fof(f32897,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30948,f1359])).
fof(f30948,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7126,f24933,f1270])).
fof(f7126,plain,(
  ~sP7('F_0','T_0',sK27('const.arithmetic.ZERO_0'))),
  inference(unit_resulting_resolution,[],[f1359,f842,f6174,f1275])).
fof(f32896,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30949,f1359])).
fof(f30949,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7127,f24933,f1270])).
fof(f7127,plain,(
  ~sP7('F_0',sK27('const.arithmetic.ZERO_0'),sK27('const.arithmetic.ZERO_0'))),
  inference(unit_resulting_resolution,[],[f1359,f6174,f6174,f1275])).
fof(f32895,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30950,f24933])).
fof(f30950,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1832,f24933,f1270])).
fof(f1832,plain,(
  ( ! [X0] : (~sP7(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1359,f1723,f1260])).
fof(f32894,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30951,f24933])).
fof(f30951,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1833,f24933,f1270])).
fof(f1833,plain,(
  ( ! [X0] : (~sP7(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1359,f1723,f1261])).
fof(f32893,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30952,f24933])).
fof(f30952,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1896,f24933,f1270])).
fof(f1896,plain,(
  ( ! [X0] : (~sP7(sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1359,f1724,f1260])).
fof(f32892,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30953,f24933])).
fof(f30953,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1897,f24933,f1270])).
fof(f1897,plain,(
  ( ! [X0] : (~sP7(sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1359,f1724,f1261])).
fof(f32891,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30954,f24933])).
fof(f30954,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7029,f24933,f1270])).
fof(f7029,plain,(
  ( ! [X0] : (~sP7(sK27('const.arithmetic.ZERO_0'),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1359,f6174,f1260])).
fof(f32890,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30955,f24933])).
fof(f30955,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7030,f24933,f1270])).
fof(f7030,plain,(
  ( ! [X0] : (~sP7(sK27('const.arithmetic.ZERO_0'),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1359,f6174,f1261])).
fof(f32889,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30956,f24933])).
fof(f30956,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6345,f24933,f1270])).
fof(f6345,plain,(
  ( ! [X0] : (~sP7(sK28('const.arithmetic.ZERO_0'),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1359,f6173,f1260])).
fof(f32888,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30957,f24933])).
fof(f30957,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6346,f24933,f1270])).
fof(f6346,plain,(
  ( ! [X0] : (~sP7(sK28('const.arithmetic.ZERO_0'),'F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1359,f6173,f1261])).
fof(f32887,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30962,f1359])).
fof(f30962,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1793,f24933,f1270])).
fof(f1793,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f842,f1359,f1260])).
fof(f32886,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30963,f24933])).
fof(f30963,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1796,f24933,f1270])).
fof(f1796,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f842,f1359,f1261])).
fof(f32885,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30967,f24933])).
fof(f30967,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6442,f24933,f1270])).
fof(f32884,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30968,f24933])).
fof(f30968,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6443,f24933,f1270])).
fof(f32883,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30969,f24933])).
fof(f30969,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7126,f24933,f1270])).
fof(f32882,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30970,f24933])).
fof(f30970,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7127,f24933,f1270])).
fof(f32881,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30971,f1359])).
fof(f30971,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1832,f24933,f1270])).
fof(f32880,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30972,f24933])).
fof(f30972,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1833,f24933,f1270])).
fof(f32879,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30973,f1359])).
fof(f30973,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1896,f24933,f1270])).
fof(f32878,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30974,f24933])).
fof(f30974,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1897,f24933,f1270])).
fof(f32877,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30975,f1359])).
fof(f30975,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7029,f24933,f1270])).
fof(f32876,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30976,f24933])).
fof(f30976,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7030,f24933,f1270])).
fof(f32875,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30977,f1359])).
fof(f30977,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6345,f24933,f1270])).
fof(f32874,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30978,f24933])).
fof(f30978,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6346,f24933,f1270])).
fof(f32873,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30985,f1359])).
fof(f30985,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f5149,f24933,f1271])).
fof(f1271,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f821])).
fof(f32872,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30986,f1359])).
fof(f30986,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6434,f24933,f1271])).
fof(f32871,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30987,f1359])).
fof(f30987,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7118,f24933,f1271])).
fof(f32870,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30988,f1359])).
fof(f30988,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6442,f24933,f1271])).
fof(f32869,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30989,f1359])).
fof(f30989,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6443,f24933,f1271])).
fof(f32868,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30990,f1359])).
fof(f30990,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7126,f24933,f1271])).
fof(f32867,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30991,f1359])).
fof(f30991,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7127,f24933,f1271])).
fof(f32866,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30992,f24933])).
fof(f30992,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1832,f24933,f1271])).
fof(f32865,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30993,f24933])).
fof(f30993,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1833,f24933,f1271])).
fof(f32864,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30994,f24933])).
fof(f30994,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1896,f24933,f1271])).
fof(f32863,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30995,f24933])).
fof(f30995,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1897,f24933,f1271])).
fof(f32862,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30996,f24933])).
fof(f30996,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7029,f24933,f1271])).
fof(f32861,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30997,f24933])).
fof(f30997,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7030,f24933,f1271])).
fof(f32860,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30998,f24933])).
fof(f30998,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6345,f24933,f1271])).
fof(f32859,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f30999,f24933])).
fof(f30999,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6346,f24933,f1271])).
fof(f32858,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31004,f24933])).
fof(f31004,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1793,f24933,f1271])).
fof(f32857,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31005,f1359])).
fof(f31005,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1796,f24933,f1271])).
fof(f32856,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31007,f24933])).
fof(f31007,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6434,f24933,f1271])).
fof(f32855,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31008,f24933])).
fof(f31008,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7118,f24933,f1271])).
fof(f32854,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31010,f24933])).
fof(f31010,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6443,f24933,f1271])).
fof(f32853,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31012,f24933])).
fof(f31012,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7127,f24933,f1271])).
fof(f32852,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31013,f24933])).
fof(f31013,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1832,f24933,f1271])).
fof(f32851,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31014,f1359])).
fof(f31014,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1833,f24933,f1271])).
fof(f32850,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31015,f24933])).
fof(f31015,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1896,f24933,f1271])).
fof(f32849,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31016,f1359])).
fof(f31016,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1897,f24933,f1271])).
fof(f32848,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31017,f24933])).
fof(f31017,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7029,f24933,f1271])).
fof(f32847,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31018,f1359])).
fof(f31018,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7030,f24933,f1271])).
fof(f32846,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31019,f24933])).
fof(f31019,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6345,f24933,f1271])).
fof(f32845,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31020,f1359])).
fof(f31020,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6346,f24933,f1271])).
fof(f32844,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31080,f31663])).
fof(f31663,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f24933,f24933,f1287])).
fof(f1287,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f829])).
fof(f829,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f828])).
fof(f828,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f666])).
fof(f666,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f539,f665,f664])).
fof(f664,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f665,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f539,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f538])).
fof(f538,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X39,X40,X41] : ((p(s(bool,X41)) <=> (p(s(bool,X40)) | p(s(bool,X39)))) <=> ((p(s(bool,X41)) | ~p(s(bool,X40))) & (p(s(bool,X41)) | ~p(s(bool,X39))) & (p(s(bool,X40)) | p(s(bool,X39)) | ~p(s(bool,X41)))))),
  file('thm.extreal.le_addr_imp','thm.sat.sat_25_prove_ltb')).
fof(f31080,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'T_0')) ) | $spl14),
  inference(unit_resulting_resolution,[],[f842,f24933,f1276])).
fof(f1276,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f824])).
fof(f824,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f823])).
fof(f823,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f822])).
fof(f822,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f665])).
fof(f32843,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31107,f31663])).
fof(f31107,plain,(
  ( ! [X0,X1] : (~sP9(X0,'T_0',X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f842,f24933,f1277])).
fof(f1277,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f824])).
fof(f32842,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31226,f24933])).
fof(f31226,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f1835,f24933,f1282])).
fof(f1282,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f824])).
fof(f1835,plain,(
  ( ! [X0] : (~sP9('F_0',X0,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1359,f1723,f1276])).
fof(f1443,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1359,f1286])).
fof(f1286,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f827])).
fof(f827,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f826])).
fof(f826,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f825])).
fof(f825,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f664])).
fof(f32841,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31227,f24933])).
fof(f31227,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f7131,f24933,f1282])).
fof(f7131,plain,(
  ( ! [X0] : (~sP9('F_0',X0,sK27('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f1359,f6174,f1276])).
fof(f32840,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31228,f24933])).
fof(f31228,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f6447,f24933,f1282])).
fof(f6447,plain,(
  ( ! [X0] : (~sP9('F_0',X0,sK28('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f1359,f6173,f1276])).
fof(f32839,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31229,f24933])).
fof(f31229,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f1841,f24933,f1282])).
fof(f1841,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f842,f1359,f1277])).
fof(f32838,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31230,f24933])).
fof(f31230,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1440,f1841,f24933,f1282])).
fof(f1440,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f842,f1285])).
fof(f1285,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f827])).
fof(f32837,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31232,f24933])).
fof(f31232,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1438,f1841,f24933,f1282])).
fof(f1438,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) )),
  inference(unit_resulting_resolution,[],[f840,f1284])).
fof(f1284,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f827])).
fof(f32836,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31233,f24933])).
fof(f31233,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1463,f1841,f24933,f1282])).
fof(f1463,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f850,f1284])).
fof(f32835,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31234,f24933])).
fof(f31234,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1439,f1841,f24933,f1282])).
fof(f1439,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f844,f1284])).
fof(f32834,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31235,f24933])).
fof(f31235,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1742,f1841,f24933,f1282])).
fof(f1742,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f1721,f1284])).
fof(f32833,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31236,f24933])).
fof(f31236,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1782,f1841,f24933,f1282])).
fof(f1782,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1722,f1284])).
fof(f32832,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31237,f24933])).
fof(f31237,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1836,f1841,f24933,f1282])).
fof(f1836,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1723,f1284])).
fof(f32831,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31238,f24933])).
fof(f31238,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7165,f1841,f24933,f1282])).
fof(f7165,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,sK27('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f6174,f1284])).
fof(f32830,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31239,f24933])).
fof(f31239,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6481,f1841,f24933,f1282])).
fof(f6481,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,sK28('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f6173,f1284])).
fof(f32829,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31240,f24933])).
fof(f31240,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f7135,f24933,f1282])).
fof(f7135,plain,(
  ( ! [X0] : (~sP9('F_0',sK27('const.arithmetic.ZERO_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f1359,f6174,f1277])).
fof(f32828,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31241,f24933])).
fof(f31241,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7166,f7135,f24933,f1282])).
fof(f7166,plain,(
  ( ! [X0,X1] : (sP8(X0,sK27('const.arithmetic.ZERO_0'),X1)) )),
  inference(unit_resulting_resolution,[],[f6174,f1285])).
fof(f32827,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31243,f24933])).
fof(f31243,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1438,f7135,f24933,f1282])).
fof(f32826,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31244,f24933])).
fof(f31244,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1463,f7135,f24933,f1282])).
fof(f32825,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31245,f24933])).
fof(f31245,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1439,f7135,f24933,f1282])).
fof(f32824,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31246,f24933])).
fof(f31246,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1742,f7135,f24933,f1282])).
fof(f32823,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31247,f24933])).
fof(f31247,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1782,f7135,f24933,f1282])).
fof(f32822,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31248,f24933])).
fof(f31248,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1836,f7135,f24933,f1282])).
fof(f32821,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31249,f24933])).
fof(f31249,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7165,f7135,f24933,f1282])).
fof(f32820,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31250,f24933])).
fof(f31250,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6481,f7135,f24933,f1282])).
fof(f32819,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31251,f24933])).
fof(f31251,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f6451,f24933,f1282])).
fof(f6451,plain,(
  ( ! [X0] : (~sP9('F_0',sK28('const.arithmetic.ZERO_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f1359,f6173,f1277])).
fof(f32818,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31252,f24933])).
fof(f31252,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6482,f6451,f24933,f1282])).
fof(f6482,plain,(
  ( ! [X0,X1] : (sP8(X0,sK28('const.arithmetic.ZERO_0'),X1)) )),
  inference(unit_resulting_resolution,[],[f6173,f1285])).
fof(f32817,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31254,f24933])).
fof(f31254,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1438,f6451,f24933,f1282])).
fof(f32816,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31255,f24933])).
fof(f31255,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1463,f6451,f24933,f1282])).
fof(f32815,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31256,f24933])).
fof(f31256,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1439,f6451,f24933,f1282])).
fof(f32814,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31257,f24933])).
fof(f31257,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1742,f6451,f24933,f1282])).
fof(f32813,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31258,f24933])).
fof(f31258,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1782,f6451,f24933,f1282])).
fof(f32812,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31259,f24933])).
fof(f31259,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1836,f6451,f24933,f1282])).
fof(f32811,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31260,f24933])).
fof(f31260,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7165,f6451,f24933,f1282])).
fof(f32810,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31261,f24933])).
fof(f31261,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6481,f6451,f24933,f1282])).
fof(f32809,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31284,f31552])).
fof(f31552,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1286])).
fof(f31284,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f24933,f24933,f1282])).
fof(f1815,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f842,f1359,f1276])).
fof(f32808,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31286,f31552])).
fof(f31286,plain,(
  ( ! [X0] : (~sP8('F_0',X0,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f24933,f24933,f1282])).
fof(f32807,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31288,f31552])).
fof(f31288,plain,(
  ( ! [X0] : (~sP8('F_0',X0,sK27('const.arithmetic.ZERO_0'))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f24933,f24933,f1282])).
fof(f32806,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31290,f31552])).
fof(f31290,plain,(
  ( ! [X0] : (~sP8('F_0',X0,sK28('const.arithmetic.ZERO_0'))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f24933,f24933,f1282])).
fof(f32805,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31292,f31552])).
fof(f31292,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f24933,f24933,f1282])).
fof(f32804,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31294,f31552])).
fof(f31294,plain,(
  ( ! [X0] : (~sP8('F_0',sK27('const.arithmetic.ZERO_0'),X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f24933,f24933,f1282])).
fof(f32803,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31296,f31552])).
fof(f31296,plain,(
  ( ! [X0] : (~sP8('F_0',sK28('const.arithmetic.ZERO_0'),X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f24933,f24933,f1282])).
fof(f32802,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31313,f31552])).
fof(f31313,plain,(
  ~sP8('F_0','T_0','F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f1359,f24933,f1282])).
fof(f32801,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31315,f31552])).
fof(f31315,plain,(
  ~sP8('F_0',sK27('const.arithmetic.ZERO_0'),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f1359,f24933,f1282])).
fof(f32800,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31317,f31552])).
fof(f31317,plain,(
  ~sP8('F_0',sK28('const.arithmetic.ZERO_0'),'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f1359,f24933,f1282])).
fof(f32799,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31321,f24933])).
fof(f31321,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1437,f24933,f1282])).
fof(f1437,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f842,f1284])).
fof(f32798,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31323,f24933])).
fof(f31323,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f1437,f24933,f1282])).
fof(f32797,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31324,f24933])).
fof(f31324,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f1437,f24933,f1282])).
fof(f32796,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31326,f24933])).
fof(f31326,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f1438,f24933,f1282])).
fof(f32795,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31327,f24933])).
fof(f31327,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f1438,f24933,f1282])).
fof(f32794,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31329,f24933])).
fof(f31329,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f1463,f24933,f1282])).
fof(f32793,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31330,f24933])).
fof(f31330,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f1463,f24933,f1282])).
fof(f32792,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31332,f24933])).
fof(f31332,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f1439,f24933,f1282])).
fof(f32791,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31333,f24933])).
fof(f31333,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f1439,f24933,f1282])).
fof(f32790,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31335,f24933])).
fof(f31335,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f1742,f24933,f1282])).
fof(f32789,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31336,f24933])).
fof(f31336,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f1742,f24933,f1282])).
fof(f32788,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31338,f24933])).
fof(f31338,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f1782,f24933,f1282])).
fof(f32787,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31339,f24933])).
fof(f31339,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f1782,f24933,f1282])).
fof(f32786,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31340,f24933])).
fof(f31340,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f1836,f24933,f1282])).
fof(f32785,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31342,f24933])).
fof(f31342,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f1836,f24933,f1282])).
fof(f32784,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31343,f24933])).
fof(f31343,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f1836,f24933,f1282])).
fof(f32783,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31344,f24933])).
fof(f31344,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f7165,f24933,f1282])).
fof(f32782,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31346,f24933])).
fof(f31346,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f7165,f24933,f1282])).
fof(f32781,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31347,f24933])).
fof(f31347,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f7165,f24933,f1282])).
fof(f32780,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31348,f24933])).
fof(f31348,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f6481,f24933,f1282])).
fof(f32779,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31350,f24933])).
fof(f31350,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f6481,f24933,f1282])).
fof(f32778,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31351,f24933])).
fof(f31351,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f6481,f24933,f1282])).
fof(f32777,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31353,f24933])).
fof(f31353,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1441,f24933,f1282])).
fof(f1441,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f840,f1285])).
fof(f32776,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31354,f24933])).
fof(f31354,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1464,f24933,f1282])).
fof(f1464,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f850,f1285])).
fof(f32775,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31355,f24933])).
fof(f31355,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1442,f24933,f1282])).
fof(f1442,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f844,f1285])).
fof(f32774,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31356,f24933])).
fof(f31356,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1743,f24933,f1282])).
fof(f1743,plain,(
  ( ! [X0,X1] : (sP8(X0,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),X1)) )),
  inference(unit_resulting_resolution,[],[f1721,f1285])).
fof(f32773,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31357,f24933])).
fof(f31357,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1783,f24933,f1282])).
fof(f1783,plain,(
  ( ! [X0,X1] : (sP8(X0,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),X1)) )),
  inference(unit_resulting_resolution,[],[f1722,f1285])).
fof(f32772,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31358,f24933])).
fof(f31358,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1837,f24933,f1282])).
fof(f1837,plain,(
  ( ! [X0,X1] : (sP8(X0,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),X1)) )),
  inference(unit_resulting_resolution,[],[f1723,f1285])).
fof(f32771,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31359,f24933])).
fof(f31359,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f7166,f24933,f1282])).
fof(f32770,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31360,f24933])).
fof(f31360,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f6482,f24933,f1282])).
fof(f32769,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31361,f24933])).
fof(f31361,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1443,f24933,f1282])).
fof(f32768,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31362,f24933])).
fof(f31362,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f1443,f24933,f1282])).
fof(f32767,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31363,f24933])).
fof(f31363,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f1443,f24933,f1282])).
fof(f32766,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31364,f24933])).
fof(f31364,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f1443,f24933,f1282])).
fof(f32765,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31366,f24933])).
fof(f31366,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f1443,f24933,f1282])).
fof(f32764,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31367,f24933])).
fof(f31367,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f1443,f24933,f1282])).
fof(f32763,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31369,f31552])).
fof(f31369,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f24933,f24933,f1282])).
fof(f32762,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31370,f31552])).
fof(f31370,plain,(
  ( ! [X0] : (~sP8('F_0',X0,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f24933,f24933,f1282])).
fof(f32761,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31371,f31552])).
fof(f31371,plain,(
  ( ! [X0] : (~sP8('F_0',X0,sK27('const.arithmetic.ZERO_0'))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f24933,f24933,f1282])).
fof(f32760,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31372,f31552])).
fof(f31372,plain,(
  ( ! [X0] : (~sP8('F_0',X0,sK28('const.arithmetic.ZERO_0'))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f24933,f24933,f1282])).
fof(f32759,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31373,f31552])).
fof(f31373,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f24933,f24933,f1282])).
fof(f32758,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31374,f31552])).
fof(f31374,plain,(
  ( ! [X0] : (~sP8('F_0',sK27('const.arithmetic.ZERO_0'),X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f24933,f24933,f1282])).
fof(f32757,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31375,f31552])).
fof(f31375,plain,(
  ( ! [X0] : (~sP8('F_0',sK28('const.arithmetic.ZERO_0'),X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f24933,f24933,f1282])).
fof(f32756,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31419,f31552])).
fof(f31419,plain,(
  ~sP8('F_0','F_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1359,f24933,f1282])).
fof(f32755,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31420,f31552])).
fof(f31420,plain,(
  ~sP8('F_0','F_0',sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f1359,f24933,f1282])).
fof(f32754,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31421,f31552])).
fof(f31421,plain,(
  ~sP8('F_0','F_0',sK27('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f1359,f24933,f1282])).
fof(f32753,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31422,f31552])).
fof(f31422,plain,(
  ~sP8('F_0','F_0',sK28('const.arithmetic.ZERO_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f1359,f24933,f1282])).
fof(f32752,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31443,f24933])).
fof(f31443,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1438,f842,f24933,f1283])).
fof(f1283,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f827])).
fof(f32751,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31444,f24933])).
fof(f31444,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1463,f842,f24933,f1283])).
fof(f32750,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31445,f24933])).
fof(f31445,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1439,f842,f24933,f1283])).
fof(f32749,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31446,f24933])).
fof(f31446,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1742,f842,f24933,f1283])).
fof(f32748,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31447,f24933])).
fof(f31447,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1782,f842,f24933,f1283])).
fof(f32747,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31448,f24933])).
fof(f31448,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1836,f842,f24933,f1283])).
fof(f32746,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31449,f24933])).
fof(f31449,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7165,f842,f24933,f1283])).
fof(f32745,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31450,f24933])).
fof(f31450,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6481,f842,f24933,f1283])).
fof(f32744,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31451,f24933])).
fof(f31451,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1440,f842,f24933,f1283])).
fof(f32743,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31452,f24933])).
fof(f31452,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1441,f842,f24933,f1283])).
fof(f32742,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31453,f24933])).
fof(f31453,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1464,f842,f24933,f1283])).
fof(f32741,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31454,f24933])).
fof(f31454,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1442,f842,f24933,f1283])).
fof(f32740,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31455,f24933])).
fof(f31455,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1743,f842,f24933,f1283])).
fof(f32739,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31456,f24933])).
fof(f31456,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1783,f842,f24933,f1283])).
fof(f32738,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31457,f24933])).
fof(f31457,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1837,f842,f24933,f1283])).
fof(f32737,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31458,f24933])).
fof(f31458,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7166,f842,f24933,f1283])).
fof(f32736,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31459,f24933])).
fof(f31459,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6482,f842,f24933,f1283])).
fof(f32735,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31481,f31552])).
fof(f31481,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f842,f24933,f24933,f1283])).
fof(f32734,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31495,f31552])).
fof(f31495,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) ) | $spl14),
  inference(unit_resulting_resolution,[],[f842,f1359,f24933,f1283])).
fof(f32733,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31497,f24933])).
fof(f31497,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1437,f842,f24933,f1283])).
fof(f32732,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31498,f24933])).
fof(f31498,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1438,f842,f24933,f1283])).
fof(f32731,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31499,f24933])).
fof(f31499,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1463,f842,f24933,f1283])).
fof(f32730,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31500,f24933])).
fof(f31500,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1439,f842,f24933,f1283])).
fof(f32729,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31501,f24933])).
fof(f31501,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1742,f842,f24933,f1283])).
fof(f32728,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31502,f24933])).
fof(f31502,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1782,f842,f24933,f1283])).
fof(f32727,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31503,f24933])).
fof(f31503,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1836,f842,f24933,f1283])).
fof(f32726,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31504,f24933])).
fof(f31504,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7165,f842,f24933,f1283])).
fof(f32725,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31505,f24933])).
fof(f31505,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6481,f842,f24933,f1283])).
fof(f32724,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31507,f24933])).
fof(f31507,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1441,f842,f24933,f1283])).
fof(f32723,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31508,f24933])).
fof(f31508,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1464,f842,f24933,f1283])).
fof(f32722,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31509,f24933])).
fof(f31509,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1442,f842,f24933,f1283])).
fof(f32721,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31510,f24933])).
fof(f31510,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1743,f842,f24933,f1283])).
fof(f32720,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31511,f24933])).
fof(f31511,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1783,f842,f24933,f1283])).
fof(f32719,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31512,f24933])).
fof(f31512,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1837,f842,f24933,f1283])).
fof(f32718,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31513,f24933])).
fof(f31513,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7166,f842,f24933,f1283])).
fof(f32717,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31514,f24933])).
fof(f31514,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6482,f842,f24933,f1283])).
fof(f32716,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31536,f31552])).
fof(f31536,plain,(
  ( ! [X0,X1] : (~sP8('T_0',X0,X1)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f842,f24933,f24933,f1283])).
fof(f32715,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31550,f31552])).
fof(f31550,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f842,f1359,f24933,f1283])).
fof(f32714,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31554,f1359])).
fof(f31554,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f5237,f24933,f24933,f1287])).
fof(f5237,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f5222,f1278])).
fof(f1278,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f824])).
fof(f5222,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f842,f1359,f1359,f1283])).
fof(f32713,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31556,f24933])).
fof(f31556,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f24933,f24933,f1287])).
fof(f32712,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31557,f24933])).
fof(f31557,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f24933,f24933,f1287])).
fof(f32711,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31558,f24933])).
fof(f31558,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f24933,f24933,f1287])).
fof(f32710,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31559,f24933])).
fof(f31559,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f24933,f24933,f1287])).
fof(f32709,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31560,f24933])).
fof(f31560,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f24933,f24933,f1287])).
fof(f32708,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31561,f24933])).
fof(f31561,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f24933,f24933,f1287])).
fof(f32707,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31562,f1359])).
fof(f31562,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7178,f24933,f24933,f1287])).
fof(f7178,plain,(
  ~sP9(sK27('const.arithmetic.ZERO_0'),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1359,f6174,f1290])).
fof(f1290,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f829])).
fof(f32706,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31563,f1359])).
fof(f31563,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6494,f24933,f24933,f1287])).
fof(f6494,plain,(
  ~sP9(sK28('const.arithmetic.ZERO_0'),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1359,f1359,f6173,f1290])).
fof(f32705,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31565,f24933])).
fof(f31565,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f1359,f24933,f1287])).
fof(f32704,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31566,f24933])).
fof(f31566,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f1359,f24933,f1287])).
fof(f32703,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31567,f24933])).
fof(f31567,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f1359,f24933,f1287])).
fof(f32702,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31568,f24933])).
fof(f31568,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f1359,f24933,f1287])).
fof(f32701,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31569,f24933])).
fof(f31569,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f1359,f24933,f1287])).
fof(f32700,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31570,f24933])).
fof(f31570,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f1359,f24933,f1287])).
fof(f32699,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31572,f1359])).
fof(f31572,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f24933,f24933,f1287])).
fof(f32698,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31573,f1359])).
fof(f31573,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f24933,f24933,f1287])).
fof(f32697,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31574,f1359])).
fof(f31574,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f24933,f24933,f1287])).
fof(f32696,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31575,f1359])).
fof(f31575,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f24933,f24933,f1287])).
fof(f32695,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31576,f1359])).
fof(f31576,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f24933,f24933,f1287])).
fof(f32694,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31577,f1359])).
fof(f31577,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f24933,f24933,f1287])).
fof(f32693,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31578,f1359])).
fof(f31578,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f24933,f24933,f1287])).
fof(f32692,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31579,f24933])).
fof(f31579,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7178,f24933,f24933,f1287])).
fof(f32691,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31580,f24933])).
fof(f31580,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6494,f24933,f24933,f1287])).
fof(f32690,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31586,f1359])).
fof(f31586,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f1359,f24933,f1287])).
fof(f32689,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31587,f1359])).
fof(f31587,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f1359,f24933,f1287])).
fof(f32688,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31588,f1359])).
fof(f31588,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f1359,f24933,f1287])).
fof(f32687,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31589,f24933])).
fof(f31589,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7178,f1359,f24933,f1287])).
fof(f32686,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31590,f24933])).
fof(f31590,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6494,f1359,f24933,f1287])).
fof(f32685,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31595,f1359])).
fof(f31595,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f5237,f24933,f24933,f1287])).
fof(f32684,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31596,f24933])).
fof(f31596,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f24933,f24933,f1287])).
fof(f32683,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31597,f24933])).
fof(f31597,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f24933,f24933,f1287])).
fof(f32682,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31598,f24933])).
fof(f31598,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f24933,f24933,f1287])).
fof(f32681,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31599,f24933])).
fof(f31599,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f24933,f24933,f1287])).
fof(f32680,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31601,f24933])).
fof(f31601,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f24933,f24933,f1287])).
fof(f32679,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31602,f24933])).
fof(f31602,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f24933,f24933,f1287])).
fof(f32678,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31603,f1359])).
fof(f31603,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7178,f24933,f24933,f1287])).
fof(f32677,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31604,f1359])).
fof(f31604,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6494,f24933,f24933,f1287])).
fof(f32676,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31605,f24933])).
fof(f31605,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1359,f24933,f1287])).
fof(f32675,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31606,f24933])).
fof(f31606,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f1359,f24933,f1287])).
fof(f32674,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31607,f24933])).
fof(f31607,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f1359,f24933,f1287])).
fof(f32673,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31608,f24933])).
fof(f31608,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f1359,f24933,f1287])).
fof(f32672,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31610,f24933])).
fof(f31610,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f1359,f24933,f1287])).
fof(f32671,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31611,f24933])).
fof(f31611,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f1359,f24933,f1287])).
fof(f32670,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31613,f1359])).
fof(f31613,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f24933,f24933,f1287])).
fof(f32669,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31614,f1359])).
fof(f31614,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f24933,f24933,f1287])).
fof(f32668,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31615,f1359])).
fof(f31615,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f24933,f24933,f1287])).
fof(f32667,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31616,f1359])).
fof(f31616,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f24933,f24933,f1287])).
fof(f32666,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31617,f1359])).
fof(f31617,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f24933,f24933,f1287])).
fof(f32665,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31618,f1359])).
fof(f31618,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f24933,f24933,f1287])).
fof(f32664,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31619,f1359])).
fof(f31619,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f24933,f24933,f1287])).
fof(f32663,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31620,f24933])).
fof(f31620,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7178,f24933,f24933,f1287])).
fof(f32662,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31621,f24933])).
fof(f31621,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6494,f24933,f24933,f1287])).
fof(f32661,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31627,f1359])).
fof(f31627,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1359,f24933,f1287])).
fof(f32660,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31628,f1359])).
fof(f31628,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f1359,f24933,f1287])).
fof(f32659,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31629,f1359])).
fof(f31629,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f1359,f24933,f1287])).
fof(f32658,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31630,f1359])).
fof(f31630,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f1359,f24933,f1287])).
fof(f32657,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31631,f24933])).
fof(f31631,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7178,f1359,f24933,f1287])).
fof(f32656,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31632,f24933])).
fof(f31632,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6494,f1359,f24933,f1287])).
fof(f32655,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31637,f1359])).
fof(f31637,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f5237,f24933,f24933,f1287])).
fof(f32654,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31638,f24933])).
fof(f31638,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f24933,f24933,f1287])).
fof(f32653,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31639,f24933])).
fof(f31639,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f24933,f24933,f1287])).
fof(f32652,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31640,f24933])).
fof(f31640,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f24933,f24933,f1287])).
fof(f32651,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31641,f24933])).
fof(f31641,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f24933,f24933,f1287])).
fof(f32650,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31643,f24933])).
fof(f31643,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f24933,f24933,f1287])).
fof(f32649,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31644,f24933])).
fof(f31644,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f24933,f24933,f1287])).
fof(f32648,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31645,f1359])).
fof(f31645,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7178,f24933,f24933,f1287])).
fof(f32647,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31646,f1359])).
fof(f31646,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6494,f24933,f24933,f1287])).
fof(f32646,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31647,f1359])).
fof(f31647,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f5237,f1359,f24933,f1287])).
fof(f32645,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31649,f24933])).
fof(f31649,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f1359,f24933,f1287])).
fof(f32644,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31650,f24933])).
fof(f31650,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f1359,f24933,f1287])).
fof(f32643,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31651,f1359])).
fof(f31651,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7178,f1359,f24933,f1287])).
fof(f32642,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31652,f1359])).
fof(f31652,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6494,f1359,f24933,f1287])).
fof(f32641,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31653,f1359])).
fof(f31653,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f5237,f24933,f24933,f1287])).
fof(f32640,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31655,f24933])).
fof(f31655,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f24933,f24933,f1287])).
fof(f32639,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31656,f24933])).
fof(f31656,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f24933,f24933,f1287])).
fof(f32638,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31657,f24933])).
fof(f31657,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f24933,f24933,f1287])).
fof(f32637,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31658,f24933])).
fof(f31658,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f24933,f24933,f1287])).
fof(f32636,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31659,f24933])).
fof(f31659,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f24933,f24933,f1287])).
fof(f32635,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31660,f24933])).
fof(f31660,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f24933,f24933,f1287])).
fof(f32634,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31661,f1359])).
fof(f31661,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7178,f24933,f24933,f1287])).
fof(f32633,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31662,f1359])).
fof(f31662,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6494,f24933,f24933,f1287])).
fof(f32632,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31667,f1359])).
fof(f31667,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f5237,f1359,f24933,f1287])).
fof(f32631,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31669,f24933])).
fof(f31669,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f1359,f24933,f1287])).
fof(f32630,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31670,f24933])).
fof(f31670,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f1359,f24933,f1287])).
fof(f32629,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31671,f24933])).
fof(f31671,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f1359,f24933,f1287])).
fof(f32628,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31672,f1359])).
fof(f31672,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7178,f1359,f24933,f1287])).
fof(f32627,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31673,f1359])).
fof(f31673,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6494,f1359,f24933,f1287])).
fof(f32626,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31679,f24933])).
fof(f31679,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1959,f842,f24933,f1290])).
fof(f1959,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) )),
  inference(unit_resulting_resolution,[],[f1438,f842,f1342])).
fof(f1342,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f1279])).
fof(f1279,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f824])).
fof(f32625,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31680,f24933])).
fof(f31680,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1961,f842,f24933,f1290])).
fof(f1961,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f1439,f842,f1342])).
fof(f32624,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31681,f24933])).
fof(f31681,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7169,f842,f24933,f1290])).
fof(f7169,plain,(
  ( ! [X0] : (sP9('T_0',X0,sK27('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f842,f6174,f1288])).
fof(f1288,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f829])).
fof(f32623,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31682,f24933])).
fof(f31682,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6485,f842,f24933,f1290])).
fof(f6485,plain,(
  ( ! [X0] : (sP9('T_0',X0,sK28('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f842,f6173,f1288])).
fof(f32622,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31683,f24933])).
fof(f31683,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1854,f842,f24933,f1290])).
fof(f1854,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f842,f842,f1289])).
fof(f1289,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f829])).
fof(f32621,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31684,f24933])).
fof(f31684,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1963,f842,f24933,f1290])).
fof(f1963,plain,(
  ( ! [X0] : (sP9('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f1441,f842,f1342])).
fof(f32620,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31685,f24933])).
fof(f31685,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1965,f842,f24933,f1290])).
fof(f1965,plain,(
  ( ! [X0] : (sP9('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f1442,f842,f1342])).
fof(f32619,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31686,f24933])).
fof(f31686,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7172,f842,f24933,f1290])).
fof(f7172,plain,(
  ( ! [X0] : (sP9('T_0',sK27('const.arithmetic.ZERO_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f842,f6174,f1289])).
fof(f32618,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31687,f24933])).
fof(f31687,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6488,f842,f24933,f1290])).
fof(f6488,plain,(
  ( ! [X0] : (sP9('T_0',sK28('const.arithmetic.ZERO_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f842,f6173,f1289])).
fof(f32617,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31770,f24933])).
fof(f31770,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1846,f842,f24933,f1290])).
fof(f1846,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f842,f842,f1288])).
fof(f32616,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31771,f24933])).
fof(f31771,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1959,f842,f24933,f1290])).
fof(f32615,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31772,f24933])).
fof(f31772,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1961,f842,f24933,f1290])).
fof(f32614,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31773,f24933])).
fof(f31773,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7169,f842,f24933,f1290])).
fof(f32613,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31774,f24933])).
fof(f31774,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6485,f842,f24933,f1290])).
fof(f32612,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31776,f24933])).
fof(f31776,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1963,f842,f24933,f1290])).
fof(f32611,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31777,f24933])).
fof(f31777,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1965,f842,f24933,f1290])).
fof(f32610,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31778,f24933])).
fof(f31778,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7172,f842,f24933,f1290])).
fof(f32609,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31779,f24933])).
fof(f31779,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6488,f842,f24933,f1290])).
fof(f32608,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31784,f1359])).
fof(f31784,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1868,f24933,f1293])).
fof(f1293,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f832])).
fof(f832,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f831])).
fof(f831,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f830])).
fof(f830,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f668])).
fof(f668,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f1868,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f842,f1359,f1305])).
fof(f1305,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f837])).
fof(f837,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f836])).
fof(f836,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f669])).
fof(f669,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f635,f668,f667])).
fof(f667,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f635,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f541])).
fof(f541,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f540])).
fof(f540,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f61])).
fof(f61,axiom,(
  ! [X39,X40,X41] : ((p(s(bool,X41)) <=> (p(s(bool,X40)) => p(s(bool,X39)))) <=> ((p(s(bool,X41)) | p(s(bool,X40))) & (p(s(bool,X41)) | ~p(s(bool,X39))) & (~p(s(bool,X40)) | p(s(bool,X39)) | ~p(s(bool,X41)))))),
  file('thm.extreal.le_addr_imp','thm.sat.sat_27_prove_ltb')).
fof(f32607,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31785,f24933])).
fof(f31785,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1937,f24933,f1293])).
fof(f1937,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f1447,f842,f1341])).
fof(f1341,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f1296])).
fof(f1296,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f832])).
fof(f1447,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f870,f1301])).
fof(f1301,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f835])).
fof(f835,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f834])).
fof(f834,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f833])).
fof(f833,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f667])).
fof(f870,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f148])).
fof(f148,axiom,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.extreal.le_addr_imp','thm.numeral.numeral_distrib|split|34')).
fof(f32606,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31786,f24933])).
fof(f31786,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1938,f24933,f1293])).
fof(f1938,plain,(
  ( ! [X0] : (sP11('T_0',X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))) )),
  inference(unit_resulting_resolution,[],[f1448,f842,f1341])).
fof(f1448,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))) )),
  inference(unit_resulting_resolution,[],[f841,f1301])).
fof(f32605,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31787,f24933])).
fof(f31787,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1876,f24933,f1293])).
fof(f1876,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f842,f842,f1306])).
fof(f1306,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f837])).
fof(f32604,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31788,f24933])).
fof(f31788,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1940,f24933,f1293])).
fof(f1940,plain,(
  ( ! [X0] : (sP11('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f1450,f842,f1341])).
fof(f1450,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),X1)) )),
  inference(unit_resulting_resolution,[],[f840,f1302])).
fof(f1302,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f835])).
fof(f32603,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31789,f24933])).
fof(f31789,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1942,f24933,f1293])).
fof(f1942,plain,(
  ( ! [X0] : (sP11('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f1451,f842,f1341])).
fof(f1451,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f844,f1302])).
fof(f32602,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31790,f24933])).
fof(f31790,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7256,f24933,f1293])).
fof(f7256,plain,(
  ( ! [X0] : (sP11('T_0',sK27('const.arithmetic.ZERO_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f842,f6174,f1306])).
fof(f32601,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31791,f24933])).
fof(f31791,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6572,f24933,f1293])).
fof(f6572,plain,(
  ( ! [X0] : (sP11('T_0',sK28('const.arithmetic.ZERO_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f842,f6173,f1306])).
fof(f32600,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31793,f24933])).
fof(f31793,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6570,f24933,f1293])).
fof(f6570,plain,(
  sP11('F_0','F_0',sK28('const.arithmetic.ZERO_0'))),
  inference(unit_resulting_resolution,[],[f1359,f1359,f6173,f1304])).
fof(f1304,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f837])).
fof(f32599,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31794,f24933])).
fof(f31794,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7254,f24933,f1293])).
fof(f7254,plain,(
  sP11('F_0','F_0',sK27('const.arithmetic.ZERO_0'))),
  inference(unit_resulting_resolution,[],[f1359,f1359,f6174,f1304])).
fof(f32598,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31797,f1359])).
fof(f31797,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7273,f24933,f1293])).
fof(f7273,plain,(
  ( ! [X0] : (sP11(sK27('const.arithmetic.ZERO_0'),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1446,f6174,f1341])).
fof(f1446,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1359,f1301])).
fof(f32597,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31798,f24933])).
fof(f31798,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7276,f24933,f1293])).
fof(f7276,plain,(
  ( ! [X0] : (sP11(sK27('const.arithmetic.ZERO_0'),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1449,f6174,f1341])).
fof(f1449,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f842,f1302])).
fof(f32596,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31799,f24933])).
fof(f31799,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7277,f24933,f1293])).
fof(f7277,plain,(
  ( ! [X0] : (sP11(sK27('const.arithmetic.ZERO_0'),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f1450,f6174,f1341])).
fof(f32595,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31800,f24933])).
fof(f31800,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7260,f24933,f1293])).
fof(f7260,plain,(
  ( ! [X0] : (sP11(sK27('const.arithmetic.ZERO_0'),sK27('const.arithmetic.ZERO_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f6174,f6174,f1306])).
fof(f32594,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31801,f1359])).
fof(f31801,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6589,f24933,f1293])).
fof(f6589,plain,(
  ( ! [X0] : (sP11(sK28('const.arithmetic.ZERO_0'),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f1446,f6173,f1341])).
fof(f32593,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31802,f24933])).
fof(f31802,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6592,f24933,f1293])).
fof(f6592,plain,(
  ( ! [X0] : (sP11(sK28('const.arithmetic.ZERO_0'),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f1449,f6173,f1341])).
fof(f32592,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31803,f24933])).
fof(f31803,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6593,f24933,f1293])).
fof(f6593,plain,(
  ( ! [X0] : (sP11(sK28('const.arithmetic.ZERO_0'),'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f1450,f6173,f1341])).
fof(f32591,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31804,f24933])).
fof(f31804,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6576,f24933,f1293])).
fof(f6576,plain,(
  ( ! [X0] : (sP11(sK28('const.arithmetic.ZERO_0'),sK28('const.arithmetic.ZERO_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f6173,f6173,f1306])).
fof(f32590,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31817,f1359])).
fof(f31817,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f5294,f24933,f1293])).
fof(f5294,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1452,f842,f1359,f1299])).
fof(f1299,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f832])).
fof(f1452,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1359,f1303])).
fof(f1303,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f835])).
fof(f32589,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31818,f1359])).
fof(f31818,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6570,f24933,f1293])).
fof(f32588,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31819,f1359])).
fof(f31819,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7254,f24933,f1293])).
fof(f32587,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31820,f24933])).
fof(f31820,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))) | $spl14),
  inference(unit_resulting_resolution,[],[f5295,f24933,f1293])).
fof(f5295,plain,(
  sP11('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1454,f842,f1359,f1299])).
fof(f1454,plain,(
  ( ! [X0,X1] : (sP10('const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f841,f1303])).
fof(f32586,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31821,f24933])).
fof(f31821,plain,(
  p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f5296,f24933,f1293])).
fof(f5296,plain,(
  sP11('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1453,f842,f1359,f1299])).
fof(f1453,plain,(
  ( ! [X0,X1] : (sP10('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f870,f1303])).
fof(f32585,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31822,f24933])).
fof(f31822,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7273,f24933,f1293])).
fof(f32584,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31823,f24933])).
fof(f31823,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7276,f24933,f1293])).
fof(f32583,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31824,f24933])).
fof(f31824,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7277,f24933,f1293])).
fof(f32582,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31825,f24933])).
fof(f31825,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7260,f24933,f1293])).
fof(f32581,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31826,f24933])).
fof(f31826,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6589,f24933,f1293])).
fof(f32580,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31827,f24933])).
fof(f31827,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6592,f24933,f1293])).
fof(f32579,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31828,f24933])).
fof(f31828,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6593,f24933,f1293])).
fof(f32578,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31829,f24933])).
fof(f31829,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6576,f24933,f1293])).
fof(f32577,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31905,f31941])).
fof(f31941,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1303])).
fof(f31905,plain,(
  ~sP10('F_0','T_0','T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1863,f842,f24933,f1299])).
fof(f1863,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f842,f1359,f1294])).
fof(f1294,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f832])).
fof(f32576,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31907,f31941])).
fof(f31907,plain,(
  ~sP10('F_0',sK27('const.arithmetic.ZERO_0'),'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f7182,f842,f24933,f1299])).
fof(f7182,plain,(
  ( ! [X0] : (~sP11('F_0',sK27('const.arithmetic.ZERO_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f1359,f6174,f1294])).
fof(f32575,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31909,f31941])).
fof(f31909,plain,(
  ~sP10('F_0',sK28('const.arithmetic.ZERO_0'),'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f6498,f842,f24933,f1299])).
fof(f6498,plain,(
  ( ! [X0] : (~sP11('F_0',sK28('const.arithmetic.ZERO_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f1359,f6173,f1294])).
fof(f32574,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31937,f31941])).
fof(f31937,plain,(
  ( ! [X0] : (~sP10('T_0',X0,'T_0')) ) | $spl14),
  inference(unit_resulting_resolution,[],[f842,f842,f24933,f1300])).
fof(f1300,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f835])).
fof(f32573,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31959,f1359])).
fof(f31959,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f1863,f842,f24933,f1304])).
fof(f32572,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31960,f1359])).
fof(f31960,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7182,f842,f24933,f1304])).
fof(f32571,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31961,f1359])).
fof(f31961,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6498,f842,f24933,f1304])).
fof(f32570,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31962,f24933])).
fof(f31962,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7261,f842,f24933,f1304])).
fof(f7261,plain,(
  ~sP11(sK27('const.arithmetic.ZERO_0'),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1359,f842,f6174,f1307])).
fof(f1307,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f837])).
fof(f32569,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31963,f24933])).
fof(f31963,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6577,f842,f24933,f1304])).
fof(f6577,plain,(
  ~sP11(sK28('const.arithmetic.ZERO_0'),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1359,f842,f6173,f1307])).
fof(f32568,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31968,f1359])).
fof(f31968,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f5322,f842,f24933,f1304])).
fof(f5322,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f5308,f1295])).
fof(f1295,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f832])).
fof(f5308,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f842,f842,f1359,f1300])).
fof(f32567,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31970,f24933])).
fof(f31970,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f7182,f842,f24933,f1304])).
fof(f32566,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31971,f24933])).
fof(f31971,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6498,f842,f24933,f1304])).
fof(f32565,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31972,f1359])).
fof(f31972,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f7261,f842,f24933,f1304])).
fof(f32564,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f31973,f1359])).
fof(f31973,plain,(
  p(s(bool,'F_0')) | $spl14),
  inference(unit_resulting_resolution,[],[f6577,f842,f24933,f1304])).
fof(f32563,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32149,f1428])).
fof(f32149,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1793,f24933,f1343])).
fof(f32562,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32151,f1428])).
fof(f32151,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1796,f24933,f1343])).
fof(f32561,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32160,f1834])).
fof(f1834,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1723,f1267])).
fof(f32160,plain,(
  ( ! [X0] : (~sP6(X0,'F_0',sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1832,f24933,f1343])).
fof(f32560,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32162,f1834])).
fof(f32162,plain,(
  ( ! [X0] : (~sP6('F_0',X0,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1833,f24933,f1343])).
fof(f32559,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32164,f1898])).
fof(f1898,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f1724,f1267])).
fof(f32164,plain,(
  ( ! [X0] : (~sP6(X0,'F_0',sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1896,f24933,f1343])).
fof(f32558,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32166,f1898])).
fof(f32166,plain,(
  ( ! [X0] : (~sP6('F_0',X0,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1897,f24933,f1343])).
fof(f32557,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32168,f7103])).
fof(f7103,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,sK27('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f6174,f1267])).
fof(f32168,plain,(
  ( ! [X0] : (~sP6(X0,'F_0',sK27('const.arithmetic.ZERO_0'))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7029,f24933,f1343])).
fof(f32556,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32170,f7103])).
fof(f32170,plain,(
  ( ! [X0] : (~sP6('F_0',X0,sK27('const.arithmetic.ZERO_0'))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f7030,f24933,f1343])).
fof(f32555,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32172,f6419])).
fof(f6419,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,sK28('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f6173,f1267])).
fof(f32172,plain,(
  ( ! [X0] : (~sP6(X0,'F_0',sK28('const.arithmetic.ZERO_0'))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6345,f24933,f1343])).
fof(f32554,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32174,f6419])).
fof(f32174,plain,(
  ( ! [X0] : (~sP6('F_0',X0,sK28('const.arithmetic.ZERO_0'))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6346,f24933,f1343])).
fof(f32553,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32177,f24933])).
fof(f32177,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1396,f24933,f1392])).
fof(f1392,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1190,f1359])).
fof(f1190,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f788])).
fof(f788,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f787])).
fof(f787,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f652])).
fof(f652,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f1396,plain,(
  ( ! [X0] : (sP0(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) )),
  inference(unit_resulting_resolution,[],[f840,f1191])).
fof(f1191,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f788])).
fof(f32552,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32178,f24933])).
fof(f32178,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1466,f24933,f1392])).
fof(f1466,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f874,f1191])).
fof(f32551,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32179,f24933])).
fof(f32179,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1479,f24933,f1392])).
fof(f1479,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X1)))) )),
  inference(unit_resulting_resolution,[],[f876,f1191])).
fof(f876,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X6] : p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X6))))),
  file('thm.extreal.le_addr_imp','thm.extreal.le_refl')).
fof(f32550,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32180,f24933])).
fof(f32180,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1455,f24933,f1392])).
fof(f1455,plain,(
  ( ! [X0] : (sP0(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(unit_resulting_resolution,[],[f850,f1191])).
fof(f32549,plain,(
  $false | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f32181,f24933])).
fof(f32181,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f22301,f24933,f1392])).
fof(f22301,plain,(
  ( ! [X0] : (sP0(X0,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) ) | $spl13),
  inference(unit_resulting_resolution,[],[f20973,f1191])).
fof(f32548,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32182,f24933])).
fof(f32182,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1494,f24933,f1392])).
fof(f1494,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)))) )),
  inference(unit_resulting_resolution,[],[f878,f1191])).
fof(f878,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))),
  inference(rectify,[],[f218])).
fof(f218,axiom,(
  ! [X6] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6))))),
  file('thm.extreal.le_addr_imp','thm.real.REAL_LE_REFL')).
fof(f32547,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32183,f24933])).
fof(f32183,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1397,f24933,f1392])).
fof(f1397,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f844,f1191])).
fof(f32546,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32184,f24933])).
fof(f32184,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1883,f24933,f1392])).
fof(f1883,plain,(
  ( ! [X0] : (sP0(X0,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f1724,f1191])).
fof(f32545,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32185,f24933])).
fof(f32185,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6896,f24933,f1392])).
fof(f6896,plain,(
  ( ! [X0] : (sP0(X0,sK27('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f6174,f1191])).
fof(f32544,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32186,f24933])).
fof(f32186,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1728,f24933,f1392])).
fof(f1728,plain,(
  ( ! [X0] : (sP0(X0,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f1721,f1191])).
fof(f32543,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32187,f24933])).
fof(f32187,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1768,f24933,f1392])).
fof(f1768,plain,(
  ( ! [X0] : (sP0(X0,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1722,f1191])).
fof(f32542,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32188,f24933])).
fof(f32188,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1819,f24933,f1392])).
fof(f1819,plain,(
  ( ! [X0] : (sP0(X0,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1723,f1191])).
fof(f32541,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32189,f24933])).
fof(f32189,plain,(
  p(s(bool,sK28('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2519,f24933,f1392])).
fof(f2519,plain,(
  ( ! [X0] : (sP0(X0,sK28('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f1756,f1191])).
fof(f32540,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32190,f24933])).
fof(f32190,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6214,f24933,f1392])).
fof(f6214,plain,(
  ( ! [X0] : (sP0(X0,sK28('const.arithmetic.ZERO_0'))) )),
  inference(unit_resulting_resolution,[],[f6173,f1191])).
fof(f32539,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32191,f24933])).
fof(f32191,plain,(
  p(s(bool,sK28('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2057,f24933,f1392])).
fof(f2057,plain,(
  ( ! [X0] : (sP0(X0,sK28('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f1753,f1191])).
fof(f32538,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32192,f24933])).
fof(f32192,plain,(
  p(s(bool,sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2209,f24933,f1392])).
fof(f2209,plain,(
  ( ! [X0] : (sP0(X0,sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1754,f1191])).
fof(f32537,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32193,f24933])).
fof(f32193,plain,(
  p(s(bool,sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2347,f24933,f1392])).
fof(f2347,plain,(
  ( ! [X0] : (sP0(X0,sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1755,f1191])).
fof(f32536,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32194,f24933])).
fof(f32194,plain,(
  p(s(bool,sK29('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3251,f24933,f1392])).
fof(f3251,plain,(
  ( ! [X0] : (sP0(X0,sK29('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f3145,f1191])).
fof(f32535,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32195,f24933])).
fof(f32195,plain,(
  p(s(bool,sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3859,f24933,f1392])).
fof(f3859,plain,(
  ( ! [X0] : (sP0(X0,sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f3157,f1191])).
fof(f32534,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32196,f24933])).
fof(f32196,plain,(
  p(s(bool,sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f4046,f24933,f1392])).
fof(f4046,plain,(
  ( ! [X0] : (sP0(X0,sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f3233,f1191])).
fof(f32533,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32197,f24933])).
fof(f32197,plain,(
  p(s(bool,sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f4310,f24933,f1392])).
fof(f4310,plain,(
  ( ! [X0] : (sP0(X0,sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) )),
  inference(unit_resulting_resolution,[],[f3245,f1191])).
fof(f32532,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32198,f24933])).
fof(f32198,plain,(
  p(s(bool,sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2759,f24933,f1392])).
fof(f2759,plain,(
  ( ! [X0] : (sP0(X0,sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1761,f1191])).
fof(f32531,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32199,f24933])).
fof(f32199,plain,(
  p(s(bool,sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3079,f24933,f1392])).
fof(f3079,plain,(
  ( ! [X0] : (sP0(X0,sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1787,f1191])).
fof(f32530,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32200,f24933])).
fof(f32200,plain,(
  p(s(bool,sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2925,f24933,f1392])).
fof(f2925,plain,(
  ( ! [X0] : (sP0(X0,sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1762,f1191])).
fof(f32529,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32201,f24933])).
fof(f32201,plain,(
  p(s(bool,sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3165,f24933,f1392])).
fof(f3165,plain,(
  ( ! [X0] : (sP0(X0,sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) )),
  inference(unit_resulting_resolution,[],[f1788,f1191])).
fof(f32528,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32202,f24933])).
fof(f32202,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1398,f24933,f1392])).
fof(f1398,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f842,f1192])).
fof(f1192,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f788])).
fof(f32527,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32203,f24933])).
fof(f32203,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1399,f24933,f1392])).
fof(f1399,plain,(
  ( ! [X0] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),X0)) )),
  inference(unit_resulting_resolution,[],[f840,f1192])).
fof(f32526,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32204,f24933])).
fof(f32204,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1467,f24933,f1392])).
fof(f1467,plain,(
  ( ! [X0,X1] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f874,f1192])).
fof(f32525,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32205,f24933])).
fof(f32205,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1480,f24933,f1392])).
fof(f1480,plain,(
  ( ! [X0,X1] : (sP0('const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f876,f1192])).
fof(f32524,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32206,f24933])).
fof(f32206,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1456,f24933,f1392])).
fof(f1456,plain,(
  ( ! [X0] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f850,f1192])).
fof(f32523,plain,(
  $false | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f32207,f24933])).
fof(f32207,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f22302,f24933,f1392])).
fof(f22302,plain,(
  ( ! [X0] : (sP0('const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),X0)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f20973,f1192])).
fof(f32522,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32208,f24933])).
fof(f32208,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1495,f24933,f1392])).
fof(f1495,plain,(
  ( ! [X0,X1] : (sP0('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f878,f1192])).
fof(f32521,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32209,f24933])).
fof(f32209,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1400,f24933,f1392])).
fof(f1400,plain,(
  ( ! [X0] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f844,f1192])).
fof(f32520,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32210,f24933])).
fof(f32210,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1884,f24933,f1392])).
fof(f1884,plain,(
  ( ! [X0] : (sP0(sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f1724,f1192])).
fof(f32519,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32211,f24933])).
fof(f32211,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6897,f24933,f1392])).
fof(f6897,plain,(
  ( ! [X0] : (sP0(sK27('const.arithmetic.ZERO_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f6174,f1192])).
fof(f32518,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32212,f24933])).
fof(f32212,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1729,f24933,f1392])).
fof(f1729,plain,(
  ( ! [X0] : (sP0(sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f1721,f1192])).
fof(f32517,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32213,f24933])).
fof(f32213,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1769,f24933,f1392])).
fof(f1769,plain,(
  ( ! [X0] : (sP0(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f1722,f1192])).
fof(f32516,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32214,f24933])).
fof(f32214,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1820,f24933,f1392])).
fof(f1820,plain,(
  ( ! [X0] : (sP0(sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f1723,f1192])).
fof(f32515,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32215,f24933])).
fof(f32215,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f2520,f24933,f1392])).
fof(f2520,plain,(
  ( ! [X0] : (sP0(sK28('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f1756,f1192])).
fof(f32514,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32216,f24933])).
fof(f32216,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6215,f24933,f1392])).
fof(f6215,plain,(
  ( ! [X0] : (sP0(sK28('const.arithmetic.ZERO_0'),X0)) )),
  inference(unit_resulting_resolution,[],[f6173,f1192])).
fof(f32513,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32217,f24933])).
fof(f32217,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f2058,f24933,f1392])).
fof(f2058,plain,(
  ( ! [X0] : (sP0(sK28('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f1753,f1192])).
fof(f32512,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32218,f24933])).
fof(f32218,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f2210,f24933,f1392])).
fof(f2210,plain,(
  ( ! [X0] : (sP0(sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f1754,f1192])).
fof(f32511,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32219,f24933])).
fof(f32219,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f2348,f24933,f1392])).
fof(f2348,plain,(
  ( ! [X0] : (sP0(sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f1755,f1192])).
fof(f32510,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32220,f24933])).
fof(f32220,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f3252,f24933,f1392])).
fof(f3252,plain,(
  ( ! [X0] : (sP0(sK29('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f3145,f1192])).
fof(f32509,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32221,f24933])).
fof(f32221,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f3860,f24933,f1392])).
fof(f3860,plain,(
  ( ! [X0] : (sP0(sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f3157,f1192])).
fof(f32508,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32222,f24933])).
fof(f32222,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f4047,f24933,f1392])).
fof(f4047,plain,(
  ( ! [X0] : (sP0(sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f3233,f1192])).
fof(f32507,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32223,f24933])).
fof(f32223,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f4311,f24933,f1392])).
fof(f4311,plain,(
  ( ! [X0] : (sP0(sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f3245,f1192])).
fof(f32506,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32224,f24933])).
fof(f32224,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f2760,f24933,f1392])).
fof(f2760,plain,(
  ( ! [X0] : (sP0(sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f1761,f1192])).
fof(f32505,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32225,f24933])).
fof(f32225,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f3080,f24933,f1392])).
fof(f3080,plain,(
  ( ! [X0] : (sP0(sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f1787,f1192])).
fof(f32504,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32226,f24933])).
fof(f32226,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f2926,f24933,f1392])).
fof(f2926,plain,(
  ( ! [X0] : (sP0(sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f1762,f1192])).
fof(f32503,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32227,f24933])).
fof(f32227,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f3166,f24933,f1392])).
fof(f3166,plain,(
  ( ! [X0] : (sP0(sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))),X0)) )),
  inference(unit_resulting_resolution,[],[f1788,f1192])).
fof(f32502,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32232,f24933])).
fof(f32232,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1395,f24933,f1392])).
fof(f1395,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f842,f1191])).
fof(f32501,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32233,f24933])).
fof(f32233,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1396,f24933,f1392])).
fof(f32500,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32234,f24933])).
fof(f32234,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1466,f24933,f1392])).
fof(f32499,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32235,f24933])).
fof(f32235,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1479,f24933,f1392])).
fof(f32498,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32236,f24933])).
fof(f32236,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1455,f24933,f1392])).
fof(f32497,plain,(
  $false | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f32237,f24933])).
fof(f32237,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f22301,f24933,f1392])).
fof(f32496,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32238,f24933])).
fof(f32238,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1494,f24933,f1392])).
fof(f32495,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32239,f24933])).
fof(f32239,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1397,f24933,f1392])).
fof(f32494,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32240,f24933])).
fof(f32240,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1883,f24933,f1392])).
fof(f32493,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32241,f24933])).
fof(f32241,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6896,f24933,f1392])).
fof(f32492,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32242,f24933])).
fof(f32242,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1728,f24933,f1392])).
fof(f32491,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32243,f24933])).
fof(f32243,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1768,f24933,f1392])).
fof(f32490,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32244,f24933])).
fof(f32244,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1819,f24933,f1392])).
fof(f32489,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32245,f24933])).
fof(f32245,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f2519,f24933,f1392])).
fof(f32488,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32246,f24933])).
fof(f32246,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f6214,f24933,f1392])).
fof(f32487,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32247,f24933])).
fof(f32247,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f2057,f24933,f1392])).
fof(f32486,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32248,f24933])).
fof(f32248,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f2209,f24933,f1392])).
fof(f32485,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32249,f24933])).
fof(f32249,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f2347,f24933,f1392])).
fof(f32484,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32250,f24933])).
fof(f32250,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f3251,f24933,f1392])).
fof(f32483,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32251,f24933])).
fof(f32251,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f3859,f24933,f1392])).
fof(f32482,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32252,f24933])).
fof(f32252,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f4046,f24933,f1392])).
fof(f32481,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32253,f24933])).
fof(f32253,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f4310,f24933,f1392])).
fof(f32480,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32254,f24933])).
fof(f32254,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f2759,f24933,f1392])).
fof(f32479,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32255,f24933])).
fof(f32255,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f3079,f24933,f1392])).
fof(f32478,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32256,f24933])).
fof(f32256,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f2925,f24933,f1392])).
fof(f32477,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32257,f24933])).
fof(f32257,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f3165,f24933,f1392])).
fof(f32476,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32259,f24933])).
fof(f32259,plain,(
  p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | $spl14),
  inference(unit_resulting_resolution,[],[f1399,f24933,f1392])).
fof(f32475,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32260,f24933])).
fof(f32260,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1467,f24933,f1392])).
fof(f32474,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32261,f24933])).
fof(f32261,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X0))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1480,f24933,f1392])).
fof(f32473,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32262,f24933])).
fof(f32262,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1456,f24933,f1392])).
fof(f32472,plain,(
  $false | ($spl13 | $spl14)),
  inference(subsumption_resolution,[],[f32263,f24933])).
fof(f32263,plain,(
  p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))) | ($spl13 | $spl14)),
  inference(unit_resulting_resolution,[],[f22302,f24933,f1392])).
fof(f32471,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32264,f24933])).
fof(f32264,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f1495,f24933,f1392])).
fof(f32470,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32265,f24933])).
fof(f32265,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) | $spl14),
  inference(unit_resulting_resolution,[],[f1400,f24933,f1392])).
fof(f32469,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32266,f24933])).
fof(f32266,plain,(
  p(s(bool,sK27('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1884,f24933,f1392])).
fof(f32468,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32267,f24933])).
fof(f32267,plain,(
  p(s(bool,sK27('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6897,f24933,f1392])).
fof(f32467,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32268,f24933])).
fof(f32268,plain,(
  p(s(bool,sK27('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1729,f24933,f1392])).
fof(f32466,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32269,f24933])).
fof(f32269,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1769,f24933,f1392])).
fof(f32465,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32270,f24933])).
fof(f32270,plain,(
  p(s(bool,sK27('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f1820,f24933,f1392])).
fof(f32464,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32271,f24933])).
fof(f32271,plain,(
  p(s(bool,sK28('const.combin.I_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2520,f24933,f1392])).
fof(f32463,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32272,f24933])).
fof(f32272,plain,(
  p(s(bool,sK28('const.arithmetic.ZERO_0'))) | $spl14),
  inference(unit_resulting_resolution,[],[f6215,f24933,f1392])).
fof(f32462,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32273,f24933])).
fof(f32273,plain,(
  p(s(bool,sK28('const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2058,f24933,f1392])).
fof(f32461,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32274,f24933])).
fof(f32274,plain,(
  p(s(bool,sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2210,f24933,f1392])).
fof(f32460,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32275,f24933])).
fof(f32275,plain,(
  p(s(bool,sK28('const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2348,f24933,f1392])).
fof(f32459,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32276,f24933])).
fof(f32276,plain,(
  p(s(bool,sK29('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3252,f24933,f1392])).
fof(f32458,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32277,f24933])).
fof(f32277,plain,(
  p(s(bool,sK30('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3860,f24933,f1392])).
fof(f32457,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32278,f24933])).
fof(f32278,plain,(
  p(s(bool,sK31('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f4047,f24933,f1392])).
fof(f32456,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32279,f24933])).
fof(f32279,plain,(
  p(s(bool,sK32('const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f4311,f24933,f1392])).
fof(f32455,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32280,f24933])).
fof(f32280,plain,(
  p(s(bool,sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2760,f24933,f1392])).
fof(f32454,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32281,f24933])).
fof(f32281,plain,(
  p(s(bool,sK33('const.extreal.extreal_ainv_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3080,f24933,f1392])).
fof(f32453,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32282,f24933])).
fof(f32282,plain,(
  p(s(bool,sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.NegInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f2926,f24933,f1392])).
fof(f32452,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32283,f24933])).
fof(f32283,plain,(
  p(s(bool,sK33('const.combin.I_1'(s('type.extreal.extreal','const.extreal.PosInf_0'))))) | $spl14),
  inference(unit_resulting_resolution,[],[f3166,f24933,f1392])).
fof(f32451,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32294,f24933])).
fof(f32294,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1147])).
fof(f32450,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32297,f24933])).
fof(f32297,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1114])).
fof(f1114,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) )),
  inference(cnf_transformation,[],[f461])).
fof(f461,plain,(
  ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X6,X7] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.le_addr_imp','thm.extreal.le_total')).
fof(f32449,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32299,f24933])).
fof(f32299,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1114])).
fof(f32448,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32306,f32369])).
fof(f32369,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1085])).
fof(f32306,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1029])).
fof(f1029,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f724])).
fof(f724,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0'))))) & (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))) | s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0)))),
  inference(nnf_transformation,[],[f377])).
fof(f377,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(flattening,[],[f376])).
fof(f376,plain,(
  ! [X0] : (~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  inference(rectify,[],[f209])).
fof(f209,axiom,(
  ! [X6] : (~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.PosInf_0')))))),
  file('thm.extreal.le_addr_imp','thm.extreal.lt_infty|split|5')).
fof(f32447,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32322,f24933])).
fof(f32322,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1147])).
fof(f32446,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32324,f32371])).
fof(f32371,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1084])).
fof(f32324,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1031])).
fof(f1031,plain,(
  ( ! [X0] : (p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) )),
  inference(cnf_transformation,[],[f725])).
fof(f725,plain,(
  ! [X0] : ((s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0))))) & (~p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0)))),
  inference(nnf_transformation,[],[f379])).
fof(f379,plain,(
  ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(flattening,[],[f378])).
fof(f378,plain,(
  ! [X0] : (~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X0)))))),
  inference(rectify,[],[f208])).
fof(f208,axiom,(
  ! [X6] : (~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6) <=> p(s(bool,'const.extreal.extreal_lt_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal',X6)))))),
  file('thm.extreal.le_addr_imp','thm.extreal.lt_infty|split|6')).
fof(f32436,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32330,f24933])).
fof(f32330,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1149])).
fof(f1149,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f762])).
fof(f762,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) & (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))))),
  inference(nnf_transformation,[],[f491])).
fof(f491,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) <=> ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(flattening,[],[f490])).
fof(f490,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) <=> ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(rectify,[],[f94])).
fof(f94,axiom,(
  ! [X7,X6] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) <=> ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X6)))))),
  file('thm.extreal.le_addr_imp','thm.real.real_lt')).
fof(f32435,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32332,f32359])).
fof(f32359,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1069])).
fof(f32332,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1183])).
fof(f1183,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f782])).
fof(f782,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) & ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))))))),
  inference(flattening,[],[f781])).
fof(f781,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) | (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f507])).
fof(f507,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) <=> (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)))),
  inference(flattening,[],[f506])).
fof(f506,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) <=> (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f194])).
fof(f194,axiom,(
  ! [X22,X45] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X22))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X45)))))) <=> (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | ~s('type.num.num','const.num.|48|_0') = s('type.num.num',X45)))),
  file('thm.extreal.le_addr_imp','thm.real.lt_int|split|1')).
fof(f32434,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32334,f32359])).
fof(f32334,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1182])).
fof(f1182,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f782])).
fof(f32433,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32336,f24933])).
fof(f32336,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1149])).
fof(f32432,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32339,f24933])).
fof(f32339,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1113])).
fof(f1113,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f460])).
fof(f460,plain,(
  ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))),
  inference(rectify,[],[f224])).
fof(f224,axiom,(
  ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X6)))))),
  file('thm.extreal.le_addr_imp','thm.real.REAL_LE_TOTAL')).
fof(f32431,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32341,f24933])).
fof(f32341,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1113])).
fof(f32430,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32347,f24933])).
fof(f32347,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1175])).
fof(f1175,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f777])).
fof(f777,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f503])).
fof(f503,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f502])).
fof(f502,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f169])).
fof(f169,axiom,(
  ! [X22,X45] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X22))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X45)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X45),s('type.num.num',X22)))))),
  file('thm.extreal.le_addr_imp','thm.numeral.numeral_lte|split|5')).
fof(f32429,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32349,f24933])).
fof(f32349,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1215])).
fof(f1215,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f797])).
fof(f797,plain,(
  ! [X0,X1] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f517])).
fof(f517,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))))),
  inference(flattening,[],[f516])).
fof(f516,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))))),
  inference(rectify,[],[f178])).
fof(f178,axiom,(
  ! [X45,X22] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X45),s('type.num.num',X22)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.num.num',X45)))))),
  file('thm.extreal.le_addr_imp','thm.arithmetic.NOT_LEQ')).
fof(f32428,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32351,f24933])).
fof(f32351,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1215])).
fof(f32427,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f32353,f24933])).
fof(f32353,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1175])).
fof(f32372,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f854,f24933,f1084])).
fof(f854,plain,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  inference(cnf_transformation,[],[f86])).
fof(f86,axiom,(
  s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.NegInf_0')))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_add_def|split|4')).
fof(f32370,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f851,f24933,f1085])).
fof(f851,plain,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(cnf_transformation,[],[f87])).
fof(f87,axiom,(
  s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal','const.extreal.NegInf_0'),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_add_def|split|2')).
fof(f32368,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f915,f24933,f1081])).
fof(f915,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f219])).
fof(f219,axiom,(
  ! [X6] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.extreal.le_addr_imp','thm.real.REAL_MUL_RZERO')).
fof(f32366,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f915,f24933,f1078])).
fof(f32364,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f915,f24933,f1075])).
fof(f32362,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f915,f24933,f1072])).
fof(f32360,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6088,f24933,f1069])).
fof(f32358,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6088,f24933,f1066])).
fof(f32356,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6088,f24933,f972])).
fof(f972,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f687])).
fof(f687,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f354])).
fof(f354,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f147])).
fof(f147,axiom,(
  ! [X22] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X22))),
  file('thm.extreal.le_addr_imp','thm.numeral.numeral_distrib|split|30')).
fof(f32354,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f24933,f1175])).
fof(f32352,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f24933,f1215])).
fof(f32350,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f24933,f1215])).
fof(f32348,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f24933,f1175])).
fof(f32346,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6088,f24933,f978])).
fof(f978,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f690])).
fof(f32344,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f915,f24933,f1145])).
fof(f1145,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) != s('type.realax.real',X1)) )),
  inference(cnf_transformation,[],[f760])).
fof(f760,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & s('type.realax.real',X0) != s('type.realax.real',X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))),
  inference(flattening,[],[f759])).
fof(f759,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) = s('type.realax.real',X1))) & ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & s('type.realax.real',X0) != s('type.realax.real',X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))),
  inference(nnf_transformation,[],[f487])).
fof(f487,plain,(
  ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) <=> (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) = s('type.realax.real',X1)))),
  inference(rectify,[],[f253])).
fof(f253,axiom,(
  ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) <=> (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X6),s('type.realax.real',X7)))) | s('type.realax.real',X6) = s('type.realax.real',X7)))),
  file('thm.extreal.le_addr_imp','thm.real.REAL_LE_LT')).
fof(f32342,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f24933,f1113])).
fof(f32340,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f24933,f1113])).
fof(f32337,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f24933,f1149])).
fof(f32335,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6101,f24933,f1182])).
fof(f6101,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))) )),
  inference(backward_demodulation,[],[f6088,f1376])).
fof(f1376,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))) )),
  inference(subsumption_resolution,[],[f959,f1359])).
fof(f959,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f681])).
fof(f681,plain,(
  ! [X0] : ((s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f200])).
fof(f200,axiom,(
  ! [X22] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X22))) <=> p(s(bool,'F_0')))),
  file('thm.extreal.le_addr_imp','thm.numeral.numeral_eq|split|2')).
fof(f32333,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6101,f24933,f1183])).
fof(f32331,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f24933,f1149])).
fof(f32325,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1035,f24933,f1031])).
fof(f1035,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f383])).
fof(f383,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f382])).
fof(f382,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X22] : ~s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X22)))),
  file('thm.extreal.le_addr_imp','thm.extreal.num_not_infty|split|0')).
fof(f32323,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f24933,f1147])).
fof(f32307,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1036,f24933,f1029])).
fof(f1036,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f385])).
fof(f385,plain,(
  ! [X0] : s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f384])).
fof(f384,plain,(
  ! [X0] : ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X22] : ~s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num',X22)))),
  file('thm.extreal.le_addr_imp','thm.extreal.num_not_infty|split|1')).
fof(f32304,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1379,f24933,f1229])).
fof(f1379,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f1042,f1041])).
fof(f1042,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f397])).
fof(f397,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f396])).
fof(f396,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8] : ! [X15] : s(X8,'const.combin.I_1'(s(X8,X15))) = s(X8,happ(s(fun(X8,X8),'const.combin.I_0'),s(X8,X15)))),
  file('thm.extreal.le_addr_imp','ARITY_0_1_const.combin.I')).
fof(f32302,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f851,f24933,f1142])).
fof(f32300,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f24933,f1114])).
fof(f32298,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f24933,f1114])).
fof(f32295,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f24933,f1147])).
fof(f32293,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f854,f24933,f976])).
fof(f976,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))) )),
  inference(cnf_transformation,[],[f689])).
fof(f689,plain,(
  ! [X0] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) | s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0)) & (s('type.extreal.extreal','const.extreal.NegInf_0') != s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0'))))))),
  inference(nnf_transformation,[],[f356])).
fof(f356,plain,(
  ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X6] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal','const.extreal.NegInf_0')))) <=> s('type.extreal.extreal','const.extreal.NegInf_0') = s('type.extreal.extreal',X6))),
  file('thm.extreal.le_addr_imp','thm.extreal.le_infty|split|2')).
fof(f32291,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f851,f24933,f974])).
fof(f974,plain,(
  ( ! [X0] : (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0))))) )),
  inference(cnf_transformation,[],[f688])).
fof(f688,plain,(
  ! [X0] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0)) & (s('type.extreal.extreal','const.extreal.PosInf_0') != s('type.extreal.extreal',X0) | p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0))))))),
  inference(nnf_transformation,[],[f355])).
fof(f355,plain,(
  ! [X0] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X0)))) <=> s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X0))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X6] : (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.PosInf_0'),s('type.extreal.extreal',X6)))) <=> s('type.extreal.extreal','const.extreal.PosInf_0') = s('type.extreal.extreal',X6))),
  file('thm.extreal.le_addr_imp','thm.extreal.le_infty|split|3')).
fof(f32287,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1395,f1359,f24933,f1392])).
fof(f32285,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1395,f24933,f24933,f1392])).
fof(f32231,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1398,f1359,f24933,f1392])).
fof(f32229,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1395,f24933,f24933,f1392])).
fof(f32175,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1434,f6346,f24933,f1343])).
fof(f1434,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f1359,f1269])).
fof(f32173,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1431,f6345,f24933,f1343])).
fof(f1431,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f1359,f1268])).
fof(f1268,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f819])).
fof(f32171,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1434,f7030,f24933,f1343])).
fof(f32169,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1431,f7029,f24933,f1343])).
fof(f32167,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1434,f1897,f24933,f1343])).
fof(f32165,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1431,f1896,f24933,f1343])).
fof(f32163,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1434,f1833,f24933,f1343])).
fof(f32161,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1431,f1832,f24933,f1343])).
fof(f32152,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1434,f1796,f24933,f1343])).
fof(f32150,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1431,f1793,f24933,f1343])).
fof(f32090,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f842,f1048,f24933,f1340])).
fof(f1340,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f1339])).
fof(f1339,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f1321])).
fof(f1321,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f839])).
fof(f839,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f838])).
fof(f838,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f637])).
fof(f637,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f636])).
fof(f636,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f550])).
fof(f550,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X6,X26,X7,X27] : ((s(bool,X6) = s(bool,X26) & (p(s(bool,X26)) => s(bool,X7) = s(bool,X27))) => ((p(s(bool,X6)) => p(s(bool,X7))) <=> (p(s(bool,X26)) => p(s(bool,X27)))))),
  file('thm.extreal.le_addr_imp','thm.bool.IMP_CONG')).
fof(f32088,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f842,f1047,f24933,f1340])).
fof(f1047,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f402])).
fof(f402,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f401])).
fof(f401,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f155])).
fof(f155,axiom,(
  ! [X8] : ! [X23,X24] : s(X8,X23) = s(X8,'const.bool.COND_3'(s(bool,'T_0'),s(X8,X23),s(X8,X24)))),
  file('thm.extreal.le_addr_imp','thm.bool.COND_CLAUSES|split|0')).
fof(f32086,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f842,f1049,f24933,f1340])).
fof(f1049,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f406])).
fof(f406,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(flattening,[],[f405])).
fof(f405,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.ind_type.FCONS_3'(s(X0,X1),s(fun('type.num.num',X0),X2),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X8] : ! [X17,X12] : s(X8,X17) = s(X8,'const.ind_type.FCONS_3'(s(X8,X17),s(fun('type.num.num',X8),X12),s('type.num.num','const.num.|48|_0')))),
  file('thm.extreal.le_addr_imp','thm.ind_type.FCONS|split|0')).
fof(f32084,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f842,f1041,f24933,f1340])).
fof(f32082,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f842,f1056,f24933,f1340])).
fof(f1056,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f416])).
fof(f416,plain,(
  ! [X0,X1,X2,X3] : s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(flattening,[],[f415])).
fof(f415,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X1) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X8] : ! [X13,X14,X12] : s(X8,X13) = s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.NegInf_0'),s(X8,X13),s(X8,X14),s(fun('type.realax.real',X8),X12)))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_case_def|split|0')).
fof(f32080,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f842,f1055,f24933,f1340])).
fof(f1055,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))) )),
  inference(cnf_transformation,[],[f414])).
fof(f414,plain,(
  ! [X0,X1,X2,X3] : s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(flattening,[],[f413])).
fof(f413,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,X2) = s(X0,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X0,X1),s(X0,X2),s(fun('type.realax.real',X0),X3)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8] : ! [X13,X14,X12] : s(X8,X14) = s(X8,'const.extreal.extreal_CASE_4'(s('type.extreal.extreal','const.extreal.PosInf_0'),s(X8,X13),s(X8,X14),s(fun('type.realax.real',X8),X12)))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_case_def|split|1')).
fof(f32078,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f842,f1379,f24933,f1340])).
fof(f32075,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f842,f884,f24933,f1340])).
fof(f884,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f133])).
fof(f133,axiom,(
  ! [X22] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X22),s('type.num.num','const.num.|48|_0')))),
  file('thm.extreal.le_addr_imp','thm.numeral.numeral_distrib|split|29')).
fof(f32058,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f842,f887,f24933,f1340])).
fof(f887,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f134])).
fof(f134,axiom,(
  ! [X22] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)))),
  file('thm.extreal.le_addr_imp','thm.numeral.numeral_distrib|split|26')).
fof(f32053,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f842,f1087,f24933,f1340])).
fof(f1087,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f434])).
fof(f434,plain,(
  ! [X0,X1] : s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f156])).
fof(f156,axiom,(
  ! [X22,X45] : s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X22))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X45)))))),
  file('thm.extreal.le_addr_imp','thm.real.le_int|split|1')).
fof(f32042,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f842,f900,f24933,f1340])).
fof(f900,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))) )),
  inference(cnf_transformation,[],[f297])).
fof(f297,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  inference(rectify,[],[f108])).
fof(f108,axiom,(
  ! [X42] : s(bool,'T_0') = s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X42))),s('type.extreal.extreal','const.extreal.PosInf_0')))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_le_def|split|3')).
fof(f32040,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f842,f853,f24933,f1340])).
fof(f32038,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f842,f888,f24933,f1340])).
fof(f32036,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f842,f1603,f24933,f1340])).
fof(f32033,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1876,f842,f842,f24933,f1307])).
fof(f32009,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5322,f842,f24933,f1305])).
fof(f31996,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5322,f842,f1359,f24933,f1304])).
fof(f31986,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5322,f842,f24933,f24933,f1304])).
fof(f31967,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1863,f1359,f842,f24933,f1304])).
fof(f31965,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5322,f24933,f842,f24933,f1304])).
fof(f31942,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5308,f24933,f1303])).
fof(f31940,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5308,f24933,f1301])).
fof(f31938,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1449,f842,f842,f24933,f1300])).
fof(f31910,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1452,f6498,f842,f24933,f1299])).
fof(f31908,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1452,f7182,f842,f24933,f1299])).
fof(f31906,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1452,f1863,f842,f24933,f1299])).
fof(f31856,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1876,f842,f24933,f1294])).
fof(f31833,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5294,f1359,f24933,f1293])).
fof(f31831,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1868,f24933,f24933,f1293])).
fof(f31808,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1868,f1359,f24933,f1293])).
fof(f31806,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1868,f24933,f24933,f1293])).
fof(f31783,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1846,f1359,f842,f24933,f1290])).
fof(f31781,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1846,f24933,f842,f24933,f1290])).
fof(f31730,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1854,f842,f1359,f24933,f1290])).
fof(f31714,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1846,f842,f24933,f24933,f1290])).
fof(f31677,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5237,f1359,f1359,f24933,f1287])).
fof(f31675,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5237,f24933,f1359,f24933,f1287])).
fof(f31666,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5237,f1359,f24933,f24933,f1287])).
fof(f31664,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5237,f24933,f24933,f24933,f1287])).
fof(f31636,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1359,f1359,f24933,f1287])).
fof(f31634,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5237,f24933,f1359,f24933,f1287])).
fof(f31625,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1359,f24933,f24933,f1287])).
fof(f31623,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5237,f24933,f24933,f24933,f1287])).
fof(f31594,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f1359,f1359,f24933,f1287])).
fof(f31592,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5237,f24933,f1359,f24933,f1287])).
fof(f31584,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1359,f24933,f24933,f1287])).
fof(f31582,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5237,f24933,f24933,f24933,f1287])).
fof(f31553,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5222,f24933,f1286])).
fof(f31551,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1437,f842,f1359,f24933,f1283])).
fof(f31537,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1437,f842,f24933,f24933,f1283])).
fof(f31496,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1440,f842,f1359,f24933,f1283])).
fof(f31482,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1437,f842,f24933,f24933,f1283])).
fof(f31439,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1443,f1359,f24933,f1282])).
fof(f31437,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f6481,f1359,f24933,f1282])).
fof(f31435,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f7165,f1359,f24933,f1282])).
fof(f31433,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f1836,f1359,f24933,f1282])).
fof(f31426,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1437,f1359,f24933,f1282])).
fof(f31415,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1443,f24933,f24933,f1282])).
fof(f31413,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6451,f6482,f24933,f24933,f1282])).
fof(f31411,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f7135,f7166,f24933,f24933,f1282])).
fof(f31409,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1837,f24933,f24933,f1282])).
fof(f31407,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1783,f24933,f24933,f1282])).
fof(f31405,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1743,f24933,f24933,f1282])).
fof(f31403,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1442,f24933,f24933,f1282])).
fof(f31401,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1464,f24933,f24933,f1282])).
fof(f31399,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1441,f24933,f24933,f1282])).
fof(f31397,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f1440,f24933,f24933,f1282])).
fof(f31395,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6447,f6481,f24933,f24933,f1282])).
fof(f31393,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f7131,f7165,f24933,f24933,f1282])).
fof(f31391,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1835,f1836,f24933,f24933,f1282])).
fof(f31389,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f1782,f24933,f24933,f1282])).
fof(f31387,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f1742,f24933,f24933,f1282])).
fof(f31385,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f1439,f24933,f24933,f1282])).
fof(f31383,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f1463,f24933,f24933,f1282])).
fof(f31381,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1841,f1438,f24933,f24933,f1282])).
fof(f31379,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1815,f1437,f24933,f24933,f1282])).
fof(f31318,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f6451,f1359,f24933,f1282])).
fof(f31316,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f7135,f1359,f24933,f1282])).
fof(f31314,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f1841,f1359,f24933,f1282])).
fof(f31297,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f6451,f24933,f24933,f1282])).
fof(f31295,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f7135,f24933,f24933,f1282])).
fof(f31293,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f1841,f24933,f24933,f1282])).
fof(f31291,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f6447,f24933,f24933,f1282])).
fof(f31289,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f7131,f24933,f24933,f1282])).
fof(f31287,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f1835,f24933,f24933,f1282])).
fof(f31285,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1443,f1815,f24933,f24933,f1282])).
fof(f31108,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1854,f842,f24933,f1277])).
fof(f31081,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1846,f842,f24933,f1276])).
fof(f31054,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5120,f842,f842,f24933,f1275])).
fof(f5120,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1428,f842,f842,f1262])).
fof(f1262,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f817])).
fof(f31024,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1796,f1359,f24933,f1271])).
fof(f31022,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1793,f24933,f24933,f1271])).
fof(f31003,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5149,f1359,f24933,f1271])).
fof(f31001,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1793,f24933,f24933,f1271])).
fof(f30982,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1793,f1359,f24933,f1270])).
fof(f30980,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1793,f24933,f24933,f1270])).
fof(f30961,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5149,f1359,f24933,f1270])).
fof(f30959,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1793,f24933,f24933,f1270])).
fof(f30940,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5135,f24933,f1269])).
fof(f30938,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5135,f24933,f1268])).
fof(f30936,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1428,f842,f842,f24933,f1266])).
fof(f30908,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5120,f842,f24933,f1261])).
fof(f30879,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5120,f842,f24933,f1260])).
fof(f30847,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f3438,f1041,f24933,f1258])).
fof(f3438,plain,(
  ( ! [X0] : (sP5('T_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f842,f1041,f1256])).
fof(f30812,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f3431,f1639,f24933,f1258])).
fof(f3431,plain,(
  sP5('T_0','F_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))))),
  inference(unit_resulting_resolution,[],[f842,f1639,f1256])).
fof(f30807,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f3428,f1603,f24933,f1258])).
fof(f3428,plain,(
  sP5('T_0','T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)))),
  inference(unit_resulting_resolution,[],[f842,f1603,f1256])).
fof(f30802,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f3427,f1041,f24933,f1258])).
fof(f3427,plain,(
  ( ! [X0] : (sP5('T_0','const.combin.I_1'(s(bool,X0)),X0)) )),
  inference(unit_resulting_resolution,[],[f842,f1041,f1256])).
fof(f30767,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f3420,f1639,f24933,f1258])).
fof(f3420,plain,(
  sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal','const.extreal.extreal_add_2'(s('type.extreal.extreal',sK12),s('type.extreal.extreal',sK13)))),'F_0')),
  inference(unit_resulting_resolution,[],[f842,f1639,f1256])).
fof(f30762,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f3417,f1603,f24933,f1258])).
fof(f3417,plain,(
  sP5('T_0','const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'T_0')),
  inference(unit_resulting_resolution,[],[f842,f1603,f1256])).
fof(f32426,plain,(
  $false | $spl14),
  inference(trivial_inequality_removal,[],[f30757])).
fof(f30757,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f12416,f24933,f1258])).
fof(f12416,plain,(
  sP5(sK28('const.arithmetic.ZERO_0'),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f6302,f1404,f1413,f1359,f1241])).
fof(f32425,plain,(
  $false | $spl14),
  inference(trivial_inequality_removal,[],[f30755])).
fof(f30755,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f12241,f24933,f1258])).
fof(f12241,plain,(
  sP5(sK28('const.arithmetic.ZERO_0'),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1425,f1401,f6301,f842,f1240])).
fof(f32424,plain,(
  $false | $spl14),
  inference(trivial_inequality_removal,[],[f30751])).
fof(f30751,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f12412,f24933,f1258])).
fof(f12412,plain,(
  sP5(sK27('const.arithmetic.ZERO_0'),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f6984,f1404,f1413,f1359,f1241])).
fof(f32423,plain,(
  $false | $spl14),
  inference(trivial_inequality_removal,[],[f30749])).
fof(f30749,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f12192,f24933,f1258])).
fof(f12192,plain,(
  sP5(sK27('const.arithmetic.ZERO_0'),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1425,f1401,f6983,f842,f1240])).
fof(f32422,plain,(
  $false | $spl14),
  inference(trivial_inequality_removal,[],[f30747])).
fof(f30747,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f12410,f24933,f1258])).
fof(f12410,plain,(
  sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1421,f1404,f1413,f1359,f1241])).
fof(f32421,plain,(
  $false | $spl14),
  inference(trivial_inequality_removal,[],[f30746])).
fof(f30746,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f12167,f24933,f1258])).
fof(f12167,plain,(
  sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1425,f1401,f1412,f842,f1240])).
fof(f32420,plain,(
  $false | $spl14),
  inference(trivial_inequality_removal,[],[f30743])).
fof(f30743,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f12408,f24933,f1258])).
fof(f12408,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1420,f1404,f1413,f1359,f1241])).
fof(f32419,plain,(
  $false | $spl14),
  inference(trivial_inequality_removal,[],[f30742])).
fof(f30742,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f12143,f24933,f1258])).
fof(f12143,plain,(
  sP5('const.extreal.extreal_le_2'(s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.extreal.extreal',sK13)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1425,f1401,f1411,f842,f1240])).
fof(f32418,plain,(
  $false | $spl14),
  inference(trivial_inequality_removal,[],[f30718])).
fof(f30718,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl14),
  inference(unit_resulting_resolution,[],[f12407,f24933,f1258])).
fof(f12407,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1419,f1404,f1413,f1359,f1241])).
fof(f32417,plain,(
  $false | $spl14),
  inference(trivial_inequality_removal,[],[f30716])).
fof(f30716,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f12131,f24933,f1258])).
fof(f12131,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1425,f1401,f1410,f842,f1240])).
fof(f30713,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f3452,f1592,f24933,f1255])).
fof(f30711,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f3451,f1591,f24933,f1255])).
fof(f30709,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f3449,f1592,f24933,f1255])).
fof(f30707,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f3448,f1591,f24933,f1255])).
fof(f30673,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f1359,f1359,f24933,f1251])).
fof(f30671,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f24933,f1359,f24933,f1251])).
fof(f30649,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f1359,f24933,f24933,f1251])).
fof(f30647,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f24933,f24933,f24933,f1251])).
fof(f30565,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f1359,f1359,f24933,f1251])).
fof(f30563,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f24933,f1359,f24933,f1251])).
fof(f30541,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f1359,f24933,f24933,f1251])).
fof(f30539,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f24933,f24933,f24933,f1251])).
fof(f30457,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1425,f1359,f1359,f24933,f1251])).
fof(f30455,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f24933,f1359,f24933,f1251])).
fof(f30433,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f1359,f24933,f24933,f1251])).
fof(f30431,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1419,f24933,f24933,f24933,f1251])).
fof(f30349,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5075,f24933,f1250])).
fof(f5075,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f842,f1359,f842,f1247])).
fof(f30347,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5075,f24933,f1249])).
fof(f30345,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1410,f842,f842,f24933,f1247])).
fof(f30317,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5048,f24933,f1246])).
fof(f5048,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f842,f1359,f842,f1243])).
fof(f30315,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f5048,f24933,f1245])).
fof(f30313,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f3541,f842,f842,f24933,f1243])).
fof(f3541,plain,(
  ( ! [X0] : (sP4('T_0',X0,'const.combin.I_1'(s(bool,X0)))) )),
  inference(unit_resulting_resolution,[],[f3427,f1238])).
fof(f29765,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1401,f1415,f24933,f1242])).
fof(f29750,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1406,f1415,f24933,f1242])).
fof(f29693,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f6266,f1413,f24933,f1242])).
fof(f29691,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f6266,f1413,f24933,f1242])).
fof(f29645,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f6948,f1413,f24933,f1242])).
fof(f29643,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f6948,f1413,f24933,f1242])).
fof(f29576,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f1401,f1413,f24933,f1242])).
fof(f29574,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f1401,f1413,f24933,f1242])).
fof(f29572,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1401,f1413,f24933,f1242])).
fof(f29559,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f1404,f1413,f24933,f1242])).
fof(f29557,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f1404,f1413,f24933,f1242])).
fof(f29555,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f1404,f1413,f24933,f1242])).
fof(f29553,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f1404,f1413,f24933,f1242])).
fof(f29548,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f1404,f1413,f24933,f1242])).
fof(f29546,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f1404,f1413,f24933,f1242])).
fof(f29544,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1404,f1413,f24933,f1242])).
fof(f29493,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f1404,f6301,f24933,f1242])).
fof(f29491,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f1404,f6301,f24933,f1242])).
fof(f29478,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f6266,f6301,f24933,f1242])).
fof(f29357,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f1401,f6301,f24933,f1242])).
fof(f28825,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f1404,f6983,f24933,f1242])).
fof(f28823,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f1404,f6983,f24933,f1242])).
fof(f28762,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f6948,f6983,f24933,f1242])).
fof(f28689,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f1401,f6983,f24933,f1242])).
fof(f28004,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1406,f1410,f24933,f1242])).
fof(f27991,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f1404,f1410,f24933,f1242])).
fof(f27989,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f1404,f1410,f24933,f1242])).
fof(f27987,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1404,f1410,f24933,f1242])).
fof(f27974,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f6266,f1410,f24933,f1242])).
fof(f27924,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f6948,f1410,f24933,f1242])).
fof(f27851,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1401,f1410,f24933,f1242])).
fof(f27849,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1401,f1410,f24933,f1242])).
fof(f27223,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1401,f1415,f24933,f1241])).
fof(f27208,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1406,f1415,f24933,f1241])).
fof(f27151,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f6266,f1413,f24933,f1241])).
fof(f27149,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f6266,f1413,f24933,f1241])).
fof(f27103,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f6948,f1413,f24933,f1241])).
fof(f27101,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f6948,f1413,f24933,f1241])).
fof(f27034,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f1401,f1413,f24933,f1241])).
fof(f27032,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f1401,f1413,f24933,f1241])).
fof(f27030,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1401,f1413,f24933,f1241])).
fof(f27017,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f1404,f1413,f24933,f1241])).
fof(f27015,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f1404,f1413,f24933,f1241])).
fof(f27013,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f1404,f1413,f24933,f1241])).
fof(f27011,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f1404,f1413,f24933,f1241])).
fof(f27006,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f1404,f1413,f24933,f1241])).
fof(f27004,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f1404,f1413,f24933,f1241])).
fof(f27002,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1404,f1413,f24933,f1241])).
fof(f26951,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f1404,f6301,f24933,f1241])).
fof(f26949,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f1404,f6301,f24933,f1241])).
fof(f26936,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6225,f6266,f6301,f24933,f1241])).
fof(f26815,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f6341,f1401,f6301,f24933,f1241])).
fof(f26283,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f1404,f6983,f24933,f1241])).
fof(f26281,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f1404,f6983,f24933,f1241])).
fof(f26220,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6907,f6948,f6983,f24933,f1241])).
fof(f26147,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f7025,f1401,f6983,f24933,f1241])).
fof(f25462,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1406,f1410,f24933,f1241])).
fof(f25449,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f1404,f1410,f24933,f1241])).
fof(f25447,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f1404,f1410,f24933,f1241])).
fof(f25445,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1404,f1410,f24933,f1241])).
fof(f25432,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6303,f6224,f6266,f1410,f24933,f1241])).
fof(f25382,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f6985,f6906,f6948,f1410,f24933,f1241])).
fof(f25309,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3452,f1401,f1410,f24933,f1241])).
fof(f25307,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1422,f3451,f1401,f1410,f24933,f1241])).
fof(f25196,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f3428,f842,f842,f24933,f1239])).
fof(f25168,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1664,f842,f24933,f1210])).
fof(f1664,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f842,f842,f1206])).
fof(f1206,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f795])).
fof(f25116,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1664,f842,f24933,f1207])).
fof(f25064,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1677,f1359,f24933,f1205])).
fof(f25062,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1667,f24933,f24933,f1205])).
fof(f25012,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1667,f1359,f24933,f1205])).
fof(f25010,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1667,f24933,f24933,f1205])).
fof(f24960,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1592,f24933,f1126])).
fof(f24958,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1359,f1591,f24933,f1126])).
fof(f24956,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1592,f24933,f1126])).
fof(f24954,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1591,f24933,f1126])).
fof(f24950,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1592,f24933,f1126])).
fof(f24948,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1359,f1591,f24933,f1126])).
fof(f24946,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1592,f24933,f1126])).
fof(f24944,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f24933,f1591,f24933,f1126])).
fof(f24940,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1591,f24933,f956])).
fof(f24938,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f1603,f24933,f951])).
fof(f24936,plain,(
  $false | $spl14),
  inference(unit_resulting_resolution,[],[f842,f24933])).
fof(f24933,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | $spl14),
  inference(cnf_transformation,[],[f24933_D])).
fof(f24933_D,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f24935,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl16),
  inference(cnf_transformation,[],[f24935_D])).
fof(f24935_D,plain,(
  p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f24932,plain,(
  ( ! [X3] : (p(s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0')))) | ~p(s(bool,X3))) )),
  inference(forward_demodulation,[],[f24924,f1379])).
fof(f24924,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | p(s(bool,happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,'const.min.@_1'(s(fun(bool,bool),'const.combin.I_0'))))))) )),
  inference(superposition,[],[f1050,f1379])).
fof(f1050,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1))))))) )),
  inference(cnf_transformation,[],[f578])).
fof(f578,plain,(
  ! [X0,X1,X2] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(ennf_transformation,[],[f408])).
fof(f408,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(flattening,[],[f407])).
fof(f407,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8] : ! [X16,X6] : (p(s(bool,happ(s(fun(X8,bool),X16),s(X8,X6)))) => p(s(bool,happ(s(fun(X8,bool),X16),s(X8,'const.min.@_1'(s(fun(X8,bool),X16)))))))),
  file('thm.extreal.le_addr_imp','thm.bool.SELECT_AX')).
fof(f20968,plain,(
  s('type.extreal.extreal',sK13) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl10),
  inference(cnf_transformation,[],[f20968_D])).
fof(f20968_D,plain,(
  s('type.extreal.extreal',sK13) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f20970,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl13),
  inference(cnf_transformation,[],[f20970_D])).
fof(f20970_D,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f20966,plain,(
  ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal',sK13) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(subsumption_resolution,[],[f20815,f842])).
fof(f20815,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',sK13),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.extreal.extreal',sK13) = s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f1184,f1603])).
fof(f1184,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal',X1)) )),
  inference(cnf_transformation,[],[f784])).
fof(f784,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal',X1)) & (s('type.extreal.extreal',X0) != s('type.extreal.extreal',X1) | (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))))),
  inference(flattening,[],[f783])).
fof(f783,plain,(
  ! [X0,X1] : (((~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) | ~p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) | s('type.extreal.extreal',X0) = s('type.extreal.extreal',X1)) & (s('type.extreal.extreal',X0) != s('type.extreal.extreal',X1) | (p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0)))))))),
  inference(nnf_transformation,[],[f508])).
fof(f508,plain,(
  ! [X0,X1] : ((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X0),s('type.extreal.extreal',X1)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X1),s('type.extreal.extreal',X0))))) <=> s('type.extreal.extreal',X0) = s('type.extreal.extreal',X1))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X6,X7] : ((p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X6),s('type.extreal.extreal',X7)))) & p(s(bool,'const.extreal.extreal_le_2'(s('type.extreal.extreal',X7),s('type.extreal.extreal',X6))))) <=> s('type.extreal.extreal',X6) = s('type.extreal.extreal',X7))),
  file('thm.extreal.le_addr_imp','thm.extreal.le_antisym')).
fof(f16113,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) | $spl6),
  inference(cnf_transformation,[],[f16113_D])).
fof(f16113_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0')))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f16114,plain,(
  ( ! [X3] : (~p(s(bool,sK26(X3,'const.combin.I_0',bool)))) ) | $spl8),
  inference(cnf_transformation,[],[f16114_D])).
fof(f16114_D,plain,(
  ( ! [X3] : (~p(s(bool,sK26(X3,'const.combin.I_0',bool)))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f16104,plain,(
  ( ! [X3] : (~p(s(bool,sK26(X3,'const.combin.I_0',bool))) | p(s(bool,'|dollar|exists_1'(s(fun(bool,bool),'const.combin.I_0'))))) )),
  inference(superposition,[],[f1052,f1379])).
fof(f1052,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK26(X2,X1,X0))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f730])).
fof(f730,plain,(
  ! [X0,X1,X2] : ((s(X0,X2) = s(X0,sK26(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK26(X2,X1,X0)))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK26])],[f579])).
fof(f579,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f410])).
fof(f410,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f409])).
fof(f409,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X8] : ! [X16,X21] : (! [X6] : (s(X8,X6) = s(X8,X21) => p(s(bool,happ(s(fun(X8,bool),X16),s(X8,X6))))) => p(s(bool,'|dollar|exists_1'(s(fun(X8,bool),X16)))))),
  file('thm.extreal.le_addr_imp','thm.bool.JRH_INDUCT_UTIL')).
fof(f5606,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))))) ) | $spl4),
  inference(cnf_transformation,[],[f5606_D])).
fof(f5606_D,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f5601,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl3),
  inference(cnf_transformation,[],[f5601_D])).
fof(f5601_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f5590,plain,(
  ( ! [X9] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X9)))))) )),
  inference(superposition,[],[f967,f1367])).
fof(f1367,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) )),
  inference(forward_demodulation,[],[f911,f858])).
fof(f911,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f308])).
fof(f308,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f251])).
fof(f251,axiom,(
  ! [X17] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.extreal.extreal_size_1'(s('type.extreal.extreal','const.extreal.Normal_1'(s('type.realax.real',X17)))))),
  file('thm.extreal.le_addr_imp','thm.extreal.extreal_size_def|split|2')).
fof(f967,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f685])).
fof(f685,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f352])).
fof(f352,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f125])).
fof(f125,axiom,(
  ! [X22] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X22))),
  file('thm.extreal.le_addr_imp','thm.numeral.numeral_distrib|split|17')).
fof(f5599,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) ) | $spl0),
  inference(cnf_transformation,[],[f5599_D])).
fof(f5599_D,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f5583,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f967,f1370])).
fof(f1370,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f916,f858])).
fof(f916,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f142])).
fof(f142,axiom,(
  ! [X22] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X22),s('type.num.num','const.num.|48|_0')))),
  file('thm.extreal.le_addr_imp','thm.numeral.numeral_distrib|split|11')).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 896
% Passive clauses: 8127
% Generated clauses: 44735
% Final active clauses: 819
% Final passive clauses: 4641
% Input formulas: 261
% Initial clauses: 489
% 
% Duplicate literals: 8
% Trivial inequalities: 91
% Fw subsumption resolutions: 1331
% Fw demodulations: 11492
% Bw demodulations: 695
% Condensations: 17
% 
% Simple tautologies: 77
% Equational tautologies: 66
% Forward subsumptions: 21703
% Fw demodulations to eq. taut.: 464
% Bw demodulations to eq. taut.: 8
% 
% Unit resulting resolution: 18583
% Factoring: 3
% Forward superposition: 6636
% Backward superposition: 5010
% Self superposition: 71
% Equality factoring: 2
% Equality resolution: 27
% 
% Split clauses: 8
% Split components: 16
% SAT solver clauses: 9496
% SAT solver unit clauses: 5238
% SAT solver binary clauses: 3780
% 
% Sat splits: 8
% Sat splitting refutations: 1351
% 
% Memory used [KB]: 28400
% Time elapsed: 1.167 s
% ------------------------------
% ------------------------------
% Success in time 2.462 s
