Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.bool
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.bool
% SZS output start Proof for thm.bool
fof(f152,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f151,f24,f35,f33,f41,f23,f37,f22,f34,f39])).
fof(f39,plain,(
  s(sK0,sK2) = s(sK0,sK3) | $spl6),
  inference(cnf_transformation,[],[f39_D])).
fof(f39_D,plain,(
  s(sK0,sK2) = s(sK0,sK3) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f34,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | $spl2),
  inference(cnf_transformation,[],[f34_D])).
fof(f34_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f22,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | s(sK0,sK2) = s(sK0,sK3)),
  inference(cnf_transformation,[],[f20])).
fof(f20,plain,(
  ((s(sK0,sK2) = s(sK0,sK3) & p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK3))))) | p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))))) & (! [X3] : (s(sK0,sK2) != s(sK0,X3) | ~p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,X3))))) | ~p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f19])).
fof(f19,plain,(
  ? [X0,X1,X2] : ((? [X4] : (s(X0,X2) = s(X0,X4) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (! [X3] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f18])).
fof(f18,plain,(
  ? [X0,X1,X2] : ((? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (! [X3] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f14])).
fof(f14,plain,(
  ? [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <~> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(ennf_transformation,[],[f9])).
fof(f9,plain,(
  ~! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f8])).
fof(f8,plain,(
  ~! [X0] : ! [X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (? [X9] : (s(X6,X8) = s(X6,X9) & p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9))))) <=> p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (? [X9] : (s(X6,X8) = s(X6,X9) & p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9))))) <=> p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X8)))))),
  file('thm.bool.UNWIND_THM|49|','thm.bool.UNWIND_THM|49|')).
fof(f37,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK3)))) | $spl4),
  inference(cnf_transformation,[],[f37_D])).
fof(f37_D,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK3)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f23,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK3))))),
  inference(cnf_transformation,[],[f20])).
fof(f41,plain,(
  $false | ($spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f40,f35])).
fof(f40,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | ($spl4 | $spl6)),
  inference(forward_demodulation,[],[f37,f39])).
fof(f33,plain,(
  ( ! [X3] : (s(sK0,sK2) != s(sK0,X3) | ~p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,X3))))) ) | $spl0),
  inference(cnf_transformation,[],[f33_D])).
fof(f33_D,plain,(
  ( ! [X3] : (s(sK0,sK2) != s(sK0,X3) | ~p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,X3))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f35,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | $spl3),
  inference(cnf_transformation,[],[f35_D])).
fof(f35_D,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f24,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | ~p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,X3)))) | s(sK0,sK2) != s(sK0,X3)) )),
  inference(cnf_transformation,[],[f20])).
fof(f151,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f150,f25])).
fof(f25,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.bool.UNWIND_THM|49|',aHL_TRUTH)).
fof(f150,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f145,f49])).
fof(f49,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) | $spl2),
  inference(subsumption_resolution,[],[f46,f26])).
fof(f26,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f10])).
fof(f10,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.bool.UNWIND_THM|49|',aHL_FALSITY)).
fof(f46,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2))) | $spl2),
  inference(superposition,[],[f34,f27])).
fof(f27,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.bool.UNWIND_THM|49|',aHL_BOOL_CASES)).
fof(f145,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),sK1),s(sK0,sK2)))) | $spl0),
  inference(equality_resolution,[],[f33])).
% SZS output end Proof for thm.bool
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 19
% Passive clauses: 44
% Generated clauses: 129
% Final active clauses: 12
% Final passive clauses: 12
% Input formulas: 6
% Initial clauses: 11
% 
% Duplicate literals: 3
% Trivial inequalities: 2
% Fw subsumption resolutions: 15
% Fw demodulations: 5
% Bw demodulations: 4
% 
% Simple tautologies: 4
% Equational tautologies: 4
% Forward subsumptions: 42
% 
% Binary resolution: 19
% Factoring: 1
% Forward superposition: 24
% Backward superposition: 23
% Self superposition: 6
% Equality factoring: 2
% Equality resolution: 1
% 
% Split clauses: 6
% Split components: 12
% SAT solver clauses: 19
% SAT solver unit clauses: 6
% SAT solver binary clauses: 9
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 6
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 17
% TWLsolver calls for satisfiability: 10
% 
% Memory used [KB]: 639
% Time elapsed: 0.014 s
% ------------------------------
% ------------------------------
% Success in time 0.016 s
