Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (63498)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 124
% Passive clauses: 1848
% Generated clauses: 21842
% Final active clauses: 120
% Final passive clauses: 1639
% Input formulas: 22
% Initial clauses: 23
% 
% Duplicate literals: 1512
% Trivial inequalities: 56
% Fw subsumption resolutions: 2075
% Fw demodulations: 3378
% Bw demodulations: 2
% 
% Simple tautologies: 203
% Equational tautologies: 19
% Forward subsumptions: 11427
% 
% Binary resolution: 11
% Forward superposition: 8423
% Backward superposition: 6016
% Self superposition: 169
% Equality factoring: 36
% Equality resolution: 5
% 
% Split clauses: 770
% Split components: 1545
% SAT solver clauses: 1253
% SAT solver unit clauses: 1
% SAT solver binary clauses: 20
% 
% Sat splits: 770
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1253
% TWLsolver calls for satisfiability: 84
% 
% Memory used [KB]: 14456
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (63512)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 91
% Passive clauses: 1451
% Generated clauses: 13781
% Final active clauses: 86
% Final passive clauses: 1349
% Input formulas: 22
% Initial clauses: 23
% 
% Duplicate literals: 568
% Trivial inequalities: 53
% Fw subsumption resolutions: 573
% Fw demodulations: 1459
% Bw demodulations: 12
% Condensations: 688
% 
% Simple tautologies: 128
% Equational tautologies: 18
% Forward subsumptions: 6776
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 11
% Forward superposition: 4751
% Backward superposition: 4308
% Self superposition: 184
% Equality factoring: 45
% Equality resolution: 5
% 
% Split clauses: 119
% Split components: 243
% SAT solver clauses: 2087
% SAT solver unit clauses: 40
% SAT solver binary clauses: 26
% 
% Sat splits: 119
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2087
% TWLsolver calls for satisfiability: 112
% 
% Memory used [KB]: 18166
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% (63541)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 309
% Passive clauses: 9838
% Generated clauses: 187292
% Final active clauses: 78
% Final passive clauses: 473
% Input formulas: 22
% Initial clauses: 23
% Discarded non-redundant clauses: 17901
% 
% Duplicate literals: 4606
% Trivial inequalities: 752
% Fw subsumption resolutions: 9256
% Fw demodulations: 23379
% Bw demodulations: 3726
% Condensations: 11677
% Global subsumptions: 125
% 
% Simple tautologies: 1115
% Equational tautologies: 328
% Forward subsumptions: 121742
% Fw demodulations to eq. taut.: 14
% Bw demodulations to eq. taut.: 606
% 
% Unit resulting resolution: 1149
% Forward superposition: 42760
% Backward superposition: 36145
% Self superposition: 829
% Equality factoring: 298
% Equality resolution: 52
% 
% Split clauses: 484
% Split components: 978
% SAT solver clauses: 11063
% SAT solver unit clauses: 133
% SAT solver binary clauses: 1110
% 
% Sat splits: 484
% Sat splitting refutations: 129
% 
% Memory used [KB]: 102087
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 6
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 6
% Input formulas: 22
% Initial clauses: 8
% 
% Selected by SInE selection: 7
% SInE iterations: 1
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 1
% 
% Memory used [KB]: 767
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.ringNorm
% (63803)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 86
% Passive clauses: 1447
% Generated clauses: 14447
% Final active clauses: 67
% Final passive clauses: 1009
% Input formulas: 22
% Initial clauses: 23
% 
% Duplicate literals: 237
% Trivial inequalities: 171
% Fw subsumption resolutions: 874
% Fw demodulations: 2413
% Bw demodulations: 188
% Global subsumptions: 3
% 
% Simple tautologies: 244
% Equational tautologies: 76
% Forward subsumptions: 8440
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 22
% 
% Unit resulting resolution: 616
% Forward superposition: 4776
% Backward superposition: 4615
% Self superposition: 173
% Equality factoring: 30
% Equality resolution: 13
% 
% Split clauses: 11
% Split components: 22
% SAT solver clauses: 1509
% SAT solver unit clauses: 85
% SAT solver binary clauses: 313
% 
% Sat splits: 11
% Sat splitting refutations: 46
% 
% Memory used [KB]: 18933
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 175 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 21
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 21
% Input formulas: 22
% Initial clauses: 23
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 3
% SAT solver unit clauses: 1
% 
% Memory used [KB]: 895
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% remaining time: 175 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 21
% Passive clauses: 1
% Generated clauses: 2
% Final active clauses: 21
% Input formulas: 22
% Initial clauses: 23
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 5
% SAT solver unit clauses: 2
% SAT solver binary clauses: 1
% 
% Memory used [KB]: 895
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% remaining time: 175 next slice time: 85
lrs+1_5:4_cond=on:fsr=off:fde=none:gs=on:gsem=on:lwlo=on:nm=64:nwc=1:stl=60:sos=all:spl=off_80 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 21
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 21
% Input formulas: 22
% Initial clauses: 23
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 1
% SAT solver unit clauses: 1
% 
% TWLsolver clauses: 1
% TWLsolver calls for satisfiability: 1
% 
% Memory used [KB]: 895
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% remaining time: 175 next slice time: 115
lrs+11_28_cond=on:gs=on:gsaa=from_current:gsem=on:gsssp=full:lwlo=on:nm=64:nwc=1.7:stl=30:sac=on:sdd=off:sfr=on:ssfp=40000:ssfq=1.1:smm=off:ssnc=none:urr=on:updr=off_109 on thm.ringNorm
% (63837)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 128
% Passive clauses: 1000
% Generated clauses: 38654
% Final active clauses: 124
% Final passive clauses: 239
% Input formulas: 22
% Initial clauses: 23
% Discarded non-redundant clauses: 19234
% 
% Duplicate literals: 3647
% Trivial inequalities: 361
% Fw subsumption resolutions: 1636
% Fw demodulations: 2301
% Bw demodulations: 56
% Condensations: 6522
% Global subsumptions: 1
% 
% Simple tautologies: 538
% Equational tautologies: 1256
% Forward subsumptions: 19727
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 33
% Unit resulting resolution: 4512
% Forward superposition: 10310
% Backward superposition: 8398
% Self superposition: 548
% Equality factoring: 493
% Equality resolution: 44
% 
% Split clauses: 21
% Split components: 43
% SAT solver clauses: 1128
% SAT solver unit clauses: 67
% SAT solver binary clauses: 79
% SAT solver learnt clauses: 61
% SAT solver learnt literals: 226
% 
% Sat splits: 21
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1024
% TWLsolver calls for satisfiability: 41650
% 
% Memory used [KB]: 18933
% Time elapsed: 11.600 s
% ------------------------------
% ------------------------------
% remaining time: 59 next slice time: 59
ott+11_3:1_br=off:gs=on:gsem=on:nwc=1:sas=minisat:sos=all:spl=off:urr=on:updr=off_135 on thm.ringNorm
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 21
% Passive clauses: 1
% Generated clauses: 1
% Final active clauses: 21
% Input formulas: 22
% Initial clauses: 23
% 
% Simple tautologies: 2
% 
% SAT solver clauses: 2
% SAT solver unit clauses: 1
% 
% Memory used [KB]: 1023
% Time elapsed: 0.003 s
% ------------------------------
% ------------------------------
% remaining time: 59 next slice time: 59
lrs+3_5_bd=preordered:bsr=unit_only:gsp=input_only:gs=on:gsem=on:gsssp=full:lwlo=on:nm=64:nwc=1:sas=minisat:stl=90:spl=off:sp=occurrence:urr=ec_only:updr=off_815 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f29622,plain,(
  $false),
  inference(subsumption_resolution,[],[f68,f29621])).
fof(f29621,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))))) = s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))),s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X4)))))) )),
  inference(forward_demodulation,[],[f29620,f16734])).
fof(f16734,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3)))) )),
  inference(forward_demodulation,[],[f16524,f83])).
fof(f83,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X20,X21,X22] : s(X6,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X6),X20),s('type.quote.varmap'(X6),X21),s('type.canonical.spolynom'(X6),X22))) = s(X6,happ(s(fun('type.canonical.spolynom'(X6),X6),happ(s(fun('type.quote.varmap'(X6),fun('type.canonical.spolynom'(X6),X6)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X6),X20))),s('type.quote.varmap'(X6),X21))),s('type.canonical.spolynom'(X6),X22)))),
  file('thm.ringNorm.interp_sp_def|split|3','ARITY_1_3_const.ringNorm.r_interp_sp')).
fof(f16524,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))) )),
  inference(superposition,[],[f82,f72])).
fof(f72,plain,(
  ( ! [X0,X1] : (s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1] : s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0] : ! [X1] : s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X7] : s(fun('type.quote.varmap'(X6),fun('type.canonical.spolynom'(X6),X6)),'const.ringNorm.r_interp_sp_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.quote.varmap'(X6),fun('type.canonical.spolynom'(X6),X6)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.interp_sp_def|split|3','thm.ringNorm.r_interp_sp_def')).
fof(f82,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))) = s(X0,happ(s(fun('type.canonical.spolynom'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.spolynom'(X0),X0)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.spolynom'(X0),X3)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X20,X21,X22] : s(X6,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X6),X20),s('type.quote.varmap'(X6),X21),s('type.canonical.spolynom'(X6),X22))) = s(X6,happ(s(fun('type.canonical.spolynom'(X6),X6),happ(s(fun('type.quote.varmap'(X6),fun('type.canonical.spolynom'(X6),X6)),'const.canonical.interp_sp_1'(s('type.semi_ring.semi_ring'(X6),X20))),s('type.quote.varmap'(X6),X21))),s('type.canonical.spolynom'(X6),X22)))),
  file('thm.ringNorm.interp_sp_def|split|3','ARITY_1_3_const.canonical.interp_sp')).
fof(f29620,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))))) = s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))),s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X4)))))) )),
  inference(forward_demodulation,[],[f29619,f16734])).
fof(f29619,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))))) = s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))),s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X4)))))) )),
  inference(forward_demodulation,[],[f29350,f16734])).
fof(f29350,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))))) = s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))),s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X4)))))) )),
  inference(superposition,[],[f88,f27729])).
fof(f27729,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s(X0,X3)))) )),
  inference(forward_demodulation,[],[f27702,f80])).
fof(f80,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X20,X21,X22] : s(X6,'const.ring.ring_RM_3'(s('type.ring.ring'(X6),X20),s(X6,X21),s(X6,X22))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X20))),s(X6,X21))),s(X6,X22)))),
  file('thm.ringNorm.interp_sp_def|split|3','ARITY_1_3_const.ring.ring_RM')).
fof(f27702,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(superposition,[],[f78,f26515])).
fof(f26515,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1))) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(superposition,[],[f87,f73])).
fof(f73,plain,(
  ( ! [X0,X1] : (s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : ! [X1] : s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))) = s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X0),X1))),s(X0,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RP_1'(s('type.ring.ring'(X0),X1))),s(fun(X0,fun(X0,X0)),'const.ring.ring_RM_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7] : s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7))) = s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,'const.ring.ring_R|48|_1'(s('type.ring.ring'(X6),X7))),s(X6,'const.ring.ring_R|49|_1'(s('type.ring.ring'(X6),X7))),s(fun(X6,fun(X6,X6)),'const.ring.ring_RP_1'(s('type.ring.ring'(X6),X7))),s(fun(X6,fun(X6,X6)),'const.ring.ring_RM_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.interp_sp_def|split|3','thm.ring.semi_ring_of_def')).
fof(f87,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3,X4] : s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(fun(X0,fun(X0,X0)),X4) = s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X16,X17,X18,X19] : s(fun(X6,fun(X6,X6)),X19) = s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),'const.semi_ring.recordtype|46|semi_ring_4'(s(X6,X16),s(X6,X17),s(fun(X6,fun(X6,X6)),X18),s(fun(X6,fun(X6,X6)),X19)))))),
  file('thm.ringNorm.interp_sp_def|split|3','thm.semi_ring.semi_ring_SRM')).
fof(f78,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,X2),s(X0,X3))) = s(X0,happ(s(fun(X0,X0),happ(s(fun(X0,fun(X0,X0)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X0),X1))),s(X0,X2))),s(X0,X3)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X20,X21,X22] : s(X6,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X6),X20),s(X6,X21),s(X6,X22))) = s(X6,happ(s(fun(X6,X6),happ(s(fun(X6,fun(X6,X6)),'const.semi_ring.semi_ring_SRM_1'(s('type.semi_ring.semi_ring'(X6),X20))),s(X6,X21))),s(X6,X22)))),
  file('thm.ringNorm.interp_sp_def|split|3','ARITY_1_3_const.semi_ring.semi_ring_SRM')).
fof(f88,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))))) = s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))),s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X4)))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))))) = s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))),s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X4)))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))))) = s(X0,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X0),X1),s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))),s(X0,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X4)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X13,X8,X9,X10] : s(X6,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X6),X13),s('type.quote.varmap'(X6),X8),s('type.canonical.spolynom'(X6),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X6),X9),s('type.canonical.spolynom'(X6),X10))))) = s(X6,'const.semi_ring.semi_ring_SRM_3'(s('type.semi_ring.semi_ring'(X6),X13),s(X6,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X6),X13),s('type.quote.varmap'(X6),X8),s('type.canonical.spolynom'(X6),X9))),s(X6,'const.canonical.interp_sp_3'(s('type.semi_ring.semi_ring'(X6),X13),s('type.quote.varmap'(X6),X8),s('type.canonical.spolynom'(X6),X10)))))),
  file('thm.ringNorm.interp_sp_def|split|3','thm.canonical.interp_sp_def|split|3')).
fof(f68,plain,(
  s(sK0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK4))))) != s(sK0,'const.ring.ring_RM_3'(s('type.ring.ring'(sK0),sK1),s(sK0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK3))),s(sK0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK4)))))),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  s(sK0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.canonical.spolynom'(sK0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(sK0),sK3),s('type.canonical.spolynom'(sK0),sK4))))) != s(sK0,'const.ring.ring_RM_3'(s('type.ring.ring'(sK0),sK1),s(sK0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK3))),s(sK0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(sK0),sK1),s('type.quote.varmap'(sK0),sK2),s('type.canonical.spolynom'(sK0),sK4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f63])).
fof(f63,plain,(
  ? [X0,X1,X2,X3,X4] : s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))))) != s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))),s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X4)))))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1,X2,X3,X4] : s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))))) = s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))),s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X4)))))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X0),X3),s('type.canonical.spolynom'(X0),X4))))) = s(X0,'const.ring.ring_RM_3'(s('type.ring.ring'(X0),X1),s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X3))),s(X0,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.spolynom'(X0),X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : s(X6,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.canonical.spolynom'(X6),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X6),X9),s('type.canonical.spolynom'(X6),X10))))) = s(X6,'const.ring.ring_RM_3'(s('type.ring.ring'(X6),X7),s(X6,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.canonical.spolynom'(X6),X9))),s(X6,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.canonical.spolynom'(X6),X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : s(X6,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.canonical.spolynom'(X6),'const.canonical.SPmult_2'(s('type.canonical.spolynom'(X6),X9),s('type.canonical.spolynom'(X6),X10))))) = s(X6,'const.ring.ring_RM_3'(s('type.ring.ring'(X6),X7),s(X6,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.canonical.spolynom'(X6),X9))),s(X6,'const.ringNorm.r_interp_sp_3'(s('type.ring.ring'(X6),X7),s('type.quote.varmap'(X6),X8),s('type.canonical.spolynom'(X6),X10)))))),
  file('thm.ringNorm.interp_sp_def|split|3','thm.ringNorm.interp_sp_def|split|3')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 80
% Passive clauses: 1329
% Generated clauses: 28642
% Final active clauses: 79
% Final passive clauses: 1249
% Input formulas: 22
% Initial clauses: 23
% 
% Duplicate literals: 1969
% Trivial inequalities: 318
% Fw subsumption resolutions: 2984
% Bw subsumption resolutions: 1
% Fw demodulations: 283
% 
% Simple tautologies: 487
% Equational tautologies: 196
% Forward subsumptions: 20811
% 
% Binary resolution: 48
% Forward superposition: 10510
% Backward superposition: 11111
% Self superposition: 739
% Equality factoring: 638
% Equality resolution: 18
% 
% SAT solver clauses: 1330
% SAT solver unit clauses: 36
% SAT solver binary clauses: 43
% 
% Memory used [KB]: 13816
% Time elapsed: 5.500 s
% ------------------------------
% ------------------------------
% Success in time 29.555 s
