Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f1732,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1723,f110,f175,f947,f177,f949,f181,f109,f179,f1684,f176,f1685,f180])).
fof(f180,plain,(
  $false | ($spl1 | $spl2)),
  inference(resolution,[],[f176,f175])).
fof(f1685,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,X3,X2,'const.pred_set.EMPTY_0',X1)) ) | $spl30),
  inference(cnf_transformation,[],[f1685_D])).
fof(f1685_D,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X0,X3,X2,'const.pred_set.EMPTY_0',X1)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f176,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl2),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1684,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X4)) | sP0(X0,X3,X2,'const.pred_set.EMPTY_0',X1)) ) | $spl1),
  inference(subsumption_resolution,[],[f1683,f175])).
fof(f1683,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,X4)) | sP0(X0,X3,X2,'const.pred_set.EMPTY_0',X1)) ) | $spl1),
  inference(forward_demodulation,[],[f1682,f188])).
fof(f188,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(resolution,[],[f153,f118])).
fof(f118,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.pred_set.INJ_EMPTY|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f153,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X13] : ~p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.INJ_EMPTY|split|0','thm.pred_set.NOT_IN_EMPTY')).
fof(f1682,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X4)) | sP0(X0,X3,X2,'const.pred_set.EMPTY_0',X1) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X1,'const.pred_set.EMPTY_0',X2,X3,X0)),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) ) | $spl1),
  inference(subsumption_resolution,[],[f1668,f113])).
fof(f113,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.pred_set.INJ_EMPTY|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f1668,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,X4)) | sP0(X0,X3,X2,'const.pred_set.EMPTY_0',X1) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X1,'const.pred_set.EMPTY_0',X2,X3,X0)),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) | s(bool,'T_0') = s(bool,X4)) ) | $spl1),
  inference(superposition,[],[f163,f280])).
fof(f280,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl1),
  inference(subsumption_resolution,[],[f277,f175])).
fof(f277,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'F_0')) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f215,f188])).
fof(f215,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(resolution,[],[f156,f114])).
fof(f114,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f73])).
fof(f156,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.pred_set.INJ_EMPTY|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f163,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK8(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) | sP0(X0,X1,X2,X3,X4) | p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP0(X0,X1,X2,X3,X4) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X8),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,X8))) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))) | s(X0,X8) = s(X0,X9)))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK7(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK7(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK8(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X0,sK9(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & s(X1,happ(s(fun(X0,X1),X2),s(X0,sK8(X4,X3,X2,X1,X0)))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK9(X4,X3,X2,X1,X0)))) & s(X0,sK8(X4,X3,X2,X1,X0)) != s(X0,sK9(X4,X3,X2,X1,X0))) | sP0(X0,X1,X2,X3,X4)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8,sK9])],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP0(X0,X1,X2,X3,X4) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8,X9] : (~p(s(bool,'const.bool.IN_2'(s(X0,X8),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) | s(X1,happ(s(fun(X0,X1),X2),s(X0,X8))) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))) | s(X0,X8) = s(X0,X9)))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,happ(s(fun(X0,X1),X2),s(X0,X5))) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))) & s(X0,X5) != s(X0,X6)) | sP0(X0,X1,X2,X3,X4)))),
  inference(rectify,[],[f104])).
fof(f104,plain,(
  ! [X1,X0,X2,X3,X4] : ((~sP0(X1,X0,X2,X3,X4) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) & s(X1,X5) != s(X1,X6)) | sP0(X1,X0,X2,X3,X4)))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X1,X0,X2,X3,X4] : ((~sP0(X1,X0,X2,X3,X4) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))) & ((? [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) | ? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) & s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) & s(X1,X5) != s(X1,X6))) | sP0(X1,X0,X2,X3,X4)))),
  inference(nnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X1,X0,X2,X3,X4] : (sP0(X1,X0,X2,X3,X4) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f179,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f179_D])).
fof(f179_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f109,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.INJ_EMPTY|split|0',aHL_TRUTH)).
fof(f181,plain,(
  $false | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f179,f177])).
fof(f949,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK1,sK2),sK3),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK2,bool),sK4))) | $spl8),
  inference(cnf_transformation,[],[f949_D])).
fof(f949_D,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK1,sK2),sK3),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK2,bool),sK4))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f177,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl4),
  inference(cnf_transformation,[],[f177_D])).
fof(f177_D,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f947,plain,(
  ( ! [X75] : (~p(s(bool,X75)) | s(bool,'F_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK1,sK2),sK3),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK2,bool),sK4)))) )),
  inference(subsumption_resolution,[],[f713,f117])).
fof(f117,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f75])).
fof(f713,plain,(
  ( ! [X75] : (~p(s(bool,X75)) | s(bool,'F_0') = s(bool,X75) | s(bool,'F_0') = s(bool,'const.pred_set.INJ_3'(s(fun(sK1,sK2),sK3),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK2,bool),sK4)))) )),
  inference(superposition,[],[f108,f285])).
fof(f285,plain,(
  ( ! [X10,X11] : (s(bool,X10) = s(bool,X11) | s(bool,'F_0') = s(bool,X11) | s(bool,'F_0') = s(bool,X10)) )),
  inference(resolution,[],[f230,f118])).
fof(f230,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | s(bool,X7) = s(bool,X8) | s(bool,'F_0') = s(bool,X8)) )),
  inference(resolution,[],[f159,f118])).
fof(f159,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f66])).
fof(f108,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK1,sK2),sK3),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK2,bool),sK4))))),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'const.pred_set.INJ_3'(s(fun(sK1,sK2),sK3),s(fun(sK1,bool),'const.pred_set.EMPTY_0'),s(fun(sK2,bool),sK4))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3,sK4])],[f58])).
fof(f58,plain,(
  ? [X0,X1,X2,X3] : ~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3))))),
  inference(ennf_transformation,[],[f28])).
fof(f28,plain,(
  ~! [X0,X1,X2,X3] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ~! [X0,X1] : ! [X2,X3] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X3))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X7,bool),X9))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : p(s(bool,'const.pred_set.INJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(X7,bool),X9))))),
  file('thm.pred_set.INJ_EMPTY|split|0','thm.pred_set.INJ_EMPTY|split|0')).
fof(f175,plain,(
  ~p(s(bool,'F_0')) | $spl1),
  inference(cnf_transformation,[],[f175_D])).
fof(f175_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f110,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.INJ_EMPTY|split|0',aHL_FALSITY)).
fof(f1723,plain,(
  $false | ($spl1 | $spl8 | $spl30)),
  inference(subsumption_resolution,[],[f1712,f175])).
fof(f1712,plain,(
  p(s(bool,'F_0')) | ($spl8 | $spl30)),
  inference(superposition,[],[f1701,f949])).
fof(f1701,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X5,X6),X7),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),X8))))) ) | $spl30),
  inference(resolution,[],[f1685,f172])).
fof(f172,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP0(X1,X0,X2,X3,X4) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4))))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) | sP0(X1,X0,X2,X3,X4)) & (~sP0(X1,X0,X2,X3,X4) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4))))))),
  inference(nnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> sP0(X1,X0,X2,X3,X4))),
  inference(definition_folding,[],[f69,f70])).
fof(f69,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) | (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)))))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) => (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) => s(X1,X5) = s(X1,X6)))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X1,X0),X2),s(fun(X1,bool),X3),s(fun(X0,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X1,X7),s(fun(X1,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))),s(fun(X0,bool),X4))))) & ! [X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(X1,X6),s(fun(X1,bool),X3))))) => (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) => s(X1,X5) = s(X1,X6)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X7,X6] : ! [X8,X9,X12] : (p(s(bool,'const.pred_set.INJ_3'(s(fun(X6,X7),X8),s(fun(X6,bool),X9),s(fun(X7,bool),X12)))) <=> (! [X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X9)))) => p(s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,X13))),s(fun(X7,bool),X12))))) & ! [X13,X14] : ((p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X9)))) & p(s(bool,'const.bool.IN_2'(s(X6,X14),s(fun(X6,bool),X9))))) => (s(X7,happ(s(fun(X6,X7),X8),s(X6,X13))) = s(X7,happ(s(fun(X6,X7),X8),s(X6,X14))) => s(X6,X13) = s(X6,X14)))))),
  file('thm.pred_set.INJ_EMPTY|split|0','thm.pred_set.INJ_DEF')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 45
% Passive clauses: 334
% Generated clauses: 1574
% Final active clauses: 41
% Final passive clauses: 248
% Input formulas: 25
% Initial clauses: 65
% 
% Duplicate literals: 122
% Trivial inequalities: 26
% Fw subsumption resolutions: 210
% Fw demodulations: 45
% Bw demodulations: 2
% 
% Simple tautologies: 95
% Equational tautologies: 31
% Forward subsumptions: 687
% 
% Binary resolution: 107
% Factoring: 1
% Forward superposition: 445
% Backward superposition: 424
% Self superposition: 50
% Equality factoring: 52
% Equality resolution: 8
% 
% Split clauses: 41
% Split components: 83
% SAT solver clauses: 65
% SAT solver unit clauses: 7
% SAT solver binary clauses: 48
% SAT solver learnt clauses: 1
% 
% Sat splits: 41
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 61
% TWLsolver calls for satisfiability: 24
% 
% Memory used [KB]: 1407
% Time elapsed: 0.043 s
% ------------------------------
% ------------------------------
% Success in time 0.045 s
