Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.quotient
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quotient
% SZS output start Proof for thm.quotient
fof(f12432,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f12430,f350,f562,f8980,f9031,f9029,f6025,f9028])).
fof(f9028,plain,(
  p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,sK18(sK13,sK12,sK10))))) | $spl160),
  inference(cnf_transformation,[],[f9028_D])).
fof(f9028_D,plain,(
  p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,sK18(sK13,sK12,sK10))))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f6025,plain,(
  p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,sK18(sK13,sK12,sK10)))))),
  inference(resolution,[],[f5915,f342])).
fof(f342,plain,(
  p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK10,bool),sK12),s(fun(sK10,bool),sK13))))),
  inference(cnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X4] : (~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,X4)))) | ~p(s(bool,happ(s(fun(sK10,bool),sK13),s(sK10,X4)))) | p(s(bool,happ(s(fun(sK10,bool),sK11),s(sK10,X4))))) & p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK10,bool),sK12),s(fun(sK10,bool),sK13)))) & ~p(s(bool,'|dollar|exists_1'(s(fun(sK10,bool),sK11))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f185])).
fof(f185,plain,(
  ? [X0,X1,X2,X3] : (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) & ~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ? [X0,X1,X2,X3] : (! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) & ~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : (! [X4] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) => (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3] : (! [X4] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) => (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))) => p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (! [X10] : (p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10)))) => (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X10)))))) => (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9)))) => p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X7))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (! [X10] : (p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10)))) => (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X10)))))) => (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9)))) => p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X7))))))),
  file('thm.quotient.LEFT_RES_EXISTS_REGULAR','thm.quotient.LEFT_RES_EXISTS_REGULAR')).
fof(f5915,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X2,X1,X0)))))) )),
  inference(forward_demodulation,[],[f418,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0] : ! [X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X7,X10] : s(bool,happ(s(fun(X6,bool),X7),s(X6,X10))) = s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7)))),
  file('thm.quotient.LEFT_RES_EXISTS_REGULAR','thm.pred_set.SPECIFICATION')).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK18(X2,X1,X0)),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0))))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X7,X16] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X16)))) <=> ? [X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7)))) & p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X10))))))),
  file('thm.quotient.LEFT_RES_EXISTS_REGULAR','thm.bool.RES_EXISTS_THM')).
fof(f9029,plain,(
  ~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,sK18(sK13,sK12,sK10))))) | $spl161),
  inference(cnf_transformation,[],[f9029_D])).
fof(f9029_D,plain,(
  ~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,sK18(sK13,sK12,sK10))))) <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f9031,plain,(
  p(s(bool,happ(s(fun(sK10,bool),sK11),s(sK10,sK18(sK13,sK12,sK10))))) | $spl162),
  inference(cnf_transformation,[],[f9031_D])).
fof(f9031_D,plain,(
  p(s(bool,happ(s(fun(sK10,bool),sK11),s(sK10,sK18(sK13,sK12,sK10))))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f8980,plain,(
  p(s(bool,happ(s(fun(sK10,bool),sK11),s(sK10,sK18(sK13,sK12,sK10))))) | ~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,sK18(sK13,sK12,sK10)))))),
  inference(resolution,[],[f5993,f341])).
fof(f341,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(sK10,bool),sK13),s(sK10,X4)))) | p(s(bool,happ(s(fun(sK10,bool),sK11),s(sK10,X4)))) | ~p(s(bool,happ(s(fun(sK10,bool),sK12),s(sK10,X4))))) )),
  inference(cnf_transformation,[],[f237])).
fof(f5993,plain,(
  p(s(bool,happ(s(fun(sK10,bool),sK13),s(sK10,sK18(sK13,sK12,sK10)))))),
  inference(resolution,[],[f419,f342])).
fof(f419,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f276])).
fof(f562,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f562_D])).
fof(f562_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f350,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quotient.LEFT_RES_EXISTS_REGULAR',aHL_FALSITY)).
fof(f12430,plain,(
  $false | ($spl3 | $spl162)),
  inference(subsumption_resolution,[],[f12429,f562])).
fof(f12429,plain,(
  p(s(bool,'F_0')) | $spl162),
  inference(forward_demodulation,[],[f12382,f592])).
fof(f592,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK10,bool),sK11)))),
  inference(resolution,[],[f362,f343])).
fof(f343,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK10,bool),sK11))))),
  inference(cnf_transformation,[],[f237])).
fof(f362,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.quotient.LEFT_RES_EXISTS_REGULAR','thm.bool.EQ_CLAUSES|split|3')).
fof(f12382,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK10,bool),sK11)))) | $spl162),
  inference(resolution,[],[f9031,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X1,X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14])],[f267])).
fof(f267,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f266])).
fof(f266,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f111])).
fof(f111,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X16] : (p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X16)))) <=> ? [X10] : p(s(bool,happ(s(fun(X6,bool),X16),s(X6,X10)))))),
  file('thm.quotient.LEFT_RES_EXISTS_REGULAR','thm.bool.EXISTS_THM')).
% SZS output end Proof for thm.quotient
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 600
% Passive clauses: 4502
% Generated clauses: 12069
% Final active clauses: 592
% Final passive clauses: 3696
% Input formulas: 69
% Initial clauses: 212
% 
% Duplicate literals: 117
% Trivial inequalities: 20
% Fw subsumption resolutions: 849
% Fw demodulations: 1459
% Bw demodulations: 7
% 
% Simple tautologies: 167
% Equational tautologies: 22
% Forward subsumptions: 4840
% 
% Binary resolution: 4508
% Factoring: 541
% Forward superposition: 1276
% Backward superposition: 2959
% Self superposition: 28
% Equality factoring: 2
% Equality resolution: 6
% 
% Split clauses: 171
% Split components: 345
% SAT solver clauses: 304
% SAT solver unit clauses: 11
% SAT solver binary clauses: 46
% SAT solver learnt clauses: 1
% 
% Sat splits: 171
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 303
% TWLsolver calls for satisfiability: 267
% 
% Memory used [KB]: 9083
% Time elapsed: 0.318 s
% ------------------------------
% ------------------------------
% Success in time 0.32 s
