Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.arithmetic
% (121131)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 454
% Passive clauses: 6350
% Generated clauses: 41051
% Final active clauses: 361
% Final passive clauses: 5641
% Input formulas: 133
% Initial clauses: 275
% 
% Duplicate literals: 106
% Trivial inequalities: 252
% Fw subsumption resolutions: 4550
% Fw demodulations: 7332
% Bw demodulations: 89
% 
% Simple tautologies: 366
% Equational tautologies: 24
% Forward subsumptions: 21752
% Fw demodulations to eq. taut.: 10
% 
% Binary resolution: 1876
% Factoring: 2
% Forward superposition: 10983
% Backward superposition: 15365
% Self superposition: 66
% Equality factoring: 14
% Equality resolution: 13
% 
% Split clauses: 160
% Split components: 321
% SAT solver clauses: 307
% SAT solver unit clauses: 9
% SAT solver binary clauses: 99
% 
% Sat splits: 160
% Sat splitting refutations: 12
% 
% TWLsolver clauses: 304
% TWLsolver calls for satisfiability: 102
% 
% Memory used [KB]: 19445
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.arithmetic
% (121199)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 370
% Passive clauses: 3311
% Generated clauses: 27933
% Final active clauses: 162
% Final passive clauses: 1296
% Input formulas: 133
% Initial clauses: 275
% 
% Duplicate literals: 110
% Trivial inequalities: 157
% Fw subsumption resolutions: 2054
% Fw demodulations: 7008
% Bw demodulations: 1723
% Condensations: 280
% Global subsumptions: 20
% 
% Simple tautologies: 397
% Equational tautologies: 12
% Forward subsumptions: 10362
% Bw demodulations to eq. taut.: 12
% 
% Binary resolution: 1102
% Factoring: 1
% Forward superposition: 5210
% Backward superposition: 4435
% Self superposition: 84
% Equality factoring: 2
% Equality resolution: 13
% 
% Split clauses: 680
% Split components: 1425
% SAT solver clauses: 3882
% SAT solver unit clauses: 333
% SAT solver binary clauses: 125
% SAT solver learnt clauses: 48
% SAT solver learnt literals: 33
% 
% Sat splits: 680
% Sat splitting refutations: 96
% 
% TWLsolver clauses: 3796
% TWLsolver calls for satisfiability: 1904
% 
% Memory used [KB]: 16886
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.arithmetic
% (121237)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 512
% Passive clauses: 2738
% Generated clauses: 398614
% Final active clauses: 448
% Final passive clauses: 1543
% Input formulas: 133
% Initial clauses: 275
% Discarded non-redundant clauses: 460
% 
% Duplicate literals: 8
% Trivial inequalities: 15
% Fw subsumption resolutions: 200
% Fw demodulations: 272467
% Bw demodulations: 372
% Condensations: 23
% Global subsumptions: 3
% 
% Simple tautologies: 85
% Equational tautologies: 25
% Forward subsumptions: 122885
% Fw demodulations to eq. taut.: 84
% 
% Unit resulting resolution: 115420
% Factoring: 1
% Forward superposition: 4709
% Backward superposition: 4758
% Self superposition: 92
% Equality factoring: 2
% Equality resolution: 8
% 
% Split clauses: 11
% Split components: 22
% SAT solver clauses: 2787
% SAT solver unit clauses: 1200
% SAT solver binary clauses: 1119
% 
% Sat splits: 11
% Sat splitting refutations: 22
% 
% Memory used [KB]: 208269
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f6150,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f621,f622,f735,f623,f624,f736,f737,f740,f741,f744,f745,f748,f751,f743,f747,f752,f739,f750,f754,f759,f762,f757,f769,f777,f783,f772,f787,f791,f793,f797,f802,f805,f800,f688,f809,f689,f690,f812,f691,f692,f813,f693,f814,f811,f694,f695,f818,f820,f696,f822,f703,f826,f704,f827,f707,f790,f829,f718,f830,f719,f833,f722,f832,f835,f727,f836,f733,f729,f837,f734,f910,f911,f914,f785,f944,f945,f948,f954,f1143,f1144,f1181,f1316,f1319,f1320,f1322,f1326,f1324,f1287,f788,f1288,f792,f1291,f807,f795,f1524,f1525,f1562,f1577,f1578,f1600,f1590,f1591,f1601,f1606,f1607,f1682,f1684,f764,f1689,f1692,f1691,f1641,f1642,f1693,f1643,f1644,f1694,f1645,f1695,f1663,f1671,f1703,f1672,f1706,f1675,f1705,f1708,f1848,f1849,f1872,f1861,f1862,f1873,f1882,f1883,f1906,f1895,f1896,f1907,f1916,f1917,f1940,f1929,f1930,f1941,f1946,f1947,f1681,f2023,f2026,f2025,f2027,f2028,f816,f1981,f1982,f2029,f1983,f1984,f2030,f1985,f2031,f2003,f2041,f1318,f2050,f2051,f2077,f2063,f2064,f2078,f2087,f2088,f2112,f2100,f2101,f2113,f2122,f770,f2123,f2147,f2135,f2136,f2148,f2296,f2386,f2395,f2410,f2455,f2501,f2547,f2593,f778,f2678,f2740,f2720,f2750,f2941,f2942,f2943,f3318,f3319,f3362,f798,f3363,f3431,f3520,f3529,f3543,f3588,f3633,f3681,f3728,f803,f3811,f3873,f3853,f3881,f4008,f4009,f2435,f4052,f4053,f2481,f4097,f4098,f2527,f4142,f755,f4143,f2573,f4211,f4300,f4309,f2396,f4323,f2622,f4371,f3708,f4417,f3755,f760,f4482,f4870,f4932,f2384,f4912,f4940,f2748,f5218,f5220,f5424,f5426,f5764,f5762,f5460,f5766,f5466,f5771,f5467,f5772,f5470,f5768,f5774,f5555,f5909,f5559,f5910,f5908,f5710,f6033,f5711,f6035,f5712,f6037,f5713,f6039,f5907,f6140,f6141,f6143,f6144,f6146,f1702,f6148,f6149])).
fof(f6149,plain,(
  $false | ($spl7 | $spl86)),
  inference(subsumption_resolution,[],[f6131,f1317])).
fof(f1317,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl86),
  inference(cnf_transformation,[],[f1317_D])).
fof(f1317_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f6131,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl7),
  inference(trivial_inequality_removal,[],[f6114])).
fof(f6114,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl7),
  inference(superposition,[],[f739,f441])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ! [X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(ennf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) => ! [X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X13,X14] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num',X14)))) => ! [X15] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X14))),s('type.num.num',X13))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X14),s('type.num.num',X13))))))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.LESS_EQ_ADD_SUB')).
fof(f6148,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) | $spl223),
  inference(cnf_transformation,[],[f6148_D])).
fof(f6148_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl223),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl223])])).
fof(f1702,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) | $spl109),
  inference(cnf_transformation,[],[f1702_D])).
fof(f1702_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f6146,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) | $spl7),
  inference(forward_demodulation,[],[f6145,f477])).
fof(f477,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.ADD_CLAUSES|split|0')).
fof(f6145,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) | $spl7),
  inference(forward_demodulation,[],[f6109,f450])).
fof(f450,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f122])).
fof(f122,axiom,(
  ! [X6,X7] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num',X6)))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.ADD_SYM')).
fof(f6109,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) | $spl7),
  inference(superposition,[],[f739,f465])).
fof(f465,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f333])).
fof(f333,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))))),
  inference(nnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f107])).
fof(f107,axiom,(
  ! [X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X7))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.LESS_EQ_|48|')).
fof(f6144,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X18)))) ) | $spl220),
  inference(cnf_transformation,[],[f6144_D])).
fof(f6144_D,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X18)))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f6143,plain,(
  ( ! [X18] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X18)))) ) | $spl7),
  inference(forward_demodulation,[],[f6142,f477])).
fof(f6142,plain,(
  ( ! [X18] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X18)))) ) | $spl7),
  inference(forward_demodulation,[],[f6106,f450])).
fof(f6106,plain,(
  ( ! [X18] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',X18)))) ) | $spl7),
  inference(superposition,[],[f739,f430])).
fof(f430,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(flattening,[],[f326])).
fof(f326,plain,(
  ! [X0,X1] : ((s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f105])).
fof(f105,axiom,(
  ! [X6,X7] : (s('type.num.num',X6) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X7)))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.SUB_EQ_EQ_|48|')).
fof(f6141,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) ) | $spl218),
  inference(cnf_transformation,[],[f6141_D])).
fof(f6141_D,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f6140,plain,(
  ( ! [X17] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) ) | $spl7),
  inference(forward_demodulation,[],[f6139,f477])).
fof(f6139,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) ) | $spl7),
  inference(forward_demodulation,[],[f6105,f450])).
fof(f6105,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))) ) | $spl7),
  inference(superposition,[],[f739,f430])).
fof(f5907,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl204),
  inference(cnf_transformation,[],[f5907_D])).
fof(f5907_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f6039,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))) | $spl217),
  inference(cnf_transformation,[],[f6039_D])).
fof(f6039_D,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))) <=> ~$spl217),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl217])])).
fof(f5713,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK4))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ($spl24 | $spl42)),
  inference(superposition,[],[f803,f771])).
fof(f771,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl24),
  inference(cnf_transformation,[],[f771_D])).
fof(f771_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f6037,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | $spl215),
  inference(cnf_transformation,[],[f6037_D])).
fof(f6037_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))) <=> ~$spl215),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl215])])).
fof(f5712,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ($spl24 | $spl38)),
  inference(superposition,[],[f798,f771])).
fof(f6035,plain,(
  s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK4))) | $spl213),
  inference(cnf_transformation,[],[f6035_D])).
fof(f6035_D,plain,(
  s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK4))) <=> ~$spl213),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl213])])).
fof(f5711,plain,(
  s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK4))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ($spl24 | $spl26)),
  inference(superposition,[],[f778,f771])).
fof(f6033,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl211),
  inference(cnf_transformation,[],[f6033_D])).
fof(f6033_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) <=> ~$spl211),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl211])])).
fof(f5710,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ($spl22 | $spl24)),
  inference(superposition,[],[f770,f771])).
fof(f5908,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl205),
  inference(cnf_transformation,[],[f5908_D])).
fof(f5908_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl205),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl205])])).
fof(f5910,plain,(
  ( ! [X118] : (s('type.num.num',X118) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X118),s('type.num.num',sK4)))) ) | $spl208),
  inference(cnf_transformation,[],[f5910_D])).
fof(f5910_D,plain,(
  ( ! [X118] : (s('type.num.num',X118) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X118),s('type.num.num',sK4)))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f5559,plain,(
  ( ! [X118] : (s('type.num.num',X118) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X118),s('type.num.num',sK4))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) ) | $spl24),
  inference(superposition,[],[f432,f771])).
fof(f432,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f327])).
fof(f5909,plain,(
  ( ! [X114] : (s('type.num.num',sK4) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X114)))) ) | $spl206),
  inference(cnf_transformation,[],[f5909_D])).
fof(f5909_D,plain,(
  ( ! [X114] : (s('type.num.num',sK4) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X114)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f5555,plain,(
  ( ! [X114] : (s('type.num.num',sK4) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X114))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) ) | $spl24),
  inference(superposition,[],[f431,f771])).
fof(f431,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f327])).
fof(f5774,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) | $spl203),
  inference(cnf_transformation,[],[f5774_D])).
fof(f5774_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl203),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl203])])).
fof(f5768,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | $spl196),
  inference(cnf_transformation,[],[f5768_D])).
fof(f5768_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f5470,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) | $spl24),
  inference(superposition,[],[f771,f465])).
fof(f5772,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X18)))) ) | $spl200),
  inference(cnf_transformation,[],[f5772_D])).
fof(f5772_D,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X18)))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f5467,plain,(
  ( ! [X18] : (s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))),s('type.num.num',X18)))) ) | $spl24),
  inference(superposition,[],[f771,f430])).
fof(f5771,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) ) | $spl198),
  inference(cnf_transformation,[],[f5771_D])).
fof(f5771_D,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f5466,plain,(
  ( ! [X17] : (s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6)))))) ) | $spl24),
  inference(superposition,[],[f771,f430])).
fof(f5766,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) | $spl195),
  inference(cnf_transformation,[],[f5766_D])).
fof(f5766_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) <=> ~$spl195),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl195])])).
fof(f5460,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num',sK6)))) | $spl24),
  inference(superposition,[],[f771,f434])).
fof(f434,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f328])).
fof(f328,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X6,X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7)))))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.SUB_EQ_|48|')).
fof(f5762,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | $spl190),
  inference(cnf_transformation,[],[f5762_D])).
fof(f5762_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f5764,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl192),
  inference(cnf_transformation,[],[f5764_D])).
fof(f5764_D,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f5426,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | $spl24),
  inference(superposition,[],[f771,f476])).
fof(f476,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.arithmetic.SUB_LEFT_ADD',aHL_BOOL_CASES)).
fof(f5424,plain,(
  s('type.num.num',sK4) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | $spl24),
  inference(superposition,[],[f771,f475])).
fof(f475,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X28] : (s(bool,'T_0') = s(bool,X28) | s(bool,'F_0') = s(bool,X28))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.bool.BOOL_CASES_AX')).
fof(f5220,plain,(
  $false | $spl25),
  inference(subsumption_resolution,[],[f5219,f519])).
fof(f519,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))),
  inference(rectify,[],[f119])).
fof(f119,axiom,(
  ! [X7] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.ZERO_LESS_EQ')).
fof(f5219,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | $spl25),
  inference(subsumption_resolution,[],[f5131,f461])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X9] : ! [X31,X32] : s(X9,X31) = s(X9,'const.bool.COND_3'(s(bool,'T_0'),s(X9,X31),s(X9,X32)))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.bool.COND_CLAUSES|split|0')).
fof(f5131,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | $spl25),
  inference(superposition,[],[f772,f474])).
fof(f474,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f337])).
fof(f337,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X28] : (s(bool,'T_0') = s(bool,X28) <=> p(s(bool,X28)))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.bool.EQ_CLAUSES|split|0')).
fof(f5218,plain,(
  $false | $spl25),
  inference(subsumption_resolution,[],[f5217,f519])).
fof(f5217,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | $spl25),
  inference(subsumption_resolution,[],[f5130,f461])).
fof(f5130,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | $spl25),
  inference(superposition,[],[f772,f472])).
fof(f472,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X28] : (s(bool,'T_0') = s(bool,X28) <=> p(s(bool,X28)))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.bool.EQ_CLAUSES|split|1')).
fof(f2748,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl170),
  inference(cnf_transformation,[],[f2748_D])).
fof(f2748_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f4940,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK4))) | $spl189),
  inference(cnf_transformation,[],[f4940_D])).
fof(f4940_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK4))) <=> ~$spl189),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl189])])).
fof(f4912,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK4))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl18),
  inference(superposition,[],[f760,f605])).
fof(f605,plain,(
  ( ! [X0] : (s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(cnf_transformation,[],[f252])).
fof(f252,plain,(
  ! [X0] : s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f126])).
fof(f126,axiom,(
  ! [X6] : s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X6))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.PRE_SUB|49|')).
fof(f2384,plain,(
  ( ! [X71] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X71))),s('type.num.num',X71)))) ) | $spl152),
  inference(cnf_transformation,[],[f2384_D])).
fof(f2384_D,plain,(
  ( ! [X71] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X71))),s('type.num.num',X71)))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f4932,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl187),
  inference(cnf_transformation,[],[f4932_D])).
fof(f4932_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl187),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl187])])).
fof(f4870,plain,(
  ( ! [X71] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X71))),s('type.num.num',X71)))) ) | $spl18),
  inference(superposition,[],[f760,f607])).
fof(f607,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X0)))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X15] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.num.num',X15)))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.SUC_SUB')).
fof(f4482,plain,(
  $false | ($spl0 | $spl9)),
  inference(subsumption_resolution,[],[f2952,f734])).
fof(f2952,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | $spl9),
  inference(subsumption_resolution,[],[f1648,f480])).
fof(f480,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.ADD_|48|')).
fof(f1648,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | $spl9),
  inference(superposition,[],[f743,f434])).
fof(f760,plain,(
  ( ! [X13] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) | $spl18),
  inference(cnf_transformation,[],[f760_D])).
fof(f760_D,plain,(
  ( ! [X13] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3755,plain,(
  ( ! [X24,X23] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24)) ) | $spl178),
  inference(cnf_transformation,[],[f3755_D])).
fof(f3755_D,plain,(
  ( ! [X24,X23] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24)) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f4417,plain,(
  ( ! [X24,X23,X25] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | $spl21),
  inference(superposition,[],[f1742,f430])).
fof(f1742,plain,(
  ( ! [X23,X22] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')),X22,X23)) ) | $spl21),
  inference(resolution,[],[f764,f410])).
fof(f410,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f320])).
fof(f320,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f319])).
fof(f319,plain,(
  ! [X0,X2,X3] : ((~sP0(X0,X2,X3) | p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))) & ((~p(s(bool,X3)) & ~p(s(bool,X2)) & p(s(bool,X0))) | sP0(X0,X2,X3)))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ! [X0,X2,X3] : ((~sP0(X0,X2,X3) | (p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0)))) & ((~p(s(bool,X3)) & ~p(s(bool,X2)) & p(s(bool,X0))) | sP0(X0,X2,X3)))),
  inference(nnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0,X2,X3] : (sP0(X0,X2,X3) <=> (p(s(bool,X3)) | p(s(bool,X2)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f3708,plain,(
  ( ! [X24,X23] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24)) ) | $spl176),
  inference(cnf_transformation,[],[f3708_D])).
fof(f3708_D,plain,(
  ( ! [X24,X23] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f4371,plain,(
  ( ! [X24,X23,X25] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | $spl21),
  inference(superposition,[],[f1739,f430])).
fof(f1739,plain,(
  ( ! [X17,X16] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')),X16,X17)) ) | $spl21),
  inference(resolution,[],[f764,f406])).
fof(f406,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f316])).
fof(f316,plain,(
  ! [X2,X1,X3] : ((~sP1(X2,X1,X3) | p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X2))) | sP1(X2,X1,X3)))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ! [X2,X1,X3] : ((~sP1(X2,X1,X3) | (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X2))) | sP1(X2,X1,X3)))),
  inference(nnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X2,X1,X3] : (sP1(X2,X1,X3) <=> (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f2622,plain,(
  ( ! [X24,X23] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24)) ) | $spl166),
  inference(cnf_transformation,[],[f2622_D])).
fof(f2622_D,plain,(
  ( ! [X24,X23] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f4323,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | $spl21),
  inference(superposition,[],[f1738,f430])).
fof(f1738,plain,(
  ( ! [X14,X15] : (sP1(X14,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')),X15)) ) | $spl21),
  inference(resolution,[],[f764,f405])).
fof(f405,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f317])).
fof(f2396,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18)) ) | $spl156),
  inference(cnf_transformation,[],[f2396_D])).
fof(f2396_D,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f4309,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X19)))) ) | $spl14),
  inference(subsumption_resolution,[],[f4248,f478])).
fof(f478,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.SUB_|48||split|1')).
fof(f4248,plain,(
  ( ! [X19,X18] : (s('type.num.num',X18) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X19)))) ) | $spl14),
  inference(superposition,[],[f755,f430])).
fof(f4300,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK4))) | $spl185),
  inference(cnf_transformation,[],[f4300_D])).
fof(f4300_D,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK4))) <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f4211,plain,(
  ( ! [X71] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK4))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X71))),s('type.num.num',X71)))) ) | $spl14),
  inference(superposition,[],[f755,f607])).
fof(f2573,plain,(
  ( ! [X24,X23] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24)) ) | $spl164),
  inference(cnf_transformation,[],[f2573_D])).
fof(f2573_D,plain,(
  ( ! [X24,X23] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f4143,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | $spl21),
  inference(superposition,[],[f1736,f430])).
fof(f1736,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl21),
  inference(resolution,[],[f764,f402])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X3] : ((~sP2(X0,X1,X3) | p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X3)))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X0,X1,X3] : ((~sP2(X0,X1,X3) | (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X3)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X3)))),
  inference(nnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0,X1,X3] : (sP2(X0,X1,X3) <=> (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f755,plain,(
  ( ! [X12] : (s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) ) | $spl14),
  inference(cnf_transformation,[],[f755_D])).
fof(f755_D,plain,(
  ( ! [X12] : (s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f4142,plain,(
  ( ! [X21,X22,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | $spl21),
  inference(superposition,[],[f1736,f430])).
fof(f2527,plain,(
  ( ! [X24,X23] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24)) ) | $spl162),
  inference(cnf_transformation,[],[f2527_D])).
fof(f2527_D,plain,(
  ( ! [X24,X23] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f4098,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | $spl21),
  inference(superposition,[],[f1735,f430])).
fof(f1735,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')),X9)) ) | $spl21),
  inference(resolution,[],[f764,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f314])).
fof(f4097,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | $spl21),
  inference(superposition,[],[f1735,f430])).
fof(f2481,plain,(
  ( ! [X24,X23] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24)) ) | $spl160),
  inference(cnf_transformation,[],[f2481_D])).
fof(f2481_D,plain,(
  ( ! [X24,X23] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f4053,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | $spl21),
  inference(superposition,[],[f1733,f430])).
fof(f1733,plain,(
  ( ! [X4,X5] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')),X4,X5)) ) | $spl21),
  inference(resolution,[],[f764,f398])).
fof(f398,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f310])).
fof(f310,plain,(
  ! [X3,X1,X2] : ((~sP3(X3,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X3))) | sP3(X3,X1,X2)))),
  inference(flattening,[],[f309])).
fof(f309,plain,(
  ! [X3,X1,X2] : ((~sP3(X3,X1,X2) | (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3)))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X3))) | sP3(X3,X1,X2)))),
  inference(nnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X3,X1,X2] : (sP3(X3,X1,X2) <=> (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X3))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f4052,plain,(
  ( ! [X21,X22,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | $spl21),
  inference(superposition,[],[f1733,f430])).
fof(f2435,plain,(
  ( ! [X24,X23] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) | $spl158),
  inference(cnf_transformation,[],[f2435_D])).
fof(f2435_D,plain,(
  ( ! [X24,X23] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f4009,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X25)))) ) | $spl21),
  inference(superposition,[],[f1732,f430])).
fof(f1732,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) ) | $spl21),
  inference(resolution,[],[f764,f396])).
fof(f396,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f311])).
fof(f4008,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK4)))) ) | $spl21),
  inference(superposition,[],[f1732,f430])).
fof(f3881,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK6))) | $spl183),
  inference(cnf_transformation,[],[f3881_D])).
fof(f3881_D,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK6))) <=> ~$spl183),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl183])])).
fof(f3853,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl42),
  inference(superposition,[],[f803,f605])).
fof(f3873,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl181),
  inference(cnf_transformation,[],[f3873_D])).
fof(f3873_D,plain,(
  s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl181),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl181])])).
fof(f3811,plain,(
  ( ! [X71] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X71))),s('type.num.num',X71)))) ) | $spl42),
  inference(superposition,[],[f803,f607])).
fof(f803,plain,(
  ( ! [X13] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) | $spl42),
  inference(cnf_transformation,[],[f803_D])).
fof(f803_D,plain,(
  ( ! [X13] : (s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f3728,plain,(
  ( ! [X24,X23,X25] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl45),
  inference(superposition,[],[f1350,f430])).
fof(f1350,plain,(
  ( ! [X23,X22] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X22,X23)) ) | $spl45),
  inference(resolution,[],[f807,f410])).
fof(f3681,plain,(
  ( ! [X24,X23,X25] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl45),
  inference(superposition,[],[f1347,f430])).
fof(f1347,plain,(
  ( ! [X17,X16] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X16,X17)) ) | $spl45),
  inference(resolution,[],[f807,f406])).
fof(f3633,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl45),
  inference(superposition,[],[f1346,f430])).
fof(f1346,plain,(
  ( ! [X14,X15] : (sP1(X14,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X15)) ) | $spl45),
  inference(resolution,[],[f807,f405])).
fof(f3588,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl45),
  inference(superposition,[],[f1344,f430])).
fof(f1344,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl45),
  inference(resolution,[],[f807,f402])).
fof(f3543,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl45),
  inference(superposition,[],[f1343,f430])).
fof(f1343,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X9)) ) | $spl45),
  inference(resolution,[],[f807,f401])).
fof(f3529,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X19)))) ) | $spl38),
  inference(subsumption_resolution,[],[f3468,f478])).
fof(f3468,plain,(
  ( ! [X19,X18] : (s('type.num.num',X18) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X19)))) ) | $spl38),
  inference(superposition,[],[f798,f430])).
fof(f3520,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK6))) | $spl175),
  inference(cnf_transformation,[],[f3520_D])).
fof(f3520_D,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK6))) <=> ~$spl175),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl175])])).
fof(f3431,plain,(
  ( ! [X71] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X71))),s('type.num.num',X71)))) ) | $spl38),
  inference(superposition,[],[f798,f607])).
fof(f3363,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl45),
  inference(superposition,[],[f1341,f430])).
fof(f1341,plain,(
  ( ! [X4,X5] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')),X4,X5)) ) | $spl45),
  inference(resolution,[],[f807,f398])).
fof(f798,plain,(
  ( ! [X12] : (s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) ) | $spl38),
  inference(cnf_transformation,[],[f798_D])).
fof(f798_D,plain,(
  ( ! [X12] : (s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f3362,plain,(
  ( ! [X21,X22,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl45),
  inference(superposition,[],[f1341,f430])).
fof(f3319,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl45),
  inference(superposition,[],[f1340,f430])).
fof(f1340,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) ) | $spl45),
  inference(resolution,[],[f807,f396])).
fof(f3318,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl45),
  inference(superposition,[],[f1340,f430])).
fof(f2943,plain,(
  $false | ($spl1 | $spl87)),
  inference(subsumption_resolution,[],[f2849,f838])).
fof(f838,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | $spl1),
  inference(resolution,[],[f733,f515])).
fof(f515,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f133])).
fof(f133,axiom,(
  ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7)))))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.LESS_IMP_LESS_OR_EQ')).
fof(f2849,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | $spl87),
  inference(resolution,[],[f1318,f508])).
fof(f508,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f204])).
fof(f204,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X6,X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X6)))))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.NOT_LESS')).
fof(f2942,plain,(
  $false | ($spl1 | $spl87)),
  inference(subsumption_resolution,[],[f2847,f841])).
fof(f841,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl1),
  inference(resolution,[],[f733,f508])).
fof(f2847,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl87),
  inference(resolution,[],[f1318,f511])).
fof(f511,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f346])).
fof(f346,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(flattening,[],[f345])).
fof(f345,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1))) & ((~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) & s('type.num.num',X0) != s('type.num.num',X1)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num',X1)))),
  inference(rectify,[],[f131])).
fof(f131,axiom,(
  ! [X6,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X7)))) <=> (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) | s('type.num.num',X6) = s('type.num.num',X7)))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.LESS_OR_EQ')).
fof(f2941,plain,(
  $false | ($spl1 | $spl87)),
  inference(subsumption_resolution,[],[f2846,f841])).
fof(f2846,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl87),
  inference(resolution,[],[f1318,f515])).
fof(f2750,plain,(
  s('type.num.num',sK5) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK5))) | $spl173),
  inference(cnf_transformation,[],[f2750_D])).
fof(f2750_D,plain,(
  s('type.num.num',sK5) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK5))) <=> ~$spl173),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl173])])).
fof(f2720,plain,(
  s('type.num.num',sK5) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl26),
  inference(superposition,[],[f778,f605])).
fof(f2740,plain,(
  s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl169),
  inference(cnf_transformation,[],[f2740_D])).
fof(f2740_D,plain,(
  s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl169),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl169])])).
fof(f2678,plain,(
  ( ! [X71] : (s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X71))),s('type.num.num',X71)))) ) | $spl26),
  inference(superposition,[],[f778,f607])).
fof(f778,plain,(
  ( ! [X13] : (s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) | $spl26),
  inference(cnf_transformation,[],[f778_D])).
fof(f778_D,plain,(
  ( ! [X13] : (s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2593,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl29),
  inference(superposition,[],[f984,f430])).
fof(f984,plain,(
  ( ! [X14,X15] : (sP1(X14,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X15)) ) | $spl29),
  inference(resolution,[],[f785,f405])).
fof(f2547,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl29),
  inference(superposition,[],[f982,f430])).
fof(f982,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X10,X11)) ) | $spl29),
  inference(resolution,[],[f785,f402])).
fof(f2501,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl29),
  inference(superposition,[],[f981,f430])).
fof(f981,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X9)) ) | $spl29),
  inference(resolution,[],[f785,f401])).
fof(f2455,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl29),
  inference(superposition,[],[f979,f430])).
fof(f979,plain,(
  ( ! [X4,X5] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X4,X5)) ) | $spl29),
  inference(resolution,[],[f785,f398])).
fof(f2410,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl29),
  inference(superposition,[],[f978,f430])).
fof(f978,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) ) | $spl29),
  inference(resolution,[],[f785,f396])).
fof(f2395,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X19)))) ) | $spl22),
  inference(subsumption_resolution,[],[f2333,f478])).
fof(f2333,plain,(
  ( ! [X19,X18] : (s('type.num.num',X18) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X19) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X19)))) ) | $spl22),
  inference(superposition,[],[f770,f430])).
fof(f2386,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK5))) | $spl155),
  inference(cnf_transformation,[],[f2386_D])).
fof(f2386_D,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK5))) <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f2296,plain,(
  ( ! [X71] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X71))),s('type.num.num',X71)))) ) | $spl22),
  inference(superposition,[],[f770,f607])).
fof(f2148,plain,(
  ( ! [X21,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21)) ) | $spl150),
  inference(cnf_transformation,[],[f2148_D])).
fof(f2148_D,plain,(
  ( ! [X21,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f2136,plain,(
  ( ! [X24,X23,X25] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f854,f430])).
fof(f854,plain,(
  ( ! [X23,X22] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)),X22,X23)) ) | $spl1),
  inference(resolution,[],[f733,f410])).
fof(f2135,plain,(
  ( ! [X21,X22,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f854,f430])).
fof(f2147,plain,(
  ( ! [X21,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) | $spl148),
  inference(cnf_transformation,[],[f2147_D])).
fof(f2147_D,plain,(
  ( ! [X21,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f2123,plain,(
  ( ! [X24,X23,X25] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f854,f430])).
fof(f770,plain,(
  ( ! [X12] : (s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) ) | $spl22),
  inference(cnf_transformation,[],[f770_D])).
fof(f770_D,plain,(
  ( ! [X12] : (s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2122,plain,(
  ( ! [X21,X22,X20] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f854,f430])).
fof(f2113,plain,(
  ( ! [X21,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21)) ) | $spl146),
  inference(cnf_transformation,[],[f2113_D])).
fof(f2113_D,plain,(
  ( ! [X21,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f2101,plain,(
  ( ! [X24,X23,X25] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f851,f430])).
fof(f851,plain,(
  ( ! [X17,X16] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)),X16,X17)) ) | $spl1),
  inference(resolution,[],[f733,f406])).
fof(f2100,plain,(
  ( ! [X21,X22,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f851,f430])).
fof(f2112,plain,(
  ( ! [X21,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) | $spl144),
  inference(cnf_transformation,[],[f2112_D])).
fof(f2112_D,plain,(
  ( ! [X21,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f2088,plain,(
  ( ! [X24,X23,X25] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f851,f430])).
fof(f2087,plain,(
  ( ! [X21,X22,X20] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f851,f430])).
fof(f2078,plain,(
  ( ! [X21,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X21)) ) | $spl142),
  inference(cnf_transformation,[],[f2078_D])).
fof(f2078_D,plain,(
  ( ! [X21,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X21)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f2064,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f850,f430])).
fof(f850,plain,(
  ( ! [X14,X15] : (sP1(X14,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)),X15)) ) | $spl1),
  inference(resolution,[],[f733,f405])).
fof(f2063,plain,(
  ( ! [X21,X22,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f850,f430])).
fof(f2077,plain,(
  ( ! [X21,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21)) ) | $spl140),
  inference(cnf_transformation,[],[f2077_D])).
fof(f2077_D,plain,(
  ( ! [X21,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f2051,plain,(
  ( ! [X24,X23,X25] : (sP1(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f850,f430])).
fof(f2050,plain,(
  ( ! [X21,X22,X20] : (sP1(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f850,f430])).
fof(f1318,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl87),
  inference(cnf_transformation,[],[f1318_D])).
fof(f1318_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f2041,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) | $spl139),
  inference(cnf_transformation,[],[f2041_D])).
fof(f2041_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f2003,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl73),
  inference(superposition,[],[f832,f441])).
fof(f2031,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6)))) ) | $spl136),
  inference(cnf_transformation,[],[f2031_D])).
fof(f2031_D,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6)))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f1985,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) ) | $spl73),
  inference(superposition,[],[f832,f423])).
fof(f423,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1))) | s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f325])).
fof(f325,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | ((s('type.num.num',X2) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1)))) & (s('type.num.num',X0) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1))) | s('type.num.num',X2) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | (s('type.num.num',X2) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) => (s('type.num.num',X2) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X1)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X6,X7,X8] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))) => (s('type.num.num',X8) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X7)))))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.ADD_EQ_SUB')).
fof(f2030,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK5)))) | s('type.num.num',sK6) != s('type.num.num',X2)) ) | $spl134),
  inference(cnf_transformation,[],[f2030_D])).
fof(f2030_D,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK5)))) | s('type.num.num',sK6) != s('type.num.num',X2)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f1984,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X3))))) | s('type.num.num',sK6) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',sK5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) ) | $spl73),
  inference(superposition,[],[f832,f422])).
fof(f422,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) | s('type.num.num',X1) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f324])).
fof(f324,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0)))) | ((s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) | s('type.num.num',X1) = s('type.num.num',X2)) & (s('type.num.num',X1) != s('type.num.num',X2) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))))))),
  inference(nnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0)))) | (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(flattening,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0))))) | (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(ennf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X0))))) => (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X2))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X8,X7,X6] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num',X8))))) => (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X6))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X8),s('type.num.num',X7))) <=> s('type.num.num',X6) = s('type.num.num',X7)))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.SUB_CANCEL')).
fof(f1983,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X2))))) | s('type.num.num',sK6) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK5))))) ) | $spl73),
  inference(superposition,[],[f832,f422])).
fof(f2029,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X0)))) | s('type.num.num',sK5) != s('type.num.num',X0)) ) | $spl132),
  inference(cnf_transformation,[],[f2029_D])).
fof(f2029_D,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X0)))) | s('type.num.num',sK5) != s('type.num.num',X0)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f1982,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',sK6))))) | s('type.num.num',sK5) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) ) | $spl73),
  inference(superposition,[],[f832,f420])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) | s('type.num.num',X1) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))) | ((s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) | s('type.num.num',X1) = s('type.num.num',X2)) & (s('type.num.num',X1) != s('type.num.num',X2) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))) | (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(flattening,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2))))) | (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(ennf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2))))) => (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num',X0))) <=> s('type.num.num',X1) = s('type.num.num',X2)))),
  inference(rectify,[],[f84])).
fof(f84,axiom,(
  ! [X8,X7,X6] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X6))))) => (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X8))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X8))) <=> s('type.num.num',X6) = s('type.num.num',X7)))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.CANCEL_SUB')).
fof(f1981,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',sK6))))) | s('type.num.num',sK5) != s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X0))))) ) | $spl73),
  inference(superposition,[],[f832,f420])).
fof(f816,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl57),
  inference(cnf_transformation,[],[f816_D])).
fof(f816_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f2028,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl73),
  inference(forward_demodulation,[],[f1973,f478])).
fof(f1973,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl73),
  inference(superposition,[],[f832,f430])).
fof(f2027,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl73),
  inference(forward_demodulation,[],[f1972,f478])).
fof(f1972,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl73),
  inference(superposition,[],[f832,f430])).
fof(f2025,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | $spl131),
  inference(cnf_transformation,[],[f2025_D])).
fof(f2025_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) <=> ~$spl131),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl131])])).
fof(f2026,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) ) | $spl73),
  inference(forward_demodulation,[],[f1960,f481])).
fof(f481,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.SUB|split|0')).
fof(f1960,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) ) | $spl73),
  inference(superposition,[],[f832,f430])).
fof(f2023,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl73),
  inference(forward_demodulation,[],[f1959,f481])).
fof(f1959,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl73),
  inference(superposition,[],[f832,f430])).
fof(f1681,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | $spl97),
  inference(cnf_transformation,[],[f1681_D])).
fof(f1681_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f1947,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl73),
  inference(superposition,[],[f832,f430])).
fof(f1946,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl73),
  inference(superposition,[],[f832,f430])).
fof(f1941,plain,(
  ( ! [X21,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21)) ) | $spl128),
  inference(cnf_transformation,[],[f1941_D])).
fof(f1941_D,plain,(
  ( ! [X21,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f1930,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f848,f430])).
fof(f848,plain,(
  ( ! [X10,X11] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)),X10,X11)) ) | $spl1),
  inference(resolution,[],[f733,f402])).
fof(f1929,plain,(
  ( ! [X21,X22,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f848,f430])).
fof(f1940,plain,(
  ( ! [X21,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) | $spl126),
  inference(cnf_transformation,[],[f1940_D])).
fof(f1940_D,plain,(
  ( ! [X21,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f1917,plain,(
  ( ! [X24,X23,X25] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f848,f430])).
fof(f1916,plain,(
  ( ! [X21,X22,X20] : (sP2('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f848,f430])).
fof(f1907,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X21)) ) | $spl124),
  inference(cnf_transformation,[],[f1907_D])).
fof(f1907_D,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X21)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f1896,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f847,f430])).
fof(f847,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)),X9)) ) | $spl1),
  inference(resolution,[],[f733,f401])).
fof(f1895,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f847,f430])).
fof(f1906,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21)) ) | $spl122),
  inference(cnf_transformation,[],[f1906_D])).
fof(f1906_D,plain,(
  ( ! [X21,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f1883,plain,(
  ( ! [X24,X23,X25] : (sP2(X23,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f847,f430])).
fof(f1882,plain,(
  ( ! [X21,X22,X20] : (sP2(X20,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f847,f430])).
fof(f1873,plain,(
  ( ! [X21,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21)) ) | $spl120),
  inference(cnf_transformation,[],[f1873_D])).
fof(f1873_D,plain,(
  ( ! [X21,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f1862,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f845,f430])).
fof(f845,plain,(
  ( ! [X4,X5] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)),X4,X5)) ) | $spl1),
  inference(resolution,[],[f733,f398])).
fof(f1861,plain,(
  ( ! [X21,X22,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f845,f430])).
fof(f1872,plain,(
  ( ! [X21,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) | $spl118),
  inference(cnf_transformation,[],[f1872_D])).
fof(f1872_D,plain,(
  ( ! [X21,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f1849,plain,(
  ( ! [X24,X23,X25] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X23,X24) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f845,f430])).
fof(f1848,plain,(
  ( ! [X21,X22,X20] : (sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)),X20,X21) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f845,f430])).
fof(f1708,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num','const.num.|48|_0')))) | $spl117),
  inference(cnf_transformation,[],[f1708_D])).
fof(f1708_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f1705,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) | $spl113),
  inference(cnf_transformation,[],[f1705_D])).
fof(f1705_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f1675,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num','const.num.|48|_0')))) | $spl9),
  inference(superposition,[],[f743,f465])).
fof(f1706,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num',X14)))) ) | $spl114),
  inference(cnf_transformation,[],[f1706_D])).
fof(f1706_D,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num',X14)))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f1672,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),s('type.num.num',X14)))) ) | $spl9),
  inference(superposition,[],[f743,f430])).
fof(f1703,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))))) ) | $spl110),
  inference(cnf_transformation,[],[f1703_D])).
fof(f1703_D,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f1671,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK4) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6)))))))) ) | $spl9),
  inference(superposition,[],[f743,f430])).
fof(f1663,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | $spl9),
  inference(superposition,[],[f743,f441])).
fof(f1695,plain,(
  ( ! [X4] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6)))) ) | $spl106),
  inference(cnf_transformation,[],[f1695_D])).
fof(f1695_D,plain,(
  ( ! [X4] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6)))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f1645,plain,(
  ( ! [X4] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) ) | $spl9),
  inference(superposition,[],[f743,f423])).
fof(f1694,plain,(
  ( ! [X2] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK5)))) | s('type.num.num',sK6) != s('type.num.num',X2)) ) | $spl104),
  inference(cnf_transformation,[],[f1694_D])).
fof(f1694_D,plain,(
  ( ! [X2] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK5)))) | s('type.num.num',sK6) != s('type.num.num',X2)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f1644,plain,(
  ( ! [X3] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X3))))) | s('type.num.num',sK6) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',sK5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) ) | $spl9),
  inference(superposition,[],[f743,f422])).
fof(f1643,plain,(
  ( ! [X2] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X2))))) | s('type.num.num',sK6) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',sK5))))) ) | $spl9),
  inference(superposition,[],[f743,f422])).
fof(f1693,plain,(
  ( ! [X0] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X0)))) | s('type.num.num',sK5) != s('type.num.num',X0)) ) | $spl102),
  inference(cnf_transformation,[],[f1693_D])).
fof(f1693_D,plain,(
  ( ! [X0] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X0)))) | s('type.num.num',sK5) != s('type.num.num',X0)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1642,plain,(
  ( ! [X1] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',sK6))))) | s('type.num.num',sK5) != s('type.num.num',X1) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5))))) ) | $spl9),
  inference(superposition,[],[f743,f420])).
fof(f1641,plain,(
  ( ! [X0] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',sK6))))) | s('type.num.num',sK5) != s('type.num.num',X0) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X0))))) ) | $spl9),
  inference(superposition,[],[f743,f420])).
fof(f1691,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | $spl101),
  inference(cnf_transformation,[],[f1691_D])).
fof(f1691_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f1692,plain,(
  ( ! [X13] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl9),
  inference(forward_demodulation,[],[f1633,f478])).
fof(f1633,plain,(
  ( ! [X13] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl9),
  inference(superposition,[],[f743,f430])).
fof(f1689,plain,(
  ( ! [X12] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl9),
  inference(forward_demodulation,[],[f1632,f478])).
fof(f1632,plain,(
  ( ! [X12] : (s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl9),
  inference(superposition,[],[f743,f430])).
fof(f764,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl21),
  inference(cnf_transformation,[],[f764_D])).
fof(f764_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1684,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl99),
  inference(cnf_transformation,[],[f1684_D])).
fof(f1684_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f1682,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl9),
  inference(forward_demodulation,[],[f1610,f477])).
fof(f1610,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl9),
  inference(superposition,[],[f743,f465])).
fof(f1607,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl9),
  inference(superposition,[],[f743,f430])).
fof(f1606,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl9),
  inference(superposition,[],[f743,f430])).
fof(f1601,plain,(
  ( ! [X21,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) ) | $spl94),
  inference(cnf_transformation,[],[f1601_D])).
fof(f1601_D,plain,(
  ( ! [X21,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f1591,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f844,f430])).
fof(f844,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) ) | $spl1),
  inference(resolution,[],[f733,f396])).
fof(f1590,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f844,f430])).
fof(f1600,plain,(
  ( ! [X21,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) ) | $spl92),
  inference(cnf_transformation,[],[f1600_D])).
fof(f1600_D,plain,(
  ( ! [X21,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f1578,plain,(
  ( ! [X24,X23,X25] : (sP3(X23,X24,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X25) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X25)))) ) | $spl1),
  inference(superposition,[],[f844,f430])).
fof(f1577,plain,(
  ( ! [X21,X22,X20] : (sP3(X20,X21,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X22) | s('type.num.num',X22) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X22),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f844,f430])).
fof(f1562,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5)))) | $spl90),
  inference(cnf_transformation,[],[f1562_D])).
fof(f1562_D,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5)))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f1525,plain,(
  ( ! [X13] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl1),
  inference(superposition,[],[f841,f430])).
fof(f1524,plain,(
  ( ! [X12] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f841,f430])).
fof(f795,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | $spl37),
  inference(cnf_transformation,[],[f795_D])).
fof(f795_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f807,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) | $spl45),
  inference(cnf_transformation,[],[f807_D])).
fof(f807_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f1291,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0')))) | $spl49),
  inference(superposition,[],[f811,f465])).
fof(f792,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',X14)))) ) | $spl34),
  inference(cnf_transformation,[],[f792_D])).
fof(f792_D,plain,(
  ( ! [X14] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',X14)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1288,plain,(
  ( ! [X14] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',X14)))) ) | $spl49),
  inference(superposition,[],[f811,f430])).
fof(f788,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5)))))) ) | $spl30),
  inference(cnf_transformation,[],[f788_D])).
fof(f788_D,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5)))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1287,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5)))))) ) | $spl49),
  inference(superposition,[],[f811,f430])).
fof(f1324,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) | $spl89),
  inference(cnf_transformation,[],[f1324_D])).
fof(f1324_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f1326,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) ) | $spl49),
  inference(forward_demodulation,[],[f1325,f477])).
fof(f1325,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) ) | $spl49),
  inference(forward_demodulation,[],[f1272,f450])).
fof(f1272,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) ) | $spl49),
  inference(superposition,[],[f811,f430])).
fof(f1322,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK4)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl49),
  inference(forward_demodulation,[],[f1321,f477])).
fof(f1321,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl49),
  inference(forward_demodulation,[],[f1271,f450])).
fof(f1271,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl49),
  inference(superposition,[],[f811,f430])).
fof(f1320,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl49),
  inference(forward_demodulation,[],[f1262,f477])).
fof(f1262,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0')))) | $spl49),
  inference(superposition,[],[f811,f465])).
fof(f1319,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl49),
  inference(forward_demodulation,[],[f1259,f477])).
fof(f1259,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) ) | $spl49),
  inference(superposition,[],[f811,f430])).
fof(f1316,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',sK5)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl49),
  inference(forward_demodulation,[],[f1258,f477])).
fof(f1258,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) ) | $spl49),
  inference(superposition,[],[f811,f430])).
fof(f1181,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | $spl85),
  inference(cnf_transformation,[],[f1181_D])).
fof(f1181_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f1144,plain,(
  ( ! [X13] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) ) | $spl1),
  inference(superposition,[],[f838,f430])).
fof(f1143,plain,(
  ( ! [X12] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) ) | $spl1),
  inference(superposition,[],[f838,f430])).
fof(f954,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK5) | $spl83),
  inference(cnf_transformation,[],[f954_D])).
fof(f954_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK5) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f948,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK5) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) | $spl1),
  inference(superposition,[],[f840,f465])).
fof(f840,plain,(
  s('type.num.num',sK5) != s('type.num.num',sK6) | $spl1),
  inference(resolution,[],[f733,f512])).
fof(f512,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f346])).
fof(f945,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK5) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl1),
  inference(superposition,[],[f840,f430])).
fof(f944,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK5) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f840,f430])).
fof(f785,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | $spl29),
  inference(cnf_transformation,[],[f785_D])).
fof(f785_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f914,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0')))) | $spl1),
  inference(superposition,[],[f733,f465])).
fof(f911,plain,(
  ( ! [X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) ) | $spl1),
  inference(superposition,[],[f733,f430])).
fof(f910,plain,(
  ( ! [X12] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) ) | $spl1),
  inference(superposition,[],[f733,f430])).
fof(f734,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | $spl0),
  inference(cnf_transformation,[],[f734_D])).
fof(f734_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f837,plain,(
  ( ! [X28,X26,X27] : (s('type.num.num','const.bool.COND_3'(s(bool,X26),s('type.num.num',X27),s('type.num.num',X28))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,X26) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | s('type.num.num',X28) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6)))) ) | $spl80),
  inference(cnf_transformation,[],[f837_D])).
fof(f837_D,plain,(
  ( ! [X28,X26,X27] : (s('type.num.num','const.bool.COND_3'(s(bool,X26),s('type.num.num',X27),s('type.num.num',X28))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,X26) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | s('type.num.num',X28) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f729,plain,(
  ( ! [X28,X26,X27] : (s('type.num.num','const.bool.COND_3'(s(bool,X26),s('type.num.num',X27),s('type.num.num',X28))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num',X28) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | s(bool,X26) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) )),
  inference(superposition,[],[f394,f453])).
fof(f453,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(X0,X5) != s(X0,X6) | p(s(bool,X2)) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6))) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(ennf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X9] : ! [X30,X40,X29,X41,X42,X43] : ((s(bool,X30) = s(bool,X40) & (p(s(bool,X40)) => s(X9,X29) = s(X9,X41)) & (~p(s(bool,X40)) => s(X9,X42) = s(X9,X43))) => s(X9,'const.bool.COND_3'(s(bool,X30),s(X9,X29),s(X9,X42))) = s(X9,'const.bool.COND_3'(s(bool,X40),s(X9,X41),s(X9,X43))))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.bool.COND_CONG')).
fof(f394,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6)))))),
  inference(cnf_transformation,[],[f308])).
fof(f308,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5,sK6])],[f264])).
fof(f264,plain,(
  ? [X0,X1,X2] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X2))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1,X2] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num',X0),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X8))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X8))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8))),s('type.num.num',X6),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num',X8)))))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.SUB_LEFT_ADD')).
fof(f733,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | $spl1),
  inference(cnf_transformation,[],[f733_D])).
fof(f733_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f836,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,X20) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | s('type.num.num',sK4) != s('type.num.num',X21)) ) | $spl78),
  inference(cnf_transformation,[],[f836_D])).
fof(f836_D,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,X20) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | s('type.num.num',sK4) != s('type.num.num',X21)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f727,plain,(
  ( ! [X21,X22,X20] : (s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | s('type.num.num',sK4) != s('type.num.num',X21) | s(bool,X20) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) )),
  inference(superposition,[],[f394,f452])).
fof(f452,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | ~p(s(bool,X2)) | s(X0,X3) != s(X0,X4) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f277])).
fof(f835,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))),s('type.num.num','const.num.|48|_0')))) | $spl77),
  inference(cnf_transformation,[],[f835_D])).
fof(f835_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f832,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | $spl73),
  inference(cnf_transformation,[],[f832_D])).
fof(f832_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f722,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f394,f465])).
fof(f833,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))),s('type.num.num',X13)))) ) | $spl74),
  inference(cnf_transformation,[],[f833_D])).
fof(f833_D,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))),s('type.num.num',X13)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f719,plain,(
  ( ! [X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))),s('type.num.num',X13)))) )),
  inference(superposition,[],[f394,f430])).
fof(f830,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6)))))))) ) | $spl70),
  inference(cnf_transformation,[],[f830_D])).
fof(f830_D,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6)))))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f718,plain,(
  ( ! [X12] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6)))))))) )),
  inference(superposition,[],[f394,f430])).
fof(f829,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) | $spl69),
  inference(cnf_transformation,[],[f829_D])).
fof(f829_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f790,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | $spl33),
  inference(cnf_transformation,[],[f790_D])).
fof(f790_D,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f707,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f394,f465])).
fof(f827,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))),s('type.num.num',X18)))) ) | $spl66),
  inference(cnf_transformation,[],[f827_D])).
fof(f827_D,plain,(
  ( ! [X18] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))),s('type.num.num',X18)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f704,plain,(
  ( ! [X18] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X18) | s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))),s('type.num.num',X18)))) )),
  inference(superposition,[],[f394,f430])).
fof(f826,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6)))))) ) | $spl64),
  inference(cnf_transformation,[],[f826_D])).
fof(f826_D,plain,(
  ( ! [X17] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6)))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f703,plain,(
  ( ! [X17] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X17) | s('type.num.num',X17) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X17),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6)))))) )),
  inference(superposition,[],[f394,f430])).
fof(f822,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6)))) | $spl63),
  inference(cnf_transformation,[],[f822_D])).
fof(f822_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6)))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f696,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))),
  inference(superposition,[],[f394,f434])).
fof(f820,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK6) | $spl61),
  inference(cnf_transformation,[],[f820_D])).
fof(f820_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK6) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f818,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | $spl59),
  inference(cnf_transformation,[],[f818_D])).
fof(f818_D,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f695,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK6)),
  inference(superposition,[],[f394,f432])).
fof(f694,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5)))),
  inference(superposition,[],[f394,f431])).
fof(f811,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5)))))) | $spl49),
  inference(cnf_transformation,[],[f811_D])).
fof(f811_D,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5)))))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f814,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6)))) ) | $spl54),
  inference(cnf_transformation,[],[f814_D])).
fof(f814_D,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f693,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num',X4))) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',sK6))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))))) )),
  inference(superposition,[],[f394,f423])).
fof(f813,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',X2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5)))))) | s('type.num.num',sK6) != s('type.num.num',X2)) ) | $spl52),
  inference(cnf_transformation,[],[f813_D])).
fof(f813_D,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',X2))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5)))))) | s('type.num.num',sK6) != s('type.num.num',X2)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f692,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',X3))))) | s('type.num.num',sK6) != s('type.num.num',X3) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))))) )),
  inference(superposition,[],[f394,f422])).
fof(f691,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',X2))))) | s('type.num.num',sK6) != s('type.num.num',X2) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))))) )),
  inference(superposition,[],[f394,f422])).
fof(f812,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X0)))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) ) | $spl50),
  inference(cnf_transformation,[],[f812_D])).
fof(f812_D,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X0)))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f690,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',sK6))))) | s('type.num.num',X1) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))))) )),
  inference(superposition,[],[f394,f420])).
fof(f689,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',sK6))))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num',X0))))) )),
  inference(superposition,[],[f394,f420])).
fof(f809,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))))) | $spl47),
  inference(cnf_transformation,[],[f809_D])).
fof(f809_D,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f688,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',sK5))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(superposition,[],[f394,f443])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ! [X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))))),
  inference(ennf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) => ! [X2] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1))),s('type.num.num',X0))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X14,X13] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num',X14)))) => ! [X15] : s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X15),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X14),s('type.num.num',X13))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X13))),s('type.num.num',X14))))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.arithmetic.SUB_SUB')).
fof(f800,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | $spl41),
  inference(cnf_transformation,[],[f800_D])).
fof(f800_D,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f805,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f804,f478])).
fof(f804,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f683,f478])).
fof(f683,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK6),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f394,f465])).
fof(f802,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f801,f478])).
fof(f801,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f680,f478])).
fof(f680,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK6) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK6),s('type.num.num',X13)))) )),
  inference(superposition,[],[f394,f430])).
fof(f797,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) )),
  inference(forward_demodulation,[],[f796,f478])).
fof(f796,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) )),
  inference(forward_demodulation,[],[f679,f478])).
fof(f679,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK6)))) )),
  inference(superposition,[],[f394,f430])).
fof(f793,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f670,f481])).
fof(f670,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f394,f465])).
fof(f791,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',X14)))) )),
  inference(forward_demodulation,[],[f667,f481])).
fof(f667,plain,(
  ( ! [X14] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X14) | s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',X14)))) )),
  inference(superposition,[],[f394,f430])).
fof(f787,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5)))))) )),
  inference(forward_demodulation,[],[f666,f481])).
fof(f666,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X13),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5)))))) )),
  inference(superposition,[],[f394,f430])).
fof(f772,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | $spl25),
  inference(cnf_transformation,[],[f772_D])).
fof(f772_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f783,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f782,f477])).
fof(f782,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f781,f450])).
fof(f781,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f780,f481])).
fof(f780,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f779,f477])).
fof(f779,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f654,f450])).
fof(f654,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f394,f465])).
fof(f777,plain,(
  ( ! [X13] : (s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f776,f477])).
fof(f776,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f775,f450])).
fof(f775,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f774,f481])).
fof(f774,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f773,f477])).
fof(f773,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f651,f450])).
fof(f651,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK5) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',X13)))) )),
  inference(superposition,[],[f394,f430])).
fof(f769,plain,(
  ( ! [X12] : (s('type.num.num',sK4) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(forward_demodulation,[],[f768,f477])).
fof(f768,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(forward_demodulation,[],[f767,f450])).
fof(f767,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(forward_demodulation,[],[f766,f481])).
fof(f766,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(forward_demodulation,[],[f765,f477])).
fof(f765,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK4))),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(forward_demodulation,[],[f650,f450])).
fof(f650,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK6))),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK5)))) )),
  inference(superposition,[],[f394,f430])).
fof(f757,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | $spl17),
  inference(cnf_transformation,[],[f757_D])).
fof(f757_D,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f762,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f761,f477])).
fof(f761,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(forward_demodulation,[],[f641,f477])).
fof(f641,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK4),s('type.num.num','const.num.|48|_0'))))),
  inference(superposition,[],[f394,f465])).
fof(f759,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f758,f477])).
fof(f758,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) )),
  inference(forward_demodulation,[],[f638,f477])).
fof(f638,plain,(
  ( ! [X13] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X13) | s('type.num.num',sK4) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK4),s('type.num.num',X13)))) )),
  inference(superposition,[],[f394,f430])).
fof(f754,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) )),
  inference(forward_demodulation,[],[f753,f477])).
fof(f753,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) )),
  inference(forward_demodulation,[],[f637,f477])).
fof(f637,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK5))),s('type.num.num',sK6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X12) | s('type.num.num',X12) != s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X12),s('type.num.num',sK4)))) )),
  inference(superposition,[],[f394,f430])).
fof(f750,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl12),
  inference(cnf_transformation,[],[f750_D])).
fof(f750_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f739,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | $spl7),
  inference(cnf_transformation,[],[f739_D])).
fof(f739_D,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f752,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(forward_demodulation,[],[f632,f460])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9] : ! [X31,X32] : s(X9,X32) = s(X9,'const.bool.COND_3'(s(bool,'F_0'),s(X9,X31),s(X9,X32)))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.bool.COND_CLAUSES|split|1')).
fof(f632,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f394,f476])).
fof(f747,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) | $spl10),
  inference(cnf_transformation,[],[f747_D])).
fof(f747_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f743,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | $spl9),
  inference(cnf_transformation,[],[f743_D])).
fof(f743_D,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f751,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(forward_demodulation,[],[f631,f461])).
fof(f631,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f394,f476])).
fof(f748,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(forward_demodulation,[],[f630,f460])).
fof(f630,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f394,f475])).
fof(f745,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(forward_demodulation,[],[f629,f461])).
fof(f629,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))),
  inference(superposition,[],[f394,f475])).
fof(f744,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(forward_demodulation,[],[f628,f461])).
fof(f628,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(superposition,[],[f394,f474])).
fof(f741,plain,(
  s('type.num.num',sK4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(forward_demodulation,[],[f627,f461])).
fof(f627,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(superposition,[],[f394,f472])).
fof(f740,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(forward_demodulation,[],[f626,f460])).
fof(f626,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(superposition,[],[f394,f470])).
fof(f470,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f335])).
fof(f335,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X28] : (s(bool,'F_0') = s(bool,X28) <=> ~p(s(bool,X28)))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.bool.EQ_CLAUSES|split|2')).
fof(f737,plain,(
  s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(forward_demodulation,[],[f625,f460])).
fof(f625,plain,(
  s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))),
  inference(superposition,[],[f394,f468])).
fof(f468,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f334])).
fof(f334,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X28] : (s(bool,'F_0') = s(bool,X28) <=> ~p(s(bool,X28)))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.bool.EQ_CLAUSES|split|3')).
fof(f736,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | ~p(s(bool,X11))) ) | $spl4),
  inference(cnf_transformation,[],[f736_D])).
fof(f736_D,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | ~p(s(bool,X11))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f624,plain,(
  ( ! [X12] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X12),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | ~p(s(bool,X12)) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) )),
  inference(superposition,[],[f394,f440])).
fof(f440,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X31,X32] : ((p(s(bool,X31)) => p(s(bool,X32))) => ((p(s(bool,X32)) => p(s(bool,X31))) => s(bool,X31) = s(bool,X32)))),
  file('thm.arithmetic.SUB_LEFT_ADD','thm.bool.IMP_ANTISYM_AX')).
fof(f623,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f394,f440])).
fof(f735,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X9),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | p(s(bool,X9))) ) | $spl2),
  inference(cnf_transformation,[],[f735_D])).
fof(f735_D,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X9),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | p(s(bool,X9))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f622,plain,(
  ( ! [X10] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X10),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6)))) | p(s(bool,X10))) )),
  inference(superposition,[],[f394,f437])).
fof(f437,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f271])).
fof(f621,plain,(
  ( ! [X9] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) != s('type.num.num','const.bool.COND_3'(s(bool,X9),s('type.num.num',sK4),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK4),s('type.num.num',sK5))),s('type.num.num',sK6))))) | p(s(bool,X9)) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5),s('type.num.num',sK6))))) )),
  inference(superposition,[],[f394,f437])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 231
% Passive clauses: 3071
% Generated clauses: 5494
% Final active clauses: 216
% Final passive clauses: 2447
% Input formulas: 133
% Initial clauses: 222
% 
% Selected by SInE selection: 104
% SInE iterations: 1
% 
% Duplicate literals: 38
% Trivial inequalities: 20
% Fw subsumption resolutions: 423
% Fw demodulations: 545
% 
% Simple tautologies: 119
% Equational tautologies: 2
% Forward subsumptions: 1048
% 
% Binary resolution: 1105
% Forward superposition: 2958
% Backward superposition: 253
% 
% Split clauses: 165
% Split components: 330
% SAT solver clauses: 173
% SAT solver unit clauses: 2
% SAT solver binary clauses: 46
% 
% Sat splits: 165
% Sat splitting refutations: 7
% 
% Memory used [KB]: 3837
% Time elapsed: 0.150 s
% ------------------------------
% ------------------------------
% Success in time 11.9 s
