Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.realax
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.realax
% SZS output start Proof for thm.realax
fof(f1704,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1703,f1586,f1173,f1584,f1193,f1154,f1535])).
fof(f1535,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f1526,f1483])).
fof(f1483,plain,(
  ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0')))) | $spl22),
  inference(backward_demodulation,[],[f1476,f33])).
fof(f33,plain,(
  ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))))),
  inference(cnf_transformation,[],[f14])).
fof(f14,plain,(
  ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))))),
  inference(flattening,[],[f6])).
fof(f6,negated_conjecture,(
  ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))))),
  file('thm.realax.TREAL_INV_|48|','thm.realax.TREAL_INV_|48|')).
fof(f1476,plain,(
  s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0') = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))) | $spl22),
  inference(superposition,[],[f1397,f36])).
fof(f36,plain,(
  s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0') = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_|49|_0')))),
  inference(cnf_transformation,[],[f11])).
fof(f11,axiom,(
  s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0') = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_|49|_0')))),
  file('thm.realax.TREAL_INV_|48|','thm.realax.treal_|48|')).
fof(f1397,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0') = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1)))))) ) | $spl22),
  inference(forward_demodulation,[],[f1359,f41])).
fof(f41,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f18])).
fof(f18,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X8,X9] : s(X6,X8) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X8),s(X6,X9)))),
  file('thm.realax.TREAL_INV_|48|','thm.bool.COND_CLAUSES|split|0')).
fof(f1359,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,'T_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0')))))))))) ) | $spl22),
  inference(backward_demodulation,[],[f1154,f46])).
fof(f46,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,sK1(X1,X0)),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0')))))))))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1] : ((~p(s(bool,sK1(X1,X0))) | s('type.hreal.hreal',X0) = s('type.hreal.hreal',X1)) & (s('type.hreal.hreal',X0) != s('type.hreal.hreal',X1) | p(s(bool,sK1(X1,X0)))) & s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,sK1(X1,X0)),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ? [X2] : ((~p(s(bool,X2)) | s('type.hreal.hreal',X0) = s('type.hreal.hreal',X1)) & (s('type.hreal.hreal',X0) != s('type.hreal.hreal',X1) | p(s(bool,X2))) & s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,X2),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))))))))))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : ? [X2] : (((~p(s(bool,X2)) | s('type.hreal.hreal',X0) = s('type.hreal.hreal',X1)) & (s('type.hreal.hreal',X0) != s('type.hreal.hreal',X1) | p(s(bool,X2)))) & s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,X2),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))))))))))),
  inference(nnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : ? [X2] : ((p(s(bool,X2)) <=> s('type.hreal.hreal',X0) = s('type.hreal.hreal',X1)) & s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,X2),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X1),s('type.hreal.hreal',X0))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))))))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X7,X10] : ? [X11] : ((p(s(bool,X11)) <=> s('type.hreal.hreal',X7) = s('type.hreal.hreal',X10)) & s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X7),s('type.hreal.hreal',X10))))) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,X11),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X10),s('type.hreal.hreal',X7))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X7),s('type.hreal.hreal',X10))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X10),s('type.hreal.hreal',X7))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))))))))))),
  file('thm.realax.TREAL_INV_|48|','thm.realax.treal_inv')).
fof(f1526,plain,(
  p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))))),
  inference(superposition,[],[f785,f36])).
fof(f785,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))))))) )),
  inference(equality_resolution,[],[f48])).
fof(f48,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X3))) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X1))) | p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X3))))))) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X3)))))) | s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X3))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X1)))) & (s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X3))) != s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X1))) | p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X3))))))))),
  inference(nnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X1))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X3)))))) <=> s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X3))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X2),s('type.hreal.hreal',X1))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X12,X13,X14,X15] : (p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X12),s('type.hreal.hreal',X13))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X14),s('type.hreal.hreal',X15)))))) <=> s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X12),s('type.hreal.hreal',X15))) = s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal',X14),s('type.hreal.hreal',X13))))),
  file('thm.realax.TREAL_INV_|48|','thm.realax.treal_eq')).
fof(f1154,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,sK1(X4,X5))) ) | $spl22),
  inference(cnf_transformation,[],[f1154_D])).
fof(f1154_D,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,sK1(X4,X5))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1193,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,sK1(X2,X3))) ) | $spl32),
  inference(resolution,[],[f1173,f53])).
fof(f53,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f37,f38])).
fof(f38,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.realax.TREAL_INV_|48|',aHL_BOOL_CASES)).
fof(f37,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f15])).
fof(f15,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.realax.TREAL_INV_|48|',aHL_FALSITY)).
fof(f1584,plain,(
  s(bool,'T_0') = s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))) | $spl48),
  inference(cnf_transformation,[],[f1584_D])).
fof(f1584_D,plain,(
  s(bool,'T_0') = s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1173,plain,(
  ( ! [X45,X46] : (p(s(bool,sK1(X45,X46)))) ) | $spl32),
  inference(cnf_transformation,[],[f1173_D])).
fof(f1173_D,plain,(
  ( ! [X45,X46] : (p(s(bool,sK1(X45,X46)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1586,plain,(
  ( ! [X4,X5] : (p(s(bool,sK1(X4,X5))) | s(bool,'T_0') = s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0')))) )),
  inference(subsumption_resolution,[],[f1580,f45])).
fof(f45,plain,(
  ( ! [X0,X1] : (p(s(bool,sK1(X1,X0))) | s('type.hreal.hreal',X0) != s('type.hreal.hreal',X1)) )),
  inference(cnf_transformation,[],[f31])).
fof(f1580,plain,(
  ( ! [X4,X5] : (p(s(bool,sK1(X4,X5))) | s('type.hreal.hreal',X4) = s('type.hreal.hreal',X5) | s(bool,'T_0') = s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0')))) )),
  inference(superposition,[],[f1526,f337])).
fof(f337,plain,(
  ( ! [X23,X21,X22] : (s(bool,X23) = s(bool,sK1(X21,X22)) | s('type.hreal.hreal',X21) = s('type.hreal.hreal',X22) | s(bool,'T_0') = s(bool,X23)) )),
  inference(subsumption_resolution,[],[f235,f34])).
fof(f34,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.realax.TREAL_INV_|48|',aHL_TRUTH)).
fof(f235,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'T_0')) | s('type.hreal.hreal',X21) = s('type.hreal.hreal',X22) | s(bool,X23) = s(bool,sK1(X21,X22)) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f44,f51])).
fof(f51,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f38,f38])).
fof(f44,plain,(
  ( ! [X0,X1] : (~p(s(bool,sK1(X1,X0))) | s('type.hreal.hreal',X0) = s('type.hreal.hreal',X1)) )),
  inference(cnf_transformation,[],[f31])).
fof(f1703,plain,(
  $false | $spl48),
  inference(subsumption_resolution,[],[f1702,f34])).
fof(f1702,plain,(
  ~p(s(bool,'T_0')) | $spl48),
  inference(forward_demodulation,[],[f1695,f1584])).
fof(f1695,plain,(
  ~p(s(bool,'const.realax.treal_eq_2'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'))))),
  inference(backward_demodulation,[],[f1688,f33])).
fof(f1688,plain,(
  s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0') = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0')))),
  inference(superposition,[],[f1026,f36])).
fof(f1026,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0') = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X0)))))) )),
  inference(forward_demodulation,[],[f992,f41])).
fof(f992,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_inv_1'(s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X0))))) = s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,'T_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.realax.treal_|48|_0'),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.bool.COND_3'(s(bool,'const.hreal.hreal_lt_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X0))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X0))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.hreal.hreal','const.hreal.hreal_|49|_0'))),s('type.pair.prod'('type.hreal.hreal','type.hreal.hreal'),'const.pair.,_2'(s('type.hreal.hreal','const.hreal.hreal_|49|_0'),s('type.hreal.hreal','const.hreal.hreal_add_2'(s('type.hreal.hreal','const.hreal.hreal_inv_1'(s('type.hreal.hreal','const.hreal.hreal_sub_2'(s('type.hreal.hreal',X0),s('type.hreal.hreal',X0))))),s('type.hreal.hreal','const.hreal.hreal_|49|_0')))))))))) )),
  inference(superposition,[],[f46,f120])).
fof(f120,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK1(X0,X0))) )),
  inference(equality_resolution,[],[f59])).
fof(f59,plain,(
  ( ! [X0,X1] : (s('type.hreal.hreal',X0) != s('type.hreal.hreal',X1) | s(bool,'T_0') = s(bool,sK1(X0,X1))) )),
  inference(subsumption_resolution,[],[f55,f37])).
fof(f55,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s('type.hreal.hreal',X0) != s('type.hreal.hreal',X1) | s(bool,'T_0') = s(bool,sK1(X0,X1))) )),
  inference(superposition,[],[f45,f38])).
% SZS output end Proof for thm.realax
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 57
% Passive clauses: 319
% Generated clauses: 1627
% Final active clauses: 36
% Final passive clauses: 171
% Input formulas: 12
% Initial clauses: 16
% 
% Duplicate literals: 128
% Trivial inequalities: 18
% Fw subsumption resolutions: 127
% Fw demodulations: 44
% Bw demodulations: 10
% 
% Simple tautologies: 30
% Equational tautologies: 17
% Forward subsumptions: 894
% 
% Binary resolution: 20
% Forward superposition: 538
% Backward superposition: 614
% Self superposition: 55
% Equality factoring: 28
% Equality resolution: 6
% 
% Split clauses: 46
% Split components: 92
% SAT solver clauses: 71
% SAT solver unit clauses: 5
% SAT solver binary clauses: 32
% SAT solver learnt clauses: 1
% 
% Sat splits: 46
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 70
% TWLsolver calls for satisfiability: 12
% 
% Memory used [KB]: 1791
% Time elapsed: 0.051 s
% ------------------------------
% ------------------------------
% Success in time 0.052 s
