Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.util_prob
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.util_prob
% SZS output start Proof for thm.util_prob
fof(f8584,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f7878,f7916,f8542,f337,f591,f6253,f7915,f340,f585,f8543,f338,f589,f588,f5943,f339,f586,f584,f398,f594,f592,f347,f593])).
fof(f593,plain,(
  ~p(s(bool,'F_0')) | $spl13),
  inference(cnf_transformation,[],[f593_D])).
fof(f593_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f347,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.util_prob.MINIMAL_EXISTS|48|',aHL_FALSITY)).
fof(f592,plain,(
  p(s(bool,'F_0')) | $spl12),
  inference(cnf_transformation,[],[f592_D])).
fof(f592_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f594,plain,(
  p(s(bool,'T_0')) | $spl14),
  inference(cnf_transformation,[],[f594_D])).
fof(f594_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f398,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f259])).
fof(f259,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X16] : ((p(s(bool,'F_0')) => p(s(bool,X16))) <=> p(s(bool,'T_0')))),
  file('thm.util_prob.MINIMAL_EXISTS|48|','thm.bool.IMP_CLAUSES|split|2')).
fof(f584,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X3))))) ) | $spl0),
  inference(cnf_transformation,[],[f584_D])).
fof(f584_D,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X3))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f586,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16(X1)),s('type.num.num',X1))))) ) | $spl4),
  inference(cnf_transformation,[],[f586_D])).
fof(f586_D,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16(X1)),s('type.num.num',X1))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f339,plain,(
  ( ! [X3,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16(X1)),s('type.num.num',X1)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X1)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X3))))) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) & ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK15)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X5))))))) & (! [X3] : ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X3)))) | ! [X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X1)))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16(X1)),s('type.num.num',X1)))) & p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(X1))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16])],[f231])).
fof(f231,plain,(
  ? [X0] : ((? [X6] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X6)))) | ? [X4] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X4)))) & ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X4)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X5))))))) & (! [X3] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3)))) | ! [X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) | ? [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X1)))) & p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))))))),
  inference(rectify,[],[f230])).
fof(f230,plain,(
  ? [X0] : ((? [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) | ? [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) & ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3))))))) & (! [X1] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) | ! [X2] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | ? [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) & p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3))))))))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ? [X0] : (? [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) <~> ? [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) & ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3)))))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0] : (? [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) <=> ? [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) & ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) => ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3)))))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : (? [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) <=> ? [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) & ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) => ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (? [X7] : p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X7)))) <=> ? [X7] : (p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X7)))) & ! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) => ~p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X8)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (? [X7] : p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X7)))) <=> ? [X7] : (p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X7)))) & ! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) => ~p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X8)))))))),
  file('thm.util_prob.MINIMAL_EXISTS|48|','thm.util_prob.MINIMAL_EXISTS|48|')).
fof(f5943,plain,(
  $false | ($spl2 | $spl4 | $spl8 | $spl10 | $spl14)),
  inference(subsumption_resolution,[],[f5942,f5717])).
fof(f5717,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK15))))) | ($spl4 | $spl8 | $spl10 | $spl14)),
  inference(resolution,[],[f5716,f589])).
fof(f5716,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16(sK15)),s('type.num.num',sK15)))) | ($spl4 | $spl10 | $spl14)),
  inference(subsumption_resolution,[],[f5709,f594])).
fof(f5709,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16(sK15)),s('type.num.num',sK15)))) | ($spl4 | $spl10)),
  inference(superposition,[],[f586,f928])).
fof(f928,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15))) | $spl10),
  inference(resolution,[],[f591,f367])).
fof(f367,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X16] : (s(bool,'T_0') = s(bool,X16) <=> p(s(bool,X16)))),
  file('thm.util_prob.MINIMAL_EXISTS|48|','thm.bool.EQ_CLAUSES|split|1')).
fof(f5942,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK15))))) | ($spl2 | $spl10 | $spl14)),
  inference(subsumption_resolution,[],[f5940,f594])).
fof(f5940,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK15))))) | ($spl2 | $spl10)),
  inference(superposition,[],[f585,f928])).
fof(f588,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) | $spl6),
  inference(cnf_transformation,[],[f588_D])).
fof(f588_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f589,plain,(
  ( ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK15)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X5))))) ) | $spl8),
  inference(cnf_transformation,[],[f589_D])).
fof(f589_D,plain,(
  ( ! [X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK15)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X5))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f338,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X5)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK15)))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))) )),
  inference(cnf_transformation,[],[f232])).
fof(f8543,plain,(
  $false | ($spl0 | $spl10 | $spl14)),
  inference(resolution,[],[f584,f4937])).
fof(f4937,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK22(sK13))))) | ($spl10 | $spl14)),
  inference(subsumption_resolution,[],[f4935,f594])).
fof(f4935,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK22(sK13))))) | ($spl10 | $spl14)),
  inference(superposition,[],[f420,f4676])).
fof(f4676,plain,(
  s(bool,'T_0') = s(bool,sK21(sK13)) | ($spl10 | $spl14)),
  inference(resolution,[],[f4675,f367])).
fof(f4675,plain,(
  p(s(bool,sK21(sK13))) | ($spl10 | $spl14)),
  inference(subsumption_resolution,[],[f4673,f594])).
fof(f4673,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,sK21(sK13))) | $spl10),
  inference(superposition,[],[f421,f928])).
fof(f421,plain,(
  ( ! [X2,X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,sK21(X0)))) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0] : ((~p(s(bool,sK21(X0))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK22(X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,sK21(X0)))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,sK21(X0)),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21,sK22])],[f273])).
fof(f273,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,X1))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(rectify,[],[f272])).
fof(f272,plain,(
  ! [X0] : ? [X1] : ((~p(s(bool,X1)) | ? [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,X1))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(flattening,[],[f271])).
fof(f271,plain,(
  ! [X0] : ? [X1] : (((~p(s(bool,X1)) | ? [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) & (! [X2] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | p(s(bool,X1)))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0] : ? [X1] : ((p(s(bool,X1)) <=> ? [X2] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2))))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X0))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X1),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X0))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X6] : ? [X24] : ((p(s(bool,X24)) <=> ? [X7] : p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X7))))) & s('type.option.option'('type.num.num'),'const.while.OLEAST_1'(s(fun('type.num.num',bool),X6))) = s('type.option.option'('type.num.num'),'const.bool.COND_3'(s(bool,X24),s('type.option.option'('type.num.num'),'const.option.SOME_1'(s('type.num.num','const.while.LEAST_1'(s(fun('type.num.num',bool),X6))))),s('type.option.option'('type.num.num'),'const.option.NONE_0'))))),
  file('thm.util_prob.MINIMAL_EXISTS|48|','thm.while.OLEAST_def')).
fof(f420,plain,(
  ( ! [X0] : (~p(s(bool,sK21(X0))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK22(X0)))))) )),
  inference(cnf_transformation,[],[f274])).
fof(f585,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X1)))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(X1)))))) ) | $spl2),
  inference(cnf_transformation,[],[f585_D])).
fof(f585_D,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X1)))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(X1)))))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f340,plain,(
  ( ! [X3,X1] : (p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(X1))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X1)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X3))))) )),
  inference(cnf_transformation,[],[f232])).
fof(f7915,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK18(sK13)))))) | $spl142),
  inference(cnf_transformation,[],[f7915_D])).
fof(f7915_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK18(sK13)))))) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f6253,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK18(sK13)))))) | ($spl2 | $spl6)),
  inference(resolution,[],[f6043,f585])).
fof(f6043,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK18(sK13))))) | $spl6),
  inference(resolution,[],[f364,f588])).
fof(f364,plain,(
  ( ! [X0,X3] : (~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3)))) | p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK18(X0)))))) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0] : (! [X3] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3)))) | (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK18(X0))))) & ! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',sK18(X0))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f237])).
fof(f237,plain,(
  ! [X0] : (! [X3] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3)))) | ? [X1] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) & ! [X2] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X1)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))))))),
  inference(rectify,[],[f177])).
fof(f177,plain,(
  ! [X0] : (! [X1] : ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) | ? [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) & ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3)))))))),
  inference(ennf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : (? [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) => ? [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) & ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) => ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3)))))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : (? [X1] : p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X1)))) => ? [X2] : (p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) & ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2)))) => ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3)))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : (? [X7] : p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X7)))) => ? [X7] : (p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X7)))) & ! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X8),s('type.num.num',X7)))) => ~p(s(bool,happ(s(fun('type.num.num',bool),X6),s('type.num.num',X8)))))))),
  file('thm.util_prob.MINIMAL_EXISTS|48|','thm.arithmetic.WOP')).
fof(f591,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | $spl10),
  inference(cnf_transformation,[],[f591_D])).
fof(f591_D,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f337,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK15)))) | p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK14))))),
  inference(cnf_transformation,[],[f232])).
fof(f8542,plain,(
  $false | ($spl0 | $spl6)),
  inference(resolution,[],[f584,f6878])).
fof(f6878,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK20(sK13,'type.num.num'))))) | $spl6),
  inference(resolution,[],[f412,f588])).
fof(f412,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X1,X0)))))) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK19,sK20])],[f266])).
fof(f266,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f265])).
fof(f265,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X17] : ! [X6] : (~? [X18] : p(s(bool,happ(s(fun(X17,bool),X6),s(X17,X18)))) <=> ! [X18] : ~p(s(bool,happ(s(fun(X17,bool),X6),s(X17,X18)))))),
  file('thm.util_prob.MINIMAL_EXISTS|48|','thm.bool.NOT_EXISTS_THM')).
fof(f7916,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK18(sK13)))))) | $spl143),
  inference(cnf_transformation,[],[f7916_D])).
fof(f7916_D,plain,(
  ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK18(sK13)))))) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f7878,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',sK16(sK18(sK13)))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),sK13),s('type.num.num',X0))))) ) | ($spl4 | $spl6)),
  inference(resolution,[],[f365,f6254])).
fof(f6254,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK16(sK18(sK13))),s('type.num.num',sK18(sK13))))) | ($spl4 | $spl6)),
  inference(resolution,[],[f6043,f586])).
fof(f365,plain,(
  ( ! [X2,X0,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',sK18(X0))))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X2)))) | ~p(s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',X3))))) )),
  inference(cnf_transformation,[],[f238])).
% SZS output end Proof for thm.util_prob
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 565
% Passive clauses: 3348
% Generated clauses: 8235
% Final active clauses: 553
% Final passive clauses: 2584
% Input formulas: 69
% Initial clauses: 239
% 
% Duplicate literals: 148
% Trivial inequalities: 30
% Fw subsumption resolutions: 522
% Fw demodulations: 333
% Bw demodulations: 7
% 
% Simple tautologies: 159
% Equational tautologies: 23
% Forward subsumptions: 3559
% Fw demodulations to eq. taut.: 6
% 
% Binary resolution: 4505
% Factoring: 694
% Forward superposition: 811
% Backward superposition: 859
% Self superposition: 13
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 134
% Split components: 270
% SAT solver clauses: 261
% SAT solver unit clauses: 8
% SAT solver binary clauses: 50
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 1
% 
% Sat splits: 134
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 259
% TWLsolver calls for satisfiability: 207
% 
% Memory used [KB]: 4477
% Time elapsed: 0.185 s
% ------------------------------
% ------------------------------
% Success in time 0.181 s
