Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f4824,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f4813,f518,f150,f250,f151,f252,f248,f4711,f556,f152,f253,f251,f249])).
fof(f249,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | $spl1),
  inference(cnf_transformation,[],[f249_D])).
fof(f249_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK1),s('type.num.num',sK0))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f251,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK1) | $spl3),
  inference(cnf_transformation,[],[f251_D])).
fof(f251_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f253,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK0) | $spl5),
  inference(cnf_transformation,[],[f253_D])).
fof(f253_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK0) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f152,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK1) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK1),s('type.num.num',sK0)))),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) & s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0))) & (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK1) | s('type.num.num','const.num.|48|_0') != s('type.num.num',sK0))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f109])).
fof(f109,plain,(
  ? [X0,X1] : ((s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))) & (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ? [X0,X1] : ((s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))) & (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))) | (s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0))))),
  inference(nnf_transformation,[],[f91])).
fof(f91,plain,(
  ? [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))) <~> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ~! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X0))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X1) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X7),s('type.num.num',X6))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X7),s('type.num.num',X6))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X7) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X6)))),
  file('thm.arithmetic.MAX_EQ_|48|','thm.arithmetic.MAX_EQ_|48|')).
fof(f556,plain,(
  $false | ($spl0 | $spl3 | $spl4)),
  inference(subsumption_resolution,[],[f543,f251])).
fof(f543,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | ($spl0 | $spl4)),
  inference(superposition,[],[f496,f521])).
fof(f521,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f248,f252])).
fof(f496,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f472,f210])).
fof(f210,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X11] : ! [X8,X9] : s(X11,X9) = s(X11,'const.bool.COND_3'(s(bool,'F_0'),s(X11,X8),s(X11,X9)))),
  file('thm.arithmetic.MAX_EQ_|48|','thm.bool.COND_CLAUSES|split|1')).
fof(f472,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(superposition,[],[f219,f269])).
fof(f269,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(resolution,[],[f168,f205])).
fof(f205,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0'))))),
  file('thm.arithmetic.MAX_EQ_|48|','thm.prim_rec.NOT_LESS_|48|')).
fof(f168,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.arithmetic.MAX_EQ_|48|','thm.bool.EQ_CLAUSES|split|3')).
fof(f219,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7,X6] : s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X7),s('type.num.num',X6))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num',X6),s('type.num.num',X7)))),
  file('thm.arithmetic.MAX_EQ_|48|','thm.arithmetic.MAX_DEF')).
fof(f4711,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f4710,f253])).
fof(f4710,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f4709,f248])).
fof(f4709,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f4628,f211])).
fof(f211,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X11] : ! [X8,X9] : s(X11,X8) = s(X11,'const.bool.COND_3'(s(bool,'T_0'),s(X11,X8),s(X11,X9)))),
  file('thm.arithmetic.MAX_EQ_|48|','thm.bool.COND_CLAUSES|split|0')).
fof(f4628,plain,(
  s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',sK0),s('type.num.num',sK1))) | ($spl0 | $spl3)),
  inference(superposition,[],[f219,f4618])).
fof(f4618,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4572,f251])).
fof(f4572,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | $spl0),
  inference(superposition,[],[f497,f248])).
fof(f497,plain,(
  ( ! [X2,X1] : (s('type.num.num',X1) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f473,f210])).
fof(f473,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X1),s('type.num.num',X2))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X2),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(superposition,[],[f219,f161])).
fof(f161,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.arithmetic.MAX_EQ_|48|',aHL_BOOL_CASES)).
fof(f248,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | $spl0),
  inference(cnf_transformation,[],[f248_D])).
fof(f248_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK1),s('type.num.num',sK0))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f252,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | $spl4),
  inference(cnf_transformation,[],[f252_D])).
fof(f252_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f151,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK1),s('type.num.num',sK0)))),
  inference(cnf_transformation,[],[f110])).
fof(f250,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | $spl2),
  inference(cnf_transformation,[],[f250_D])).
fof(f250_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f150,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK1),s('type.num.num',sK0)))),
  inference(cnf_transformation,[],[f110])).
fof(f518,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(trivial_inequality_removal,[],[f507])).
fof(f507,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | ($spl1 | $spl2 | $spl4)),
  inference(superposition,[],[f273,f486])).
fof(f486,plain,(
  ( ! [X2] : (s('type.num.num',X2) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',X2),s('type.num.num',X2)))) )),
  inference(superposition,[],[f219,f212])).
fof(f212,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X2)))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X2)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X11] : ! [X14,X10] : s(X11,X10) = s(X11,'const.bool.COND_3'(s(bool,X14),s(X11,X10),s(X11,X10)))),
  file('thm.arithmetic.MAX_EQ_|48|','thm.bool.COND_ID')).
fof(f273,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f272,f250])).
fof(f272,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num',sK1),s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f249,f252])).
fof(f4813,plain,(
  $false | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f4812,f253])).
fof(f4812,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | ($spl0 | $spl2)),
  inference(duplicate_literal_removal,[],[f4778])).
fof(f4778,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK0) | ($spl0 | $spl2)),
  inference(superposition,[],[f4739,f713])).
fof(f713,plain,(
  ( ! [X3] : (s('type.num.num',X3) = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3)) )),
  inference(forward_demodulation,[],[f689,f211])).
fof(f689,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X3),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X3)) )),
  inference(superposition,[],[f219,f441])).
fof(f441,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X5)) )),
  inference(resolution,[],[f201,f164])).
fof(f164,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.arithmetic.MAX_EQ_|48|','thm.bool.EQ_CLAUSES|split|0')).
fof(f201,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))))),
  file('thm.arithmetic.MAX_EQ_|48|','thm.arithmetic.NOT_ZERO_LT_ZERO')).
fof(f4739,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MAX_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f250,f248])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 80
% Passive clauses: 920
% Generated clauses: 4543
% Final active clauses: 69
% Final passive clauses: 765
% Input formulas: 39
% Initial clauses: 95
% 
% Duplicate literals: 212
% Trivial inequalities: 128
% Fw subsumption resolutions: 437
% Fw demodulations: 326
% Bw demodulations: 4
% 
% Simple tautologies: 153
% Equational tautologies: 30
% Forward subsumptions: 2277
% 
% Binary resolution: 119
% Factoring: 1
% Forward superposition: 1601
% Backward superposition: 1439
% Self superposition: 81
% Equality factoring: 52
% Equality resolution: 15
% 
% Split clauses: 69
% Split components: 141
% SAT solver clauses: 116
% SAT solver unit clauses: 11
% SAT solver binary clauses: 59
% SAT solver learnt clauses: 3
% SAT solver learnt literals: 2
% 
% Sat splits: 69
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 112
% TWLsolver calls for satisfiability: 38
% 
% Memory used [KB]: 2942
% Time elapsed: 0.116 s
% ------------------------------
% ------------------------------
% Success in time 0.118 s
