Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (48762)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 697
% Passive clauses: 9024
% Generated clauses: 19849
% Final active clauses: 682
% Final passive clauses: 7891
% Input formulas: 35
% Initial clauses: 153
% 
% Duplicate literals: 426
% Trivial inequalities: 102
% Fw subsumption resolutions: 1947
% Fw demodulations: 332
% Bw demodulations: 10
% 
% Simple tautologies: 206
% Equational tautologies: 56
% Forward subsumptions: 7055
% 
% Binary resolution: 10618
% Factoring: 1269
% Forward superposition: 2827
% Backward superposition: 1873
% Self superposition: 24
% Equality factoring: 28
% Equality resolution: 27
% 
% Split clauses: 555
% Split components: 1118
% SAT solver clauses: 975
% SAT solver unit clauses: 25
% SAT solver binary clauses: 160
% 
% Sat splits: 555
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 974
% TWLsolver calls for satisfiability: 392
% 
% Memory used [KB]: 9722
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (48790)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 350
% Passive clauses: 2465
% Generated clauses: 10029
% Final active clauses: 338
% Final passive clauses: 1990
% Input formulas: 35
% Initial clauses: 153
% 
% Duplicate literals: 181
% Trivial inequalities: 35
% Fw subsumption resolutions: 835
% Fw demodulations: 404
% Bw demodulations: 135
% Condensations: 214
% Global subsumptions: 13
% 
% Simple tautologies: 208
% Equational tautologies: 17
% Forward subsumptions: 2792
% 
% Binary resolution: 3317
% Factoring: 191
% Forward superposition: 1249
% Backward superposition: 704
% Self superposition: 24
% Equality factoring: 28
% Equality resolution: 9
% 
% Split clauses: 52
% Split components: 109
% SAT solver clauses: 2955
% SAT solver unit clauses: 161
% SAT solver binary clauses: 75
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 52
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 2903
% TWLsolver calls for satisfiability: 1476
% 
% Memory used [KB]: 12025
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f22819,plain,(
  $false),
  inference(subsumption_resolution,[],[f22818,f2308])).
fof(f2308,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f198,f197,f199,f206,f213,f212,f250,f247,f255,f254,f370,f273,f270,f269,f268,f278,f291,f290,f289,f288,f287,f286,f285,f295,f294,f293,f292,f299,f298,f297,f296,f303,f302,f301,f300,f307,f306,f305,f304,f363,f311,f310,f309,f308,f318,f317,f316,f315,f324,f321,f320,f319,f331,f330,f329,f362,f327,f326,f325,f335,f334,f333,f332,f339,f338,f337,f336,f361,f371,f372,f374,f373,f376,f375,f377,f378,f380,f379,f381,f382,f384,f383,f386,f385,f388,f387,f390,f389,f391,f392,f394,f393,f396,f395,f397,f404,f403,f402,f401,f400,f399,f398,f406,f405,f409,f417,f430,f433,f435,f438,f441,f443,f452,f457,f458,f461,f464,f470,f474,f475,f478,f480,f484,f486,f489,f491,f510,f508,f512,f522,f520,f523,f525,f533,f583,f633,f632,f631,f630,f643,f642,f641,f640,f657,f656,f654,f653,f669,f667,f666,f664,f704,f723,f731,f744,f752,f756,f762,f771,f776,f775,f787,f795,f799,f826,f829,f836,f849,f856,f855,f867,f875,f879,f890,f916,f917,f927,f943,f941,f939,f938,f935,f933,f932,f950,f986,f984,f982,f979,f972,f971,f1029,f1027,f1025,f1010,f1089,f1101,f1142,f1312,f1387,f1438,f1747,f1751,f1754,f1758,f1761,f1764,f1767,f1770,f1773,f1776,f1779,f1782,f1785,f1787,f1788,f1789,f1790,f1791,f1792,f1793,f1794,f1795,f1796,f1797,f1798,f1799,f1800,f1801,f1802,f1803,f1804,f1805,f1516,f1817,f1819,f1821,f1824,f1826,f1828,f1830,f1832,f1834,f1836,f1838,f1840,f1842,f1125,f1810,f1514,f1515,f1110,f1511,f1083,f1385,f1386,f1467,f1466,f1382,f1223,f1109,f1383,f1512,f245,f244,f1808,f1809,f241,f246,f1806,f1811,f242,f1044,f243,f1807,f1384,f1513,f1274,f1111,f1108,f1084,f1079,f1078,f1067,f1066,f1054,f1042,f414,f369,f217,f216,f2307])).
fof(f2307,plain,(
  ( ! [X9] : (~p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X9)) )),
  inference(subsumption_resolution,[],[f2240,f206])).
fof(f2240,plain,(
  ( ! [X9] : (~p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f216,f206])).
fof(f216,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.pred_set.SUBSET_BIGUNION_I','thm.bool.EQ_CLAUSES|split|2')).
fof(f217,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f134])).
fof(f369,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f202,f199])).
fof(f202,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f16])).
fof(f16,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.pred_set.SUBSET_BIGUNION_I','thm.bool.NOT_CLAUSES|split|2')).
fof(f414,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))))),
  inference(unit_resulting_resolution,[],[f198,f217])).
fof(f1042,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,sK13(X0,X1,bool))) )),
  inference(superposition,[],[f242,f206])).
fof(f1054,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK13(X0,X1,bool))) )),
  inference(superposition,[],[f243,f206])).
fof(f1066,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,X2,bool))) )),
  inference(superposition,[],[f278,f206])).
fof(f1067,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK15(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK15(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f278,f206])).
fof(f1078,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X1)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),sK14(X1,X0,bool))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f244,f206])).
fof(f1079,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X3)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),sK14(X3,'F_0',bool))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f244,f206])).
fof(f1084,plain,(
  ( ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),sK14(X5,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),bool)))))) )),
  inference(forward_demodulation,[],[f1081,f414])).
fof(f1081,plain,(
  ( ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(bool,bool),sK14(X5,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),bool)))))) )),
  inference(superposition,[],[f244,f414])).
fof(f1108,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,bool),sK14(X1,X0,bool)),s(fun(fun(bool,bool),bool),X1)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f245,f206])).
fof(f1111,plain,(
  ( ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,bool),sK14(X5,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),bool)),s(fun(fun(bool,bool),bool),X5))))) )),
  inference(superposition,[],[f245,f414])).
fof(f1274,plain,(
  ( ! [X66,X67] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),sK14(X67,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),bool))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X66),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X67)))))) | s(bool,'T_0') = s(bool,X66)) )),
  inference(forward_demodulation,[],[f1224,f414])).
fof(f1224,plain,(
  ( ! [X66,X67] : (~p(s(bool,'const.bool.IN_2'(s(bool,X66),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X67)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12))))),s(fun(bool,bool),sK14(X67,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),bool))))) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f244,f1142])).
fof(f1513,plain,(
  ( ! [X140,X138,X139,X137] : (s(bool,'T_0') != s(bool,f(s(X137,sK15(X138,X139,bool,X137)))) | s(fun(X137,bool),X138) = s(fun(X137,bool),X139) | s(bool,X140) = s(bool,g(s(X137,sK15(X138,X139,bool,X137)))) | s(bool,'T_0') = s(bool,X140)) )),
  inference(superposition,[],[f278,f533])).
fof(f1384,plain,(
  ( ! [X144,X142,X143,X141] : (s(bool,X144) != s(bool,f(s(X141,sK15(X142,X143,bool,X141)))) | s(fun(X141,bool),X142) = s(fun(X141,bool),X143) | s(bool,'T_0') = s(bool,X144) | s(bool,'T_0') = s(bool,g(s(X141,sK15(X142,X143,bool,X141))))) )),
  inference(superposition,[],[f278,f533])).
fof(f1807,plain,(
  ( ! [X107,X105,X108,X106] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X105,bool),X107),s(fun(X105,bool),X106)))) | s(bool,X108) = s(bool,'const.bool.IN_2'(s(X105,sK13(X106,X107,X105)),s(fun(X105,bool),X106))) | s(bool,'T_0') = s(bool,X108)) )),
  inference(subsumption_resolution,[],[f1503,f199])).
fof(f1503,plain,(
  ( ! [X107,X105,X108,X106] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X105,bool),X107),s(fun(X105,bool),X106)))) | s(bool,X108) = s(bool,'const.bool.IN_2'(s(X105,sK13(X106,X107,X105)),s(fun(X105,bool),X106))) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f243,f533])).
fof(f243,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK13(X2,X1,X0)),s(fun(X0,bool),X2)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK13(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK13(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f149])).
fof(f149,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X16,X11] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X16),s(fun(X6,bool),X11)))) <=> ! [X7] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X16)))) => p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X11))))))),
  file('thm.pred_set.SUBSET_BIGUNION_I','thm.pred_set.SUBSET_DEF')).
fof(f1044,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK13(X1,X2,X0)),s(fun(X0,bool),X2)))) )),
  inference(subsumption_resolution,[],[f1043,f369])).
fof(f1043,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK13(X1,X2,X0)),s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f242,f206])).
fof(f242,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK13(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f1811,plain,(
  ( ! [X127,X125,X123,X126,X124] : (~p(s(bool,'const.bool.IN_2'(s(X123,X126),s(fun(X123,bool),X124)))) | p(s(bool,'const.bool.IN_2'(s(X123,X126),s(fun(X123,bool),X125)))) | s(bool,X127) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X123,bool),X124),s(fun(X123,bool),X125))) | s(bool,'T_0') = s(bool,X127)) )),
  inference(subsumption_resolution,[],[f1510,f199])).
fof(f1510,plain,(
  ( ! [X127,X125,X123,X126,X124] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X123,X126),s(fun(X123,bool),X124)))) | p(s(bool,'const.bool.IN_2'(s(X123,X126),s(fun(X123,bool),X125)))) | s(bool,X127) = s(bool,'const.pred_set.SUBSET_2'(s(fun(X123,bool),X124),s(fun(X123,bool),X125))) | s(bool,'T_0') = s(bool,X127)) )),
  inference(superposition,[],[f241,f533])).
fof(f1806,plain,(
  ( ! [X99,X97,X100,X98,X96] : (~p(s(bool,'const.bool.IN_2'(s(X96,X99),s(fun(X96,bool),X97)))) | p(s(bool,'const.bool.IN_2'(s(X96,X99),s(fun(X96,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X96,bool),bool),X98)))))) | s(bool,X100) = s(bool,'const.bool.IN_2'(s(fun(X96,bool),X97),s(fun(fun(X96,bool),bool),X98))) | s(bool,'T_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f1501,f199])).
fof(f1501,plain,(
  ( ! [X99,X97,X100,X98,X96] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X96,X99),s(fun(X96,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X96,bool),bool),X98)))))) | ~p(s(bool,'const.bool.IN_2'(s(X96,X99),s(fun(X96,bool),X97)))) | s(bool,X100) = s(bool,'const.bool.IN_2'(s(fun(X96,bool),X97),s(fun(fun(X96,bool),bool),X98))) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f246,f533])).
fof(f246,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),sK14(X2,X1,X0))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK14(X2,X1,X0)),s(fun(fun(X0,bool),bool),X2)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14])],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X4)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),X2)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))))))))),
  inference(rectify,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2))))))))),
  inference(nnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2))))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),X2))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X7,X17] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),X17)))))) <=> ? [X16] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X16)))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X16),s(fun(fun(X6,bool),bool),X17))))))),
  file('thm.pred_set.SUBSET_BIGUNION_I','thm.pred_set.IN_BIGUNION')).
fof(f241,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f1809,plain,(
  ( ! [X116,X114,X115,X113] : (p(s(bool,'const.bool.IN_2'(s(X113,X114),s(fun(X113,bool),sK14(X115,X114,X113))))) | s(bool,X116) = s(bool,'const.bool.IN_2'(s(X113,X114),s(fun(X113,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X113,bool),bool),X115))))) | s(bool,'T_0') = s(bool,X116)) )),
  inference(subsumption_resolution,[],[f1505,f199])).
fof(f1505,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X113,X114),s(fun(X113,bool),sK14(X115,X114,X113))))) | s(bool,X116) = s(bool,'const.bool.IN_2'(s(X113,X114),s(fun(X113,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X113,bool),bool),X115))))) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f244,f533])).
fof(f1808,plain,(
  ( ! [X111,X109,X112,X110] : (p(s(bool,'const.bool.IN_2'(s(fun(X109,bool),sK14(X111,X110,X109)),s(fun(fun(X109,bool),bool),X111)))) | s(bool,X112) = s(bool,'const.bool.IN_2'(s(X109,X110),s(fun(X109,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X109,bool),bool),X111))))) | s(bool,'T_0') = s(bool,X112)) )),
  inference(subsumption_resolution,[],[f1504,f199])).
fof(f1504,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(X109,bool),sK14(X111,X110,X109)),s(fun(fun(X109,bool),bool),X111)))) | s(bool,X112) = s(bool,'const.bool.IN_2'(s(X109,X110),s(fun(X109,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X109,bool),bool),X111))))) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f245,f533])).
fof(f244,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),sK14(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f154])).
fof(f245,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK14(X2,X1,X0)),s(fun(fun(X0,bool),bool),X2))))) )),
  inference(cnf_transformation,[],[f154])).
fof(f1512,plain,(
  ( ! [X132,X133,X131] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X132)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X131),s(fun(bool,bool),sK14(X132,X131,bool))))) | s(bool,X131) = s(bool,X133) | s(bool,'T_0') = s(bool,X133)) )),
  inference(superposition,[],[f244,f533])).
fof(f1383,plain,(
  ( ! [X134,X136,X135] : (~p(s(bool,'const.bool.IN_2'(s(bool,X135),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X136)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X134),s(fun(bool,bool),sK14(X136,X134,bool))))) | s(bool,'T_0') = s(bool,X135) | s(bool,'T_0') = s(bool,X134)) )),
  inference(superposition,[],[f244,f533])).
fof(f1109,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X3)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,bool),sK14(X3,'F_0',bool)),s(fun(fun(bool,bool),bool),X3)))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f245,f206])).
fof(f1223,plain,(
  ( ! [X64,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,X64),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X65)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,bool),sK14(X65,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),bool)),s(fun(fun(bool,bool),bool),X65)))) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f245,f1142])).
fof(f1382,plain,(
  ( ! [X132,X133,X131] : (~p(s(bool,'const.bool.IN_2'(s(bool,X132),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X133)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,bool),sK14(X133,X131,bool)),s(fun(fun(bool,bool),bool),X133)))) | s(bool,'T_0') = s(bool,X132) | s(bool,'T_0') = s(bool,X131)) )),
  inference(superposition,[],[f245,f533])).
fof(f1466,plain,(
  ( ! [X94,X95,X93] : (~p(s(bool,'const.bool.IN_2'(s(bool,X93),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X94)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,bool),sK14(X94,'T_0',bool)),s(fun(fun(bool,bool),bool),X94)))) | s(bool,X93) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f245,f533])).
fof(f1467,plain,(
  ( ! [X97,X98,X96] : (~p(s(bool,'const.bool.IN_2'(s(bool,X96),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X97)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),sK14(X97,'T_0',bool))))) | s(bool,X96) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f244,f533])).
fof(f1386,plain,(
  ( ! [X149,X150,X148] : (p(s(bool,'const.bool.IN_2'(s(bool,X150),s(fun(bool,bool),X149)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X149),s(fun(bool,bool),X148)))) | s(bool,'T_0') = s(bool,X150) | s(bool,'T_0') = s(bool,sK13(X148,X149,bool))) )),
  inference(superposition,[],[f242,f533])).
fof(f1385,plain,(
  ( ! [X146,X147,X145] : (~p(s(bool,'const.bool.IN_2'(s(bool,X147),s(fun(bool,bool),X145)))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X146),s(fun(bool,bool),X145)))) | s(bool,'T_0') = s(bool,X147) | s(bool,'T_0') = s(bool,sK13(X145,X146,bool))) )),
  inference(superposition,[],[f243,f533])).
fof(f1083,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X4)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),sK14(X4,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),bool)))))) )),
  inference(forward_demodulation,[],[f1080,f409])).
fof(f1080,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X4)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12))),s(fun(bool,bool),sK14(X4,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),bool)))))) )),
  inference(superposition,[],[f244,f409])).
fof(f1511,plain,(
  ( ! [X130,X128,X129] : (p(s(bool,'const.bool.IN_2'(s(fun(bool,bool),sK14(X129,X128,bool)),s(fun(fun(bool,bool),bool),X129)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X129)))))) | s(bool,X128) = s(bool,X130) | s(bool,'T_0') = s(bool,X130)) )),
  inference(superposition,[],[f245,f533])).
fof(f1110,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(bool,bool),bool),X4)))))) | p(s(bool,'const.bool.IN_2'(s(fun(bool,bool),sK14(X4,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),bool)),s(fun(fun(bool,bool),bool),X4))))) )),
  inference(superposition,[],[f245,f409])).
fof(f1515,plain,(
  ( ! [X146,X144,X145] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X145),s(fun(bool,bool),X144)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X145)))) | s(bool,X146) = s(bool,sK13(X144,X145,bool)) | s(bool,'T_0') = s(bool,X146)) )),
  inference(superposition,[],[f242,f533])).
fof(f1514,plain,(
  ( ! [X142,X143,X141] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(bool,bool),X142),s(fun(bool,bool),X141)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X141)))) | s(bool,X143) = s(bool,sK13(X141,X142,bool)) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f243,f533])).
fof(f1810,plain,(
  ( ! [X118,X117] : (p(s(bool,'const.bool.IN_2'(s(sK10,X117),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))))) | s(bool,X118) = s(bool,'const.bool.IN_2'(s(sK10,X117),s(fun(sK10,bool),sK11))) | s(bool,'T_0') = s(bool,X118)) )),
  inference(subsumption_resolution,[],[f1506,f199])).
fof(f1506,plain,(
  ( ! [X118,X117] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK10,X117),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))))) | s(bool,X118) = s(bool,'const.bool.IN_2'(s(sK10,X117),s(fun(sK10,bool),sK11))) | s(bool,'T_0') = s(bool,X118)) )),
  inference(superposition,[],[f1125,f533])).
fof(f1125,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),sK11)))) | p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12))))))) )),
  inference(subsumption_resolution,[],[f1123,f199])).
fof(f1123,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),sK11))))) )),
  inference(superposition,[],[f246,f409])).
fof(f1842,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1841])).
fof(f1841,plain,(
  ( ! [X80,X81,X79,X82] : (~sP8('T_0',X80,X81) | sP9('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1541,f583])).
fof(f1541,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~sP8('T_0',X80,X81) | sP9('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f362,f533])).
fof(f1840,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1839])).
fof(f1839,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f1539,f583])).
fof(f1539,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f339,f533])).
fof(f1838,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1837])).
fof(f1837,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f1538,f583])).
fof(f1538,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | ~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f338,f533])).
fof(f1836,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1835])).
fof(f1835,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f1537,f583])).
fof(f1537,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f337,f533])).
fof(f1834,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f1833])).
fof(f1833,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f1536,f583])).
fof(f1536,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f332,f533])).
fof(f1832,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f1831])).
fof(f1831,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f1535,f583])).
fof(f1535,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f324,f533])).
fof(f1830,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1829])).
fof(f1829,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f1534,f583])).
fof(f1534,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f315,f533])).
fof(f1828,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1827])).
fof(f1827,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f1533,f583])).
fof(f1533,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f296,f533])).
fof(f1826,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1825])).
fof(f1825,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f1532,f583])).
fof(f1532,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f292,f533])).
fof(f1824,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f1823,f387])).
fof(f1823,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f1822])).
fof(f1822,plain,(
  ( ! [X43,X41,X42,X40] : (sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f1531,f583])).
fof(f1531,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f289,f533])).
fof(f1821,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f1820])).
fof(f1820,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f1530,f583])).
fof(f1530,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f288,f533])).
fof(f1819,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1818])).
fof(f1818,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f1529,f583])).
fof(f1529,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f273,f533])).
fof(f1817,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1816])).
fof(f1816,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f1528,f583])).
fof(f1528,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f270,f533])).
fof(f1516,plain,(
  ( ! [X149,X147,X150,X148] : (s(X149,f(s(bool,'T_0'))) != s(X149,g(s(bool,'T_0'))) | s(fun(bool,X149),X147) = s(fun(bool,X149),X148) | s(bool,X150) = s(bool,sK15(X147,X148,X149,bool)) | s(bool,'T_0') = s(bool,X150)) )),
  inference(superposition,[],[f278,f533])).
fof(f1805,plain,(
  ( ! [X83,X81,X82] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X81,X82) | s(bool,X81) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(subsumption_resolution,[],[f1493,f199])).
fof(f1493,plain,(
  ( ! [X83,X81,X82] : (~p(s(bool,'T_0')) | sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X81,X82) | s(bool,X81) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f491,f533])).
fof(f1804,plain,(
  ( ! [X80,X78,X79] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X79,X78) | s(bool,X78) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(subsumption_resolution,[],[f1492,f199])).
fof(f1492,plain,(
  ( ! [X80,X78,X79] : (~p(s(bool,'T_0')) | sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X79,X78) | s(bool,X78) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f480,f533])).
fof(f1803,plain,(
  ( ! [X76,X74,X77,X75] : (sP9(X74,X75,X76) | ~sP8(X74,X75,X76) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f1491,f199])).
fof(f1491,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,'T_0')) | ~sP8(X74,X75,X76) | sP9(X74,X75,X76) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f362,f533])).
fof(f1802,plain,(
  ( ! [X70,X68,X69,X67] : (~sP9(X67,X69,X68) | p(s(bool,X69)) | p(s(bool,X68)) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1489,f199])).
fof(f1489,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | p(s(bool,X69)) | ~sP9(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f339,f533])).
fof(f1801,plain,(
  ( ! [X66,X64,X65,X63] : (sP9(X63,X64,X65) | ~p(s(bool,X64)) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1488,f199])).
fof(f1488,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X64)) | sP9(X63,X64,X65) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f338,f533])).
fof(f1800,plain,(
  ( ! [X61,X59,X62,X60] : (sP9(X59,X61,X60) | ~p(s(bool,X60)) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1487,f199])).
fof(f1487,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP9(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f337,f533])).
fof(f1799,plain,(
  ( ! [X57,X58,X56,X55] : (~sP8(X55,X56,X57) | p(s(bool,X57)) | p(s(bool,X56)) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1486,f199])).
fof(f1486,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP8(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f332,f533])).
fof(f1798,plain,(
  ( ! [X54,X52,X53,X51] : (~sP7(X53,X52,X51) | p(s(bool,X53)) | ~p(s(bool,X52)) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1485,f199])).
fof(f1485,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP7(X53,X52,X51) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f324,f533])).
fof(f1797,plain,(
  ( ! [X47,X50,X48,X49] : (~sP6(X48,X47,X49) | p(s(bool,X49)) | ~p(s(bool,X48)) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1484,f199])).
fof(f1484,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X48,X47,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f315,f533])).
fof(f1796,plain,(
  ( ! [X45,X43,X46,X44] : (~sP3(X44,X43,X45) | p(s(bool,X45)) | ~p(s(bool,X44)) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1483,f199])).
fof(f1483,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f296,f533])).
fof(f1795,plain,(
  ( ! [X39,X41,X42,X40] : (~sP4(X40,X39,X41) | p(s(bool,X41)) | ~p(s(bool,X40)) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1482,f199])).
fof(f1482,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f292,f533])).
fof(f1794,plain,(
  ( ! [X37,X35,X38,X36] : (~sP3(X37,X35,X36) | ~sP4(X36,X35,X37) | sP5(X36,X37,X35) | ~sP2(X35,X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1481,f199])).
fof(f1481,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | sP5(X36,X37,X35) | ~sP4(X36,X35,X37) | ~sP3(X37,X35,X36) | ~sP2(X35,X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f289,f533])).
fof(f1793,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1480,f199])).
fof(f1480,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f288,f533])).
fof(f1792,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f1479,f199])).
fof(f1479,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f273,f533])).
fof(f1791,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f1478,f199])).
fof(f1478,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f270,f533])).
fof(f1790,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1477,f199])).
fof(f1477,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f255,f533])).
fof(f1789,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1476,f199])).
fof(f1476,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f254,f533])).
fof(f1788,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1475,f199])).
fof(f1475,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f250,f533])).
fof(f1787,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1471,f199])).
fof(f1471,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f216,f533])).
fof(f1785,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1784])).
fof(f1784,plain,(
  ( ! [X80,X81,X79,X82] : (~sP8('T_0',X80,X81) | sP9('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1462,f583])).
fof(f1462,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~sP8('T_0',X80,X81) | sP9('T_0',X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f362,f533])).
fof(f1782,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1781])).
fof(f1781,plain,(
  ( ! [X74,X72,X75,X73] : (p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f1460,f583])).
fof(f1460,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | p(s(bool,X74)) | ~sP9('T_0',X74,X73) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f339,f533])).
fof(f1779,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f1778])).
fof(f1778,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(subsumption_resolution,[],[f1459,f583])).
fof(f1459,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,X68)) | ~p(s(bool,X69)) | sP9('T_0',X69,X70) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f338,f533])).
fof(f1776,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f1775])).
fof(f1775,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f1458,f583])).
fof(f1458,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,X64)) | ~p(s(bool,X65)) | sP9('T_0',X66,X65) | s(bool,X64) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f337,f533])).
fof(f1773,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f1772])).
fof(f1772,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f1457,f583])).
fof(f1457,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,X60)) | p(s(bool,X61)) | p(s(bool,X62)) | ~sP8('T_0',X61,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f332,f533])).
fof(f1770,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f1769])).
fof(f1769,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(subsumption_resolution,[],[f1456,f583])).
fof(f1456,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X56)) | ~p(s(bool,X57)) | p(s(bool,X58)) | ~sP7(X58,X57,'T_0') | s(bool,X56) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f324,f533])).
fof(f1767,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f1766])).
fof(f1766,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(subsumption_resolution,[],[f1455,f583])).
fof(f1455,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP6(X53,'T_0',X54) | s(bool,X52) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) )),
  inference(superposition,[],[f315,f533])).
fof(f1764,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f1763])).
fof(f1763,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f1454,f583])).
fof(f1454,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP3(X49,'T_0',X50) | s(bool,X48) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f296,f533])).
fof(f1761,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f1760])).
fof(f1760,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f1453,f583])).
fof(f1453,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'T_0',X46) | s(bool,X44) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f292,f533])).
fof(f1758,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f1757,f387])).
fof(f1757,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f1756])).
fof(f1756,plain,(
  ( ! [X43,X41,X42,X40] : (sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f1452,f583])).
fof(f1452,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X40)) | sP5(X41,X42,'T_0') | ~sP4(X41,'T_0',X42) | ~sP3(X42,'T_0',X41) | ~sP2('T_0',X42,X41) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f289,f533])).
fof(f1754,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f1753])).
fof(f1753,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f1451,f583])).
fof(f1451,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X36)) | ~p(s(bool,X37)) | p(s(bool,X38)) | ~sP5(X37,'T_0',X38) | s(bool,X36) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f288,f533])).
fof(f1751,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f1750])).
fof(f1750,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(subsumption_resolution,[],[f1450,f583])).
fof(f1450,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | sP1(X34,'T_0') | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f273,f533])).
fof(f1747,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f1746])).
fof(f1746,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f1449,f583])).
fof(f1449,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | sP1('T_0',X31) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f270,f533])).
fof(f1438,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f212,f533])).
fof(f1387,plain,(
  ( ! [X154,X152,X151,X153] : (s(X153,f(s(bool,X154))) != s(X153,g(s(bool,X154))) | s(fun(bool,X153),X151) = s(fun(bool,X153),X152) | s(bool,'T_0') = s(bool,X154) | s(bool,'T_0') = s(bool,sK15(X151,X152,X153,bool))) )),
  inference(superposition,[],[f278,f533])).
fof(f1312,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f533,f533])).
fof(f1142,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f1140,f369])).
fof(f1140,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f583,f414])).
fof(f1101,plain,(
  ~sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1089,f327])).
fof(f1089,plain,(
  ~sP8('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f369,f369,f875])).
fof(f1010,plain,(
  sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))),
  inference(unit_resulting_resolution,[],[f388,f401,f378,f369,f291])).
fof(f1025,plain,(
  sP5('F_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f386,f381,f377,f369,f291])).
fof(f1027,plain,(
  ~sP2('F_0','const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),'F_0')),
  inference(unit_resulting_resolution,[],[f401,f666,f377,f369,f291])).
fof(f1029,plain,(
  sP5('F_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))))),
  inference(unit_resulting_resolution,[],[f386,f400,f398,f369,f291])).
fof(f971,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f383,f379,f377,f369,f290])).
fof(f972,plain,(
  sP5('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f384,f380,f377,f369,f290])).
fof(f979,plain,(
  sP5('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'F_0','const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))))),
  inference(unit_resulting_resolution,[],[f384,f380,f398,f369,f290])).
fof(f982,plain,(
  sP5('F_0','F_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))),
  inference(unit_resulting_resolution,[],[f388,f382,f378,f369,f290])).
fof(f984,plain,(
  ~sP2('const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f382,f669,f378,f369,f290])).
fof(f986,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),'F_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))),
  inference(unit_resulting_resolution,[],[f388,f382,f399,f369,f290])).
fof(f950,plain,(
  ( ! [X6,X7] : (~sP4(X6,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X7) | sP5(X6,X7,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12))) | ~sP3(X7,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X6)) )),
  inference(subsumption_resolution,[],[f949,f388])).
fof(f949,plain,(
  ( ! [X6,X7] : (sP5(X6,X7,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12))) | ~sP4(X6,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X7) | ~sP3(X7,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X6) | ~sP2('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X7,X6)) )),
  inference(subsumption_resolution,[],[f946,f199])).
fof(f946,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | sP5(X6,X7,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12))) | ~sP4(X6,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X7) | ~sP3(X7,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X6) | ~sP2('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X7,X6)) )),
  inference(superposition,[],[f289,f409])).
fof(f932,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f387,f379,f375,f199,f289])).
fof(f933,plain,(
  sP5('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f387,f380,f375,f199,f289])).
fof(f935,plain,(
  sP5('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'T_0')),
  inference(unit_resulting_resolution,[],[f387,f380,f376,f199,f289])).
fof(f938,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f387,f382,f399,f199,f289])).
fof(f939,plain,(
  sP5('const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),'T_0')),
  inference(unit_resulting_resolution,[],[f387,f401,f399,f199,f289])).
fof(f941,plain,(
  ~sP4('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))))),
  inference(unit_resulting_resolution,[],[f387,f401,f654,f199,f289])).
fof(f943,plain,(
  ~sP3('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f387,f378,f664,f199,f289])).
fof(f927,plain,(
  ~sP7('F_0','T_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))),
  inference(unit_resulting_resolution,[],[f917,f308])).
fof(f917,plain,(
  ~sP6('T_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f199,f369,f795])).
fof(f916,plain,(
  ( ! [X6,X7] : (~sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X7,X6) | p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f911,f199])).
fof(f911,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | p(s(bool,X7)) | ~sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X7,X6)) )),
  inference(superposition,[],[f339,f409])).
fof(f890,plain,(
  ~sP3('T_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f199,f369,f752])).
fof(f879,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f867,f327])).
fof(f875,plain,(
  ( ! [X6,X7] : (~sP8('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X6,X7) | p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f870,f199])).
fof(f870,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | p(s(bool,X7)) | ~sP8('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X6,X7)) )),
  inference(superposition,[],[f332,f409])).
fof(f867,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f369,f369,f199,f332])).
fof(f855,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f397,f369,f369,f331])).
fof(f856,plain,(
  sP9('const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f404,f369,f369,f331])).
fof(f849,plain,(
  ~sP5('T_0','F_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))),
  inference(unit_resulting_resolution,[],[f836,f287])).
fof(f836,plain,(
  ~sP4('T_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f199,f369,f731])).
fof(f829,plain,(
  ~sP5('T_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f199,f369,f704])).
fof(f826,plain,(
  ( ! [X6,X7] : (~sP7(X7,X6,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f821,f199])).
fof(f821,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP7(X7,X6,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(superposition,[],[f324,f409])).
fof(f799,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f787,f308])).
fof(f795,plain,(
  ( ! [X6,X7] : (~sP6(X6,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f790,f199])).
fof(f790,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP6(X6,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X7)) )),
  inference(superposition,[],[f315,f409])).
fof(f787,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f369,f199,f199,f315])).
fof(f775,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f389,f199,f199,f311])).
fof(f776,plain,(
  sP7('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f390,f199,f199,f311])).
fof(f771,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f762,f285])).
fof(f762,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f369,f369,f369,f300])).
fof(f756,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f744,f286])).
fof(f752,plain,(
  ( ! [X6,X7] : (~sP3(X6,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f747,f199])).
fof(f747,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP3(X6,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X7)) )),
  inference(superposition,[],[f296,f409])).
fof(f744,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f369,f199,f199,f296])).
fof(f731,plain,(
  ( ! [X6,X7] : (~sP4(X6,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f726,f199])).
fof(f726,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP4(X6,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X7)) )),
  inference(superposition,[],[f292,f409])).
fof(f723,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f369,f199,f199,f292])).
fof(f704,plain,(
  ( ! [X6,X7] : (~sP5(X6,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f701,f199])).
fof(f701,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP5(X6,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X7)) )),
  inference(superposition,[],[f288,f409])).
fof(f664,plain,(
  ~sP5('F_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'T_0')),
  inference(unit_resulting_resolution,[],[f369,f409,f307])).
fof(f666,plain,(
  ~sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),'F_0')),
  inference(unit_resulting_resolution,[],[f369,f414,f307])).
fof(f667,plain,(
  ~sP5('F_0','T_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))),
  inference(unit_resulting_resolution,[],[f369,f409,f307])).
fof(f669,plain,(
  ~sP5('F_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))))),
  inference(unit_resulting_resolution,[],[f369,f414,f307])).
fof(f653,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f199,f405,f306])).
fof(f654,plain,(
  ~sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),'T_0')),
  inference(unit_resulting_resolution,[],[f199,f406,f306])).
fof(f656,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f199,f405,f306])).
fof(f657,plain,(
  ~sP5('T_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))))),
  inference(unit_resulting_resolution,[],[f199,f406,f306])).
fof(f640,plain,(
  sP5('T_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'T_0')),
  inference(unit_resulting_resolution,[],[f199,f409,f305])).
fof(f641,plain,(
  sP5('T_0','const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),'F_0')),
  inference(unit_resulting_resolution,[],[f199,f414,f305])).
fof(f642,plain,(
  sP5('T_0','T_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))),
  inference(unit_resulting_resolution,[],[f199,f409,f305])).
fof(f643,plain,(
  sP5('T_0','F_0','const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))))),
  inference(unit_resulting_resolution,[],[f199,f414,f305])).
fof(f630,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f369,f405,f304])).
fof(f631,plain,(
  sP5('F_0','const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),'T_0')),
  inference(unit_resulting_resolution,[],[f369,f406,f304])).
fof(f632,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f369,f405,f304])).
fof(f633,plain,(
  sP5('F_0','T_0','const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))))),
  inference(unit_resulting_resolution,[],[f369,f406,f304])).
fof(f583,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f579,f369])).
fof(f579,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f247,f206])).
fof(f533,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f206,f206])).
fof(f525,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0)) )),
  inference(subsumption_resolution,[],[f524,f199])).
fof(f524,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0)) )),
  inference(superposition,[],[f491,f409])).
fof(f523,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f199,f491])).
fof(f520,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))))) )),
  inference(unit_resulting_resolution,[],[f402,f369,f363])).
fof(f522,plain,(
  ( ! [X0] : (sP7('F_0','const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),X0)) )),
  inference(unit_resulting_resolution,[],[f403,f369,f363])).
fof(f512,plain,(
  ( ! [X0,X1] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0,X1) | ~sP8('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0,X1)) )),
  inference(subsumption_resolution,[],[f511,f199])).
fof(f511,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP8('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0,X1) | sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0,X1)) )),
  inference(superposition,[],[f362,f409])).
fof(f508,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(unit_resulting_resolution,[],[f394,f199,f362])).
fof(f510,plain,(
  ( ! [X0] : (sP9('T_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0)) )),
  inference(unit_resulting_resolution,[],[f396,f199,f362])).
fof(f491,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0,X1)) )),
  inference(subsumption_resolution,[],[f490,f199])).
fof(f490,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0,X1)) )),
  inference(superposition,[],[f338,f409])).
fof(f489,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f199,f199,f338])).
fof(f486,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(subsumption_resolution,[],[f485,f199])).
fof(f485,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(superposition,[],[f480,f409])).
fof(f484,plain,(
  ( ! [X0] : (sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f199,f480])).
fof(f480,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X1,X0)) )),
  inference(subsumption_resolution,[],[f479,f199])).
fof(f479,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X1,X0)) )),
  inference(superposition,[],[f337,f409])).
fof(f478,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f199,f199,f337])).
fof(f475,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f199,f369,f326])).
fof(f474,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f199,f369,f325])).
fof(f470,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f369,f369,f320])).
fof(f464,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f369,f369,f319])).
fof(f461,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f199,f369,f310])).
fof(f458,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f199,f369,f309])).
fof(f457,plain,(
  sP1('F_0','const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))),
  inference(unit_resulting_resolution,[],[f369,f443])).
fof(f452,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f369,f369,f370])).
fof(f443,plain,(
  ( ! [X0] : (sP1(X0,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f442,f199])).
fof(f442,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(superposition,[],[f273,f409])).
fof(f441,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f369,f199,f273])).
fof(f438,plain,(
  sP1('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f369,f435])).
fof(f435,plain,(
  ( ! [X0] : (sP1('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f434,f199])).
fof(f434,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0)) )),
  inference(superposition,[],[f270,f409])).
fof(f433,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f369,f199,f270])).
fof(f430,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f199,f199,f269])).
fof(f417,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f369,f369,f268])).
fof(f409,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))),
  inference(unit_resulting_resolution,[],[f197,f213])).
fof(f405,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f369,f212])).
fof(f406,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))))),
  inference(unit_resulting_resolution,[],[f198,f212])).
fof(f398,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),X1)) )),
  inference(unit_resulting_resolution,[],[f198,f294])).
fof(f399,plain,(
  ( ! [X0,X1] : (sP4('const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f198,f295])).
fof(f400,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),X1)) )),
  inference(unit_resulting_resolution,[],[f198,f298])).
fof(f401,plain,(
  ( ! [X0,X1] : (sP3('const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f198,f299])).
fof(f402,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),X1)) )),
  inference(unit_resulting_resolution,[],[f198,f317])).
fof(f403,plain,(
  ( ! [X0,X1] : (sP6('const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f198,f318])).
fof(f404,plain,(
  ( ! [X0,X1] : (sP8('const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)))),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f198,f335])).
fof(f397,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f369,f335])).
fof(f395,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f199,f334])).
fof(f396,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X1)) )),
  inference(unit_resulting_resolution,[],[f197,f334])).
fof(f393,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f199,f333])).
fof(f394,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(unit_resulting_resolution,[],[f197,f333])).
fof(f392,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f369,f318])).
fof(f391,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f369,f317])).
fof(f389,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f199,f316])).
fof(f390,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(unit_resulting_resolution,[],[f197,f316])).
fof(f387,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f199,f303])).
fof(f388,plain,(
  ( ! [X0,X1] : (sP2('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f197,f303])).
fof(f385,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f199,f302])).
fof(f386,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X1)) )),
  inference(unit_resulting_resolution,[],[f197,f302])).
fof(f383,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f199,f301])).
fof(f384,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(unit_resulting_resolution,[],[f197,f301])).
fof(f382,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f369,f299])).
fof(f381,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f369,f298])).
fof(f379,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f199,f297])).
fof(f380,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(unit_resulting_resolution,[],[f197,f297])).
fof(f378,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f369,f295])).
fof(f377,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f369,f294])).
fof(f375,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f199,f293])).
fof(f376,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(unit_resulting_resolution,[],[f197,f293])).
fof(f373,plain,(
  ( ! [X0] : (sP0('const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)),X0)) )),
  inference(unit_resulting_resolution,[],[f197,f255])).
fof(f374,plain,(
  ( ! [X0] : (sP0(X0,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))) )),
  inference(unit_resulting_resolution,[],[f197,f254])).
fof(f372,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f199,f255])).
fof(f371,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f199,f254])).
fof(f361,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f349])).
fof(f349,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X7,X14,X12,X15] : ((s(bool,X7) = s(bool,X14) & (p(s(bool,X14)) => s(bool,X12) = s(bool,X15))) => ((p(s(bool,X7)) => p(s(bool,X12))) <=> (p(s(bool,X14)) => p(s(bool,X15)))))),
  file('thm.pred_set.SUBSET_BIGUNION_I','thm.bool.IMP_CONG')).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f86,f127,f126])).
fof(f126,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f127,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f86,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X19,X20,X21] : ((p(s(bool,X21)) <=> (p(s(bool,X20)) | p(s(bool,X19)))) <=> ((p(s(bool,X21)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | ~p(s(bool,X19))) & (p(s(bool,X20)) | p(s(bool,X19)) | ~p(s(bool,X21)))))),
  file('thm.pred_set.SUBSET_BIGUNION_I','thm.sat.sat_25_prove_ltb')).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f194])).
fof(f338,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f194])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f194])).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f191])).
fof(f191,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f126])).
fof(f333,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f192])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f192])).
fof(f325,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f188])).
fof(f188,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f187])).
fof(f187,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f127])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f189])).
fof(f327,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f189])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f328])).
fof(f328,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f189])).
fof(f329,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f189])).
fof(f330,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f189])).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f189])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f84,f124,f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f124,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f84,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X19,X20,X21] : ((p(s(bool,X21)) <=> (p(s(bool,X20)) & p(s(bool,X19)))) <=> ((p(s(bool,X21)) | ~p(s(bool,X20)) | ~p(s(bool,X19))) & (p(s(bool,X20)) | ~p(s(bool,X21))) & (p(s(bool,X19)) | ~p(s(bool,X21)))))),
  file('thm.pred_set.SUBSET_BIGUNION_I','thm.sat.sat_23_prove_ltb')).
fof(f320,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f186])).
fof(f321,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f186])).
fof(f324,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f186])).
fof(f315,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f123])).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f184])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f184])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f184])).
fof(f308,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f181])).
fof(f181,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f124])).
fof(f309,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f182])).
fof(f310,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f182])).
fof(f311,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f182])).
fof(f363,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f314])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f182])).
fof(f304,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f82,f121,f120,f119,f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f119,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f120,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f121,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f82,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X19,X20,X21] : ((p(s(bool,X21)) <=> s(bool,X19) = s(bool,X20)) <=> ((p(s(bool,X21)) | p(s(bool,X20)) | p(s(bool,X19))) & (p(s(bool,X21)) | ~p(s(bool,X19)) | ~p(s(bool,X20))) & (p(s(bool,X20)) | ~p(s(bool,X19)) | ~p(s(bool,X21))) & (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X21)))))),
  file('thm.pred_set.SUBSET_BIGUNION_I','thm.sat.sat_2|49|_prove_ltb')).
fof(f305,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f179])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f179])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f179])).
fof(f300,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f118])).
fof(f301,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f178])).
fof(f302,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f178])).
fof(f303,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f178])).
fof(f296,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f175])).
fof(f175,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f119])).
fof(f297,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f176])).
fof(f298,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f176])).
fof(f299,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f176])).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f172])).
fof(f172,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f120])).
fof(f293,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f173])).
fof(f294,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f173])).
fof(f295,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f173])).
fof(f285,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f169])).
fof(f169,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f121])).
fof(f286,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f170])).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f170])).
fof(f288,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f170])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f170])).
fof(f290,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f170])).
fof(f291,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f170])).
fof(f278,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK15(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK15(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK15(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK15(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.SUBSET_BIGUNION_I',aHL_EXT)).
fof(f268,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f162])).
fof(f162,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f269,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f163])).
fof(f270,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f163])).
fof(f273,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f163])).
fof(f370,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f253,f369])).
fof(f253,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f254,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f156])).
fof(f255,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f156])).
fof(f247,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.pred_set.SUBSET_BIGUNION_I','thm.bool.IMP_ANTISYM_AX')).
fof(f250,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f100])).
fof(f212,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.pred_set.SUBSET_BIGUNION_I','thm.bool.EQ_CLAUSES|split|0')).
fof(f213,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f132])).
fof(f206,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.SUBSET_BIGUNION_I',aHL_BOOL_CASES)).
fof(f199,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.SUBSET_BIGUNION_I','thm.bool.TRUTH')).
fof(f197,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12))))),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK10,bool),sK11),s(fun(fun(sK10,bool),bool),sK12)))) & ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12])],[f88])).
fof(f88,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))) & ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X1),s(fun(fun(X0,bool),bool),X2)))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),X2)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X7),s(fun(fun(X6,bool),bool),X8)))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),X8)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X7),s(fun(fun(X6,bool),bool),X8)))) => p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X7),s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),X8)))))))),
  file('thm.pred_set.SUBSET_BIGUNION_I','thm.pred_set.SUBSET_BIGUNION_I')).
fof(f198,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12))))))),
  inference(cnf_transformation,[],[f129])).
fof(f22818,plain,(
  p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f22365,f414])).
fof(f22365,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK10,bool),sK11),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12))))))),
  inference(unit_resulting_resolution,[],[f3392,f243])).
fof(f3392,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK10,sK13('const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)),sK11,sK10)),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12))))))),
  inference(subsumption_resolution,[],[f3385,f199])).
fof(f3385,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK10,sK13('const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)),sK11,sK10)),s(fun(sK10,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12))))))),
  inference(superposition,[],[f1125,f2893])).
fof(f2893,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK10,sK13('const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)),sK11,sK10)),s(fun(sK10,bool),sK11)))),
  inference(subsumption_resolution,[],[f2885,f2308])).
fof(f2885,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK10,sK13('const.pred_set.BIGUNION_1'(s(fun(fun(sK10,bool),bool),sK12)),sK11,sK10)),s(fun(sK10,bool),sK11)))),
  inference(superposition,[],[f1044,f414])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 581
% Passive clauses: 1974
% Generated clauses: 22538
% Final active clauses: 429
% Final passive clauses: 707
% Input formulas: 35
% Initial clauses: 153
% 
% Duplicate literals: 2011
% Trivial inequalities: 89
% Fw subsumption resolutions: 1105
% Fw demodulations: 1591
% Bw demodulations: 353
% Condensations: 2622
% Global subsumptions: 48
% 
% Simple tautologies: 141
% Equational tautologies: 576
% Forward subsumptions: 12129
% Bw demodulations to eq. taut.: 33
% 
% Unit resulting resolution: 5429
% Factoring: 1
% Forward superposition: 5015
% Backward superposition: 3042
% Self superposition: 147
% Equality factoring: 124
% Equality resolution: 4
% 
% Split clauses: 251
% Split components: 515
% SAT solver clauses: 2567
% SAT solver unit clauses: 230
% SAT solver binary clauses: 493
% 
% Sat splits: 251
% Sat splitting refutations: 41
% 
% Memory used [KB]: 10490
% Time elapsed: 1.460 s
% ------------------------------
% ------------------------------
% Success in time 2.777 s
