Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% (88795)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 675
% Passive clauses: 10083
% Generated clauses: 18502
% Final active clauses: 668
% Final passive clauses: 9033
% Input formulas: 51
% Initial clauses: 204
% 
% Duplicate literals: 298
% Trivial inequalities: 44
% Fw subsumption resolutions: 1446
% Fw demodulations: 288
% Bw demodulations: 7
% 
% Simple tautologies: 234
% Equational tautologies: 57
% Forward subsumptions: 5816
% 
% Binary resolution: 11504
% Factoring: 1094
% Forward superposition: 1290
% Backward superposition: 2312
% Self superposition: 29
% Equality factoring: 26
% Equality resolution: 19
% 
% Split clauses: 302
% Split components: 606
% SAT solver clauses: 620
% SAT solver unit clauses: 26
% SAT solver binary clauses: 166
% 
% Sat splits: 302
% 
% TWLsolver clauses: 620
% TWLsolver calls for satisfiability: 383
% 
% Memory used [KB]: 9210
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quantHeuristics
% (88819)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 372
% Passive clauses: 3320
% Generated clauses: 14729
% Final active clauses: 270
% Final passive clauses: 1043
% Input formulas: 51
% Initial clauses: 204
% 
% Duplicate literals: 119
% Trivial inequalities: 36
% Fw subsumption resolutions: 1754
% Fw demodulations: 1192
% Bw demodulations: 1888
% Condensations: 99
% Global subsumptions: 18
% 
% Simple tautologies: 212
% Equational tautologies: 7
% Forward subsumptions: 4549
% Bw demodulations to eq. taut.: 80
% 
% Binary resolution: 2412
% Factoring: 115
% Forward superposition: 947
% Backward superposition: 1045
% Self superposition: 17
% Equality factoring: 6
% Equality resolution: 12
% 
% Split clauses: 89
% Split components: 179
% SAT solver clauses: 2652
% SAT solver unit clauses: 218
% SAT solver binary clauses: 94
% SAT solver learnt clauses: 37
% SAT solver learnt literals: 11
% 
% Sat splits: 89
% Sat splitting refutations: 72
% 
% TWLsolver clauses: 2579
% TWLsolver calls for satisfiability: 1856
% 
% Memory used [KB]: 10490
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f53106,plain,(
  $false),
  inference(subsumption_resolution,[],[f52788,f38098])).
fof(f38098,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK13,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15))),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X0))))))) )),
  inference(forward_demodulation,[],[f38044,f28740])).
fof(f28740,plain,(
  ( ! [X212,X210,X213,X211,X209] : (s(X210,happ(s(fun(sK13,X210),'_dst_x0x1x1_3'(s(fun(bool,fun(X211,X210)),X212),s(fun(sK13,bool),sK16),s(fun(sK13,X211),X213))),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X209))))) = s(X210,happ(s(fun(X211,X210),happ(s(fun(bool,fun(X211,X210)),X212),s(bool,'T_0'))),s(X211,happ(s(fun(sK13,X211),X213),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X209)))))))) )),
  inference(superposition,[],[f483,f1689])).
fof(f1689,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X0)))))) )),
  inference(unit_resulting_resolution,[],[f283,f1683,f348])).
fof(f348,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.quantHeuristics.GUESS_RULES_EQUIV|split|2','thm.bool.IMP_ANTISYM_AX')).
fof(f1683,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X0))))))) )),
  inference(subsumption_resolution,[],[f1681,f283])).
fof(f1681,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK13,bool),sK16),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X0))))))) )),
  inference(superposition,[],[f377,f553])).
fof(f553,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f280,f298])).
fof(f298,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.quantHeuristics.GUESS_RULES_EQUIV|split|2','thm.bool.EQ_CLAUSES|split|1')).
fof(f280,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK16))))),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK16)))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK15)))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16])],[f129])).
fof(f129,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X1,bool),X4),s(fun(X1,bool),X3)))))))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))))) & ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X1,bool),X4),s(fun(X1,bool),X3)))))))),
  inference(ennf_transformation,[],[f54])).
fof(f54,plain,(
  ~! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X1,bool),X4),s(fun(X1,bool),X3)))))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : ((p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X1,bool),X4),s(fun(X1,bool),X3)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : ((p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X10)))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X7,bool),X10),s(fun(X7,bool),X9)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : ((p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X10)))) & p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9))))) => p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X8),s(fun(X7,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(X7,bool),X10),s(fun(X7,bool),X9)))))))),
  file('thm.quantHeuristics.GUESS_RULES_EQUIV|split|2','thm.quantHeuristics.GUESS_RULES_EQUIV|split|2')).
fof(f377,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK18(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18])],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X7,X6] : ! [X21,X22] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X7),X21),s(fun(X7,bool),X22)))) <=> ! [X20] : p(s(bool,happ(s(fun(X7,bool),X22),s(X7,happ(s(fun(X6,X7),X21),s(X6,X20)))))))),
  file('thm.quantHeuristics.GUESS_RULES_EQUIV|split|2','thm.quantHeuristics.GUESS_REWRITES|split|2')).
fof(f283,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_EQUIV|split|2',aHL_TRUTH)).
fof(f483,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2,X3] : ! [X4,X5,X6,X7] : s(X0,happ(s(fun(X3,X0),'_dst_x0x1x1_3'(s(fun(X1,fun(X2,X0)),X4),s(fun(X3,X1),X5),s(fun(X3,X2),X6))),s(X3,X7))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X4),s(X1,happ(s(fun(X3,X1),X5),s(X3,X7))))),s(X2,happ(s(fun(X3,X2),X6),s(X3,X7)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X29,X30,X31,X32] : ! [X33,X34,X35,X36] : s(X29,happ(s(fun(X32,X29),'_dst_x0x1x1_3'(s(fun(X30,fun(X31,X29)),X33),s(fun(X32,X30),X34),s(fun(X32,X31),X35))),s(X32,X36))) = s(X29,happ(s(fun(X31,X29),happ(s(fun(X30,fun(X31,X29)),X33),s(X30,happ(s(fun(X32,X30),X34),s(X32,X36))))),s(X31,happ(s(fun(X32,X31),X35),s(X32,X36)))))),
  file('thm.quantHeuristics.GUESS_RULES_EQUIV|split|2','_dst_x0x1x1_ax')).
fof(f38044,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),happ(s(fun(bool,fun(bool,bool)),'=_0'),s(bool,'T_0'))),s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X0))))))))) )),
  inference(unit_resulting_resolution,[],[f6949,f393])).
fof(f393,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X26,X27,X28] : (p(s(bool,happ(s(fun(X26,bool),happ(s(fun(X26,fun(X26,bool)),'=_0'),s(X26,X27))),s(X26,X28)))) <=> s(X26,X27) = s(X26,X28))),
  file('thm.quantHeuristics.GUESS_RULES_EQUIV|split|2',arity_0_HL_EQ)).
fof(f6949,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X0)))))) )),
  inference(unit_resulting_resolution,[],[f283,f2329,f495])).
fof(f495,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f494])).
fof(f494,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f482])).
fof(f482,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X11,X17,X15,X18] : ((s(bool,X11) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X15) = s(bool,X18))) => ((p(s(bool,X11)) => p(s(bool,X15))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.quantHeuristics.GUESS_RULES_EQUIV|split|2','thm.bool.IMP_CONG')).
fof(f2329,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X0))))))) )),
  inference(subsumption_resolution,[],[f2327,f283])).
fof(f2327,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK13,bool),sK15),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,X0))))))) )),
  inference(superposition,[],[f391,f554])).
fof(f554,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK15)))),
  inference(unit_resulting_resolution,[],[f281,f298])).
fof(f281,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),sK15))))),
  inference(cnf_transformation,[],[f179])).
fof(f391,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27])],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X7,X6] : ! [X21,X22] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X7),X21),s(fun(X7,bool),X22)))) <=> ! [X20] : ~p(s(bool,happ(s(fun(X7,bool),X22),s(X7,happ(s(fun(X6,X7),X21),s(X6,X20)))))))),
  file('thm.quantHeuristics.GUESS_RULES_EQUIV|split|2','thm.quantHeuristics.GUESS_REWRITES|split|3')).
fof(f52788,plain,(
  p(s(bool,happ(s(fun(sK13,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15))),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK27('_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15)),sK14,sK12,sK13)))))))),
  inference(unit_resulting_resolution,[],[f6949,f6949,f12377,f3465])).
fof(f3465,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f297,f758])).
fof(f758,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f290,f290])).
fof(f290,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quantHeuristics.GUESS_RULES_EQUIV|split|2',aHL_BOOL_CASES)).
fof(f297,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f182])).
fof(f12377,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15))),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK27('_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15)),sK14,sK12,sK13))))))),
  inference(subsumption_resolution,[],[f12369,f505])).
fof(f505,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f286,f283])).
fof(f286,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f180])).
fof(f180,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f25])).
fof(f25,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESS_RULES_EQUIV|split|2','thm.bool.NOT_CLAUSES|split|2')).
fof(f12369,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15))),s(sK13,happ(s(fun(sK12,sK13),sK14),s(sK12,sK27('_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15)),sK14,sK12,sK13))))))),
  inference(superposition,[],[f2352,f912])).
fof(f912,plain,(
  s(bool,'F_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15)))))),
  inference(unit_resulting_resolution,[],[f505,f282,f345])).
fof(f345,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f142])).
fof(f282,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(sK12,sK13),sK14),s(fun(sK13,bool),'_dst_x0x1x1_3'(s(fun(bool,fun(bool,bool)),'=_0'),s(fun(sK13,bool),sK16),s(fun(sK13,bool),sK15))))))),
  inference(cnf_transformation,[],[f179])).
fof(f2352,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X2,X0),X3),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,sK27(X1,X3,X2,X0))))))) )),
  inference(subsumption_resolution,[],[f2351,f505])).
fof(f2351,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X2,X0),X3),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,sK27(X1,X3,X2,X0))))))) )),
  inference(superposition,[],[f392,f290])).
fof(f392,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f239])).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1328
% Passive clauses: 10356
% Generated clauses: 52607
% Final active clauses: 1286
% Final passive clauses: 5551
% Input formulas: 51
% Initial clauses: 204
% Discarded non-redundant clauses: 3885
% 
% Duplicate literals: 382
% Trivial inequalities: 62
% Fw subsumption resolutions: 1421
% Fw demodulations: 8609
% Bw demodulations: 221
% Condensations: 829
% Global subsumptions: 52
% 
% Simple tautologies: 146
% Equational tautologies: 648
% Forward subsumptions: 28980
% Fw demodulations to eq. taut.: 18
% Bw demodulations to eq. taut.: 15
% 
% Unit resulting resolution: 19712
% Factoring: 1
% Forward superposition: 11315
% Backward superposition: 9368
% Self superposition: 141
% Equality factoring: 30
% Equality resolution: 10
% 
% Split clauses: 85
% Split components: 175
% SAT solver clauses: 11274
% SAT solver unit clauses: 3301
% SAT solver binary clauses: 2887
% 
% Sat splits: 85
% Sat splitting refutations: 694
% 
% Memory used [KB]: 98633
% Time elapsed: 6.208 s
% ------------------------------
% ------------------------------
% Success in time 7.51 s
