Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.sum
% (101928)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 89
% Passive clauses: 3200
% Generated clauses: 18887
% Final active clauses: 86
% Final passive clauses: 3009
% Input formulas: 8
% Initial clauses: 11
% 
% Duplicate literals: 1188
% Trivial inequalities: 242
% Fw subsumption resolutions: 2335
% Fw demodulations: 895
% 
% Simple tautologies: 372
% Equational tautologies: 104
% Forward subsumptions: 7195
% 
% Binary resolution: 42
% Forward superposition: 7975
% Backward superposition: 5313
% Self superposition: 417
% Equality factoring: 305
% Equality resolution: 16
% 
% Split clauses: 923
% Split components: 1860
% SAT solver clauses: 2335
% SAT solver unit clauses: 275
% SAT solver binary clauses: 636
% 
% Sat splits: 923
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2325
% TWLsolver calls for satisfiability: 92
% 
% Memory used [KB]: 10234
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.sum
% (101939)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit

% Active clauses: 105
% Passive clauses: 739
% Generated clauses: 8643
% Final active clauses: 85
% Final passive clauses: 566
% Input formulas: 8
% Initial clauses: 11
% 
% Duplicate literals: 650
% Trivial inequalities: 89
% Fw subsumption resolutions: 348
% Fw demodulations: 607
% Bw demodulations: 73
% Condensations: 846
% 
% Simple tautologies: 88
% Equational tautologies: 41
% Forward subsumptions: 3675
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 5
% Forward superposition: 2720
% Backward superposition: 2585
% Self superposition: 214
% Equality factoring: 120
% Equality resolution: 6
% 
% Split clauses: 105
% Split components: 210
% SAT solver clauses: 1594
% SAT solver unit clauses: 18
% SAT solver binary clauses: 79
% SAT solver learnt clauses: 2
% SAT solver learnt literals: 4
% 
% Sat splits: 105
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 1592
% TWLsolver calls for satisfiability: 111
% 
% Memory used [KB]: 14328
% Time elapsed: 0.741 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.sum
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.sum
% SZS output start Proof for thm.sum
fof(f4646,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f46,f49,f47,f1648,f1651,f1650,f1579,f1656,f1580,f1581,f1582,f1583,f1584,f1585,f1659,f1657,f1586,f1587,f1588,f1589,f1653,f1660,f1590,f1658,f1661,f1591,f1662,f1654,f2575,f2531,f2513,f2608,f2505,f2504,f2487,f2456,f2602,f2455,f2400,f2603,f2620,f2609,f2394,f2610,f2393,f2613,f2388,f2623,f2601,f3643,f3646,f3633,f3656,f3856,f3857,f3829,f3864,f3862,f3883,f3884,f3887,f3888,f3889,f3894,f3933,f3936,f3934,f3938,f3941,f3939,f3946,f3948,f3973,f3976,f3977,f3981,f3984,f3987,f3990,f3988,f3994,f4010,f4012,f4014,f4016,f4037,f4038,f4053,f4054,f3654,f4201,f4203,f4202,f4204,f4206,f4205,f4212,f4214,f4213,f4215,f4217,f4216,f4238,f4239,f4345,f4346,f4347,f4348,f4349,f4350,f4351,f4352,f4353,f4354,f4355,f4356,f4357,f4476,f4369,f4370,f4371,f4372,f4373,f4374,f4375,f4376,f4377,f4378,f4379,f4380,f4381,f4478,f4388,f4389,f4390,f4391,f4392,f4393,f4479,f4394,f4395,f4396,f4397,f4398,f4399,f4400,f4401,f4402,f4480,f4403,f4481,f4482,f4468,f4405,f4483,f4422,f4485,f4484,f4423,f4487,f4486,f4428,f4429,f4430,f4431,f4432,f4433,f4488,f4434,f4435,f4474,f4436,f4437,f4438,f4439,f4467,f4475,f4440,f4469,f4472,f4441,f4473,f4442,f4489,f4471,f4443,f4490,f4491,f4477,f4445,f4492,f4470,f4462,f4494,f4493,f4463,f4496,f4495,f4497,f4283,f4267,f4093,f3892,f4529,f4530,f4601,f4537,f4618,f4623,f4241,f4094,f4626,f4634,f4637,f4639,f4532,f4640,f4643,f4645,f3859,f3644])).
fof(f3644,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK1,X0),X1),s(sK1,sK10(sK6,sK1,sK0)))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X0),X2),s(fun(sK1,X0),X1)))) ) | $spl40),
  inference(cnf_transformation,[],[f3644_D])).
fof(f3644_D,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK1,X0),X1),s(sK1,sK10(sK6,sK1,sK0)))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X0),X2),s(fun(sK1,X0),X1)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f3859,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | $spl50),
  inference(cnf_transformation,[],[f3859_D])).
fof(f3859_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4645,plain,(
  ( ! [X12,X13,X11] : (s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10(sK6,sK1,sK0)))) = s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X11),X13),s(fun(sK1,X11),X12)))) ) | ($spl42 | $spl48)),
  inference(forward_demodulation,[],[f3911,f3646])).
fof(f3911,plain,(
  ( ! [X12,X13,X11] : (s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10(sK6,sK1,sK0)))) = s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X11),X13),s(fun(sK1,X11),X12))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) | ($spl42 | $spl48)),
  inference(forward_demodulation,[],[f3904,f3857])).
fof(f3904,plain,(
  ( ! [X12,X13,X11] : (s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X11),X13),s(fun(sK1,X11),X12))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) | $spl42),
  inference(superposition,[],[f328,f3646])).
fof(f328,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X4,happ(s(fun(X2,X4),X5),s(X2,sK10(X3,X2,X1)))) = s(X4,'const.sum.sum_CASE_3'(s('type.sum.sum'(X1,X2),X3),s(fun(X1,X4),X6),s(fun(X2,X4),X5))) | s('type.sum.sum'(X1,X2),X3) = s('type.sum.sum'(X1,X2),'const.sum.INL_1'(s(X1,sK9(X3,X2,X1))))) )),
  inference(superposition,[],[f37,f34])).
fof(f34,plain,(
  ( ! [X2,X0,X1] : (s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK10(X2,X1,X0)))) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK9(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2] : (s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK9(X2,X1,X0)))) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK10(X2,X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9,sK10])],[f14])).
fof(f14,plain,(
  ! [X0,X1,X2] : (? [X4] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4))) | ? [X3] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))))),
  inference(flattening,[],[f13])).
fof(f13,plain,(
  ! [X0,X1] : ! [X2] : (? [X4] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X4))) | ? [X3] : s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X3))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7] : ! [X17] : (? [X15] : s('type.sum.sum'(X6,X7),X17) = s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,X15))) | ? [X16] : s('type.sum.sum'(X6,X7),X17) = s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X16))))),
  file('thm.sum.sum_case_cong','thm.sum.sum_CASES')).
fof(f37,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X2,X1),X5),s(X2,X3))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X2),'const.sum.INR_1'(s(X2,X3))),s(fun(X0,X1),X4),s(fun(X2,X1),X5)))) )),
  inference(cnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X5),s(X2,X3))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X2),'const.sum.INR_1'(s(X2,X3))),s(fun(X0,X1),X4),s(fun(X2,X1),X5)))),
  inference(flattening,[],[f19])).
fof(f19,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X5),s(X2,X3))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X2),'const.sum.INR_1'(s(X2,X3))),s(fun(X0,X1),X4),s(fun(X2,X1),X5)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X8,X7] : ! [X16,X13,X14] : s(X8,happ(s(fun(X7,X8),X14),s(X7,X16))) = s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X16))),s(fun(X6,X8),X13),s(fun(X7,X8),X14)))),
  file('thm.sum.sum_case_cong','thm.sum.sum_case_def|split|1')).
fof(f4643,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f4642,f3895])).
fof(f3895,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,sK9(sK6,sK1,sK0)))) | $spl42),
  inference(unit_resulting_resolution,[],[f3646,f38])).
fof(f38,plain,(
  ( ! [X10] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,X10))) | s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,X10))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,X10)))) )),
  inference(backward_demodulation,[],[f27,f28])).
fof(f28,plain,(
  ( ! [X10] : (s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,X10))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,X10))) | s('type.sum.sum'(sK0,sK1),sK6) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,X10)))) )),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),sK6) & ! [X10] : (s('type.sum.sum'(sK0,sK1),sK6) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,X10))) | s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,X10))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,X10)))) & ! [X9] : (s('type.sum.sum'(sK0,sK1),sK6) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X9))) | s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,X9))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,X9)))) & s(sK2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK2),sK7),s(fun(sK1,sK2),sK8))) != s(sK2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK6),s(fun(sK0,sK2),sK4),s(fun(sK1,sK2),sK3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8])],[f22])).
fof(f22,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (s('type.sum.sum'(X0,X1),X5) = s('type.sum.sum'(X0,X1),X6) & ! [X10] : (s('type.sum.sum'(X0,X1),X6) != s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X10))) | s(X2,happ(s(fun(X0,X2),X4),s(X0,X10))) = s(X2,happ(s(fun(X0,X2),X7),s(X0,X10)))) & ! [X9] : (s('type.sum.sum'(X0,X1),X6) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X9))) | s(X2,happ(s(fun(X1,X2),X3),s(X1,X9))) = s(X2,happ(s(fun(X1,X2),X8),s(X1,X9)))) & s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X5),s(fun(X0,X2),X7),s(fun(X1,X2),X8))) != s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X6),s(fun(X0,X2),X4),s(fun(X1,X2),X3))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.sum.sum'(X0,X1),X5) = s('type.sum.sum'(X0,X1),X6) & ! [X10] : (s('type.sum.sum'(X0,X1),X6) != s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X10))) | s(X2,happ(s(fun(X0,X2),X4),s(X0,X10))) = s(X2,happ(s(fun(X0,X2),X7),s(X0,X10)))) & ! [X9] : (s('type.sum.sum'(X0,X1),X6) != s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X9))) | s(X2,happ(s(fun(X1,X2),X3),s(X1,X9))) = s(X2,happ(s(fun(X1,X2),X8),s(X1,X9))))) & s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X5),s(fun(X0,X2),X7),s(fun(X1,X2),X8))) != s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X6),s(fun(X0,X2),X4),s(fun(X1,X2),X3))))),
  inference(ennf_transformation,[],[f11])).
fof(f11,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : ((s('type.sum.sum'(X0,X1),X5) = s('type.sum.sum'(X0,X1),X6) & ! [X10] : (s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X10))) => s(X2,happ(s(fun(X0,X2),X4),s(X0,X10))) = s(X2,happ(s(fun(X0,X2),X7),s(X0,X10)))) & ! [X9] : (s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X9))) => s(X2,happ(s(fun(X1,X2),X3),s(X1,X9))) = s(X2,happ(s(fun(X1,X2),X8),s(X1,X9))))) => s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X5),s(fun(X0,X2),X7),s(fun(X1,X2),X8))) = s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X6),s(fun(X0,X2),X4),s(fun(X1,X2),X3))))),
  inference(flattening,[],[f10])).
fof(f10,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5,X6,X7,X8] : ((s('type.sum.sum'(X0,X1),X5) = s('type.sum.sum'(X0,X1),X6) & ! [X10] : (s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X10))) => s(X2,happ(s(fun(X0,X2),X4),s(X0,X10))) = s(X2,happ(s(fun(X0,X2),X7),s(X0,X10)))) & ! [X9] : (s('type.sum.sum'(X0,X1),X6) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X9))) => s(X2,happ(s(fun(X1,X2),X3),s(X1,X9))) = s(X2,happ(s(fun(X1,X2),X8),s(X1,X9))))) => s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X5),s(fun(X0,X2),X7),s(fun(X1,X2),X8))) = s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X6),s(fun(X0,X2),X4),s(fun(X1,X2),X3))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11,X12,X13,X14] : ((s('type.sum.sum'(X6,X7),X11) = s('type.sum.sum'(X6,X7),X12) & ! [X15] : (s('type.sum.sum'(X6,X7),X12) = s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,X15))) => s(X8,happ(s(fun(X6,X8),X10),s(X6,X15))) = s(X8,happ(s(fun(X6,X8),X13),s(X6,X15)))) & ! [X16] : (s('type.sum.sum'(X6,X7),X12) = s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X16))) => s(X8,happ(s(fun(X7,X8),X9),s(X7,X16))) = s(X8,happ(s(fun(X7,X8),X14),s(X7,X16))))) => s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X7),X11),s(fun(X6,X8),X13),s(fun(X7,X8),X14))) = s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X7),X12),s(fun(X6,X8),X10),s(fun(X7,X8),X9))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11,X12,X13,X14] : ((s('type.sum.sum'(X6,X7),X11) = s('type.sum.sum'(X6,X7),X12) & ! [X15] : (s('type.sum.sum'(X6,X7),X12) = s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,X15))) => s(X8,happ(s(fun(X6,X8),X10),s(X6,X15))) = s(X8,happ(s(fun(X6,X8),X13),s(X6,X15)))) & ! [X16] : (s('type.sum.sum'(X6,X7),X12) = s('type.sum.sum'(X6,X7),'const.sum.INR_1'(s(X7,X16))) => s(X8,happ(s(fun(X7,X8),X9),s(X7,X16))) = s(X8,happ(s(fun(X7,X8),X14),s(X7,X16))))) => s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X7),X11),s(fun(X6,X8),X13),s(fun(X7,X8),X14))) = s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X7),X12),s(fun(X6,X8),X10),s(fun(X7,X8),X9))))),
  file('thm.sum.sum_case_cong','thm.sum.sum_case_cong')).
fof(f27,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),sK6)),
  inference(cnf_transformation,[],[f24])).
fof(f4642,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) != s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,sK9(sK6,sK1,sK0)))) | $spl42),
  inference(forward_demodulation,[],[f4641,f3902])).
fof(f3902,plain,(
  ( ! [X6,X7,X5] : (s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9(sK6,sK1,sK0)))) = s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X5),X6),s(fun(sK1,X5),X7)))) ) | $spl42),
  inference(superposition,[],[f36,f3646])).
fof(f36,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X2,X1),X4),s(X2,X3))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X2,X0),'const.sum.INL_1'(s(X2,X3))),s(fun(X2,X1),X4),s(fun(X0,X1),X5)))) )),
  inference(cnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X4),s(X2,X3))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X2,X0),'const.sum.INL_1'(s(X2,X3))),s(fun(X2,X1),X4),s(fun(X0,X1),X5)))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X1,happ(s(fun(X2,X1),X4),s(X2,X3))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X2,X0),'const.sum.INL_1'(s(X2,X3))),s(fun(X2,X1),X4),s(fun(X0,X1),X5)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7,X8,X6] : ! [X15,X13,X14] : s(X8,happ(s(fun(X6,X8),X13),s(X6,X15))) = s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X7),'const.sum.INL_1'(s(X6,X15))),s(fun(X6,X8),X13),s(fun(X7,X8),X14)))),
  file('thm.sum.sum_case_cong','thm.sum.sum_case_def|split|0')).
fof(f4641,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,sK9(sK6,sK1,sK0)))) != s(sK2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK2),sK4),s(fun(sK1,sK2),sK3))) | $spl42),
  inference(forward_demodulation,[],[f40,f3902])).
fof(f40,plain,(
  s(sK2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK2),sK4),s(fun(sK1,sK2),sK3))) != s(sK2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK2),sK7),s(fun(sK1,sK2),sK8)))),
  inference(backward_demodulation,[],[f27,f30])).
fof(f30,plain,(
  s(sK2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK2),sK7),s(fun(sK1,sK2),sK8))) != s(sK2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK6),s(fun(sK0,sK2),sK4),s(fun(sK1,sK2),sK3)))),
  inference(cnf_transformation,[],[f24])).
fof(f4640,plain,(
  ( ! [X6,X7,X5] : (s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9(sK5,sK1,sK0)))) = s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X5),X6),s(fun(sK1,X5),X7)))) ) | $spl160),
  inference(cnf_transformation,[],[f4640_D])).
fof(f4640_D,plain,(
  ( ! [X6,X7,X5] : (s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9(sK5,sK1,sK0)))) = s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X5),X6),s(fun(sK1,X5),X7)))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f4532,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))),sK1,sK0)))) | $spl158),
  inference(cnf_transformation,[],[f4532_D])).
fof(f4532_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))),sK1,sK0)))) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f4639,plain,(
  ( ! [X6,X7,X5] : (s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))),sK1,sK0)))) | s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9(sK5,sK1,sK0)))) = s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X5),X6),s(fun(sK1,X5),X7)))) ) | ($spl58 | $spl156)),
  inference(forward_demodulation,[],[f4638,f3892])).
fof(f4638,plain,(
  ( ! [X6,X7,X5] : (s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9(sK5,sK1,sK0)))) = s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X5),X6),s(fun(sK1,X5),X7))) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))),sK1,sK0))))) ) | ($spl58 | $spl156)),
  inference(forward_demodulation,[],[f4502,f4530])).
fof(f4502,plain,(
  ( ! [X6,X7,X5] : (s(X5,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X5),X6),s(fun(sK1,X5),X7))) = s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))),sK1,sK0))))) ) | $spl58),
  inference(superposition,[],[f332,f3892])).
fof(f332,plain,(
  ( ! [X6,X4,X2,X5,X3,X1] : (s(X4,happ(s(fun(X1,X4),X5),s(X1,sK9(X3,X2,X1)))) = s(X4,'const.sum.sum_CASE_3'(s('type.sum.sum'(X1,X2),X3),s(fun(X1,X4),X5),s(fun(X2,X4),X6))) | s('type.sum.sum'(X1,X2),X3) = s('type.sum.sum'(X1,X2),'const.sum.INR_1'(s(X2,sK10(X3,X2,X1))))) )),
  inference(superposition,[],[f36,f34])).
fof(f4637,plain,(
  $false | $spl58),
  inference(subsumption_resolution,[],[f4636,f4498])).
fof(f4498,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK5,sK1,sK0)))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK5,sK1,sK0)))) | $spl58),
  inference(unit_resulting_resolution,[],[f3892,f39])).
fof(f39,plain,(
  ( ! [X9] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X9))) | s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,X9))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,X9)))) )),
  inference(backward_demodulation,[],[f27,f29])).
fof(f29,plain,(
  ( ! [X9] : (s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,X9))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,X9))) | s('type.sum.sum'(sK0,sK1),sK6) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,X9)))) )),
  inference(cnf_transformation,[],[f24])).
fof(f4636,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK5,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK5,sK1,sK0)))) | $spl58),
  inference(forward_demodulation,[],[f4635,f4501])).
fof(f4501,plain,(
  ( ! [X4,X2,X3] : (s(X2,happ(s(fun(sK1,X2),X3),s(sK1,sK10(sK5,sK1,sK0)))) = s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X2),X4),s(fun(sK1,X2),X3)))) ) | $spl58),
  inference(superposition,[],[f37,f3892])).
fof(f4635,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK5,sK1,sK0)))) != s(sK2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK2),sK4),s(fun(sK1,sK2),sK3))) | $spl58),
  inference(forward_demodulation,[],[f40,f4501])).
fof(f4634,plain,(
  $false | ($spl47 | $spl58)),
  inference(subsumption_resolution,[],[f4633,f4498])).
fof(f4633,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK5,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK5,sK1,sK0)))) | ($spl47 | $spl58)),
  inference(forward_demodulation,[],[f3656,f4501])).
fof(f4626,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f4625,f4218])).
fof(f4218,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK5,sK1,sK0)))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,sK9(sK5,sK1,sK0)))) | $spl44),
  inference(unit_resulting_resolution,[],[f3654,f38])).
fof(f4625,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK5,sK1,sK0)))) != s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,sK9(sK5,sK1,sK0)))) | $spl44),
  inference(forward_demodulation,[],[f4624,f4221])).
fof(f4221,plain,(
  ( ! [X4,X2,X3] : (s(X2,happ(s(fun(sK0,X2),X3),s(sK0,sK9(sK5,sK1,sK0)))) = s(X2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X2),X3),s(fun(sK1,X2),X4)))) ) | $spl44),
  inference(superposition,[],[f36,f3654])).
fof(f4624,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,sK9(sK5,sK1,sK0)))) != s(sK2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK2),sK4),s(fun(sK1,sK2),sK3))) | $spl44),
  inference(forward_demodulation,[],[f40,f4221])).
fof(f4094,plain,(
  ( ! [X6,X8] : (s(X6,happ(s(fun(sK1,X6),X8),s(sK1,sK10(sK5,sK1,sK0)))) = s(X6,happ(s(fun(sK1,X6),X8),s(sK1,sK10(sK6,sK1,sK0))))) ) | $spl82),
  inference(cnf_transformation,[],[f4094_D])).
fof(f4094_D,plain,(
  ( ! [X6,X8] : (s(X6,happ(s(fun(sK1,X6),X8),s(sK1,sK10(sK5,sK1,sK0)))) = s(X6,happ(s(fun(sK1,X6),X8),s(sK1,sK10(sK6,sK1,sK0))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f4241,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))),sK1,sK0)))) | $spl102),
  inference(cnf_transformation,[],[f4241_D])).
fof(f4241_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))),sK1,sK0)))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f4623,plain,(
  ( ! [X8,X9] : (s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))),sK1,sK0)))) | s(X8,happ(s(fun(sK1,X8),X9),s(sK1,sK10(sK5,sK1,sK0)))) = s(X8,happ(s(fun(sK1,X8),X9),s(sK1,sK10(sK6,sK1,sK0))))) ) | ($spl44 | $spl58 | $spl100)),
  inference(forward_demodulation,[],[f4622,f3654])).
fof(f4622,plain,(
  ( ! [X8,X9] : (s(X8,happ(s(fun(sK1,X8),X9),s(sK1,sK10(sK5,sK1,sK0)))) = s(X8,happ(s(fun(sK1,X8),X9),s(sK1,sK10(sK6,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))),sK1,sK0))))) ) | ($spl44 | $spl58 | $spl100)),
  inference(forward_demodulation,[],[f4621,f4501])).
fof(f4621,plain,(
  ( ! [X10,X8,X9] : (s(X8,happ(s(fun(sK1,X8),X9),s(sK1,sK10(sK6,sK1,sK0)))) = s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X8),X10),s(fun(sK1,X8),X9))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))),sK1,sK0))))) ) | ($spl44 | $spl100)),
  inference(forward_demodulation,[],[f4223,f4239])).
fof(f4223,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X8),X10),s(fun(sK1,X8),X9))) = s(X8,happ(s(fun(sK1,X8),X9),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))),sK1,sK0))))) ) | $spl44),
  inference(superposition,[],[f328,f3654])).
fof(f4618,plain,(
  $false | ($spl40 | $spl58)),
  inference(subsumption_resolution,[],[f4598,f4498])).
fof(f4598,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK5,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK5,sK1,sK0)))) | ($spl40 | $spl58)),
  inference(superposition,[],[f4525,f4510])).
fof(f4510,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun(sK1,X0),X1),s(sK1,sK10(sK5,sK1,sK0)))) = s(X0,happ(s(fun(sK1,X0),X1),s(sK1,sK10(sK6,sK1,sK0))))) ) | ($spl40 | $spl58)),
  inference(backward_demodulation,[],[f4501,f3644])).
fof(f4525,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK5,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK6,sK1,sK0)))) | ($spl40 | $spl58)),
  inference(backward_demodulation,[],[f4510,f3993])).
fof(f3993,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK6,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK6,sK1,sK0)))) | $spl40),
  inference(forward_demodulation,[],[f3992,f3644])).
fof(f3992,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK6,sK1,sK0)))) != s(sK2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK2),sK4),s(fun(sK1,sK2),sK3))) | $spl40),
  inference(forward_demodulation,[],[f40,f3644])).
fof(f4537,plain,(
  $false | ($spl40 | $spl58 | $spl79)),
  inference(unit_resulting_resolution,[],[f4016,f4510])).
fof(f4601,plain,(
  $false | ($spl40 | $spl58 | $spl79)),
  inference(trivial_inequality_removal,[],[f4600])).
fof(f4600,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK5,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK5,sK1,sK0)))) | ($spl40 | $spl58 | $spl79)),
  inference(superposition,[],[f4016,f4510])).
fof(f4530,plain,(
  ( ! [X6,X5] : (s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9(sK5,sK1,sK0)))) = s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))),sK1,sK0))))) ) | $spl156),
  inference(cnf_transformation,[],[f4530_D])).
fof(f4530_D,plain,(
  ( ! [X6,X5] : (s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9(sK5,sK1,sK0)))) = s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))),sK1,sK0))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f4529,plain,(
  ( ! [X6,X5] : (s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))),sK1,sK0)))) | s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9(sK5,sK1,sK0)))) = s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))),sK1,sK0))))) ) | ($spl44 | $spl58)),
  inference(forward_demodulation,[],[f4528,f3892])).
fof(f4528,plain,(
  ( ! [X6,X5] : (s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9(sK5,sK1,sK0)))) = s(X5,happ(s(fun(sK0,X5),X6),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))),sK1,sK0))))) ) | ($spl44 | $spl58)),
  inference(forward_demodulation,[],[f4502,f4221])).
fof(f3892,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0)))) | $spl58),
  inference(cnf_transformation,[],[f3892_D])).
fof(f3892_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0)))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f4093,plain,(
  ( ! [X4,X5,X3] : (s(X3,happ(s(fun(sK0,X3),X4),s(sK0,sK9(sK5,sK1,sK0)))) = s(X3,happ(s(fun(sK1,X3),X5),s(sK1,sK10(sK6,sK1,sK0))))) ) | $spl80),
  inference(cnf_transformation,[],[f4093_D])).
fof(f4093_D,plain,(
  ( ! [X4,X5,X3] : (s(X3,happ(s(fun(sK0,X3),X4),s(sK0,sK9(sK5,sK1,sK0)))) = s(X3,happ(s(fun(sK1,X3),X5),s(sK1,sK10(sK6,sK1,sK0))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f4267,plain,(
  ( ! [X4,X5,X3] : (s(X3,happ(s(fun(sK0,X3),X4),s(sK0,sK9(sK5,sK1,sK0)))) = s(X3,happ(s(fun(sK1,X3),X5),s(sK1,sK10(sK6,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))))) ) | $spl40),
  inference(superposition,[],[f3644,f332])).
fof(f4283,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK0,X0),X1),s(sK0,sK9(sK5,sK1,sK0)))) = s(X0,happ(s(fun(sK1,X0),X2),s(sK1,sK10(sK6,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))))) ) | $spl40),
  inference(superposition,[],[f332,f3644])).
fof(f4497,plain,(
  $false | ($spl40 | $spl80)),
  inference(subsumption_resolution,[],[f4464,f4093])).
fof(f4464,plain,(
  ( ! [X201] : (s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK6,sK1,sK0)))) != s(sK2,happ(s(fun(sK0,sK2),X201),s(sK0,sK9(sK5,sK1,sK0))))) ) | ($spl40 | $spl80)),
  inference(superposition,[],[f3993,f4093])).
fof(f4495,plain,(
  ( ! [X199] : (s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,happ(s(fun(sK1,sK1),X199),s(sK1,sK10(sK6,sK1,sK0)))))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,happ(s(fun(sK1,sK1),X199),s(sK1,sK10(sK6,sK1,sK0))))))) ) | $spl152),
  inference(cnf_transformation,[],[f4495_D])).
fof(f4495_D,plain,(
  ( ! [X199] : (s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,happ(s(fun(sK1,sK1),X199),s(sK1,sK10(sK6,sK1,sK0)))))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,happ(s(fun(sK1,sK1),X199),s(sK1,sK10(sK6,sK1,sK0))))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f4496,plain,(
  ( ! [X200] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,happ(s(fun(sK0,sK1),X200),s(sK0,sK9(sK5,sK1,sK0))))))) ) | $spl154),
  inference(cnf_transformation,[],[f4496_D])).
fof(f4496_D,plain,(
  ( ! [X200] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,happ(s(fun(sK0,sK1),X200),s(sK0,sK9(sK5,sK1,sK0))))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f4463,plain,(
  ( ! [X200,X199] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,happ(s(fun(sK0,sK1),X200),s(sK0,sK9(sK5,sK1,sK0)))))) | s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,happ(s(fun(sK1,sK1),X199),s(sK1,sK10(sK6,sK1,sK0)))))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,happ(s(fun(sK1,sK1),X199),s(sK1,sK10(sK6,sK1,sK0))))))) ) | $spl80),
  inference(superposition,[],[f39,f4093])).
fof(f4493,plain,(
  ( ! [X197] : (s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,happ(s(fun(sK1,sK0),X197),s(sK1,sK10(sK6,sK1,sK0)))))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,happ(s(fun(sK1,sK0),X197),s(sK1,sK10(sK6,sK1,sK0))))))) ) | $spl148),
  inference(cnf_transformation,[],[f4493_D])).
fof(f4493_D,plain,(
  ( ! [X197] : (s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,happ(s(fun(sK1,sK0),X197),s(sK1,sK10(sK6,sK1,sK0)))))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,happ(s(fun(sK1,sK0),X197),s(sK1,sK10(sK6,sK1,sK0))))))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f4494,plain,(
  ( ! [X198] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,happ(s(fun(sK0,sK0),X198),s(sK0,sK9(sK5,sK1,sK0))))))) ) | $spl150),
  inference(cnf_transformation,[],[f4494_D])).
fof(f4494_D,plain,(
  ( ! [X198] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,happ(s(fun(sK0,sK0),X198),s(sK0,sK9(sK5,sK1,sK0))))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f4462,plain,(
  ( ! [X198,X197] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,happ(s(fun(sK0,sK0),X198),s(sK0,sK9(sK5,sK1,sK0)))))) | s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,happ(s(fun(sK1,sK0),X197),s(sK1,sK10(sK6,sK1,sK0)))))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,happ(s(fun(sK1,sK0),X197),s(sK1,sK10(sK6,sK1,sK0))))))) ) | $spl80),
  inference(superposition,[],[f38,f4093])).
fof(f4470,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X10),s(sK1,sK10(sK6,sK1,sK0))))) ) | $spl110),
  inference(cnf_transformation,[],[f4470_D])).
fof(f4470_D,plain,(
  ( ! [X10] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X10),s(sK1,sK10(sK6,sK1,sK0))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f4492,plain,(
  ( ! [X94] : (s(bool,'T_0') != s(bool,f(s(bool,happ(s(fun(sK0,bool),X94),s(sK0,sK9(sK5,sK1,sK0))))))) ) | $spl146),
  inference(cnf_transformation,[],[f4492_D])).
fof(f4492_D,plain,(
  ( ! [X94] : (s(bool,'T_0') != s(bool,f(s(bool,happ(s(fun(sK0,bool),X94),s(sK0,sK9(sK5,sK1,sK0))))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f4445,plain,(
  ( ! [X94,X93] : (s(bool,'T_0') != s(bool,f(s(bool,happ(s(fun(sK0,bool),X94),s(sK0,sK9(sK5,sK1,sK0)))))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X93),s(sK1,sK10(sK6,sK1,sK0))))) ) | ($spl3 | $spl5 | $spl16 | $spl80)),
  inference(superposition,[],[f2982,f4093])).
fof(f2982,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl5 | $spl16)),
  inference(superposition,[],[f2629,f33])).
fof(f33,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.sum.sum_case_cong',aHL_BOOL_CASES)).
fof(f2629,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | ($spl3 | $spl5 | $spl16)),
  inference(forward_demodulation,[],[f1650,f1747])).
fof(f1747,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | ($spl3 | $spl16)),
  inference(unit_resulting_resolution,[],[f49,f1658])).
fof(f4477,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X8),s(sK0,sK9(sK5,sK1,sK0))))) ) | $spl122),
  inference(cnf_transformation,[],[f4477_D])).
fof(f4477_D,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X8),s(sK0,sK9(sK5,sK1,sK0))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f4491,plain,(
  ( ! [X92,X91] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X92),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X91),s(sK1,sK10(sK6,sK1,sK0))))) ) | ($spl3 | $spl9 | $spl80)),
  inference(forward_demodulation,[],[f4444,f2799])).
fof(f2799,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl3 | $spl9)),
  inference(unit_resulting_resolution,[],[f49,f1652,f1652,f1652,f78])).
fof(f78,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f41,f41])).
fof(f41,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f33,f33])).
fof(f1652,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | $spl9),
  inference(cnf_transformation,[],[f1652_D])).
fof(f1652_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f4444,plain,(
  ( ! [X92,X91] : (s(bool,g(s(bool,'T_0'))) = s(bool,happ(s(fun(sK0,bool),X92),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X91),s(sK1,sK10(sK6,sK1,sK0))))) ) | ($spl9 | $spl80)),
  inference(superposition,[],[f2837,f4093])).
fof(f2837,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(trivial_inequality_removal,[],[f2822])).
fof(f2822,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl9),
  inference(superposition,[],[f1652,f41])).
fof(f4490,plain,(
  ( ! [X90] : (s(bool,'T_0') = s(bool,g(s(bool,happ(s(fun(sK0,bool),X90),s(sK0,sK9(sK5,sK1,sK0))))))) ) | $spl144),
  inference(cnf_transformation,[],[f4490_D])).
fof(f4490_D,plain,(
  ( ! [X90] : (s(bool,'T_0') = s(bool,g(s(bool,happ(s(fun(sK0,bool),X90),s(sK0,sK9(sK5,sK1,sK0))))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f4443,plain,(
  ( ! [X90,X89] : (s(bool,'T_0') = s(bool,g(s(bool,happ(s(fun(sK0,bool),X90),s(sK0,sK9(sK5,sK1,sK0)))))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X89),s(sK1,sK10(sK6,sK1,sK0))))) ) | ($spl16 | $spl80)),
  inference(superposition,[],[f1658,f4093])).
fof(f4471,plain,(
  ( ! [X14,X13] : (s(bool,X13) = s(bool,happ(s(fun(sK1,bool),X14),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X13)) ) | $spl112),
  inference(cnf_transformation,[],[f4471_D])).
fof(f4471_D,plain,(
  ( ! [X14,X13] : (s(bool,X13) = s(bool,happ(s(fun(sK1,bool),X14),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X13)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f4489,plain,(
  ( ! [X87] : (p(s(bool,happ(s(fun(sK0,bool),X87),s(sK0,sK9(sK5,sK1,sK0)))))) ) | $spl142),
  inference(cnf_transformation,[],[f4489_D])).
fof(f4489_D,plain,(
  ( ! [X87] : (p(s(bool,happ(s(fun(sK0,bool),X87),s(sK0,sK9(sK5,sK1,sK0)))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f4442,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,happ(s(fun(sK0,bool),X87),s(sK0,sK9(sK5,sK1,sK0))))) | s(bool,X88) = s(bool,happ(s(fun(sK1,bool),X86),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X88)) ) | $spl80),
  inference(superposition,[],[f110,f4093])).
fof(f110,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f31,f41])).
fof(f31,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.sum.sum_case_cong',aHL_TRUTH)).
fof(f4473,plain,(
  ( ! [X35,X33,X31,X34] : (s(bool,X33) = s(bool,X34) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X35) | s(bool,X34) = s(bool,happ(s(fun(sK0,bool),X31),s(sK0,sK9(sK5,sK1,sK0))))) ) | $spl116),
  inference(cnf_transformation,[],[f4473_D])).
fof(f4473_D,plain,(
  ( ! [X35,X33,X31,X34] : (s(bool,X33) = s(bool,X34) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X35) | s(bool,X34) = s(bool,happ(s(fun(sK0,bool),X31),s(sK0,sK9(sK5,sK1,sK0))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f4441,plain,(
  ( ! [X85,X83,X81,X84,X82] : (s(bool,X83) = s(bool,happ(s(fun(sK0,bool),X82),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X83) = s(bool,X84) | s(bool,X84) = s(bool,X85) | s(bool,'T_0') = s(bool,X85) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X81),s(sK1,sK10(sK6,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4472,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) = s(bool,happ(s(fun(sK1,bool),X23),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X24) | s(bool,X22) = s(bool,X24)) ) | $spl114),
  inference(cnf_transformation,[],[f4472_D])).
fof(f4472_D,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) = s(bool,happ(s(fun(sK1,bool),X23),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X24) | s(bool,X22) = s(bool,X24)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f4469,plain,(
  ( ! [X11,X9] : (s(bool,X11) = s(bool,happ(s(fun(sK0,bool),X9),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X11)) ) | $spl108),
  inference(cnf_transformation,[],[f4469_D])).
fof(f4469_D,plain,(
  ( ! [X11,X9] : (s(bool,X11) = s(bool,happ(s(fun(sK0,bool),X9),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X11)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f4440,plain,(
  ( ! [X80,X78,X76,X79,X77] : (s(bool,X78) = s(bool,happ(s(fun(sK0,bool),X77),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X79) = s(bool,happ(s(fun(sK1,bool),X76),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X79) = s(bool,X80) | s(bool,'T_0') = s(bool,X80) | s(bool,'T_0') = s(bool,X78)) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4475,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X42) = s(bool,happ(s(fun(sK1,bool),X43),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X45) | s(bool,X44) = s(bool,X45) | s(bool,X42) = s(bool,X44)) ) | $spl120),
  inference(cnf_transformation,[],[f4475_D])).
fof(f4475_D,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X42) = s(bool,happ(s(fun(sK1,bool),X43),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X45) | s(bool,X44) = s(bool,X45) | s(bool,X42) = s(bool,X44)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f4467,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X7),s(sK0,sK9(sK5,sK1,sK0))))) ) | $spl104),
  inference(cnf_transformation,[],[f4467_D])).
fof(f4467_D,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X7),s(sK0,sK9(sK5,sK1,sK0))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f4439,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X72),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X73) = s(bool,X74) | s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X75) | s(bool,X73) = s(bool,happ(s(fun(sK1,bool),X71),s(sK1,sK10(sK6,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4438,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,X68) = s(bool,happ(s(fun(sK0,bool),X67),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X68) = s(bool,X69) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X66),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X69) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4437,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,X63) = s(bool,happ(s(fun(sK0,bool),X62),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X64) = s(bool,happ(s(fun(sK1,bool),X61),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X63) | s(bool,X64) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4436,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X57),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X58) = s(bool,X59) | s(bool,X59) = s(bool,happ(s(fun(sK1,bool),X56),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X58) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4474,plain,(
  ( ! [X39,X36,X40] : (s(bool,X39) = s(bool,happ(s(fun(sK0,bool),X36),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X40) | s(bool,X39) = s(bool,X40)) ) | $spl118),
  inference(cnf_transformation,[],[f4474_D])).
fof(f4474_D,plain,(
  ( ! [X39,X36,X40] : (s(bool,X39) = s(bool,happ(s(fun(sK0,bool),X36),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X40) | s(bool,X39) = s(bool,X40)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f4435,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,X53) = s(bool,happ(s(fun(sK0,bool),X52),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X54) = s(bool,happ(s(fun(sK1,bool),X51),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X54) | s(bool,X53) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4434,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X48) = s(bool,happ(s(fun(sK0,bool),X47),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X48) = s(bool,X49) | s(bool,'T_0') = s(bool,X49) | s(bool,X50) = s(bool,happ(s(fun(sK1,bool),X46),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X50)) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4488,plain,(
  ( ! [X45] : (~p(s(bool,happ(s(fun(sK0,bool),X45),s(sK0,sK9(sK5,sK1,sK0)))))) ) | $spl140),
  inference(cnf_transformation,[],[f4488_D])).
fof(f4488_D,plain,(
  ( ! [X45] : (~p(s(bool,happ(s(fun(sK0,bool),X45),s(sK0,sK9(sK5,sK1,sK0)))))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f4433,plain,(
  ( ! [X45,X44] : (~p(s(bool,happ(s(fun(sK0,bool),X45),s(sK0,sK9(sK5,sK1,sK0))))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X44),s(sK1,sK10(sK6,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f43,f4093])).
fof(f43,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f32,f33])).
fof(f32,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f12])).
fof(f12,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.sum.sum_case_cong',aHL_FALSITY)).
fof(f4432,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X42),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X43) = s(bool,happ(s(fun(sK1,bool),X41),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X43)) ) | $spl80),
  inference(superposition,[],[f41,f4093])).
fof(f4431,plain,(
  ( ! [X39,X38,X40] : (s(bool,X40) = s(bool,happ(s(fun(sK0,bool),X39),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X38),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X40)) ) | $spl80),
  inference(superposition,[],[f41,f4093])).
fof(f4430,plain,(
  ( ! [X37,X35,X36] : (s(bool,X37) = s(bool,happ(s(fun(sK0,bool),X36),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X35),s(sK1,sK10(sK6,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f41,f4093])).
fof(f4429,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X33),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),X32),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X34)) ) | $spl80),
  inference(superposition,[],[f41,f4093])).
fof(f4428,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X31),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X30),s(sK1,sK10(sK6,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f33,f4093])).
fof(f4486,plain,(
  ( ! [X199] : (s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,happ(s(fun(sK0,sK1),X199),s(sK0,sK9(sK5,sK1,sK0)))))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,happ(s(fun(sK0,sK1),X199),s(sK0,sK9(sK5,sK1,sK0))))))) ) | $spl136),
  inference(cnf_transformation,[],[f4486_D])).
fof(f4486_D,plain,(
  ( ! [X199] : (s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,happ(s(fun(sK0,sK1),X199),s(sK0,sK9(sK5,sK1,sK0)))))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,happ(s(fun(sK0,sK1),X199),s(sK0,sK9(sK5,sK1,sK0))))))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f4487,plain,(
  ( ! [X200] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,happ(s(fun(sK1,sK1),X200),s(sK1,sK10(sK6,sK1,sK0))))))) ) | $spl138),
  inference(cnf_transformation,[],[f4487_D])).
fof(f4487_D,plain,(
  ( ! [X200] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,happ(s(fun(sK1,sK1),X200),s(sK1,sK10(sK6,sK1,sK0))))))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f4423,plain,(
  ( ! [X200,X199] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,happ(s(fun(sK1,sK1),X200),s(sK1,sK10(sK6,sK1,sK0)))))) | s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,happ(s(fun(sK0,sK1),X199),s(sK0,sK9(sK5,sK1,sK0)))))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,happ(s(fun(sK0,sK1),X199),s(sK0,sK9(sK5,sK1,sK0))))))) ) | $spl80),
  inference(superposition,[],[f39,f4093])).
fof(f4484,plain,(
  ( ! [X197] : (s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,happ(s(fun(sK0,sK0),X197),s(sK0,sK9(sK5,sK1,sK0)))))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,happ(s(fun(sK0,sK0),X197),s(sK0,sK9(sK5,sK1,sK0))))))) ) | $spl132),
  inference(cnf_transformation,[],[f4484_D])).
fof(f4484_D,plain,(
  ( ! [X197] : (s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,happ(s(fun(sK0,sK0),X197),s(sK0,sK9(sK5,sK1,sK0)))))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,happ(s(fun(sK0,sK0),X197),s(sK0,sK9(sK5,sK1,sK0))))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f4485,plain,(
  ( ! [X198] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,happ(s(fun(sK1,sK0),X198),s(sK1,sK10(sK6,sK1,sK0))))))) ) | $spl134),
  inference(cnf_transformation,[],[f4485_D])).
fof(f4485_D,plain,(
  ( ! [X198] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,happ(s(fun(sK1,sK0),X198),s(sK1,sK10(sK6,sK1,sK0))))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f4422,plain,(
  ( ! [X198,X197] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,happ(s(fun(sK1,sK0),X198),s(sK1,sK10(sK6,sK1,sK0)))))) | s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,happ(s(fun(sK0,sK0),X197),s(sK0,sK9(sK5,sK1,sK0)))))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,happ(s(fun(sK0,sK0),X197),s(sK0,sK9(sK5,sK1,sK0))))))) ) | $spl80),
  inference(superposition,[],[f38,f4093])).
fof(f4483,plain,(
  ( ! [X94] : (s(bool,'T_0') != s(bool,f(s(bool,happ(s(fun(sK1,bool),X94),s(sK1,sK10(sK6,sK1,sK0))))))) ) | $spl130),
  inference(cnf_transformation,[],[f4483_D])).
fof(f4483_D,plain,(
  ( ! [X94] : (s(bool,'T_0') != s(bool,f(s(bool,happ(s(fun(sK1,bool),X94),s(sK1,sK10(sK6,sK1,sK0))))))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f4405,plain,(
  ( ! [X94,X93] : (s(bool,'T_0') != s(bool,f(s(bool,happ(s(fun(sK1,bool),X94),s(sK1,sK10(sK6,sK1,sK0)))))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X93),s(sK0,sK9(sK5,sK1,sK0))))) ) | ($spl3 | $spl5 | $spl16 | $spl80)),
  inference(superposition,[],[f2982,f4093])).
fof(f4468,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),X8),s(sK1,sK10(sK6,sK1,sK0))))) ) | $spl106),
  inference(cnf_transformation,[],[f4468_D])).
fof(f4468_D,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),X8),s(sK1,sK10(sK6,sK1,sK0))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f4482,plain,(
  ( ! [X92,X91] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),X92),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X91),s(sK0,sK9(sK5,sK1,sK0))))) ) | ($spl3 | $spl9 | $spl80)),
  inference(forward_demodulation,[],[f4404,f2799])).
fof(f4404,plain,(
  ( ! [X92,X91] : (s(bool,g(s(bool,'T_0'))) = s(bool,happ(s(fun(sK1,bool),X92),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X91),s(sK0,sK9(sK5,sK1,sK0))))) ) | ($spl9 | $spl80)),
  inference(superposition,[],[f2837,f4093])).
fof(f4481,plain,(
  ( ! [X90] : (s(bool,'T_0') = s(bool,g(s(bool,happ(s(fun(sK1,bool),X90),s(sK1,sK10(sK6,sK1,sK0))))))) ) | $spl128),
  inference(cnf_transformation,[],[f4481_D])).
fof(f4481_D,plain,(
  ( ! [X90] : (s(bool,'T_0') = s(bool,g(s(bool,happ(s(fun(sK1,bool),X90),s(sK1,sK10(sK6,sK1,sK0))))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f4403,plain,(
  ( ! [X90,X89] : (s(bool,'T_0') = s(bool,g(s(bool,happ(s(fun(sK1,bool),X90),s(sK1,sK10(sK6,sK1,sK0)))))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X89),s(sK0,sK9(sK5,sK1,sK0))))) ) | ($spl16 | $spl80)),
  inference(superposition,[],[f1658,f4093])).
fof(f4480,plain,(
  ( ! [X87] : (p(s(bool,happ(s(fun(sK1,bool),X87),s(sK1,sK10(sK6,sK1,sK0)))))) ) | $spl126),
  inference(cnf_transformation,[],[f4480_D])).
fof(f4480_D,plain,(
  ( ! [X87] : (p(s(bool,happ(s(fun(sK1,bool),X87),s(sK1,sK10(sK6,sK1,sK0)))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f4402,plain,(
  ( ! [X88,X87,X86] : (p(s(bool,happ(s(fun(sK1,bool),X87),s(sK1,sK10(sK6,sK1,sK0))))) | s(bool,X88) = s(bool,happ(s(fun(sK0,bool),X86),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X88)) ) | $spl80),
  inference(superposition,[],[f110,f4093])).
fof(f4401,plain,(
  ( ! [X85,X83,X81,X84,X82] : (s(bool,X83) = s(bool,happ(s(fun(sK1,bool),X82),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X83) = s(bool,X84) | s(bool,X84) = s(bool,X85) | s(bool,'T_0') = s(bool,X85) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X81),s(sK0,sK9(sK5,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4400,plain,(
  ( ! [X80,X78,X76,X79,X77] : (s(bool,X78) = s(bool,happ(s(fun(sK1,bool),X77),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X79) = s(bool,happ(s(fun(sK0,bool),X76),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X79) = s(bool,X80) | s(bool,'T_0') = s(bool,X80) | s(bool,'T_0') = s(bool,X78)) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4399,plain,(
  ( ! [X74,X72,X71,X75,X73] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X72),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X73) = s(bool,X74) | s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X75) | s(bool,X73) = s(bool,happ(s(fun(sK0,bool),X71),s(sK0,sK9(sK5,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4398,plain,(
  ( ! [X70,X68,X66,X69,X67] : (s(bool,X68) = s(bool,happ(s(fun(sK1,bool),X67),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X68) = s(bool,X69) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X66),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X69) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4397,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,X63) = s(bool,happ(s(fun(sK1,bool),X62),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X64) = s(bool,happ(s(fun(sK0,bool),X61),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X63) | s(bool,X64) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4396,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X57),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X58) = s(bool,X59) | s(bool,X59) = s(bool,happ(s(fun(sK0,bool),X56),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X58) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4395,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,X53) = s(bool,happ(s(fun(sK1,bool),X52),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X54) = s(bool,happ(s(fun(sK0,bool),X51),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X54) | s(bool,X53) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4394,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X48) = s(bool,happ(s(fun(sK1,bool),X47),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X48) = s(bool,X49) | s(bool,'T_0') = s(bool,X49) | s(bool,X50) = s(bool,happ(s(fun(sK0,bool),X46),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X50)) ) | $spl80),
  inference(superposition,[],[f78,f4093])).
fof(f4479,plain,(
  ( ! [X45] : (~p(s(bool,happ(s(fun(sK1,bool),X45),s(sK1,sK10(sK6,sK1,sK0)))))) ) | $spl124),
  inference(cnf_transformation,[],[f4479_D])).
fof(f4479_D,plain,(
  ( ! [X45] : (~p(s(bool,happ(s(fun(sK1,bool),X45),s(sK1,sK10(sK6,sK1,sK0)))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f4393,plain,(
  ( ! [X45,X44] : (~p(s(bool,happ(s(fun(sK1,bool),X45),s(sK1,sK10(sK6,sK1,sK0))))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X44),s(sK0,sK9(sK5,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f43,f4093])).
fof(f4392,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X42),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X43) = s(bool,happ(s(fun(sK0,bool),X41),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X43)) ) | $spl80),
  inference(superposition,[],[f41,f4093])).
fof(f4391,plain,(
  ( ! [X39,X38,X40] : (s(bool,X40) = s(bool,happ(s(fun(sK1,bool),X39),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X38),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X40)) ) | $spl80),
  inference(superposition,[],[f41,f4093])).
fof(f4390,plain,(
  ( ! [X37,X35,X36] : (s(bool,X37) = s(bool,happ(s(fun(sK1,bool),X36),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X35),s(sK0,sK9(sK5,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f41,f4093])).
fof(f4389,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X33),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X34) = s(bool,happ(s(fun(sK0,bool),X32),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X34)) ) | $spl80),
  inference(superposition,[],[f41,f4093])).
fof(f4388,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),X31),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X30),s(sK0,sK9(sK5,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f33,f4093])).
fof(f4478,plain,(
  ( ! [X61,X62] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X62),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X61),s(sK1,sK10(sK6,sK1,sK0))))) ) | ($spl3 | $spl9 | $spl80)),
  inference(forward_demodulation,[],[f4382,f2799])).
fof(f4382,plain,(
  ( ! [X61,X62] : (s(bool,g(s(bool,'T_0'))) = s(bool,happ(s(fun(sK0,bool),X62),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X61),s(sK1,sK10(sK6,sK1,sK0))))) ) | ($spl9 | $spl80)),
  inference(superposition,[],[f4093,f2837])).
fof(f4381,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,X57) = s(bool,happ(s(fun(sK0,bool),X58),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X57) = s(bool,X59) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X56),s(sK1,sK10(sK6,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4380,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,X52) = s(bool,happ(s(fun(sK0,bool),X53),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X54) = s(bool,happ(s(fun(sK1,bool),X51),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X52)) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4379,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X47),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X48) = s(bool,X49) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,X48) = s(bool,happ(s(fun(sK1,bool),X46),s(sK1,sK10(sK6,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4378,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X42) = s(bool,happ(s(fun(sK0,bool),X43),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X41),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4377,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,X37) = s(bool,happ(s(fun(sK0,bool),X38),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X39) = s(bool,happ(s(fun(sK1,bool),X36),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X37) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4376,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X32),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X33) = s(bool,X34) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),X31),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4375,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,X27) = s(bool,happ(s(fun(sK0,bool),X28),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X29) = s(bool,happ(s(fun(sK1,bool),X26),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X29) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4374,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,X22) = s(bool,happ(s(fun(sK0,bool),X23),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,happ(s(fun(sK1,bool),X21),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X25)) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4373,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X19),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X20) = s(bool,happ(s(fun(sK1,bool),X18),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X20)) ) | $spl80),
  inference(superposition,[],[f4093,f41])).
fof(f4372,plain,(
  ( ! [X17,X15,X16] : (s(bool,X16) = s(bool,happ(s(fun(sK0,bool),X17),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X15),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X16)) ) | $spl80),
  inference(superposition,[],[f4093,f41])).
fof(f4371,plain,(
  ( ! [X14,X12,X13] : (s(bool,X13) = s(bool,happ(s(fun(sK0,bool),X14),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X12),s(sK1,sK10(sK6,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f4093,f41])).
fof(f4370,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X10),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X11) = s(bool,happ(s(fun(sK1,bool),X9),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X11)) ) | $spl80),
  inference(superposition,[],[f4093,f41])).
fof(f4369,plain,(
  ( ! [X8,X7] : (s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X8),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X7),s(sK1,sK10(sK6,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f4093,f33])).
fof(f4476,plain,(
  ( ! [X61,X62] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),X62),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X61),s(sK0,sK9(sK5,sK1,sK0))))) ) | ($spl3 | $spl9 | $spl80)),
  inference(forward_demodulation,[],[f4358,f2799])).
fof(f4358,plain,(
  ( ! [X61,X62] : (s(bool,g(s(bool,'T_0'))) = s(bool,happ(s(fun(sK1,bool),X62),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X61),s(sK0,sK9(sK5,sK1,sK0))))) ) | ($spl9 | $spl80)),
  inference(superposition,[],[f4093,f2837])).
fof(f4357,plain,(
  ( ! [X59,X57,X60,X58,X56] : (s(bool,X57) = s(bool,happ(s(fun(sK1,bool),X58),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X57) = s(bool,X59) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X56),s(sK0,sK9(sK5,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4356,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,X52) = s(bool,happ(s(fun(sK1,bool),X53),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X54) = s(bool,happ(s(fun(sK0,bool),X51),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X52)) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4355,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X47),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X48) = s(bool,X49) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,X48) = s(bool,happ(s(fun(sK0,bool),X46),s(sK0,sK9(sK5,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4354,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,X42) = s(bool,happ(s(fun(sK1,bool),X43),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X42) = s(bool,X44) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X41),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4353,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,X37) = s(bool,happ(s(fun(sK1,bool),X38),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X39) = s(bool,happ(s(fun(sK0,bool),X36),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X37) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4352,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X32),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X33) = s(bool,X34) | s(bool,X34) = s(bool,happ(s(fun(sK0,bool),X31),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,X33) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4351,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,X27) = s(bool,happ(s(fun(sK1,bool),X28),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X29) = s(bool,happ(s(fun(sK0,bool),X26),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X29) | s(bool,X27) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4350,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,X22) = s(bool,happ(s(fun(sK1,bool),X23),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,happ(s(fun(sK0,bool),X21),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X25)) ) | $spl80),
  inference(superposition,[],[f4093,f78])).
fof(f4349,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X19),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X20) = s(bool,happ(s(fun(sK0,bool),X18),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X20)) ) | $spl80),
  inference(superposition,[],[f4093,f41])).
fof(f4348,plain,(
  ( ! [X17,X15,X16] : (s(bool,X16) = s(bool,happ(s(fun(sK1,bool),X17),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X15),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X16)) ) | $spl80),
  inference(superposition,[],[f4093,f41])).
fof(f4347,plain,(
  ( ! [X14,X12,X13] : (s(bool,X13) = s(bool,happ(s(fun(sK1,bool),X14),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X12),s(sK0,sK9(sK5,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f4093,f41])).
fof(f4346,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X10),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,X11) = s(bool,happ(s(fun(sK0,bool),X9),s(sK0,sK9(sK5,sK1,sK0)))) | s(bool,'T_0') = s(bool,X11)) ) | $spl80),
  inference(superposition,[],[f4093,f41])).
fof(f4345,plain,(
  ( ! [X8,X7] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),X8),s(sK1,sK10(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X7),s(sK0,sK9(sK5,sK1,sK0))))) ) | $spl80),
  inference(superposition,[],[f4093,f33])).
fof(f4239,plain,(
  ( ! [X8,X9] : (s(X8,happ(s(fun(sK1,X8),X9),s(sK1,sK10(sK6,sK1,sK0)))) = s(X8,happ(s(fun(sK1,X8),X9),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))),sK1,sK0))))) ) | $spl100),
  inference(cnf_transformation,[],[f4239_D])).
fof(f4239_D,plain,(
  ( ! [X8,X9] : (s(X8,happ(s(fun(sK1,X8),X9),s(sK1,sK10(sK6,sK1,sK0)))) = s(X8,happ(s(fun(sK1,X8),X9),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))),sK1,sK0))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f4238,plain,(
  ( ! [X8,X9] : (s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))),sK1,sK0)))) | s(X8,happ(s(fun(sK1,X8),X9),s(sK1,sK10(sK6,sK1,sK0)))) = s(X8,happ(s(fun(sK1,X8),X9),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))),sK1,sK0))))) ) | ($spl40 | $spl44)),
  inference(forward_demodulation,[],[f4237,f3654])).
fof(f4237,plain,(
  ( ! [X8,X9] : (s(X8,happ(s(fun(sK1,X8),X9),s(sK1,sK10(sK6,sK1,sK0)))) = s(X8,happ(s(fun(sK1,X8),X9),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))),sK1,sK0))))) ) | ($spl40 | $spl44)),
  inference(forward_demodulation,[],[f4223,f3644])).
fof(f4216,plain,(
  ( ! [X194] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,happ(s(fun(sK1,sK1),X194),s(sK1,sK10(sK5,sK1,sK0))))))) ) | $spl96),
  inference(cnf_transformation,[],[f4216_D])).
fof(f4216_D,plain,(
  ( ! [X194] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,happ(s(fun(sK1,sK1),X194),s(sK1,sK10(sK5,sK1,sK0))))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f4217,plain,(
  ( ! [X193] : (s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,happ(s(fun(sK0,sK1),X193),s(sK0,sK9(sK6,sK1,sK0)))))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,happ(s(fun(sK0,sK1),X193),s(sK0,sK9(sK6,sK1,sK0))))))) ) | $spl98),
  inference(cnf_transformation,[],[f4217_D])).
fof(f4217_D,plain,(
  ( ! [X193] : (s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,happ(s(fun(sK0,sK1),X193),s(sK0,sK9(sK6,sK1,sK0)))))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,happ(s(fun(sK0,sK1),X193),s(sK0,sK9(sK6,sK1,sK0))))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f4215,plain,(
  ( ! [X194,X193] : (s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,happ(s(fun(sK0,sK1),X193),s(sK0,sK9(sK6,sK1,sK0)))))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,happ(s(fun(sK0,sK1),X193),s(sK0,sK9(sK6,sK1,sK0)))))) | s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,happ(s(fun(sK1,sK1),X194),s(sK1,sK10(sK5,sK1,sK0))))))) ) | ($spl40 | $spl42 | $spl82)),
  inference(forward_demodulation,[],[f4196,f3902])).
fof(f4196,plain,(
  ( ! [X194,X193] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,happ(s(fun(sK1,sK1),X194),s(sK1,sK10(sK5,sK1,sK0)))))) | s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK1),X193),s(fun(sK1,sK1),X194))))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK1),X193),s(fun(sK1,sK1),X194)))))) ) | ($spl40 | $spl82)),
  inference(superposition,[],[f39,f4108])).
fof(f4108,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK1,X0),X1),s(sK1,sK10(sK5,sK1,sK0)))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X0),X2),s(fun(sK1,X0),X1)))) ) | ($spl40 | $spl82)),
  inference(backward_demodulation,[],[f4094,f3644])).
fof(f4213,plain,(
  ( ! [X192] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,happ(s(fun(sK1,sK0),X192),s(sK1,sK10(sK5,sK1,sK0))))))) ) | $spl92),
  inference(cnf_transformation,[],[f4213_D])).
fof(f4213_D,plain,(
  ( ! [X192] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,happ(s(fun(sK1,sK0),X192),s(sK1,sK10(sK5,sK1,sK0))))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f4214,plain,(
  ( ! [X191] : (s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,happ(s(fun(sK0,sK0),X191),s(sK0,sK9(sK6,sK1,sK0)))))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,happ(s(fun(sK0,sK0),X191),s(sK0,sK9(sK6,sK1,sK0))))))) ) | $spl94),
  inference(cnf_transformation,[],[f4214_D])).
fof(f4214_D,plain,(
  ( ! [X191] : (s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,happ(s(fun(sK0,sK0),X191),s(sK0,sK9(sK6,sK1,sK0)))))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,happ(s(fun(sK0,sK0),X191),s(sK0,sK9(sK6,sK1,sK0))))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f4212,plain,(
  ( ! [X191,X192] : (s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,happ(s(fun(sK0,sK0),X191),s(sK0,sK9(sK6,sK1,sK0)))))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,happ(s(fun(sK0,sK0),X191),s(sK0,sK9(sK6,sK1,sK0)))))) | s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,happ(s(fun(sK1,sK0),X192),s(sK1,sK10(sK5,sK1,sK0))))))) ) | ($spl40 | $spl42 | $spl82)),
  inference(forward_demodulation,[],[f4195,f3902])).
fof(f4195,plain,(
  ( ! [X191,X192] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,happ(s(fun(sK1,sK0),X192),s(sK1,sK10(sK5,sK1,sK0)))))) | s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK0),X191),s(fun(sK1,sK0),X192))))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK0),X191),s(fun(sK1,sK0),X192)))))) ) | ($spl40 | $spl82)),
  inference(superposition,[],[f38,f4108])).
fof(f4205,plain,(
  ( ! [X89] : (p(s(bool,happ(s(fun(sK1,bool),X89),s(sK1,sK10(sK5,sK1,sK0)))))) ) | $spl88),
  inference(cnf_transformation,[],[f4205_D])).
fof(f4205_D,plain,(
  ( ! [X89] : (p(s(bool,happ(s(fun(sK1,bool),X89),s(sK1,sK10(sK5,sK1,sK0)))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f4206,plain,(
  ( ! [X90,X88] : (s(bool,X90) = s(bool,happ(s(fun(sK0,bool),X88),s(sK0,sK9(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X90)) ) | $spl90),
  inference(cnf_transformation,[],[f4206_D])).
fof(f4206_D,plain,(
  ( ! [X90,X88] : (s(bool,X90) = s(bool,happ(s(fun(sK0,bool),X88),s(sK0,sK9(sK6,sK1,sK0)))) | s(bool,'T_0') = s(bool,X90)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f4204,plain,(
  ( ! [X90,X88,X89] : (s(bool,X90) = s(bool,happ(s(fun(sK0,bool),X88),s(sK0,sK9(sK6,sK1,sK0)))) | p(s(bool,happ(s(fun(sK1,bool),X89),s(sK1,sK10(sK5,sK1,sK0))))) | s(bool,'T_0') = s(bool,X90)) ) | ($spl40 | $spl42 | $spl82)),
  inference(forward_demodulation,[],[f4177,f3902])).
fof(f4177,plain,(
  ( ! [X90,X88,X89] : (p(s(bool,happ(s(fun(sK1,bool),X89),s(sK1,sK10(sK5,sK1,sK0))))) | s(bool,X90) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,bool),X88),s(fun(sK1,bool),X89))) | s(bool,'T_0') = s(bool,X90)) ) | ($spl40 | $spl82)),
  inference(superposition,[],[f110,f4108])).
fof(f4202,plain,(
  ( ! [X47] : (~p(s(bool,happ(s(fun(sK1,bool),X47),s(sK1,sK10(sK5,sK1,sK0)))))) ) | $spl84),
  inference(cnf_transformation,[],[f4202_D])).
fof(f4202_D,plain,(
  ( ! [X47] : (~p(s(bool,happ(s(fun(sK1,bool),X47),s(sK1,sK10(sK5,sK1,sK0)))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f4203,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X46),s(sK0,sK9(sK6,sK1,sK0))))) ) | $spl86),
  inference(cnf_transformation,[],[f4203_D])).
fof(f4203_D,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X46),s(sK0,sK9(sK6,sK1,sK0))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f4201,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),X46),s(sK0,sK9(sK6,sK1,sK0)))) | ~p(s(bool,happ(s(fun(sK1,bool),X47),s(sK1,sK10(sK5,sK1,sK0)))))) ) | ($spl40 | $spl42 | $spl82)),
  inference(forward_demodulation,[],[f4168,f3902])).
fof(f4168,plain,(
  ( ! [X47,X46] : (~p(s(bool,happ(s(fun(sK1,bool),X47),s(sK1,sK10(sK5,sK1,sK0))))) | s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,bool),X46),s(fun(sK1,bool),X47)))) ) | ($spl40 | $spl82)),
  inference(superposition,[],[f43,f4108])).
fof(f3654,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0)))) | $spl44),
  inference(cnf_transformation,[],[f3654_D])).
fof(f3654_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4054,plain,(
  ( ! [X5,X3] : (s(X3,happ(s(fun(sK1,X3),X5),s(sK1,sK10(sK5,sK1,sK0)))) = s(X3,happ(s(fun(sK1,X3),X5),s(sK1,sK10(sK6,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))))) ) | $spl40),
  inference(superposition,[],[f328,f3644])).
fof(f4053,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK0,X0),X1),s(sK0,sK9(sK5,sK1,sK0)))) = s(X0,happ(s(fun(sK1,X0),X2),s(sK1,sK10(sK6,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))))) ) | $spl40),
  inference(superposition,[],[f332,f3644])).
fof(f4038,plain,(
  ( ! [X6,X8] : (s(X6,happ(s(fun(sK1,X6),X8),s(sK1,sK10(sK5,sK1,sK0)))) = s(X6,happ(s(fun(sK1,X6),X8),s(sK1,sK10(sK6,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))))) ) | $spl40),
  inference(superposition,[],[f3644,f328])).
fof(f4037,plain,(
  ( ! [X4,X5,X3] : (s(X3,happ(s(fun(sK0,X3),X4),s(sK0,sK9(sK5,sK1,sK0)))) = s(X3,happ(s(fun(sK1,X3),X5),s(sK1,sK10(sK6,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0))))) ) | $spl40),
  inference(superposition,[],[f3644,f332])).
fof(f4016,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK5,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK6,sK1,sK0)))) | $spl79),
  inference(cnf_transformation,[],[f4016_D])).
fof(f4016_D,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK5,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK6,sK1,sK0)))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f4014,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK5,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK6,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0)))) | ($spl40 | $spl47)),
  inference(trivial_inequality_removal,[],[f4013])).
fof(f4013,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK5,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK6,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),sK5) | s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0)))) | ($spl40 | $spl47)),
  inference(superposition,[],[f3991,f327])).
fof(f327,plain,(
  ( ! [X0] : (s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(X0,sK1,sK0)))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(X0,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),X0) | s('type.sum.sum'(sK0,sK1),X0) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(X0,sK1,sK0))))) )),
  inference(superposition,[],[f39,f34])).
fof(f3991,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK6,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK5,sK1,sK0)))) | ($spl40 | $spl47)),
  inference(forward_demodulation,[],[f3656,f3644])).
fof(f4012,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK5,sK1,sK0)))) != s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) | $spl77),
  inference(cnf_transformation,[],[f4012_D])).
fof(f4012_D,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK5,sK1,sK0)))) != s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f4010,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK5,sK1,sK0)))) != s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0)))) | $spl61),
  inference(trivial_inequality_removal,[],[f4009])).
fof(f4009,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK5,sK1,sK0)))) != s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),sK5) | s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0)))) | $spl61),
  inference(superposition,[],[f3894,f331])).
fof(f331,plain,(
  ( ! [X0] : (s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(X0,sK1,sK0)))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,sK9(X0,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),X0) | s('type.sum.sum'(sK0,sK1),X0) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(X0,sK1,sK0))))) )),
  inference(superposition,[],[f38,f34])).
fof(f3994,plain,(
  $false | ($spl40 | $spl54)),
  inference(global_subsumption,[],[f30,f29,f28,f27,f38,f39,f31,f32,f33,f34,f35,f36,f37,f41,f43,f46,f78,f110,f212,f214,f216,f217,f219,f221,f255,f257,f259,f292,f294,f296,f317,f319,f321,f322,f324,f326,f327,f328,f331,f332,f350,f1511,f383,f384,f1371,f425,f426,f1369,f446,f447,f448,f474,f475,f562,f565,f566,f567,f606,f609,f610,f611,f981,f628,f631,f632,f633,f636,f656,f659,f660,f661,f855,f1573,f772,f1648,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1586,f1587,f1588,f1589,f1590,f1591,f1643,f1615,f1848,f1849,f1850,f1853,f1855,f1856,f1857,f1861,f1863,f1865,f1866,f1941,f1870,f1872,f1873,f1874,f1875,f1876,f1933,f2632,f2633,f2634,f2637,f2639,f2640,f2641,f2645,f2647,f2649,f2650,f2727,f2655,f2657,f2658,f2659,f2660,f2661,f2719,f3044,f3045,f3046,f3049,f3051,f3052,f3053,f3057,f3058,f3059,f3060,f3177,f3178,f3181,f3183,f3184,f3185,f3189,f3190,f3191,f3192,f3641,f3501,f3502,f3503,f3506,f3508,f3509,f3510,f3643,f3516,f3517,f3518,f3519,f3520,f3647,f3648,f3523,f3524,f3533,f3534,f3535,f3536,f3552,f3553,f3556,f3558,f3559,f3560,f3565,f3566,f3574,f3583,f3591,f3594,f3595,f3596,f3599,f3600,f3610,f3619,f3627,f3630,f3631,f3632,f3633,f3634,f3635,f3636,f3637,f3638,f3659,f3661,f3664,f3827,f3671,f3672,f3673,f3676,f3678,f3679,f3680,f3830,f3686,f3687,f3688,f3689,f3690,f3691,f3831,f3832,f3694,f3695,f3696,f3697,f3710,f3711,f3713,f3714,f3730,f3731,f3734,f3736,f3737,f3738,f3743,f3744,f3752,f3761,f3771,f3774,f3775,f3776,f3777,f3781,f3782,f3792,f3801,f3811,f3814,f3815,f3816,f3817,f3819,f3820,f3822,f3823,f3824,f3839,f3845,f3856,f3829,f40,f3644,f3993,f3864])).
fof(f3845,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,bool),X3),s(fun(sK1,bool),X4))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK9(sK6,sK1,sK0))))) ) | $spl40),
  inference(superposition,[],[f259,f3664])).
fof(f3839,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,sK9(sK6,sK1,sK0)))) | $spl40),
  inference(unit_resulting_resolution,[],[f3664,f38])).
fof(f3824,plain,(
  ( ! [X47,X54,X52,X50,X48,X55,X53,X51,X49] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'('type.sum.sum'(X47,X48),X50),'const.sum.INL_1'(s('type.sum.sum'(X47,X48),X49))),s(fun('type.sum.sum'(X47,X48),bool),X51),s(fun(X50,bool),X52))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(X47,X48),bool),X51),s('type.sum.sum'(X47,X48),'const.sum.INR_1'(s(X48,sK10(X49,X48,X47)))))) | s(X53,happ(s(fun(X47,X53),X54),s(X47,sK9(X49,X48,X47)))) = s(X53,'const.sum.sum_CASE_3'(s('type.sum.sum'(X47,X48),X49),s(fun(X47,X53),X54),s(fun(X48,X53),X55)))) )),
  inference(superposition,[],[f259,f332])).
fof(f3823,plain,(
  ( ! [X39,X45,X43,X41,X38,X46,X44,X42,X40] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X41,'type.sum.sum'(X38,X39)),'const.sum.INR_1'(s('type.sum.sum'(X38,X39),X40))),s(fun(X41,bool),X42),s(fun('type.sum.sum'(X38,X39),bool),X43))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(X38,X39),bool),X43),s('type.sum.sum'(X38,X39),'const.sum.INR_1'(s(X39,sK10(X40,X39,X38)))))) | s(X44,happ(s(fun(X38,X44),X45),s(X38,sK9(X40,X39,X38)))) = s(X44,'const.sum.sum_CASE_3'(s('type.sum.sum'(X38,X39),X40),s(fun(X38,X44),X45),s(fun(X39,X44),X46)))) )),
  inference(superposition,[],[f296,f332])).
fof(f3822,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X34,X32] : (s(X30,'const.sum.sum_CASE_3'(s('type.sum.sum'(X27,X28),X29),s(fun(X27,X30),X32),s(fun(X28,X30),X31))) = s(X30,happ(s(fun(X28,X30),X31),s(X28,sK10('const.sum.INR_1'(s(X28,sK10(X29,X28,X27))),X28,X27)))) | s('type.sum.sum'(X27,X28),'const.sum.INR_1'(s(X28,sK10(X29,X28,X27)))) = s('type.sum.sum'(X27,X28),'const.sum.INL_1'(s(X27,sK9('const.sum.INR_1'(s(X28,sK10(X29,X28,X27))),X28,X27)))) | s(X33,happ(s(fun(X27,X33),X34),s(X27,sK9(X29,X28,X27)))) = s(X33,'const.sum.sum_CASE_3'(s('type.sum.sum'(X27,X28),X29),s(fun(X27,X33),X34),s(fun(X28,X33),X35)))) )),
  inference(superposition,[],[f328,f332])).
fof(f3820,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X10,X11),X12),s(fun(X10,bool),X13),s(fun(X11,bool),X14))) | s(bool,'F_0') = s(bool,happ(s(fun(X11,bool),X14),s(X11,sK10(X12,X11,X10)))) | s(X15,happ(s(fun(X10,X15),X16),s(X10,sK9(X12,X11,X10)))) = s(X15,'const.sum.sum_CASE_3'(s('type.sum.sum'(X10,X11),X12),s(fun(X10,X15),X16),s(fun(X11,X15),X17)))) )),
  inference(superposition,[],[f296,f332])).
fof(f3819,plain,(
  ( ! [X6,X4,X5,X3] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),X3) | s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(X3,sK1,sK0)))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(X3,sK1,sK0)))) | s(X4,happ(s(fun(sK0,X4),X5),s(sK0,sK9(X3,sK1,sK0)))) = s(X4,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X3),s(fun(sK0,X4),X5),s(fun(sK1,X4),X6)))) )),
  inference(superposition,[],[f39,f332])).
fof(f3817,plain,(
  ( ! [X300,X296,X298,X297,X299] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,happ(s(fun(X296,sK1),X299),s(X296,sK9(X298,X297,X296)))))) | s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X296,X297),X298),s(fun(X296,sK1),X299),s(fun(X297,sK1),X300))))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X296,X297),X298),s(fun(X296,sK1),X299),s(fun(X297,sK1),X300))))) | s('type.sum.sum'(X296,X297),X298) = s('type.sum.sum'(X296,X297),'const.sum.INR_1'(s(X297,sK10(X298,X297,X296))))) )),
  inference(superposition,[],[f39,f332])).
fof(f3816,plain,(
  ( ! [X292,X294,X293,X295,X291] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,happ(s(fun(X291,sK0),X294),s(X291,sK9(X293,X292,X291)))))) | s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X291,X292),X293),s(fun(X291,sK0),X294),s(fun(X292,sK0),X295))))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X291,X292),X293),s(fun(X291,sK0),X294),s(fun(X292,sK0),X295))))) | s('type.sum.sum'(X291,X292),X293) = s('type.sum.sum'(X291,X292),'const.sum.INR_1'(s(X292,sK10(X293,X292,X291))))) )),
  inference(superposition,[],[f38,f332])).
fof(f3815,plain,(
  ( ! [X282,X288,X290,X285,X287,X281,X283,X289,X284,X286] : (s(X288,'const.sum.sum_CASE_3'(s('type.sum.sum'(X281,X282),happ(s(fun(X283,'type.sum.sum'(X281,X282)),X286),s(X283,sK9(X285,X284,X283)))),s(fun(X281,X288),X290),s(fun(X282,X288),X289))) = s(X288,happ(s(fun(X282,X288),X289),s(X282,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X283,X284),X285),s(fun(X283,'type.sum.sum'(X281,X282)),X286),s(fun(X284,'type.sum.sum'(X281,X282)),X287)),X282,X281)))) | s('type.sum.sum'(X281,X282),'const.sum.sum_CASE_3'(s('type.sum.sum'(X283,X284),X285),s(fun(X283,'type.sum.sum'(X281,X282)),X286),s(fun(X284,'type.sum.sum'(X281,X282)),X287))) = s('type.sum.sum'(X281,X282),'const.sum.INL_1'(s(X281,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X283,X284),X285),s(fun(X283,'type.sum.sum'(X281,X282)),X286),s(fun(X284,'type.sum.sum'(X281,X282)),X287)),X282,X281)))) | s('type.sum.sum'(X283,X284),X285) = s('type.sum.sum'(X283,X284),'const.sum.INR_1'(s(X284,sK10(X285,X284,X283))))) )),
  inference(superposition,[],[f328,f332])).
fof(f3814,plain,(
  ( ! [X280,X277,X279,X273,X275,X276,X278,X272,X274] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X272,X278),'const.sum.INL_1'(s(X272,X279))),s(fun(X272,bool),happ(s(fun(X273,fun(X272,bool)),X276),s(X273,sK9(X275,X274,X273)))),s(fun(X278,bool),X280))) | s(bool,'F_0') = s(bool,happ(s(fun(X272,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X273,X274),X275),s(fun(X273,fun(X272,bool)),X276),s(fun(X274,fun(X272,bool)),X277))),s(X272,X279))) | s('type.sum.sum'(X273,X274),X275) = s('type.sum.sum'(X273,X274),'const.sum.INR_1'(s(X274,sK10(X275,X274,X273))))) )),
  inference(superposition,[],[f259,f332])).
fof(f3811,plain,(
  ( ! [X253,X251,X249,X246,X252,X250,X248,X247,X245] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X251,X245),'const.sum.INR_1'(s(X245,X252))),s(fun(X251,bool),X253),s(fun(X245,bool),happ(s(fun(X246,fun(X245,bool)),X249),s(X246,sK9(X248,X247,X246)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X245,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X246,X247),X248),s(fun(X246,fun(X245,bool)),X249),s(fun(X247,fun(X245,bool)),X250))),s(X245,X252))) | s('type.sum.sum'(X246,X247),X248) = s('type.sum.sum'(X246,X247),'const.sum.INR_1'(s(X247,sK10(X248,X247,X246))))) )),
  inference(superposition,[],[f296,f332])).
fof(f3801,plain,(
  ( ! [X167,X165,X169,X166,X164,X168] : (p(s(bool,happ(s(fun(X164,bool),X167),s(X164,sK9(X166,X165,X164))))) | s(bool,X169) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X164,X165),X166),s(fun(X164,bool),X167),s(fun(X165,bool),X168))) | s(bool,'T_0') = s(bool,X169) | s('type.sum.sum'(X164,X165),X166) = s('type.sum.sum'(X164,X165),'const.sum.INR_1'(s(X165,sK10(X166,X165,X164))))) )),
  inference(superposition,[],[f110,f332])).
fof(f3792,plain,(
  ( ! [X99,X97,X95,X98,X96] : (~p(s(bool,happ(s(fun(X95,bool),X98),s(X95,sK9(X97,X96,X95))))) | s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X95,X96),X97),s(fun(X95,bool),X98),s(fun(X96,bool),X99))) | s('type.sum.sum'(X95,X96),X97) = s('type.sum.sum'(X95,X96),'const.sum.INR_1'(s(X96,sK10(X97,X96,X95))))) )),
  inference(superposition,[],[f43,f332])).
fof(f3782,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X27,X34,X32] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X27,X33),'const.sum.INL_1'(s(X27,happ(s(fun(X28,X27),X31),s(X28,sK9(X30,X29,X28)))))),s(fun(X27,bool),X34),s(fun(X33,bool),X35))) | s(bool,'F_0') = s(bool,happ(s(fun(X27,bool),X34),s(X27,'const.sum.sum_CASE_3'(s('type.sum.sum'(X28,X29),X30),s(fun(X28,X27),X31),s(fun(X29,X27),X32))))) | s('type.sum.sum'(X28,X29),X30) = s('type.sum.sum'(X28,X29),'const.sum.INR_1'(s(X29,sK10(X30,X29,X28))))) )),
  inference(superposition,[],[f259,f332])).
fof(f3781,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20,X18] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X24,X18),'const.sum.INR_1'(s(X18,happ(s(fun(X19,X18),X22),s(X19,sK9(X21,X20,X19)))))),s(fun(X24,bool),X25),s(fun(X18,bool),X26))) | s(bool,'F_0') = s(bool,happ(s(fun(X18,bool),X26),s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,X20),X21),s(fun(X19,X18),X22),s(fun(X20,X18),X23))))) | s('type.sum.sum'(X19,X20),X21) = s('type.sum.sum'(X19,X20),'const.sum.INR_1'(s(X20,sK10(X21,X20,X19))))) )),
  inference(superposition,[],[f296,f332])).
fof(f3777,plain,(
  ( ! [X269,X271,X268,X270,X272] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X268,X271),X270),s(fun(X268,sK1),X269),s(fun(X271,sK1),X272))))) | s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,happ(s(fun(X268,sK1),X269),s(X268,sK9(X270,X271,X268)))))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,happ(s(fun(X268,sK1),X269),s(X268,sK9(X270,X271,X268)))))) | s('type.sum.sum'(X268,X271),X270) = s('type.sum.sum'(X268,X271),'const.sum.INR_1'(s(X271,sK10(X270,X271,X268))))) )),
  inference(superposition,[],[f39,f332])).
fof(f3776,plain,(
  ( ! [X263,X265,X267,X264,X266] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X263,X266),X265),s(fun(X263,sK0),X264),s(fun(X266,sK0),X267))))) | s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,happ(s(fun(X263,sK0),X264),s(X263,sK9(X265,X266,X263)))))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,happ(s(fun(X263,sK0),X264),s(X263,sK9(X265,X266,X263)))))) | s('type.sum.sum'(X263,X266),X265) = s('type.sum.sum'(X263,X266),'const.sum.INR_1'(s(X266,sK10(X265,X266,X263))))) )),
  inference(superposition,[],[f38,f332])).
fof(f3775,plain,(
  ( ! [X261,X255,X253,X257,X259,X260,X254,X256,X258,X262] : (s(X260,happ(s(fun(X254,X260),X261),s(X254,sK10(happ(s(fun(X255,'type.sum.sum'(X253,X254)),X256),s(X255,sK9(X257,X258,X255))),X254,X253)))) = s(X260,'const.sum.sum_CASE_3'(s('type.sum.sum'(X253,X254),'const.sum.sum_CASE_3'(s('type.sum.sum'(X255,X258),X257),s(fun(X255,'type.sum.sum'(X253,X254)),X256),s(fun(X258,'type.sum.sum'(X253,X254)),X259))),s(fun(X253,X260),X262),s(fun(X254,X260),X261))) | s('type.sum.sum'(X253,X254),happ(s(fun(X255,'type.sum.sum'(X253,X254)),X256),s(X255,sK9(X257,X258,X255)))) = s('type.sum.sum'(X253,X254),'const.sum.INL_1'(s(X253,sK9(happ(s(fun(X255,'type.sum.sum'(X253,X254)),X256),s(X255,sK9(X257,X258,X255))),X254,X253)))) | s('type.sum.sum'(X255,X258),X257) = s('type.sum.sum'(X255,X258),'const.sum.INR_1'(s(X258,sK10(X257,X258,X255))))) )),
  inference(superposition,[],[f328,f332])).
fof(f3774,plain,(
  ( ! [X251,X249,X246,X244,X252,X250,X248,X247,X245] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X244,X250),'const.sum.INL_1'(s(X244,X251))),s(fun(X244,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X245,X248),X247),s(fun(X245,fun(X244,bool)),X246),s(fun(X248,fun(X244,bool)),X249))),s(fun(X250,bool),X252))) | s(bool,'F_0') = s(bool,happ(s(fun(X244,bool),happ(s(fun(X245,fun(X244,bool)),X246),s(X245,sK9(X247,X248,X245)))),s(X244,X251))) | s('type.sum.sum'(X245,X248),X247) = s('type.sum.sum'(X245,X248),'const.sum.INR_1'(s(X248,sK10(X247,X248,X245))))) )),
  inference(superposition,[],[f259,f332])).
fof(f3771,plain,(
  ( ! [X218,X222,X220,X225,X219,X217,X223,X221,X224] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X223,X217),'const.sum.INR_1'(s(X217,X224))),s(fun(X223,bool),X225),s(fun(X217,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X218,X221),X220),s(fun(X218,fun(X217,bool)),X219),s(fun(X221,fun(X217,bool)),X222))))) | s(bool,'F_0') = s(bool,happ(s(fun(X217,bool),happ(s(fun(X218,fun(X217,bool)),X219),s(X218,sK9(X220,X221,X218)))),s(X217,X224))) | s('type.sum.sum'(X218,X221),X220) = s('type.sum.sum'(X218,X221),'const.sum.INR_1'(s(X221,sK10(X220,X221,X218))))) )),
  inference(superposition,[],[f296,f332])).
fof(f3761,plain,(
  ( ! [X140,X138,X136,X141,X139,X137] : (p(s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X136,X139),X138),s(fun(X136,bool),X137),s(fun(X139,bool),X140)))) | s(bool,X141) = s(bool,happ(s(fun(X136,bool),X137),s(X136,sK9(X138,X139,X136)))) | s(bool,'T_0') = s(bool,X141) | s('type.sum.sum'(X136,X139),X138) = s('type.sum.sum'(X136,X139),'const.sum.INR_1'(s(X139,sK10(X138,X139,X136))))) )),
  inference(superposition,[],[f110,f332])).
fof(f3752,plain,(
  ( ! [X70,X68,X71,X69,X67] : (~p(s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X67,X70),X69),s(fun(X67,bool),X68),s(fun(X70,bool),X71)))) | s(bool,'T_0') = s(bool,happ(s(fun(X67,bool),X68),s(X67,sK9(X69,X70,X67)))) | s('type.sum.sum'(X67,X70),X69) = s('type.sum.sum'(X67,X70),'const.sum.INR_1'(s(X70,sK10(X69,X70,X67))))) )),
  inference(superposition,[],[f43,f332])).
fof(f3744,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X15),'const.sum.INL_1'(s(X9,'const.sum.sum_CASE_3'(s('type.sum.sum'(X10,X13),X12),s(fun(X10,X9),X11),s(fun(X13,X9),X14))))),s(fun(X9,bool),X16),s(fun(X15,bool),X17))) | s(bool,'F_0') = s(bool,happ(s(fun(X9,bool),X16),s(X9,happ(s(fun(X10,X9),X11),s(X10,sK9(X12,X13,X10)))))) | s('type.sum.sum'(X10,X13),X12) = s('type.sum.sum'(X10,X13),'const.sum.INR_1'(s(X13,sK10(X12,X13,X10))))) )),
  inference(superposition,[],[f259,f332])).
fof(f3743,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X0),'const.sum.INR_1'(s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X1,X4),X3),s(fun(X1,X0),X2),s(fun(X4,X0),X5))))),s(fun(X6,bool),X7),s(fun(X0,bool),X8))) | s(bool,'F_0') = s(bool,happ(s(fun(X0,bool),X8),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK9(X3,X4,X1)))))) | s('type.sum.sum'(X1,X4),X3) = s('type.sum.sum'(X1,X4),'const.sum.INR_1'(s(X4,sK10(X3,X4,X1))))) )),
  inference(superposition,[],[f296,f332])).
fof(f3738,plain,(
  ( ! [X61,X68,X66,X64,X62,X67,X65,X63] : (s('type.sum.sum'(X62,bool),X61) = s('type.sum.sum'(X62,bool),'const.sum.INR_1'(s(bool,X63))) | s(X64,happ(s(fun(X62,X64),X65),s(X62,sK9(X61,bool,X62)))) = s(X64,'const.sum.sum_CASE_3'(s('type.sum.sum'(X62,bool),X61),s(fun(X62,X64),X65),s(fun(bool,X64),X66))) | s(bool,X63) = s(bool,X67) | s(bool,'T_0') = s(bool,sK10(X61,bool,X62)) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f332,f78])).
fof(f3737,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55,X53] : (s('type.sum.sum'(X54,bool),X53) = s('type.sum.sum'(X54,bool),'const.sum.INR_1'(s(bool,X55))) | s(X56,happ(s(fun(X54,X56),X57),s(X54,sK9(X53,bool,X54)))) = s(X56,'const.sum.sum_CASE_3'(s('type.sum.sum'(X54,bool),X53),s(fun(X54,X56),X57),s(fun(bool,X56),X58))) | s(bool,X59) = s(bool,sK10(X53,bool,X54)) | s(bool,'T_0') = s(bool,X55) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f332,f78])).
fof(f3736,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X51,X49] : (s('type.sum.sum'(X46,bool),X45) = s('type.sum.sum'(X46,bool),'const.sum.INR_1'(s(bool,'T_0'))) | s(X47,happ(s(fun(X46,X47),X48),s(X46,sK9(X45,bool,X46)))) = s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,bool),X45),s(fun(X46,X47),X48),s(fun(bool,X47),X49))) | s(bool,X50) = s(bool,X51) | s(bool,X51) = s(bool,sK10(X45,bool,X46)) | s(bool,X50) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f332,f78])).
fof(f3734,plain,(
  ( ! [X30,X35,X33,X31,X29,X36,X34,X32] : (s('type.sum.sum'(X30,bool),X29) = s('type.sum.sum'(X30,bool),'const.sum.INR_1'(s(bool,X31))) | s(X32,happ(s(fun(X30,X32),X33),s(X30,sK9(X29,bool,X30)))) = s(X32,'const.sum.sum_CASE_3'(s('type.sum.sum'(X30,bool),X29),s(fun(X30,X32),X33),s(fun(bool,X32),X34))) | s(bool,X31) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,sK10(X29,bool,X30)) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f332,f78])).
fof(f3731,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s('type.sum.sum'(X12,bool),X11) = s('type.sum.sum'(X12,bool),'const.sum.INR_1'(s(bool,X13))) | s(X14,happ(s(fun(X12,X14),X15),s(X12,sK9(X11,bool,X12)))) = s(X14,'const.sum.sum_CASE_3'(s('type.sum.sum'(X12,bool),X11),s(fun(X12,X14),X15),s(fun(bool,X14),X16))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,sK10(X11,bool,X12))) )),
  inference(superposition,[],[f332,f41])).
fof(f3730,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s('type.sum.sum'(X6,bool),X5) = s('type.sum.sum'(X6,bool),'const.sum.INR_1'(s(bool,'T_0'))) | s(X7,happ(s(fun(X6,X7),X8),s(X6,sK9(X5,bool,X6)))) = s(X7,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,bool),X5),s(fun(X6,X7),X8),s(fun(bool,X7),X9))) | s(bool,X10) = s(bool,sK10(X5,bool,X6)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f332,f41])).
fof(f3714,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(X29,bool),X32),s(X29,sK9('const.sum.INR_1'(s(X30,X31)),X30,X29)))) | s('type.sum.sum'(X29,X30),'const.sum.INR_1'(s(X30,X31))) = s('type.sum.sum'(X29,X30),'const.sum.INR_1'(s(X30,sK10('const.sum.INR_1'(s(X30,X31)),X30,X29)))) | s(bool,'F_0') = s(bool,happ(s(fun(X30,bool),X33),s(X30,X31)))) )),
  inference(superposition,[],[f332,f296])).
fof(f3713,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(X24,bool),X27),s(X24,sK9('const.sum.INL_1'(s(X24,X26)),X25,X24)))) | s('type.sum.sum'(X24,X25),'const.sum.INL_1'(s(X24,X26))) = s('type.sum.sum'(X24,X25),'const.sum.INR_1'(s(X25,sK10('const.sum.INL_1'(s(X24,X26)),X25,X24)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X27),s(X24,X26)))) )),
  inference(superposition,[],[f332,f259])).
fof(f3711,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(X6,happ(s(fun(X8,X6),X11),s(X8,X9))) = s(X6,happ(s(fun(X7,X6),X10),s(X7,sK9('const.sum.INR_1'(s(X8,X9)),X8,X7)))) | s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X9))) = s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,sK10('const.sum.INR_1'(s(X8,X9)),X8,X7))))) )),
  inference(superposition,[],[f332,f37])).
fof(f3710,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X4),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X4),s(X1,sK9('const.sum.INL_1'(s(X1,X3)),X2,X1)))) | s('type.sum.sum'(X1,X2),'const.sum.INL_1'(s(X1,X3))) = s('type.sum.sum'(X1,X2),'const.sum.INR_1'(s(X2,sK10('const.sum.INL_1'(s(X1,X3)),X2,X1))))) )),
  inference(superposition,[],[f332,f36])).
fof(f3697,plain,(
  ( ! [X94,X101,X99,X97,X95,X93,X102,X100,X98,X96] : (s(X100,happ(s(fun(X93,X100),X101),s(X93,sK9(happ(s(fun(X95,'type.sum.sum'(X93,X94)),X96),s(X95,sK9(X97,X98,X95))),X94,X93)))) = s(X100,'const.sum.sum_CASE_3'(s('type.sum.sum'(X93,X94),'const.sum.sum_CASE_3'(s('type.sum.sum'(X95,X98),X97),s(fun(X95,'type.sum.sum'(X93,X94)),X96),s(fun(X98,'type.sum.sum'(X93,X94)),X99))),s(fun(X93,X100),X101),s(fun(X94,X100),X102))) | s('type.sum.sum'(X93,X94),happ(s(fun(X95,'type.sum.sum'(X93,X94)),X96),s(X95,sK9(X97,X98,X95)))) = s('type.sum.sum'(X93,X94),'const.sum.INR_1'(s(X94,sK10(happ(s(fun(X95,'type.sum.sum'(X93,X94)),X96),s(X95,sK9(X97,X98,X95))),X94,X93)))) | s('type.sum.sum'(X95,X98),X97) = s('type.sum.sum'(X95,X98),'const.sum.INR_1'(s(X98,sK10(X97,X98,X95))))) )),
  inference(superposition,[],[f332,f332])).
fof(f3696,plain,(
  ( ! [X92,X90,X88,X87,X85,X83,X91,X89,X86,X84] : (s(X90,happ(s(fun(X83,X90),X91),s(X83,sK9(happ(s(fun(X85,'type.sum.sum'(X83,X84)),X86),s(X85,sK10(X87,X85,X88))),X84,X83)))) = s(X90,'const.sum.sum_CASE_3'(s('type.sum.sum'(X83,X84),'const.sum.sum_CASE_3'(s('type.sum.sum'(X88,X85),X87),s(fun(X88,'type.sum.sum'(X83,X84)),X89),s(fun(X85,'type.sum.sum'(X83,X84)),X86))),s(fun(X83,X90),X91),s(fun(X84,X90),X92))) | s('type.sum.sum'(X83,X84),happ(s(fun(X85,'type.sum.sum'(X83,X84)),X86),s(X85,sK10(X87,X85,X88)))) = s('type.sum.sum'(X83,X84),'const.sum.INR_1'(s(X84,sK10(happ(s(fun(X85,'type.sum.sum'(X83,X84)),X86),s(X85,sK10(X87,X85,X88))),X84,X83)))) | s('type.sum.sum'(X88,X85),X87) = s('type.sum.sum'(X88,X85),'const.sum.INL_1'(s(X88,sK9(X87,X85,X88))))) )),
  inference(superposition,[],[f332,f328])).
fof(f3695,plain,(
  ( ! [X80,X78,X76,X74,X81,X79,X77,X75,X73,X82] : (s(X80,'const.sum.sum_CASE_3'(s('type.sum.sum'(X73,X74),happ(s(fun(X76,'type.sum.sum'(X73,X74)),X79),s(X76,sK10(X77,X76,X75)))),s(fun(X73,X80),X81),s(fun(X74,X80),X82))) = s(X80,happ(s(fun(X73,X80),X81),s(X73,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X75,X76),X77),s(fun(X75,'type.sum.sum'(X73,X74)),X78),s(fun(X76,'type.sum.sum'(X73,X74)),X79)),X74,X73)))) | s('type.sum.sum'(X73,X74),'const.sum.sum_CASE_3'(s('type.sum.sum'(X75,X76),X77),s(fun(X75,'type.sum.sum'(X73,X74)),X78),s(fun(X76,'type.sum.sum'(X73,X74)),X79))) = s('type.sum.sum'(X73,X74),'const.sum.INR_1'(s(X74,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X75,X76),X77),s(fun(X75,'type.sum.sum'(X73,X74)),X78),s(fun(X76,'type.sum.sum'(X73,X74)),X79)),X74,X73)))) | s('type.sum.sum'(X75,X76),X77) = s('type.sum.sum'(X75,X76),'const.sum.INL_1'(s(X75,sK9(X77,X76,X75))))) )),
  inference(superposition,[],[f332,f328])).
fof(f3694,plain,(
  ( ! [X70,X68,X66,X64,X72,X71,X69,X67,X65,X63] : (s(X70,'const.sum.sum_CASE_3'(s('type.sum.sum'(X63,X64),happ(s(fun(X65,'type.sum.sum'(X63,X64)),X68),s(X65,sK9(X67,X66,X65)))),s(fun(X63,X70),X71),s(fun(X64,X70),X72))) = s(X70,happ(s(fun(X63,X70),X71),s(X63,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X65,X66),X67),s(fun(X65,'type.sum.sum'(X63,X64)),X68),s(fun(X66,'type.sum.sum'(X63,X64)),X69)),X64,X63)))) | s('type.sum.sum'(X63,X64),'const.sum.sum_CASE_3'(s('type.sum.sum'(X65,X66),X67),s(fun(X65,'type.sum.sum'(X63,X64)),X68),s(fun(X66,'type.sum.sum'(X63,X64)),X69))) = s('type.sum.sum'(X63,X64),'const.sum.INR_1'(s(X64,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X65,X66),X67),s(fun(X65,'type.sum.sum'(X63,X64)),X68),s(fun(X66,'type.sum.sum'(X63,X64)),X69)),X64,X63)))) | s('type.sum.sum'(X65,X66),X67) = s('type.sum.sum'(X65,X66),'const.sum.INR_1'(s(X66,sK10(X67,X66,X65))))) )),
  inference(superposition,[],[f332,f332])).
fof(f3832,plain,(
  ( ! [X61,X59,X57,X54,X62,X60,X58,X56,X55,X53] : (s('type.sum.sum'(X53,X54),happ(s(fun(X56,'type.sum.sum'(X53,X54)),X59),s(X56,X57))) = s('type.sum.sum'(X53,X54),'const.sum.INR_1'(s(X54,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X55,X56),'const.sum.INR_1'(s(X56,X57))),s(fun(X55,'type.sum.sum'(X53,X54)),X58),s(fun(X56,'type.sum.sum'(X53,X54)),X59)),X54,X53)))) | s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X53,X54),happ(s(fun(X56,'type.sum.sum'(X53,X54)),X59),s(X56,X57))),s(fun(X53,X60),X61),s(fun(X54,X60),X62))) = s(X60,happ(s(fun(X53,X60),X61),s(X53,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X55,X56),'const.sum.INR_1'(s(X56,X57))),s(fun(X55,'type.sum.sum'(X53,X54)),X58),s(fun(X56,'type.sum.sum'(X53,X54)),X59)),X54,X53))))) )),
  inference(forward_demodulation,[],[f3693,f37])).
fof(f3693,plain,(
  ( ! [X61,X59,X57,X54,X62,X60,X58,X56,X55,X53] : (s(X60,'const.sum.sum_CASE_3'(s('type.sum.sum'(X53,X54),happ(s(fun(X56,'type.sum.sum'(X53,X54)),X59),s(X56,X57))),s(fun(X53,X60),X61),s(fun(X54,X60),X62))) = s(X60,happ(s(fun(X53,X60),X61),s(X53,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X55,X56),'const.sum.INR_1'(s(X56,X57))),s(fun(X55,'type.sum.sum'(X53,X54)),X58),s(fun(X56,'type.sum.sum'(X53,X54)),X59)),X54,X53)))) | s('type.sum.sum'(X53,X54),'const.sum.sum_CASE_3'(s('type.sum.sum'(X55,X56),'const.sum.INR_1'(s(X56,X57))),s(fun(X55,'type.sum.sum'(X53,X54)),X58),s(fun(X56,'type.sum.sum'(X53,X54)),X59))) = s('type.sum.sum'(X53,X54),'const.sum.INR_1'(s(X54,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X55,X56),'const.sum.INR_1'(s(X56,X57))),s(fun(X55,'type.sum.sum'(X53,X54)),X58),s(fun(X56,'type.sum.sum'(X53,X54)),X59)),X54,X53))))) )),
  inference(superposition,[],[f332,f37])).
fof(f3831,plain,(
  ( ! [X47,X45,X43,X52,X50,X48,X46,X44,X51,X49] : (s('type.sum.sum'(X43,X44),happ(s(fun(X45,'type.sum.sum'(X43,X44)),X48),s(X45,X47))) = s('type.sum.sum'(X43,X44),'const.sum.INR_1'(s(X44,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X46),'const.sum.INL_1'(s(X45,X47))),s(fun(X45,'type.sum.sum'(X43,X44)),X48),s(fun(X46,'type.sum.sum'(X43,X44)),X49)),X44,X43)))) | s(X50,'const.sum.sum_CASE_3'(s('type.sum.sum'(X43,X44),happ(s(fun(X45,'type.sum.sum'(X43,X44)),X48),s(X45,X47))),s(fun(X43,X50),X51),s(fun(X44,X50),X52))) = s(X50,happ(s(fun(X43,X50),X51),s(X43,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X46),'const.sum.INL_1'(s(X45,X47))),s(fun(X45,'type.sum.sum'(X43,X44)),X48),s(fun(X46,'type.sum.sum'(X43,X44)),X49)),X44,X43))))) )),
  inference(forward_demodulation,[],[f3692,f36])).
fof(f3692,plain,(
  ( ! [X47,X45,X43,X52,X50,X48,X46,X44,X51,X49] : (s(X50,'const.sum.sum_CASE_3'(s('type.sum.sum'(X43,X44),happ(s(fun(X45,'type.sum.sum'(X43,X44)),X48),s(X45,X47))),s(fun(X43,X50),X51),s(fun(X44,X50),X52))) = s(X50,happ(s(fun(X43,X50),X51),s(X43,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X46),'const.sum.INL_1'(s(X45,X47))),s(fun(X45,'type.sum.sum'(X43,X44)),X48),s(fun(X46,'type.sum.sum'(X43,X44)),X49)),X44,X43)))) | s('type.sum.sum'(X43,X44),'const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X46),'const.sum.INL_1'(s(X45,X47))),s(fun(X45,'type.sum.sum'(X43,X44)),X48),s(fun(X46,'type.sum.sum'(X43,X44)),X49))) = s('type.sum.sum'(X43,X44),'const.sum.INR_1'(s(X44,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X45,X46),'const.sum.INL_1'(s(X45,X47))),s(fun(X45,'type.sum.sum'(X43,X44)),X48),s(fun(X46,'type.sum.sum'(X43,X44)),X49)),X44,X43))))) )),
  inference(superposition,[],[f332,f36])).
fof(f3691,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(X40,'const.sum.sum_CASE_3'(s('type.sum.sum'(X38,bool),X39),s(fun(X38,X40),X41),s(fun(bool,X40),X42))) = s(X40,happ(s(fun(X38,X40),X41),s(X38,sK9('const.sum.INL_1'(s(X38,sK9(X39,bool,X38))),bool,X38)))) | s('type.sum.sum'(X38,bool),'const.sum.INL_1'(s(X38,sK9(X39,bool,X38)))) = s('type.sum.sum'(X38,bool),'const.sum.INR_1'(s(bool,sK10('const.sum.INL_1'(s(X38,sK9(X39,bool,X38))),bool,X38)))) | s('type.sum.sum'(X38,bool),X39) = s('type.sum.sum'(X38,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X39,bool,X38))) )),
  inference(superposition,[],[f332,f321])).
fof(f3690,plain,(
  ( ! [X37,X35,X33,X36,X34,X32] : (s(X35,'const.sum.sum_CASE_3'(s('type.sum.sum'(X32,X33),X34),s(fun(X32,X35),X36),s(fun(X33,X35),X37))) = s(X35,happ(s(fun(X32,X35),X36),s(X32,sK9('const.sum.INL_1'(s(X32,sK9(X34,X33,X32))),X33,X32)))) | s('type.sum.sum'(X32,X33),'const.sum.INL_1'(s(X32,sK9(X34,X33,X32)))) = s('type.sum.sum'(X32,X33),'const.sum.INR_1'(s(X33,sK10('const.sum.INL_1'(s(X32,sK9(X34,X33,X32))),X33,X32)))) | s('type.sum.sum'(X32,X33),X34) = s('type.sum.sum'(X32,X33),'const.sum.INR_1'(s(X33,sK10(X34,X33,X32))))) )),
  inference(superposition,[],[f332,f34])).
fof(f3689,plain,(
  ( ! [X30,X28,X26,X24,X23,X31,X29,X27,X25] : (s(X26,'const.sum.sum_CASE_3'(s('type.sum.sum'(X23,X24),X25),s(fun(X23,X26),X27),s(fun(X24,X26),X28))) = s(X26,happ(s(fun(X23,X26),X27),s(X23,sK9('const.sum.INL_1'(s(X23,sK9(X25,X24,X23))),X24,X23)))) | s('type.sum.sum'(X23,X24),'const.sum.INL_1'(s(X23,sK9(X25,X24,X23)))) = s('type.sum.sum'(X23,X24),'const.sum.INR_1'(s(X24,sK10('const.sum.INL_1'(s(X23,sK9(X25,X24,X23))),X24,X23)))) | s(X29,happ(s(fun(X24,X29),X30),s(X24,sK10(X25,X24,X23)))) = s(X29,'const.sum.sum_CASE_3'(s('type.sum.sum'(X23,X24),X25),s(fun(X23,X29),X31),s(fun(X24,X29),X30)))) )),
  inference(superposition,[],[f332,f328])).
fof(f3688,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(X20,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X18),X19),s(fun(bool,X20),X21),s(fun(X18,X20),X22))) = s(X20,happ(s(fun(bool,X20),X21),s(bool,sK9('const.sum.INR_1'(s(X18,sK10(X19,X18,bool))),X18,bool)))) | s('type.sum.sum'(bool,X18),'const.sum.INR_1'(s(X18,sK10(X19,X18,bool)))) = s('type.sum.sum'(bool,X18),'const.sum.INR_1'(s(X18,sK10('const.sum.INR_1'(s(X18,sK10(X19,X18,bool))),X18,bool)))) | s('type.sum.sum'(bool,X18),X19) = s('type.sum.sum'(bool,X18),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X19,X18,bool))) )),
  inference(superposition,[],[f332,f326])).
fof(f3687,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(X15,'const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),X14),s(fun(X12,X15),X16),s(fun(X13,X15),X17))) = s(X15,happ(s(fun(X12,X15),X16),s(X12,sK9('const.sum.INR_1'(s(X13,sK10(X14,X13,X12))),X13,X12)))) | s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,sK10(X14,X13,X12)))) = s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,sK10('const.sum.INR_1'(s(X13,sK10(X14,X13,X12))),X13,X12)))) | s('type.sum.sum'(X12,X13),X14) = s('type.sum.sum'(X12,X13),'const.sum.INL_1'(s(X12,sK9(X14,X13,X12))))) )),
  inference(superposition,[],[f332,f34])).
fof(f3686,plain,(
  ( ! [X6,X4,X10,X8,X7,X5,X3,X11,X9] : (s(X6,'const.sum.sum_CASE_3'(s('type.sum.sum'(X3,X4),X5),s(fun(X3,X6),X7),s(fun(X4,X6),X8))) = s(X6,happ(s(fun(X3,X6),X7),s(X3,sK9('const.sum.INR_1'(s(X4,sK10(X5,X4,X3))),X4,X3)))) | s('type.sum.sum'(X3,X4),'const.sum.INR_1'(s(X4,sK10(X5,X4,X3)))) = s('type.sum.sum'(X3,X4),'const.sum.INR_1'(s(X4,sK10('const.sum.INR_1'(s(X4,sK10(X5,X4,X3))),X4,X3)))) | s(X9,happ(s(fun(X3,X9),X10),s(X3,sK9(X5,X4,X3)))) = s(X9,'const.sum.sum_CASE_3'(s('type.sum.sum'(X3,X4),X5),s(fun(X3,X9),X10),s(fun(X4,X9),X11)))) )),
  inference(superposition,[],[f332,f332])).
fof(f3830,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK0,X0),X1),s(sK0,sK9(sK6,sK1,sK0)))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X0),X1),s(fun(sK1,X0),X2)))) ) | $spl40),
  inference(subsumption_resolution,[],[f3829,f3661])).
fof(f3680,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X75,X73] : (s(bool,X72) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X68,X71),X70),s(fun(X68,bool),X69),s(fun(X71,bool),X73))) | s('type.sum.sum'(X68,X71),X70) = s('type.sum.sum'(X68,X71),'const.sum.INR_1'(s(X71,sK10(X70,X71,X68)))) | s(bool,X72) = s(bool,X74) | s(bool,'T_0') = s(bool,happ(s(fun(X68,bool),X69),s(X68,sK9(X70,X71,X68)))) | s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f332,f78])).
fof(f3679,plain,(
  ( ! [X61,X66,X64,X62,X60,X67,X65,X63] : (s(bool,X64) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X60,X63),X62),s(fun(X60,bool),X61),s(fun(X63,bool),X65))) | s('type.sum.sum'(X60,X63),X62) = s('type.sum.sum'(X60,X63),'const.sum.INR_1'(s(X63,sK10(X62,X63,X60)))) | s(bool,X66) = s(bool,happ(s(fun(X60,bool),X61),s(X60,sK9(X62,X63,X60)))) | s(bool,'T_0') = s(bool,X64) | s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f332,f78])).
fof(f3678,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X52,X55),X54),s(fun(X52,bool),X53),s(fun(X55,bool),X56))) | s('type.sum.sum'(X52,X55),X54) = s('type.sum.sum'(X52,X55),'const.sum.INR_1'(s(X55,sK10(X54,X55,X52)))) | s(bool,X57) = s(bool,X58) | s(bool,X58) = s(bool,happ(s(fun(X52,bool),X53),s(X52,sK9(X54,X55,X52)))) | s(bool,X57) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f332,f78])).
fof(f3676,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (s(bool,X40) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X36,X39),X38),s(fun(X36,bool),X37),s(fun(X39,bool),X41))) | s('type.sum.sum'(X36,X39),X38) = s('type.sum.sum'(X36,X39),'const.sum.INR_1'(s(X39,sK10(X38,X39,X36)))) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(bool,X43) = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK9(X38,X39,X36)))) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f332,f78])).
fof(f3673,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X22) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X18,X21),X20),s(fun(X18,bool),X19),s(fun(X21,bool),X23))) | s('type.sum.sum'(X18,X21),X20) = s('type.sum.sum'(X18,X21),'const.sum.INR_1'(s(X21,sK10(X20,X21,X18)))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun(X18,bool),X19),s(X18,sK9(X20,X21,X18))))) )),
  inference(superposition,[],[f332,f41])).
fof(f3672,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X15),X14),s(fun(X12,bool),X13),s(fun(X15,bool),X16))) | s('type.sum.sum'(X12,X15),X14) = s('type.sum.sum'(X12,X15),'const.sum.INR_1'(s(X15,sK10(X14,X15,X12)))) | s(bool,X17) = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK9(X14,X15,X12)))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f332,f41])).
fof(f3671,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(bool,'F_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X7,X10),X9),s(fun(X7,bool),X8),s(fun(X10,bool),X11))) | s('type.sum.sum'(X7,X10),X9) = s('type.sum.sum'(X7,X10),'const.sum.INR_1'(s(X10,sK10(X9,X10,X7)))) | s(bool,'T_0') = s(bool,happ(s(fun(X7,bool),X8),s(X7,sK9(X9,X10,X7))))) )),
  inference(superposition,[],[f332,f33])).
fof(f3827,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X1,X4),X3),s(fun(X1,X0),X2),s(fun(X4,X0),X5))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X1,X4),X3),s(fun(X1,X0),X2),s(fun(X4,X0),X6))) | s('type.sum.sum'(X1,X4),X3) = s('type.sum.sum'(X1,X4),'const.sum.INR_1'(s(X4,sK10(X3,X4,X1))))) )),
  inference(duplicate_literal_removal,[],[f3670])).
fof(f3670,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X1,X4),X3),s(fun(X1,X0),X2),s(fun(X4,X0),X5))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X1,X4),X3),s(fun(X1,X0),X2),s(fun(X4,X0),X6))) | s('type.sum.sum'(X1,X4),X3) = s('type.sum.sum'(X1,X4),'const.sum.INR_1'(s(X4,sK10(X3,X4,X1)))) | s('type.sum.sum'(X1,X4),X3) = s('type.sum.sum'(X1,X4),'const.sum.INR_1'(s(X4,sK10(X3,X4,X1))))) )),
  inference(superposition,[],[f332,f332])).
fof(f3664,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) | $spl40),
  inference(forward_demodulation,[],[f3660,f27])).
fof(f3660,plain,(
  s('type.sum.sum'(sK0,sK1),sK6) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) | $spl40),
  inference(unit_resulting_resolution,[],[f27,f3659,f327])).
fof(f3661,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK6,sK1,sK0)))) | $spl40),
  inference(unit_resulting_resolution,[],[f3659,f39])).
fof(f3659,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK6,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK6,sK1,sK0)))) | $spl40),
  inference(forward_demodulation,[],[f3657,f3644])).
fof(f3657,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK6,sK1,sK0)))) != s(sK2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK2),sK4),s(fun(sK1,sK2),sK3))) | $spl40),
  inference(backward_demodulation,[],[f3644,f40])).
fof(f3638,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X42,X40] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'('type.sum.sum'(X35,X36),X38),'const.sum.INL_1'(s('type.sum.sum'(X35,X36),X37))),s(fun('type.sum.sum'(X35,X36),bool),X39),s(fun(X38,bool),X40))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(X35,X36),bool),X39),s('type.sum.sum'(X35,X36),'const.sum.INL_1'(s(X35,sK9(X37,X36,X35)))))) | s(X41,happ(s(fun(X36,X41),X42),s(X36,sK10(X37,X36,X35)))) = s(X41,'const.sum.sum_CASE_3'(s('type.sum.sum'(X35,X36),X37),s(fun(X35,X41),X43),s(fun(X36,X41),X42)))) )),
  inference(superposition,[],[f259,f328])).
fof(f3637,plain,(
  ( ! [X30,X28,X26,X33,X31,X29,X27,X34,X32] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X29,'type.sum.sum'(X26,X27)),'const.sum.INR_1'(s('type.sum.sum'(X26,X27),X28))),s(fun(X29,bool),X30),s(fun('type.sum.sum'(X26,X27),bool),X31))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(X26,X27),bool),X31),s('type.sum.sum'(X26,X27),'const.sum.INL_1'(s(X26,sK9(X28,X27,X26)))))) | s(X32,happ(s(fun(X27,X32),X33),s(X27,sK10(X28,X27,X26)))) = s(X32,'const.sum.sum_CASE_3'(s('type.sum.sum'(X26,X27),X28),s(fun(X26,X32),X34),s(fun(X27,X32),X33)))) )),
  inference(superposition,[],[f296,f328])).
fof(f3636,plain,(
  ( ! [X23,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X18,happ(s(fun(X15,X18),X19),s(X15,sK9(X17,X16,X15)))) = s(X18,'const.sum.sum_CASE_3'(s('type.sum.sum'(X15,X16),X17),s(fun(X15,X18),X19),s(fun(X16,X18),X20))) | s(X21,happ(s(fun(X16,X21),X22),s(X16,sK10(X17,X16,X15)))) = s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X15,X16),X17),s(fun(X15,X21),X23),s(fun(X16,X21),X22)))) )),
  inference(superposition,[],[f36,f328])).
fof(f3635,plain,(
  ( ! [X14,X12,X10,X8,X7,X13,X11,X9] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X7,X8),X9),s(fun(X7,bool),X10),s(fun(X8,bool),X11))) | s(bool,'F_0') = s(bool,happ(s(fun(X7,bool),X10),s(X7,sK9(X9,X8,X7)))) | s(X12,happ(s(fun(X8,X12),X13),s(X8,sK10(X9,X8,X7)))) = s(X12,'const.sum.sum_CASE_3'(s('type.sum.sum'(X7,X8),X9),s(fun(X7,X12),X14),s(fun(X8,X12),X13)))) )),
  inference(superposition,[],[f259,f328])).
fof(f3634,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),X0) | s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(X0,sK1,sK0)))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,sK9(X0,sK1,sK0)))) | s(X1,happ(s(fun(sK1,X1),X2),s(sK1,sK10(X0,sK1,sK0)))) = s(X1,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),X0),s(fun(sK0,X1),X3),s(fun(sK1,X1),X2)))) )),
  inference(superposition,[],[f38,f328])).
fof(f3632,plain,(
  ( ! [X261,X263,X260,X262,X264] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,happ(s(fun(X261,sK1),X264),s(X261,sK10(X262,X261,X260)))))) | s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X260,X261),X262),s(fun(X260,sK1),X263),s(fun(X261,sK1),X264))))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X260,X261),X262),s(fun(X260,sK1),X263),s(fun(X261,sK1),X264))))) | s('type.sum.sum'(X260,X261),X262) = s('type.sum.sum'(X260,X261),'const.sum.INL_1'(s(X260,sK9(X262,X261,X260))))) )),
  inference(superposition,[],[f39,f328])).
fof(f3631,plain,(
  ( ! [X255,X257,X259,X256,X258] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,happ(s(fun(X256,sK0),X259),s(X256,sK10(X257,X256,X255)))))) | s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X255,X256),X257),s(fun(X255,sK0),X258),s(fun(X256,sK0),X259))))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X255,X256),X257),s(fun(X255,sK0),X258),s(fun(X256,sK0),X259))))) | s('type.sum.sum'(X255,X256),X257) = s('type.sum.sum'(X255,X256),'const.sum.INL_1'(s(X255,sK9(X257,X256,X255))))) )),
  inference(superposition,[],[f38,f328])).
fof(f3630,plain,(
  ( ! [X253,X251,X249,X246,X254,X252,X250,X248,X247] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X246,X252),'const.sum.INL_1'(s(X246,X253))),s(fun(X246,bool),happ(s(fun(X248,fun(X246,bool)),X251),s(X248,sK10(X249,X248,X247)))),s(fun(X252,bool),X254))) | s(bool,'F_0') = s(bool,happ(s(fun(X246,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X247,X248),X249),s(fun(X247,fun(X246,bool)),X250),s(fun(X248,fun(X246,bool)),X251))),s(X246,X253))) | s('type.sum.sum'(X247,X248),X249) = s('type.sum.sum'(X247,X248),'const.sum.INL_1'(s(X247,sK9(X249,X248,X247))))) )),
  inference(superposition,[],[f259,f328])).
fof(f3627,plain,(
  ( ! [X222,X220,X227,X225,X219,X223,X221,X226,X224] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X225,X219),'const.sum.INR_1'(s(X219,X226))),s(fun(X225,bool),X227),s(fun(X219,bool),happ(s(fun(X221,fun(X219,bool)),X224),s(X221,sK10(X222,X221,X220)))))) | s(bool,'F_0') = s(bool,happ(s(fun(X219,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X220,X221),X222),s(fun(X220,fun(X219,bool)),X223),s(fun(X221,fun(X219,bool)),X224))),s(X219,X226))) | s('type.sum.sum'(X220,X221),X222) = s('type.sum.sum'(X220,X221),'const.sum.INL_1'(s(X220,sK9(X222,X221,X220))))) )),
  inference(superposition,[],[f296,f328])).
fof(f3619,plain,(
  ( ! [X158,X163,X161,X159,X162,X160] : (p(s(bool,happ(s(fun(X159,bool),X162),s(X159,sK10(X160,X159,X158))))) | s(bool,X163) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X158,X159),X160),s(fun(X158,bool),X161),s(fun(X159,bool),X162))) | s(bool,'T_0') = s(bool,X163) | s('type.sum.sum'(X158,X159),X160) = s('type.sum.sum'(X158,X159),'const.sum.INL_1'(s(X158,sK9(X160,X159,X158))))) )),
  inference(superposition,[],[f110,f328])).
fof(f3610,plain,(
  ( ! [X92,X90,X93,X91,X89] : (~p(s(bool,happ(s(fun(X90,bool),X93),s(X90,sK10(X91,X90,X89))))) | s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X89,X90),X91),s(fun(X89,bool),X92),s(fun(X90,bool),X93))) | s('type.sum.sum'(X89,X90),X91) = s('type.sum.sum'(X89,X90),'const.sum.INL_1'(s(X89,sK9(X91,X90,X89))))) )),
  inference(superposition,[],[f43,f328])).
fof(f3600,plain,(
  ( ! [X28,X26,X24,X23,X21,X29,X27,X25,X22] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X21,X27),'const.sum.INL_1'(s(X21,happ(s(fun(X23,X21),X26),s(X23,sK10(X24,X23,X22)))))),s(fun(X21,bool),X28),s(fun(X27,bool),X29))) | s(bool,'F_0') = s(bool,happ(s(fun(X21,bool),X28),s(X21,'const.sum.sum_CASE_3'(s('type.sum.sum'(X22,X23),X24),s(fun(X22,X21),X25),s(fun(X23,X21),X26))))) | s('type.sum.sum'(X22,X23),X24) = s('type.sum.sum'(X22,X23),'const.sum.INL_1'(s(X22,sK9(X24,X23,X22))))) )),
  inference(superposition,[],[f259,f328])).
fof(f3599,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X18,X12),'const.sum.INR_1'(s(X12,happ(s(fun(X14,X12),X17),s(X14,sK10(X15,X14,X13)))))),s(fun(X18,bool),X19),s(fun(X12,bool),X20))) | s(bool,'F_0') = s(bool,happ(s(fun(X12,bool),X20),s(X12,'const.sum.sum_CASE_3'(s('type.sum.sum'(X13,X14),X15),s(fun(X13,X12),X16),s(fun(X14,X12),X17))))) | s('type.sum.sum'(X13,X14),X15) = s('type.sum.sum'(X13,X14),'const.sum.INL_1'(s(X13,sK9(X15,X14,X13))))) )),
  inference(superposition,[],[f296,f328])).
fof(f3596,plain,(
  ( ! [X241,X239,X242,X240,X238] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,'const.sum.sum_CASE_3'(s('type.sum.sum'(X241,X238),X240),s(fun(X241,sK1),X242),s(fun(X238,sK1),X239))))) | s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,happ(s(fun(X238,sK1),X239),s(X238,sK10(X240,X238,X241)))))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,happ(s(fun(X238,sK1),X239),s(X238,sK10(X240,X238,X241)))))) | s('type.sum.sum'(X241,X238),X240) = s('type.sum.sum'(X241,X238),'const.sum.INL_1'(s(X241,sK9(X240,X238,X241))))) )),
  inference(superposition,[],[f39,f328])).
fof(f3595,plain,(
  ( ! [X237,X235,X233,X236,X234] : (s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X236,X233),X235),s(fun(X236,sK0),X237),s(fun(X233,sK0),X234))))) | s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,happ(s(fun(X233,sK0),X234),s(X233,sK10(X235,X233,X236)))))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,happ(s(fun(X233,sK0),X234),s(X233,sK10(X235,X233,X236)))))) | s('type.sum.sum'(X236,X233),X235) = s('type.sum.sum'(X236,X233),'const.sum.INL_1'(s(X236,sK9(X235,X233,X236))))) )),
  inference(superposition,[],[f38,f328])).
fof(f3594,plain,(
  ( ! [X231,X229,X227,X225,X230,X228,X226,X224,X232] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X224,X230),'const.sum.INL_1'(s(X224,X231))),s(fun(X224,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X228,X225),X227),s(fun(X228,fun(X224,bool)),X229),s(fun(X225,fun(X224,bool)),X226))),s(fun(X230,bool),X232))) | s(bool,'F_0') = s(bool,happ(s(fun(X224,bool),happ(s(fun(X225,fun(X224,bool)),X226),s(X225,sK10(X227,X225,X228)))),s(X224,X231))) | s('type.sum.sum'(X228,X225),X227) = s('type.sum.sum'(X228,X225),'const.sum.INL_1'(s(X228,sK9(X227,X225,X228))))) )),
  inference(superposition,[],[f259,f328])).
fof(f3591,plain,(
  ( ! [X198,X204,X202,X200,X199,X197,X205,X203,X201] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X203,X197),'const.sum.INR_1'(s(X197,X204))),s(fun(X203,bool),X205),s(fun(X197,bool),'const.sum.sum_CASE_3'(s('type.sum.sum'(X201,X198),X200),s(fun(X201,fun(X197,bool)),X202),s(fun(X198,fun(X197,bool)),X199))))) | s(bool,'F_0') = s(bool,happ(s(fun(X197,bool),happ(s(fun(X198,fun(X197,bool)),X199),s(X198,sK10(X200,X198,X201)))),s(X197,X204))) | s('type.sum.sum'(X201,X198),X200) = s('type.sum.sum'(X201,X198),'const.sum.INL_1'(s(X201,sK9(X200,X198,X201))))) )),
  inference(superposition,[],[f296,f328])).
fof(f3583,plain,(
  ( ! [X140,X138,X136,X141,X139,X137] : (p(s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X139,X136),X138),s(fun(X139,bool),X140),s(fun(X136,bool),X137)))) | s(bool,X141) = s(bool,happ(s(fun(X136,bool),X137),s(X136,sK10(X138,X136,X139)))) | s(bool,'T_0') = s(bool,X141) | s('type.sum.sum'(X139,X136),X138) = s('type.sum.sum'(X139,X136),'const.sum.INL_1'(s(X139,sK9(X138,X136,X139))))) )),
  inference(superposition,[],[f110,f328])).
fof(f3574,plain,(
  ( ! [X70,X68,X71,X69,X67] : (~p(s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X70,X67),X69),s(fun(X70,bool),X71),s(fun(X67,bool),X68)))) | s(bool,'T_0') = s(bool,happ(s(fun(X67,bool),X68),s(X67,sK10(X69,X67,X70)))) | s('type.sum.sum'(X70,X67),X69) = s('type.sum.sum'(X70,X67),'const.sum.INL_1'(s(X70,sK9(X69,X67,X70))))) )),
  inference(superposition,[],[f43,f328])).
fof(f3566,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X9,X16] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X9,X15),'const.sum.INL_1'(s(X9,'const.sum.sum_CASE_3'(s('type.sum.sum'(X13,X10),X12),s(fun(X13,X9),X14),s(fun(X10,X9),X11))))),s(fun(X9,bool),X16),s(fun(X15,bool),X17))) | s(bool,'F_0') = s(bool,happ(s(fun(X9,bool),X16),s(X9,happ(s(fun(X10,X9),X11),s(X10,sK10(X12,X10,X13)))))) | s('type.sum.sum'(X13,X10),X12) = s('type.sum.sum'(X13,X10),'const.sum.INL_1'(s(X13,sK9(X12,X10,X13))))) )),
  inference(superposition,[],[f259,f328])).
fof(f3565,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,X0),'const.sum.INR_1'(s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X4,X1),X3),s(fun(X4,X0),X5),s(fun(X1,X0),X2))))),s(fun(X6,bool),X7),s(fun(X0,bool),X8))) | s(bool,'F_0') = s(bool,happ(s(fun(X0,bool),X8),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK10(X3,X1,X4)))))) | s('type.sum.sum'(X4,X1),X3) = s('type.sum.sum'(X4,X1),'const.sum.INL_1'(s(X4,sK9(X3,X1,X4))))) )),
  inference(superposition,[],[f296,f328])).
fof(f3560,plain,(
  ( ! [X61,X68,X66,X64,X62,X67,X65,X63] : (s('type.sum.sum'(bool,X62),X61) = s('type.sum.sum'(bool,X62),'const.sum.INL_1'(s(bool,X63))) | s(X64,happ(s(fun(X62,X64),X65),s(X62,sK10(X61,X62,bool)))) = s(X64,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X62),X61),s(fun(bool,X64),X66),s(fun(X62,X64),X65))) | s(bool,X63) = s(bool,X67) | s(bool,'T_0') = s(bool,sK9(X61,X62,bool)) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f328,f78])).
fof(f3559,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55,X53] : (s('type.sum.sum'(bool,X54),X53) = s('type.sum.sum'(bool,X54),'const.sum.INL_1'(s(bool,X55))) | s(X56,happ(s(fun(X54,X56),X57),s(X54,sK10(X53,X54,bool)))) = s(X56,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X54),X53),s(fun(bool,X56),X58),s(fun(X54,X56),X57))) | s(bool,X59) = s(bool,sK9(X53,X54,bool)) | s(bool,'T_0') = s(bool,X55) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f328,f78])).
fof(f3558,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X51,X49] : (s('type.sum.sum'(bool,X46),X45) = s('type.sum.sum'(bool,X46),'const.sum.INL_1'(s(bool,'T_0'))) | s(X47,happ(s(fun(X46,X47),X48),s(X46,sK10(X45,X46,bool)))) = s(X47,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X46),X45),s(fun(bool,X47),X49),s(fun(X46,X47),X48))) | s(bool,X50) = s(bool,X51) | s(bool,X51) = s(bool,sK9(X45,X46,bool)) | s(bool,X50) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f328,f78])).
fof(f3556,plain,(
  ( ! [X30,X35,X33,X31,X29,X36,X34,X32] : (s('type.sum.sum'(bool,X30),X29) = s('type.sum.sum'(bool,X30),'const.sum.INL_1'(s(bool,X31))) | s(X32,happ(s(fun(X30,X32),X33),s(X30,sK10(X29,X30,bool)))) = s(X32,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X30),X29),s(fun(bool,X32),X34),s(fun(X30,X32),X33))) | s(bool,X31) = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,sK9(X29,X30,bool)) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f328,f78])).
fof(f3553,plain,(
  ( ! [X14,X12,X15,X13,X11,X16] : (s('type.sum.sum'(bool,X12),X11) = s('type.sum.sum'(bool,X12),'const.sum.INL_1'(s(bool,X13))) | s(X14,happ(s(fun(X12,X14),X15),s(X12,sK10(X11,X12,bool)))) = s(X14,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X12),X11),s(fun(bool,X14),X16),s(fun(X12,X14),X15))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,sK9(X11,X12,bool))) )),
  inference(superposition,[],[f328,f41])).
fof(f3552,plain,(
  ( ! [X6,X10,X8,X7,X5,X9] : (s('type.sum.sum'(bool,X6),X5) = s('type.sum.sum'(bool,X6),'const.sum.INL_1'(s(bool,'T_0'))) | s(X7,happ(s(fun(X6,X7),X8),s(X6,sK10(X5,X6,bool)))) = s(X7,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X6),X5),s(fun(bool,X7),X9),s(fun(X6,X7),X8))) | s(bool,X10) = s(bool,sK9(X5,X6,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f328,f41])).
fof(f3536,plain,(
  ( ! [X24,X23,X27,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(X24,bool),X27),s(X24,sK10('const.sum.INR_1'(s(X24,X25)),X24,X23)))) | s('type.sum.sum'(X23,X24),'const.sum.INR_1'(s(X24,X25))) = s('type.sum.sum'(X23,X24),'const.sum.INL_1'(s(X23,sK9('const.sum.INR_1'(s(X24,X25)),X24,X23)))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X27),s(X24,X25)))) )),
  inference(superposition,[],[f328,f296])).
fof(f3535,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X22),s(X19,sK10('const.sum.INL_1'(s(X18,X20)),X19,X18)))) | s('type.sum.sum'(X18,X19),'const.sum.INL_1'(s(X18,X20))) = s('type.sum.sum'(X18,X19),'const.sum.INL_1'(s(X18,sK9('const.sum.INL_1'(s(X18,X20)),X19,X18)))) | s(bool,'F_0') = s(bool,happ(s(fun(X18,bool),X21),s(X18,X20)))) )),
  inference(superposition,[],[f328,f259])).
fof(f3534,plain,(
  ( ! [X6,X8,X7,X11,X9] : (s(X6,happ(s(fun(X8,X6),X11),s(X8,X9))) = s(X6,happ(s(fun(X8,X6),X11),s(X8,sK10('const.sum.INR_1'(s(X8,X9)),X8,X7)))) | s('type.sum.sum'(X7,X8),'const.sum.INR_1'(s(X8,X9))) = s('type.sum.sum'(X7,X8),'const.sum.INL_1'(s(X7,sK9('const.sum.INR_1'(s(X8,X9)),X8,X7))))) )),
  inference(superposition,[],[f328,f37])).
fof(f3533,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X4),s(X1,X3))) = s(X0,happ(s(fun(X2,X0),X5),s(X2,sK10('const.sum.INL_1'(s(X1,X3)),X2,X1)))) | s('type.sum.sum'(X1,X2),'const.sum.INL_1'(s(X1,X3))) = s('type.sum.sum'(X1,X2),'const.sum.INL_1'(s(X1,sK9('const.sum.INL_1'(s(X1,X3)),X2,X1))))) )),
  inference(superposition,[],[f328,f36])).
fof(f3524,plain,(
  ( ! [X70,X68,X66,X64,X72,X71,X69,X67,X65,X73] : (s(X71,happ(s(fun(X65,X71),X72),s(X65,sK10(happ(s(fun(X66,'type.sum.sum'(X64,X65)),X67),s(X66,sK10(X68,X66,X69))),X65,X64)))) = s(X71,'const.sum.sum_CASE_3'(s('type.sum.sum'(X64,X65),'const.sum.sum_CASE_3'(s('type.sum.sum'(X69,X66),X68),s(fun(X69,'type.sum.sum'(X64,X65)),X70),s(fun(X66,'type.sum.sum'(X64,X65)),X67))),s(fun(X64,X71),X73),s(fun(X65,X71),X72))) | s('type.sum.sum'(X64,X65),happ(s(fun(X66,'type.sum.sum'(X64,X65)),X67),s(X66,sK10(X68,X66,X69)))) = s('type.sum.sum'(X64,X65),'const.sum.INL_1'(s(X64,sK9(happ(s(fun(X66,'type.sum.sum'(X64,X65)),X67),s(X66,sK10(X68,X66,X69))),X65,X64)))) | s('type.sum.sum'(X69,X66),X68) = s('type.sum.sum'(X69,X66),'const.sum.INL_1'(s(X69,sK9(X68,X66,X69))))) )),
  inference(superposition,[],[f328,f328])).
fof(f3523,plain,(
  ( ! [X61,X59,X57,X54,X62,X60,X58,X56,X55,X63] : (s(X61,'const.sum.sum_CASE_3'(s('type.sum.sum'(X54,X55),happ(s(fun(X57,'type.sum.sum'(X54,X55)),X60),s(X57,sK10(X58,X57,X56)))),s(fun(X54,X61),X63),s(fun(X55,X61),X62))) = s(X61,happ(s(fun(X55,X61),X62),s(X55,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X56,X57),X58),s(fun(X56,'type.sum.sum'(X54,X55)),X59),s(fun(X57,'type.sum.sum'(X54,X55)),X60)),X55,X54)))) | s('type.sum.sum'(X54,X55),'const.sum.sum_CASE_3'(s('type.sum.sum'(X56,X57),X58),s(fun(X56,'type.sum.sum'(X54,X55)),X59),s(fun(X57,'type.sum.sum'(X54,X55)),X60))) = s('type.sum.sum'(X54,X55),'const.sum.INL_1'(s(X54,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X56,X57),X58),s(fun(X56,'type.sum.sum'(X54,X55)),X59),s(fun(X57,'type.sum.sum'(X54,X55)),X60)),X55,X54)))) | s('type.sum.sum'(X56,X57),X58) = s('type.sum.sum'(X56,X57),'const.sum.INL_1'(s(X56,sK9(X58,X57,X56))))) )),
  inference(superposition,[],[f328,f328])).
fof(f3648,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X53,X51,X49] : (s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48))) = s('type.sum.sum'(X44,X45),'const.sum.INL_1'(s(X44,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X45,X44)))) | s(X51,'const.sum.sum_CASE_3'(s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48))),s(fun(X44,X51),X53),s(fun(X45,X51),X52))) = s(X51,happ(s(fun(X45,X51),X52),s(X45,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X45,X44))))) )),
  inference(forward_demodulation,[],[f3522,f37])).
fof(f3522,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X53,X51,X49] : (s(X51,'const.sum.sum_CASE_3'(s('type.sum.sum'(X44,X45),happ(s(fun(X47,'type.sum.sum'(X44,X45)),X50),s(X47,X48))),s(fun(X44,X51),X53),s(fun(X45,X51),X52))) = s(X51,happ(s(fun(X45,X51),X52),s(X45,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X45,X44)))) | s('type.sum.sum'(X44,X45),'const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50))) = s('type.sum.sum'(X44,X45),'const.sum.INL_1'(s(X44,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X46,X47),'const.sum.INR_1'(s(X47,X48))),s(fun(X46,'type.sum.sum'(X44,X45)),X49),s(fun(X47,'type.sum.sum'(X44,X45)),X50)),X45,X44))))) )),
  inference(superposition,[],[f328,f37])).
fof(f3647,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (s('type.sum.sum'(X34,X35),happ(s(fun(X36,'type.sum.sum'(X34,X35)),X39),s(X36,X38))) = s('type.sum.sum'(X34,X35),'const.sum.INL_1'(s(X34,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X36,X37),'const.sum.INL_1'(s(X36,X38))),s(fun(X36,'type.sum.sum'(X34,X35)),X39),s(fun(X37,'type.sum.sum'(X34,X35)),X40)),X35,X34)))) | s(X41,'const.sum.sum_CASE_3'(s('type.sum.sum'(X34,X35),happ(s(fun(X36,'type.sum.sum'(X34,X35)),X39),s(X36,X38))),s(fun(X34,X41),X43),s(fun(X35,X41),X42))) = s(X41,happ(s(fun(X35,X41),X42),s(X35,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X36,X37),'const.sum.INL_1'(s(X36,X38))),s(fun(X36,'type.sum.sum'(X34,X35)),X39),s(fun(X37,'type.sum.sum'(X34,X35)),X40)),X35,X34))))) )),
  inference(forward_demodulation,[],[f3521,f36])).
fof(f3521,plain,(
  ( ! [X39,X37,X35,X43,X41,X38,X36,X34,X42,X40] : (s(X41,'const.sum.sum_CASE_3'(s('type.sum.sum'(X34,X35),happ(s(fun(X36,'type.sum.sum'(X34,X35)),X39),s(X36,X38))),s(fun(X34,X41),X43),s(fun(X35,X41),X42))) = s(X41,happ(s(fun(X35,X41),X42),s(X35,sK10('const.sum.sum_CASE_3'(s('type.sum.sum'(X36,X37),'const.sum.INL_1'(s(X36,X38))),s(fun(X36,'type.sum.sum'(X34,X35)),X39),s(fun(X37,'type.sum.sum'(X34,X35)),X40)),X35,X34)))) | s('type.sum.sum'(X34,X35),'const.sum.sum_CASE_3'(s('type.sum.sum'(X36,X37),'const.sum.INL_1'(s(X36,X38))),s(fun(X36,'type.sum.sum'(X34,X35)),X39),s(fun(X37,'type.sum.sum'(X34,X35)),X40))) = s('type.sum.sum'(X34,X35),'const.sum.INL_1'(s(X34,sK9('const.sum.sum_CASE_3'(s('type.sum.sum'(X36,X37),'const.sum.INL_1'(s(X36,X38))),s(fun(X36,'type.sum.sum'(X34,X35)),X39),s(fun(X37,'type.sum.sum'(X34,X35)),X40)),X35,X34))))) )),
  inference(superposition,[],[f328,f36])).
fof(f3520,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(X31,'const.sum.sum_CASE_3'(s('type.sum.sum'(X29,bool),X30),s(fun(X29,X31),X33),s(fun(bool,X31),X32))) = s(X31,happ(s(fun(bool,X31),X32),s(bool,sK10('const.sum.INL_1'(s(X29,sK9(X30,bool,X29))),bool,X29)))) | s('type.sum.sum'(X29,bool),'const.sum.INL_1'(s(X29,sK9(X30,bool,X29)))) = s('type.sum.sum'(X29,bool),'const.sum.INL_1'(s(X29,sK9('const.sum.INL_1'(s(X29,sK9(X30,bool,X29))),bool,X29)))) | s('type.sum.sum'(X29,bool),X30) = s('type.sum.sum'(X29,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X30,bool,X29))) )),
  inference(superposition,[],[f328,f321])).
fof(f3519,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s(X26,'const.sum.sum_CASE_3'(s('type.sum.sum'(X23,X24),X25),s(fun(X23,X26),X28),s(fun(X24,X26),X27))) = s(X26,happ(s(fun(X24,X26),X27),s(X24,sK10('const.sum.INL_1'(s(X23,sK9(X25,X24,X23))),X24,X23)))) | s('type.sum.sum'(X23,X24),'const.sum.INL_1'(s(X23,sK9(X25,X24,X23)))) = s('type.sum.sum'(X23,X24),'const.sum.INL_1'(s(X23,sK9('const.sum.INL_1'(s(X23,sK9(X25,X24,X23))),X24,X23)))) | s('type.sum.sum'(X23,X24),X25) = s('type.sum.sum'(X23,X24),'const.sum.INR_1'(s(X24,sK10(X25,X24,X23))))) )),
  inference(superposition,[],[f328,f34])).
fof(f3518,plain,(
  ( ! [X14,X21,X19,X17,X15,X22,X20,X18,X16] : (s(X17,'const.sum.sum_CASE_3'(s('type.sum.sum'(X14,X15),X16),s(fun(X14,X17),X19),s(fun(X15,X17),X18))) = s(X17,happ(s(fun(X15,X17),X18),s(X15,sK10('const.sum.INL_1'(s(X14,sK9(X16,X15,X14))),X15,X14)))) | s('type.sum.sum'(X14,X15),'const.sum.INL_1'(s(X14,sK9(X16,X15,X14)))) = s('type.sum.sum'(X14,X15),'const.sum.INL_1'(s(X14,sK9('const.sum.INL_1'(s(X14,sK9(X16,X15,X14))),X15,X14)))) | s(X20,happ(s(fun(X15,X20),X21),s(X15,sK10(X16,X15,X14)))) = s(X20,'const.sum.sum_CASE_3'(s('type.sum.sum'(X14,X15),X16),s(fun(X14,X20),X22),s(fun(X15,X20),X21)))) )),
  inference(superposition,[],[f328,f328])).
fof(f3517,plain,(
  ( ! [X12,X10,X13,X11,X9] : (s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X9),X10),s(fun(bool,X11),X13),s(fun(X9,X11),X12))) = s(X11,happ(s(fun(X9,X11),X12),s(X9,sK10('const.sum.INR_1'(s(X9,sK10(X10,X9,bool))),X9,bool)))) | s('type.sum.sum'(bool,X9),'const.sum.INR_1'(s(X9,sK10(X10,X9,bool)))) = s('type.sum.sum'(bool,X9),'const.sum.INL_1'(s(bool,sK9('const.sum.INR_1'(s(X9,sK10(X10,X9,bool))),X9,bool)))) | s('type.sum.sum'(bool,X9),X10) = s('type.sum.sum'(bool,X9),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X10,X9,bool))) )),
  inference(superposition,[],[f328,f326])).
fof(f3516,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(X6,'const.sum.sum_CASE_3'(s('type.sum.sum'(X3,X4),X5),s(fun(X3,X6),X8),s(fun(X4,X6),X7))) = s(X6,happ(s(fun(X4,X6),X7),s(X4,sK10('const.sum.INR_1'(s(X4,sK10(X5,X4,X3))),X4,X3)))) | s('type.sum.sum'(X3,X4),'const.sum.INR_1'(s(X4,sK10(X5,X4,X3)))) = s('type.sum.sum'(X3,X4),'const.sum.INL_1'(s(X3,sK9('const.sum.INR_1'(s(X4,sK10(X5,X4,X3))),X4,X3)))) | s('type.sum.sum'(X3,X4),X5) = s('type.sum.sum'(X3,X4),'const.sum.INL_1'(s(X3,sK9(X5,X4,X3))))) )),
  inference(superposition,[],[f328,f34])).
fof(f3510,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X75,X73] : (s(bool,X72) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X68),X70),s(fun(X71,bool),X73),s(fun(X68,bool),X69))) | s('type.sum.sum'(X71,X68),X70) = s('type.sum.sum'(X71,X68),'const.sum.INL_1'(s(X71,sK9(X70,X68,X71)))) | s(bool,X72) = s(bool,X74) | s(bool,'T_0') = s(bool,happ(s(fun(X68,bool),X69),s(X68,sK10(X70,X68,X71)))) | s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f328,f78])).
fof(f3509,plain,(
  ( ! [X61,X66,X64,X62,X60,X67,X65,X63] : (s(bool,X64) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X63,X60),X62),s(fun(X63,bool),X65),s(fun(X60,bool),X61))) | s('type.sum.sum'(X63,X60),X62) = s('type.sum.sum'(X63,X60),'const.sum.INL_1'(s(X63,sK9(X62,X60,X63)))) | s(bool,X66) = s(bool,happ(s(fun(X60,bool),X61),s(X60,sK10(X62,X60,X63)))) | s(bool,'T_0') = s(bool,X64) | s(bool,X66) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f328,f78])).
fof(f3508,plain,(
  ( ! [X59,X57,X54,X52,X58,X56,X55,X53] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X55,X52),X54),s(fun(X55,bool),X56),s(fun(X52,bool),X53))) | s('type.sum.sum'(X55,X52),X54) = s('type.sum.sum'(X55,X52),'const.sum.INL_1'(s(X55,sK9(X54,X52,X55)))) | s(bool,X57) = s(bool,X58) | s(bool,X58) = s(bool,happ(s(fun(X52,bool),X53),s(X52,sK10(X54,X52,X55)))) | s(bool,X57) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f328,f78])).
fof(f3506,plain,(
  ( ! [X39,X37,X43,X41,X38,X36,X42,X40] : (s(bool,X40) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X39,X36),X38),s(fun(X39,bool),X41),s(fun(X36,bool),X37))) | s('type.sum.sum'(X39,X36),X38) = s('type.sum.sum'(X39,X36),'const.sum.INL_1'(s(X39,sK9(X38,X36,X39)))) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(bool,X43) = s(bool,happ(s(fun(X36,bool),X37),s(X36,sK10(X38,X36,X39)))) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f328,f78])).
fof(f3503,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (s(bool,X22) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X21,X18),X20),s(fun(X21,bool),X23),s(fun(X18,bool),X19))) | s('type.sum.sum'(X21,X18),X20) = s('type.sum.sum'(X21,X18),'const.sum.INL_1'(s(X21,sK9(X20,X18,X21)))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,happ(s(fun(X18,bool),X19),s(X18,sK10(X20,X18,X21))))) )),
  inference(superposition,[],[f328,f41])).
fof(f3502,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X15,X12),X14),s(fun(X15,bool),X16),s(fun(X12,bool),X13))) | s('type.sum.sum'(X15,X12),X14) = s('type.sum.sum'(X15,X12),'const.sum.INL_1'(s(X15,sK9(X14,X12,X15)))) | s(bool,X17) = s(bool,happ(s(fun(X12,bool),X13),s(X12,sK10(X14,X12,X15)))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f328,f41])).
fof(f3501,plain,(
  ( ! [X10,X8,X7,X11,X9] : (s(bool,'F_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X10,X7),X9),s(fun(X10,bool),X11),s(fun(X7,bool),X8))) | s('type.sum.sum'(X10,X7),X9) = s('type.sum.sum'(X10,X7),'const.sum.INL_1'(s(X10,sK9(X9,X7,X10)))) | s(bool,'T_0') = s(bool,happ(s(fun(X7,bool),X8),s(X7,sK10(X9,X7,X10))))) )),
  inference(superposition,[],[f328,f33])).
fof(f3641,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X4,X1),X3),s(fun(X4,X0),X5),s(fun(X1,X0),X2))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X4,X1),X3),s(fun(X4,X0),X6),s(fun(X1,X0),X2))) | s('type.sum.sum'(X4,X1),X3) = s('type.sum.sum'(X4,X1),'const.sum.INL_1'(s(X4,sK9(X3,X1,X4))))) )),
  inference(duplicate_literal_removal,[],[f3500])).
fof(f3500,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X4,X1),X3),s(fun(X4,X0),X5),s(fun(X1,X0),X2))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(X4,X1),X3),s(fun(X4,X0),X6),s(fun(X1,X0),X2))) | s('type.sum.sum'(X4,X1),X3) = s('type.sum.sum'(X4,X1),'const.sum.INL_1'(s(X4,sK9(X3,X1,X4)))) | s('type.sum.sum'(X4,X1),X3) = s('type.sum.sum'(X4,X1),'const.sum.INL_1'(s(X4,sK9(X3,X1,X4))))) )),
  inference(superposition,[],[f328,f328])).
fof(f3192,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'('type.sum.sum'(bool,X16),X18),'const.sum.INL_1'(s('type.sum.sum'(bool,X16),X17))),s(fun('type.sum.sum'(bool,X16),bool),X19),s(fun(X18,bool),X20))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(bool,X16),bool),X19),s('type.sum.sum'(bool,X16),'const.sum.INR_1'(s(X16,sK10(X17,X16,bool)))))) | s('type.sum.sum'(bool,X16),X17) = s('type.sum.sum'(bool,X16),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X17,X16,bool))) )),
  inference(superposition,[],[f259,f326])).
fof(f3191,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X13,'type.sum.sum'(bool,X11)),'const.sum.INR_1'(s('type.sum.sum'(bool,X11),X12))),s(fun(X13,bool),X14),s(fun('type.sum.sum'(bool,X11),bool),X15))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(bool,X11),bool),X15),s('type.sum.sum'(bool,X11),'const.sum.INR_1'(s(X11,sK10(X12,X11,bool)))))) | s('type.sum.sum'(bool,X11),X12) = s('type.sum.sum'(bool,X11),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X12,X11,bool))) )),
  inference(superposition,[],[f296,f326])).
fof(f3190,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(X8,happ(s(fun(X6,X8),X9),s(X6,sK10(X7,X6,bool)))) = s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X6),X7),s(fun(bool,X8),X10),s(fun(X6,X8),X9))) | s('type.sum.sum'(bool,X6),X7) = s('type.sum.sum'(bool,X6),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X7,X6,bool))) )),
  inference(superposition,[],[f37,f326])).
fof(f3189,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X2),X3),s(fun(bool,bool),X4),s(fun(X2,bool),X5))) | s(bool,'F_0') = s(bool,happ(s(fun(X2,bool),X5),s(X2,sK10(X3,X2,bool)))) | s('type.sum.sum'(bool,X2),X3) = s('type.sum.sum'(bool,X2),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X3,X2,bool))) )),
  inference(superposition,[],[f296,f326])).
fof(f3185,plain,(
  ( ! [X28,X26,X27,X25] : (s('type.sum.sum'(bool,bool),X25) = s('type.sum.sum'(bool,bool),'const.sum.INR_1'(s(bool,X26))) | s('type.sum.sum'(bool,bool),X25) = s('type.sum.sum'(bool,bool),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X25,bool,bool)) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,sK10(X25,bool,bool)) | s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f326,f78])).
fof(f3184,plain,(
  ( ! [X24,X23,X21,X22] : (s('type.sum.sum'(bool,bool),X21) = s('type.sum.sum'(bool,bool),'const.sum.INR_1'(s(bool,X22))) | s('type.sum.sum'(bool,bool),X21) = s('type.sum.sum'(bool,bool),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X21,bool,bool)) | s(bool,X23) = s(bool,sK10(X21,bool,bool)) | s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f326,f78])).
fof(f3183,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.sum.sum'(bool,bool),X17) = s('type.sum.sum'(bool,bool),'const.sum.INR_1'(s(bool,'T_0'))) | s('type.sum.sum'(bool,bool),X17) = s('type.sum.sum'(bool,bool),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X17,bool,bool)) | s(bool,X18) = s(bool,X19) | s(bool,X19) = s(bool,sK10(X17,bool,bool)) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f326,f78])).
fof(f3181,plain,(
  ( ! [X12,X10,X11,X9] : (s('type.sum.sum'(bool,bool),X9) = s('type.sum.sum'(bool,bool),'const.sum.INR_1'(s(bool,X10))) | s('type.sum.sum'(bool,bool),X9) = s('type.sum.sum'(bool,bool),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X9,bool,bool)) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,sK10(X9,bool,bool)) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f326,f78])).
fof(f3178,plain,(
  ( ! [X4,X3] : (s('type.sum.sum'(bool,bool),X3) = s('type.sum.sum'(bool,bool),'const.sum.INR_1'(s(bool,X4))) | s('type.sum.sum'(bool,bool),X3) = s('type.sum.sum'(bool,bool),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X3,bool,bool)) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,sK10(X3,bool,bool))) )),
  inference(superposition,[],[f326,f41])).
fof(f3177,plain,(
  ( ! [X2,X1] : (s('type.sum.sum'(bool,bool),X1) = s('type.sum.sum'(bool,bool),'const.sum.INR_1'(s(bool,'T_0'))) | s('type.sum.sum'(bool,bool),X1) = s('type.sum.sum'(bool,bool),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X1,bool,bool)) | s(bool,X2) = s(bool,sK10(X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f326,f41])).
fof(f3060,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'('type.sum.sum'(X16,bool),X18),'const.sum.INL_1'(s('type.sum.sum'(X16,bool),X17))),s(fun('type.sum.sum'(X16,bool),bool),X19),s(fun(X18,bool),X20))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(X16,bool),bool),X19),s('type.sum.sum'(X16,bool),'const.sum.INL_1'(s(X16,sK9(X17,bool,X16)))))) | s('type.sum.sum'(X16,bool),X17) = s('type.sum.sum'(X16,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X17,bool,X16))) )),
  inference(superposition,[],[f259,f321])).
fof(f3059,plain,(
  ( ! [X14,X12,X15,X13,X11] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X13,'type.sum.sum'(X11,bool)),'const.sum.INR_1'(s('type.sum.sum'(X11,bool),X12))),s(fun(X13,bool),X14),s(fun('type.sum.sum'(X11,bool),bool),X15))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(X11,bool),bool),X15),s('type.sum.sum'(X11,bool),'const.sum.INL_1'(s(X11,sK9(X12,bool,X11)))))) | s('type.sum.sum'(X11,bool),X12) = s('type.sum.sum'(X11,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X12,bool,X11))) )),
  inference(superposition,[],[f296,f321])).
fof(f3058,plain,(
  ( ! [X6,X10,X8,X7,X9] : (s(X8,happ(s(fun(X6,X8),X9),s(X6,sK9(X7,bool,X6)))) = s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(X6,bool),X7),s(fun(X6,X8),X9),s(fun(bool,X8),X10))) | s('type.sum.sum'(X6,bool),X7) = s('type.sum.sum'(X6,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X7,bool,X6))) )),
  inference(superposition,[],[f36,f321])).
fof(f3057,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X2,bool),X3),s(fun(X2,bool),X4),s(fun(bool,bool),X5))) | s(bool,'F_0') = s(bool,happ(s(fun(X2,bool),X4),s(X2,sK9(X3,bool,X2)))) | s('type.sum.sum'(X2,bool),X3) = s('type.sum.sum'(X2,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X3,bool,X2))) )),
  inference(superposition,[],[f259,f321])).
fof(f3053,plain,(
  ( ! [X28,X26,X27,X25] : (s('type.sum.sum'(bool,bool),X25) = s('type.sum.sum'(bool,bool),'const.sum.INL_1'(s(bool,X26))) | s('type.sum.sum'(bool,bool),X25) = s('type.sum.sum'(bool,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X25,bool,bool)) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,sK9(X25,bool,bool)) | s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f321,f78])).
fof(f3052,plain,(
  ( ! [X24,X23,X21,X22] : (s('type.sum.sum'(bool,bool),X21) = s('type.sum.sum'(bool,bool),'const.sum.INL_1'(s(bool,X22))) | s('type.sum.sum'(bool,bool),X21) = s('type.sum.sum'(bool,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X21,bool,bool)) | s(bool,X23) = s(bool,sK9(X21,bool,bool)) | s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f321,f78])).
fof(f3051,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.sum.sum'(bool,bool),X17) = s('type.sum.sum'(bool,bool),'const.sum.INL_1'(s(bool,'T_0'))) | s('type.sum.sum'(bool,bool),X17) = s('type.sum.sum'(bool,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X17,bool,bool)) | s(bool,X18) = s(bool,X19) | s(bool,X19) = s(bool,sK9(X17,bool,bool)) | s(bool,X18) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f321,f78])).
fof(f3049,plain,(
  ( ! [X12,X10,X11,X9] : (s('type.sum.sum'(bool,bool),X9) = s('type.sum.sum'(bool,bool),'const.sum.INL_1'(s(bool,X10))) | s('type.sum.sum'(bool,bool),X9) = s('type.sum.sum'(bool,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X9,bool,bool)) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X12) = s(bool,sK9(X9,bool,bool)) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f321,f78])).
fof(f3046,plain,(
  ( ! [X4,X3] : (s('type.sum.sum'(bool,bool),X3) = s('type.sum.sum'(bool,bool),'const.sum.INL_1'(s(bool,X4))) | s('type.sum.sum'(bool,bool),X3) = s('type.sum.sum'(bool,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X3,bool,bool)) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,sK9(X3,bool,bool))) )),
  inference(superposition,[],[f321,f41])).
fof(f3045,plain,(
  ( ! [X2,X1] : (s('type.sum.sum'(bool,bool),X1) = s('type.sum.sum'(bool,bool),'const.sum.INL_1'(s(bool,'T_0'))) | s('type.sum.sum'(bool,bool),X1) = s('type.sum.sum'(bool,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X1,bool,bool)) | s(bool,X2) = s(bool,sK9(X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f321,f41])).
fof(f3044,plain,(
  ( ! [X0] : (s('type.sum.sum'(bool,bool),X0) = s('type.sum.sum'(bool,bool),'const.sum.INL_1'(s(bool,'F_0'))) | s('type.sum.sum'(bool,bool),X0) = s('type.sum.sum'(bool,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X0,bool,bool)) | s(bool,'T_0') = s(bool,sK9(X0,bool,bool))) )),
  inference(superposition,[],[f321,f33])).
fof(f2719,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(bool,X4) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X5,X1),'const.sum.INR_1'(s(X1,X3))),s(fun(X5,bool),X6),s(fun(X1,bool),X2))) | s(bool,X0) = s(bool,X4)) )),
  inference(condensation,[],[f2718])).
fof(f2718,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55,X53] : (s(bool,'T_0') = s(bool,X58) | s(bool,'F_0') = s(bool,happ(s(fun(X54,bool),X57),s(X54,X55))) | s(bool,X59) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X53,X54),'const.sum.INR_1'(s(X54,X55))),s(fun(X53,bool),X56),s(fun(X54,bool),X57))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(duplicate_literal_removal,[],[f2671])).
fof(f2671,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55,X53] : (s(bool,'T_0') = s(bool,X58) | s(bool,'F_0') = s(bool,happ(s(fun(X54,bool),X57),s(X54,X55))) | s(bool,X59) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X53,X54),'const.sum.INR_1'(s(X54,X55))),s(fun(X53,bool),X56),s(fun(X54,bool),X57))) | s(bool,'T_0') = s(bool,X58) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f296,f78])).
fof(f2661,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))) | s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),X4),s(X1,sK10(X2,X1,X0)))) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,sK9(X2,X1,X0))))) )),
  inference(superposition,[],[f296,f34])).
fof(f2660,plain,(
  ( ! [X182,X180,X183,X181] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X181,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(X180,sK1,sK0)))))),s(fun(X181,bool),X182),s(fun(sK2,bool),X183))) | s(bool,'F_0') = s(bool,happ(s(fun(sK2,bool),X183),s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,sK9(X180,sK1,sK0)))))) | s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),X180) | s('type.sum.sum'(sK0,sK1),X180) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(X180,sK1,sK0))))) )),
  inference(superposition,[],[f296,f331])).
fof(f2659,plain,(
  ( ! [X177,X178,X176,X179] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X177,sK2),'const.sum.INR_1'(s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(X176,sK1,sK0)))))),s(fun(X177,bool),X178),s(fun(sK2,bool),X179))) | s(bool,'F_0') = s(bool,happ(s(fun(sK2,bool),X179),s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(X176,sK1,sK0)))))) | s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),X176) | s('type.sum.sum'(sK0,sK1),X176) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(X176,sK1,sK0))))) )),
  inference(superposition,[],[f296,f327])).
fof(f2658,plain,(
  ( ! [X175,X173,X171,X174,X172,X170] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X173,'type.sum.sum'(X170,X171)),'const.sum.INR_1'(s('type.sum.sum'(X170,X171),X172))),s(fun(X173,bool),X174),s(fun('type.sum.sum'(X170,X171),bool),X175))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(X170,X171),bool),X175),s('type.sum.sum'(X170,X171),'const.sum.INL_1'(s(X170,sK9(X172,X171,X170)))))) | s('type.sum.sum'(X170,X171),X172) = s('type.sum.sum'(X170,X171),'const.sum.INR_1'(s(X171,sK10(X172,X171,X170))))) )),
  inference(superposition,[],[f296,f34])).
fof(f2657,plain,(
  ( ! [X167,X165,X169,X166,X164,X168] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X167,'type.sum.sum'(X164,X165)),'const.sum.INR_1'(s('type.sum.sum'(X164,X165),X166))),s(fun(X167,bool),X168),s(fun('type.sum.sum'(X164,X165),bool),X169))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(X164,X165),bool),X169),s('type.sum.sum'(X164,X165),'const.sum.INR_1'(s(X165,sK10(X166,X165,X164)))))) | s('type.sum.sum'(X164,X165),X166) = s('type.sum.sum'(X164,X165),'const.sum.INL_1'(s(X164,sK9(X166,X165,X164))))) )),
  inference(superposition,[],[f296,f34])).
fof(f2655,plain,(
  ( ! [X158,X156,X159,X157,X155,X160] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X158,bool),'const.sum.INR_1'(s(bool,'T_0'))),s(fun(X158,bool),X159),s(fun(bool,bool),X160))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X160),s(bool,sK11(X155,X156,X157,bool)))) | s(fun(bool,X157),X155) = s(fun(bool,X157),X156) | s(X157,f(s(bool,'F_0'))) != s(X157,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f296,f216])).
fof(f2727,plain,(
  ( ! [X146,X144,X142,X145,X141,X139] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X146),s(bool,happ(s(fun(X139,bool),X142),s(X139,X141))))) | s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X144,bool),'const.sum.INR_1'(s(bool,'T_0'))),s(fun(X144,bool),X145),s(fun(bool,bool),X146))) | s(bool,'F_0') = s(bool,happ(s(fun(X139,bool),X142),s(X139,X141)))) )),
  inference(forward_demodulation,[],[f2653,f36])).
fof(f2653,plain,(
  ( ! [X146,X144,X142,X140,X145,X143,X141,X139] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X144,bool),'const.sum.INR_1'(s(bool,'T_0'))),s(fun(X144,bool),X145),s(fun(bool,bool),X146))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X146),s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X139,X140),'const.sum.INL_1'(s(X139,X141))),s(fun(X139,bool),X142),s(fun(X140,bool),X143))))) | s(bool,'F_0') = s(bool,happ(s(fun(X139,bool),X142),s(X139,X141)))) )),
  inference(superposition,[],[f296,f259])).
fof(f2650,plain,(
  ( ! [X127,X130,X128,X126,X131,X129] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X129,bool),'const.sum.INR_1'(s(bool,'T_0'))),s(fun(X129,bool),X130),s(fun(bool,bool),X131))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X131),s(bool,g(s(X126,sK11(X127,X128,bool,X126)))))) | s(fun(X126,bool),X127) = s(fun(X126,bool),X128) | s(bool,'F_0') != s(bool,f(s(X126,sK11(X127,X128,bool,X126))))) )),
  inference(superposition,[],[f296,f221])).
fof(f2649,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X123,bool),'const.sum.INR_1'(s(bool,X122))),s(fun(X123,bool),X124),s(fun(bool,bool),X125))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X125),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f296,f33])).
fof(f2647,plain,(
  ( ! [X116,X114,X112,X115,X113] : (s(bool,X112) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X113,bool),'const.sum.INR_1'(s(bool,X112))),s(fun(X113,bool),X114),s(fun(bool,bool),X115))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X115),s(bool,'T_0'))) | s(bool,X112) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f296,f41])).
fof(f2645,plain,(
  ( ! [X103,X101,X99,X102,X100,X98,X104] : (s(bool,X98) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X99,bool),'const.sum.INR_1'(s(bool,X98))),s(fun(X99,bool),X100),s(fun(bool,bool),X101))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X101),s(bool,'T_0'))) | s(bool,X102) = s(bool,X103) | s(bool,X103) = s(bool,X104) | s(bool,'T_0') = s(bool,X104) | s(bool,X98) = s(bool,X102)) )),
  inference(superposition,[],[f296,f78])).
fof(f2641,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X72,bool),'const.sum.INR_1'(s(bool,X71))),s(fun(X72,bool),X73),s(fun(bool,bool),X74))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X74),s(bool,X70))) | s(bool,X71) = s(bool,X75) | s(bool,'T_0') = s(bool,X70) | s(bool,X75) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f296,f78])).
fof(f2640,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X65,bool),'const.sum.INR_1'(s(bool,X64))),s(fun(X65,bool),X66),s(fun(bool,bool),X67))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X67),s(bool,X63))) | s(bool,X63) = s(bool,X68) | s(bool,'T_0') = s(bool,X64) | s(bool,X68) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f296,f78])).
fof(f2639,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X57,bool),'const.sum.INR_1'(s(bool,'T_0'))),s(fun(X57,bool),X58),s(fun(bool,bool),X59))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X59),s(bool,X56))) | s(bool,X60) = s(bool,X61) | s(bool,X56) = s(bool,X61) | s(bool,X60) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f296,f78])).
fof(f2637,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X44,bool),'const.sum.INR_1'(s(bool,X43))),s(fun(X44,bool),X45),s(fun(bool,bool),X46))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X46),s(bool,X42))) | s(bool,X43) = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,X42) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f296,f78])).
fof(f2634,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X29,bool),'const.sum.INR_1'(s(bool,X28))),s(fun(X29,bool),X30),s(fun(bool,bool),X31))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X31),s(bool,X27))) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f296,f41])).
fof(f2633,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X23,bool),'const.sum.INR_1'(s(bool,'T_0'))),s(fun(X23,bool),X24),s(fun(bool,bool),X25))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X25),s(bool,X22))) | s(bool,X22) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f296,f41])).
fof(f2632,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X19,bool),'const.sum.INR_1'(s(bool,'F_0'))),s(fun(X19,bool),X20),s(fun(bool,bool),X21))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X21),s(bool,X18))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f296,f33])).
fof(f1933,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(bool,X4) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X1,X5),'const.sum.INL_1'(s(X1,X3))),s(fun(X1,bool),X2),s(fun(X5,bool),X6))) | s(bool,X0) = s(bool,X4)) )),
  inference(condensation,[],[f1932])).
fof(f1932,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55,X53] : (s(bool,'T_0') = s(bool,X58) | s(bool,'F_0') = s(bool,happ(s(fun(X53,bool),X56),s(X53,X55))) | s(bool,X59) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X53,X54),'const.sum.INL_1'(s(X53,X55))),s(fun(X53,bool),X56),s(fun(X54,bool),X57))) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(duplicate_literal_removal,[],[f1886])).
fof(f1886,plain,(
  ( ! [X59,X57,X54,X60,X58,X56,X55,X53] : (s(bool,'T_0') = s(bool,X58) | s(bool,'F_0') = s(bool,happ(s(fun(X53,bool),X56),s(X53,X55))) | s(bool,X59) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X53,X54),'const.sum.INL_1'(s(X53,X55))),s(fun(X53,bool),X56),s(fun(X54,bool),X57))) | s(bool,'T_0') = s(bool,X58) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f259,f78])).
fof(f1876,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))) | s(bool,'F_0') = s(bool,happ(s(fun(X0,bool),X3),s(X0,sK9(X2,X1,X0)))) | s('type.sum.sum'(X0,X1),X2) = s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,sK10(X2,X1,X0))))) )),
  inference(superposition,[],[f259,f34])).
fof(f1875,plain,(
  ( ! [X175,X173,X174,X172] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK2,X173),'const.sum.INL_1'(s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(X172,sK1,sK0)))))),s(fun(sK2,bool),X174),s(fun(X173,bool),X175))) | s(bool,'F_0') = s(bool,happ(s(fun(sK2,bool),X174),s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,sK9(X172,sK1,sK0)))))) | s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),X172) | s('type.sum.sum'(sK0,sK1),X172) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(X172,sK1,sK0))))) )),
  inference(superposition,[],[f259,f331])).
fof(f1874,plain,(
  ( ! [X171,X169,X170,X168] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK2,X169),'const.sum.INL_1'(s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(X168,sK1,sK0)))))),s(fun(sK2,bool),X170),s(fun(X169,bool),X171))) | s(bool,'F_0') = s(bool,happ(s(fun(sK2,bool),X170),s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(X168,sK1,sK0)))))) | s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),X168) | s('type.sum.sum'(sK0,sK1),X168) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(X168,sK1,sK0))))) )),
  inference(superposition,[],[f259,f327])).
fof(f1873,plain,(
  ( ! [X167,X165,X163,X166,X164,X162] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'('type.sum.sum'(X162,X163),X165),'const.sum.INL_1'(s('type.sum.sum'(X162,X163),X164))),s(fun('type.sum.sum'(X162,X163),bool),X166),s(fun(X165,bool),X167))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(X162,X163),bool),X166),s('type.sum.sum'(X162,X163),'const.sum.INL_1'(s(X162,sK9(X164,X163,X162)))))) | s('type.sum.sum'(X162,X163),X164) = s('type.sum.sum'(X162,X163),'const.sum.INR_1'(s(X163,sK10(X164,X163,X162))))) )),
  inference(superposition,[],[f259,f34])).
fof(f1872,plain,(
  ( ! [X158,X156,X161,X159,X157,X160] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'('type.sum.sum'(X156,X157),X159),'const.sum.INL_1'(s('type.sum.sum'(X156,X157),X158))),s(fun('type.sum.sum'(X156,X157),bool),X160),s(fun(X159,bool),X161))) | s(bool,'F_0') = s(bool,happ(s(fun('type.sum.sum'(X156,X157),bool),X160),s('type.sum.sum'(X156,X157),'const.sum.INR_1'(s(X157,sK10(X158,X157,X156)))))) | s('type.sum.sum'(X156,X157),X158) = s('type.sum.sum'(X156,X157),'const.sum.INL_1'(s(X156,sK9(X158,X157,X156))))) )),
  inference(superposition,[],[f259,f34])).
fof(f1870,plain,(
  ( ! [X152,X151,X149,X147,X150,X148] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X150),'const.sum.INL_1'(s(bool,'T_0'))),s(fun(bool,bool),X151),s(fun(X150,bool),X152))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X151),s(bool,sK11(X147,X148,X149,bool)))) | s(fun(bool,X149),X147) = s(fun(bool,X149),X148) | s(X149,f(s(bool,'F_0'))) != s(X149,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f259,f216])).
fof(f1941,plain,(
  ( ! [X146,X144,X142,X145,X141,X139] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X145),s(bool,happ(s(fun(X139,bool),X142),s(X139,X141))))) | s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X144),'const.sum.INL_1'(s(bool,'T_0'))),s(fun(bool,bool),X145),s(fun(X144,bool),X146))) | s(bool,'F_0') = s(bool,happ(s(fun(X139,bool),X142),s(X139,X141)))) )),
  inference(forward_demodulation,[],[f1869,f36])).
fof(f1869,plain,(
  ( ! [X146,X144,X142,X140,X145,X143,X141,X139] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X144),'const.sum.INL_1'(s(bool,'T_0'))),s(fun(bool,bool),X145),s(fun(X144,bool),X146))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X145),s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X139,X140),'const.sum.INL_1'(s(X139,X141))),s(fun(X139,bool),X142),s(fun(X140,bool),X143))))) | s(bool,'F_0') = s(bool,happ(s(fun(X139,bool),X142),s(X139,X141)))) )),
  inference(superposition,[],[f259,f259])).
fof(f1866,plain,(
  ( ! [X127,X130,X128,X126,X131,X129] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X129),'const.sum.INL_1'(s(bool,'T_0'))),s(fun(bool,bool),X130),s(fun(X129,bool),X131))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X130),s(bool,g(s(X126,sK11(X127,X128,bool,X126)))))) | s(fun(X126,bool),X127) = s(fun(X126,bool),X128) | s(bool,'F_0') != s(bool,f(s(X126,sK11(X127,X128,bool,X126))))) )),
  inference(superposition,[],[f259,f221])).
fof(f1865,plain,(
  ( ! [X125,X123,X124,X122] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X123),'const.sum.INL_1'(s(bool,X122))),s(fun(bool,bool),X124),s(fun(X123,bool),X125))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X124),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f259,f33])).
fof(f1863,plain,(
  ( ! [X116,X114,X112,X115,X113] : (s(bool,X112) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X113),'const.sum.INL_1'(s(bool,X112))),s(fun(bool,bool),X114),s(fun(X113,bool),X115))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X114),s(bool,'T_0'))) | s(bool,X112) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f259,f41])).
fof(f1861,plain,(
  ( ! [X103,X101,X99,X102,X100,X98,X104] : (s(bool,X98) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X99),'const.sum.INL_1'(s(bool,X98))),s(fun(bool,bool),X100),s(fun(X99,bool),X101))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X100),s(bool,'T_0'))) | s(bool,X102) = s(bool,X103) | s(bool,X103) = s(bool,X104) | s(bool,'T_0') = s(bool,X104) | s(bool,X98) = s(bool,X102)) )),
  inference(superposition,[],[f259,f78])).
fof(f1857,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X72),'const.sum.INL_1'(s(bool,X71))),s(fun(bool,bool),X73),s(fun(X72,bool),X74))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X73),s(bool,X70))) | s(bool,X71) = s(bool,X75) | s(bool,'T_0') = s(bool,X70) | s(bool,X75) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f259,f78])).
fof(f1856,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X65),'const.sum.INL_1'(s(bool,X64))),s(fun(bool,bool),X66),s(fun(X65,bool),X67))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X66),s(bool,X63))) | s(bool,X63) = s(bool,X68) | s(bool,'T_0') = s(bool,X64) | s(bool,X68) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f259,f78])).
fof(f1855,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X57),'const.sum.INL_1'(s(bool,'T_0'))),s(fun(bool,bool),X58),s(fun(X57,bool),X59))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X58),s(bool,X56))) | s(bool,X60) = s(bool,X61) | s(bool,X56) = s(bool,X61) | s(bool,X60) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f259,f78])).
fof(f1853,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X42] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X44),'const.sum.INL_1'(s(bool,X43))),s(fun(bool,bool),X45),s(fun(X44,bool),X46))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X45),s(bool,X42))) | s(bool,X43) = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,X42) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f259,f78])).
fof(f1850,plain,(
  ( ! [X30,X28,X31,X29,X27] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X29),'const.sum.INL_1'(s(bool,X28))),s(fun(bool,bool),X30),s(fun(X29,bool),X31))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X30),s(bool,X27))) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f259,f41])).
fof(f1849,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X23),'const.sum.INL_1'(s(bool,'T_0'))),s(fun(bool,bool),X24),s(fun(X23,bool),X25))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X24),s(bool,X22))) | s(bool,X22) = s(bool,X26) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f259,f41])).
fof(f1848,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(bool,X19),'const.sum.INL_1'(s(bool,'F_0'))),s(fun(bool,bool),X20),s(fun(X19,bool),X21))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X20),s(bool,X18))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f259,f33])).
fof(f1615,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') != s(bool,f(s(X66,sK11(X67,X68,bool,X66)))) | s(fun(X66,bool),X67) = s(fun(X66,bool),X68) | s(bool,'F_0') != s(bool,f(s(X66,sK11(X67,X68,bool,X66))))) )),
  inference(duplicate_literal_removal,[],[f1612])).
fof(f1612,plain,(
  ( ! [X68,X66,X67] : (s(bool,'T_0') != s(bool,f(s(X66,sK11(X67,X68,bool,X66)))) | s(fun(X66,bool),X67) = s(fun(X66,bool),X68) | s(fun(X66,bool),X67) = s(fun(X66,bool),X68) | s(bool,'F_0') != s(bool,f(s(X66,sK11(X67,X68,bool,X66))))) )),
  inference(superposition,[],[f35,f221])).
fof(f1643,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(fun(X1,bool),X2) = s(fun(X1,bool),X3) | s(bool,'F_0') != s(bool,f(s(X1,sK11(X2,X3,bool,X1)))) | s(bool,X4) = s(bool,g(s(X1,sK11(X2,X3,bool,X1)))) | s(bool,X0) = s(bool,X4)) )),
  inference(condensation,[],[f1642])).
fof(f1642,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s(bool,'T_0') = s(bool,X38) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') != s(bool,f(s(X35,sK11(X36,X37,bool,X35)))) | s(bool,X39) = s(bool,g(s(X35,sK11(X36,X37,bool,X35)))) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(duplicate_literal_removal,[],[f1597])).
fof(f1597,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s(bool,'T_0') = s(bool,X38) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') != s(bool,f(s(X35,sK11(X36,X37,bool,X35)))) | s(bool,X39) = s(bool,g(s(X35,sK11(X36,X37,bool,X35)))) | s(bool,'T_0') = s(bool,X38) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f221,f78])).
fof(f772,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2)) )),
  inference(equality_factoring,[],[f78])).
fof(f1573,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f765,f41])).
fof(f765,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(equality_factoring,[],[f78])).
fof(f855,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) = s(bool,X2) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f738])).
fof(f738,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(equality_factoring,[],[f78])).
fof(f661,plain,(
  ( ! [X103,X107,X105,X102,X106,X104] : (s(X104,f(s(bool,'T_0'))) != s(X104,g(s(bool,'T_0'))) | s(fun(bool,X104),X102) = s(fun(bool,X104),X103) | s(bool,X105) = s(bool,X106) | s(bool,X106) = s(bool,sK11(X102,X103,X104,bool)) | s(bool,X105) = s(bool,X107) | s(bool,'T_0') = s(bool,X107)) )),
  inference(superposition,[],[f35,f78])).
fof(f660,plain,(
  ( ! [X101,X99,X97,X100,X98] : (s('type.sum.sum'(X98,bool),X97) = s('type.sum.sum'(X98,bool),'const.sum.INR_1'(s(bool,'T_0'))) | s('type.sum.sum'(X98,bool),X97) = s('type.sum.sum'(X98,bool),'const.sum.INL_1'(s(X98,sK9(X97,bool,X98)))) | s(bool,X99) = s(bool,X100) | s(bool,X100) = s(bool,sK10(X97,bool,X98)) | s(bool,X99) = s(bool,X101) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f34,f78])).
fof(f659,plain,(
  ( ! [X94,X92,X95,X93,X96] : (s('type.sum.sum'(bool,X93),X92) = s('type.sum.sum'(bool,X93),'const.sum.INL_1'(s(bool,'T_0'))) | s('type.sum.sum'(bool,X93),X92) = s('type.sum.sum'(bool,X93),'const.sum.INR_1'(s(X93,sK10(X92,X93,bool)))) | s(bool,X94) = s(bool,X95) | s(bool,X95) = s(bool,sK9(X92,X93,bool)) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f34,f78])).
fof(f656,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (s(bool,'T_0') != s(bool,f(s(X70,sK11(X71,X72,bool,X70)))) | s(fun(X70,bool),X71) = s(fun(X70,bool),X72) | s(bool,X73) = s(bool,X74) | s(bool,X74) = s(bool,g(s(X70,sK11(X71,X72,bool,X70)))) | s(bool,X73) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f35,f78])).
fof(f636,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X13)) | s(bool,X14) = s(bool,X15) | s(bool,X13) = s(bool,X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f31,f78])).
fof(f633,plain,(
  ( ! [X116,X114,X112,X117,X115,X113] : (s(X114,f(s(bool,X115))) != s(X114,g(s(bool,X115))) | s(fun(bool,X114),X112) = s(fun(bool,X114),X113) | s(bool,X115) = s(bool,X116) | s(bool,'T_0') = s(bool,sK11(X112,X113,X114,bool)) | s(bool,X116) = s(bool,X117) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f35,f78])).
fof(f632,plain,(
  ( ! [X111,X109,X107,X110,X108] : (s('type.sum.sum'(X108,bool),X107) = s('type.sum.sum'(X108,bool),'const.sum.INR_1'(s(bool,X109))) | s('type.sum.sum'(X108,bool),X107) = s('type.sum.sum'(X108,bool),'const.sum.INL_1'(s(X108,sK9(X107,bool,X108)))) | s(bool,X109) = s(bool,X110) | s(bool,'T_0') = s(bool,sK10(X107,bool,X108)) | s(bool,X110) = s(bool,X111) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f34,f78])).
fof(f631,plain,(
  ( ! [X103,X105,X102,X106,X104] : (s('type.sum.sum'(bool,X103),X102) = s('type.sum.sum'(bool,X103),'const.sum.INL_1'(s(bool,X104))) | s('type.sum.sum'(bool,X103),X102) = s('type.sum.sum'(bool,X103),'const.sum.INR_1'(s(X103,sK10(X102,X103,bool)))) | s(bool,X104) = s(bool,X105) | s(bool,'T_0') = s(bool,sK9(X102,X103,bool)) | s(bool,X105) = s(bool,X106) | s(bool,'T_0') = s(bool,X106)) )),
  inference(superposition,[],[f34,f78])).
fof(f628,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (s(bool,X83) != s(bool,f(s(X80,sK11(X81,X82,bool,X80)))) | s(fun(X80,bool),X81) = s(fun(X80,bool),X82) | s(bool,X83) = s(bool,X84) | s(bool,'T_0') = s(bool,g(s(X80,sK11(X81,X82,bool,X80)))) | s(bool,X84) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f35,f78])).
fof(f981,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f980])).
fof(f980,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X25)) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,X26) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(duplicate_literal_removal,[],[f617])).
fof(f617,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X25)) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X24) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f43,f78])).
fof(f611,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (s(X117,f(s(bool,X118))) != s(X117,g(s(bool,X118))) | s(fun(bool,X117),X115) = s(fun(bool,X117),X116) | s(bool,X119) = s(bool,sK11(X115,X116,X117,bool)) | s(bool,'T_0') = s(bool,X118) | s(bool,X119) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f35,f78])).
fof(f610,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s('type.sum.sum'(X111,bool),X110) = s('type.sum.sum'(X111,bool),'const.sum.INR_1'(s(bool,X112))) | s('type.sum.sum'(X111,bool),X110) = s('type.sum.sum'(X111,bool),'const.sum.INL_1'(s(X111,sK9(X110,bool,X111)))) | s(bool,X113) = s(bool,sK10(X110,bool,X111)) | s(bool,'T_0') = s(bool,X112) | s(bool,X113) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f34,f78])).
fof(f609,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s('type.sum.sum'(bool,X106),X105) = s('type.sum.sum'(bool,X106),'const.sum.INL_1'(s(bool,X107))) | s('type.sum.sum'(bool,X106),X105) = s('type.sum.sum'(bool,X106),'const.sum.INR_1'(s(X106,sK10(X105,X106,bool)))) | s(bool,X108) = s(bool,sK9(X105,X106,bool)) | s(bool,'T_0') = s(bool,X107) | s(bool,X108) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f34,f78])).
fof(f606,plain,(
  ( ! [X88,X87,X85,X83,X86,X84] : (s(bool,X86) != s(bool,f(s(X83,sK11(X84,X85,bool,X83)))) | s(fun(X83,bool),X84) = s(fun(X83,bool),X85) | s(bool,X87) = s(bool,g(s(X83,sK11(X84,X85,bool,X83)))) | s(bool,'T_0') = s(bool,X86) | s(bool,X87) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f35,f78])).
fof(f567,plain,(
  ( ! [X118,X116,X120,X119,X117,X115] : (s(X117,f(s(bool,X118))) != s(X117,g(s(bool,X118))) | s(fun(bool,X117),X115) = s(fun(bool,X117),X116) | s(bool,X118) = s(bool,X119) | s(bool,'T_0') = s(bool,X119) | s(bool,X120) = s(bool,sK11(X115,X116,X117,bool)) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f35,f78])).
fof(f566,plain,(
  ( ! [X111,X114,X112,X110,X113] : (s('type.sum.sum'(X111,bool),X110) = s('type.sum.sum'(X111,bool),'const.sum.INR_1'(s(bool,X112))) | s('type.sum.sum'(X111,bool),X110) = s('type.sum.sum'(X111,bool),'const.sum.INL_1'(s(X111,sK9(X110,bool,X111)))) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113) | s(bool,X114) = s(bool,sK10(X110,bool,X111)) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f34,f78])).
fof(f565,plain,(
  ( ! [X109,X107,X105,X108,X106] : (s('type.sum.sum'(bool,X106),X105) = s('type.sum.sum'(bool,X106),'const.sum.INL_1'(s(bool,X107))) | s('type.sum.sum'(bool,X106),X105) = s('type.sum.sum'(bool,X106),'const.sum.INR_1'(s(X106,sK10(X105,X106,bool)))) | s(bool,X107) = s(bool,X108) | s(bool,'T_0') = s(bool,X108) | s(bool,X109) = s(bool,sK9(X105,X106,bool)) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f34,f78])).
fof(f562,plain,(
  ( ! [X88,X87,X85,X83,X86,X84] : (s(bool,X86) != s(bool,f(s(X83,sK11(X84,X85,bool,X83)))) | s(fun(X83,bool),X84) = s(fun(X83,bool),X85) | s(bool,X86) = s(bool,X87) | s(bool,'T_0') = s(bool,X87) | s(bool,X88) = s(bool,g(s(X83,sK11(X84,X85,bool,X83)))) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f35,f78])).
fof(f475,plain,(
  ( ! [X80,X85,X83,X81,X79,X86,X84,X82] : (s(bool,'T_0') = s(bool,happ(s(fun(X80,bool),X83),s(X80,X81))) | s(bool,X84) = s(bool,X85) | s(bool,X85) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X79,X80),'const.sum.INR_1'(s(X80,X81))),s(fun(X79,bool),X82),s(fun(X80,bool),X83))) | s(bool,X84) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f78,f37])).
fof(f474,plain,(
  ( ! [X78,X76,X74,X72,X71,X77,X75,X73] : (s(bool,'T_0') = s(bool,happ(s(fun(X71,bool),X74),s(X71,X73))) | s(bool,X76) = s(bool,X77) | s(bool,X77) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X71,X72),'const.sum.INL_1'(s(X71,X73))),s(fun(X71,bool),X74),s(fun(X72,bool),X75))) | s(bool,X76) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f78,f36])).
fof(f448,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X1) = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X0) = s(bool,X5)) )),
  inference(superposition,[],[f78,f78])).
fof(f447,plain,(
  ( ! [X111,X116,X114,X112,X110,X117,X115,X113] : (s(bool,X115) = s(bool,happ(s(fun(X111,bool),X114),s(X111,X112))) | s(bool,X115) = s(bool,X116) | s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X110,X111),'const.sum.INR_1'(s(X111,X112))),s(fun(X110,bool),X113),s(fun(X111,bool),X114))) | s(bool,X116) = s(bool,X117) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f78,f37])).
fof(f446,plain,(
  ( ! [X103,X109,X107,X105,X102,X108,X106,X104] : (s(bool,X107) = s(bool,happ(s(fun(X102,bool),X105),s(X102,X104))) | s(bool,X107) = s(bool,X108) | s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X102,X103),'const.sum.INL_1'(s(X102,X104))),s(fun(X102,bool),X105),s(fun(X103,bool),X106))) | s(bool,X108) = s(bool,X109) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f78,f36])).
fof(f1369,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f1368])).
fof(f1368,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(duplicate_literal_removal,[],[f427])).
fof(f427,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f78,f33])).
fof(f426,plain,(
  ( ! [X127,X125,X123,X128,X126,X124,X122,X129] : (s(bool,X127) = s(bool,happ(s(fun(X123,bool),X126),s(X123,X124))) | s(bool,X128) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X122,X123),'const.sum.INR_1'(s(X123,X124))),s(fun(X122,bool),X125),s(fun(X123,bool),X126))) | s(bool,'T_0') = s(bool,X127) | s(bool,X128) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f78,f37])).
fof(f425,plain,(
  ( ! [X121,X118,X116,X114,X120,X119,X117,X115] : (s(bool,X119) = s(bool,happ(s(fun(X114,bool),X117),s(X114,X116))) | s(bool,X120) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X114,X115),'const.sum.INL_1'(s(X114,X116))),s(fun(X114,bool),X117),s(fun(X115,bool),X118))) | s(bool,'T_0') = s(bool,X119) | s(bool,X120) = s(bool,X121) | s(bool,'T_0') = s(bool,X121)) )),
  inference(superposition,[],[f78,f36])).
fof(f1371,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,sK11(X2,X3,X4,bool)) | s(bool,X0) = s(bool,X1) | s(fun(bool,X4),X2) = s(fun(bool,X4),X3) | s(X4,f(s(bool,'F_0'))) != s(X4,g(s(bool,'F_0')))) )),
  inference(condensation,[],[f1370])).
fof(f1370,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,sK11(X108,X109,X110,bool)) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113) | s(fun(bool,X110),X108) = s(fun(bool,X110),X109) | s(X110,f(s(bool,'F_0'))) != s(X110,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f424])).
fof(f424,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,sK11(X108,X109,X110,bool)) | s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113) | s(fun(bool,X110),X108) = s(fun(bool,X110),X109) | s(X110,f(s(bool,'F_0'))) != s(X110,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f78,f216])).
fof(f384,plain,(
  ( ! [X127,X125,X123,X128,X126,X124,X122,X129] : (s(bool,X127) = s(bool,happ(s(fun(X123,bool),X126),s(X123,X124))) | s(bool,X127) = s(bool,X128) | s(bool,'T_0') = s(bool,X128) | s(bool,X129) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X122,X123),'const.sum.INR_1'(s(X123,X124))),s(fun(X122,bool),X125),s(fun(X123,bool),X126))) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f78,f37])).
fof(f383,plain,(
  ( ! [X121,X118,X116,X114,X120,X119,X117,X115] : (s(bool,X119) = s(bool,happ(s(fun(X114,bool),X117),s(X114,X116))) | s(bool,X119) = s(bool,X120) | s(bool,'T_0') = s(bool,X120) | s(bool,X121) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X114,X115),'const.sum.INL_1'(s(X114,X116))),s(fun(X114,bool),X117),s(fun(X115,bool),X118))) | s(bool,'T_0') = s(bool,X121)) )),
  inference(superposition,[],[f78,f36])).
fof(f1511,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X1) = s(bool,X3)) )),
  inference(condensation,[],[f1510])).
fof(f1510,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s(bool,X80) = s(bool,X81) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X83) | s(bool,X84) = s(bool,X85) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,X80) = s(bool,X84)) )),
  inference(duplicate_literal_removal,[],[f377])).
fof(f377,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s(bool,X80) = s(bool,X81) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X83) | s(bool,'T_0') = s(bool,X83) | s(bool,X84) = s(bool,X85) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,X80) = s(bool,X84)) )),
  inference(superposition,[],[f78,f78])).
fof(f350,plain,(
  ( ! [X26,X24,X25] : (s(X26,f(s(bool,'T_0'))) != s(X26,g(s(bool,'T_0'))) | s(fun(bool,X26),X24) = s(fun(bool,X26),X25) | s(X26,f(s(bool,'F_0'))) != s(X26,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f347])).
fof(f347,plain,(
  ( ! [X26,X24,X25] : (s(X26,f(s(bool,'T_0'))) != s(X26,g(s(bool,'T_0'))) | s(fun(bool,X26),X24) = s(fun(bool,X26),X25) | s(fun(bool,X26),X24) = s(fun(bool,X26),X25) | s(X26,f(s(bool,'F_0'))) != s(X26,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f35,f216])).
fof(f326,plain,(
  ( ! [X12,X13] : (s('type.sum.sum'(bool,X13),X12) = s('type.sum.sum'(bool,X13),'const.sum.INR_1'(s(X13,sK10(X12,X13,bool)))) | s('type.sum.sum'(bool,X13),X12) = s('type.sum.sum'(bool,X13),'const.sum.INL_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK9(X12,X13,bool))) )),
  inference(superposition,[],[f34,f33])).
fof(f324,plain,(
  ( ! [X6,X8,X7] : (s('type.sum.sum'(bool,X7),X6) = s('type.sum.sum'(bool,X7),'const.sum.INL_1'(s(bool,X8))) | s('type.sum.sum'(bool,X7),X6) = s('type.sum.sum'(bool,X7),'const.sum.INR_1'(s(X7,sK10(X6,X7,bool)))) | s(bool,'T_0') = s(bool,sK9(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f34,f41])).
fof(f322,plain,(
  ( ! [X2,X0,X1] : (s('type.sum.sum'(bool,X1),X0) = s('type.sum.sum'(bool,X1),'const.sum.INL_1'(s(bool,'T_0'))) | s('type.sum.sum'(bool,X1),X0) = s('type.sum.sum'(bool,X1),'const.sum.INR_1'(s(X1,sK10(X0,X1,bool)))) | s(bool,X2) = s(bool,sK9(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f34,f41])).
fof(f321,plain,(
  ( ! [X12,X13] : (s('type.sum.sum'(X13,bool),X12) = s('type.sum.sum'(X13,bool),'const.sum.INL_1'(s(X13,sK9(X12,bool,X13)))) | s('type.sum.sum'(X13,bool),X12) = s('type.sum.sum'(X13,bool),'const.sum.INR_1'(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK10(X12,bool,X13))) )),
  inference(superposition,[],[f34,f33])).
fof(f319,plain,(
  ( ! [X6,X8,X7] : (s('type.sum.sum'(X7,bool),X6) = s('type.sum.sum'(X7,bool),'const.sum.INR_1'(s(bool,X8))) | s('type.sum.sum'(X7,bool),X6) = s('type.sum.sum'(X7,bool),'const.sum.INL_1'(s(X7,sK9(X6,bool,X7)))) | s(bool,'T_0') = s(bool,sK10(X6,bool,X7)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f34,f41])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (s('type.sum.sum'(X1,bool),X0) = s('type.sum.sum'(X1,bool),'const.sum.INR_1'(s(bool,'T_0'))) | s('type.sum.sum'(X1,bool),X0) = s('type.sum.sum'(X1,bool),'const.sum.INL_1'(s(X1,sK9(X0,bool,X1)))) | s(bool,X2) = s(bool,sK10(X0,bool,X1)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f34,f41])).
fof(f296,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X24,X25),'const.sum.INR_1'(s(X25,X26))),s(fun(X24,bool),X27),s(fun(X25,bool),X28))) | s(bool,'F_0') = s(bool,happ(s(fun(X25,bool),X28),s(X25,X26)))) )),
  inference(superposition,[],[f37,f33])).
fof(f294,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X17) = s(bool,happ(s(fun(X13,bool),X16),s(X13,X14))) | s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INR_1'(s(X13,X14))),s(fun(X12,bool),X15),s(fun(X13,bool),X16))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f37,f41])).
fof(f292,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X1,bool),X4),s(X1,X2))) | s(bool,X5) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),'const.sum.INR_1'(s(X1,X2))),s(fun(X0,bool),X3),s(fun(X1,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f37,f41])).
fof(f259,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X24,X25),'const.sum.INL_1'(s(X24,X26))),s(fun(X24,bool),X27),s(fun(X25,bool),X28))) | s(bool,'F_0') = s(bool,happ(s(fun(X24,bool),X27),s(X24,X26)))) )),
  inference(superposition,[],[f36,f33])).
fof(f257,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,X17) = s(bool,happ(s(fun(X12,bool),X15),s(X12,X14))) | s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X12,X13),'const.sum.INL_1'(s(X12,X14))),s(fun(X12,bool),X15),s(fun(X13,bool),X16))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f36,f41])).
fof(f255,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X3),s(X0,X2))) | s(bool,X5) = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X0,X1),'const.sum.INL_1'(s(X0,X2))),s(fun(X0,bool),X3),s(fun(X1,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f36,f41])).
fof(f221,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,g(s(X16,sK11(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'F_0') != s(bool,f(s(X16,sK11(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f35,f33])).
fof(f219,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK11(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK11(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f35,f41])).
fof(f217,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK11(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK11(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f35,f41])).
fof(f216,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,sK11(X16,X17,X18,bool)) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f35,f33])).
fof(f214,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK11(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f35,f41])).
fof(f212,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK11(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f35,f41])).
fof(f35,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK11(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK11(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK11(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK11(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11])],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.sum.sum_case_cong',aHL_EXT)).
fof(f3988,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X8),X9),s(fun(sK1,X8),X10))) = s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) | $spl72),
  inference(cnf_transformation,[],[f3988_D])).
fof(f3988_D,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X8),X9),s(fun(sK1,X8),X10))) = s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3990,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | $spl74),
  inference(cnf_transformation,[],[f3990_D])).
fof(f3990_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3987,plain,(
  ( ! [X10,X8,X9] : (s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X8),X9),s(fun(sK1,X8),X10))) = s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) | $spl42),
  inference(forward_demodulation,[],[f3903,f3646])).
fof(f3903,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X8),X9),s(fun(sK1,X8),X10))) = s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) | $spl42),
  inference(superposition,[],[f332,f3646])).
fof(f3984,plain,(
  $false | ($spl40 | $spl42 | $spl48 | $spl51)),
  inference(global_subsumption,[],[f30,f29,f28,f27,f38,f39,f31,f32,f33,f34,f35,f36,f37,f41,f43,f46,f78,f110,f212,f214,f216,f217,f219,f221,f255,f257,f259,f292,f294,f296,f317,f319,f321,f322,f324,f326,f327,f328,f331,f332,f350,f1511,f383,f384,f1371,f425,f426,f1369,f446,f447,f448,f474,f475,f562,f565,f566,f567,f606,f609,f610,f611,f981,f628,f631,f632,f633,f636,f656,f659,f660,f661,f855,f1573,f772,f1648,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1586,f1587,f1588,f1589,f1590,f1591,f1643,f1615,f1848,f1849,f1850,f1853,f1855,f1856,f1857,f1861,f1863,f1865,f1866,f1941,f1870,f1872,f1873,f1874,f1875,f1876,f1933,f2632,f2633,f2634,f2637,f2639,f2640,f2641,f2645,f2647,f2649,f2650,f2727,f2655,f2657,f2658,f2659,f2660,f2661,f2719,f3044,f3045,f3046,f3049,f3051,f3052,f3053,f3057,f3058,f3059,f3060,f3177,f3178,f3181,f3183,f3184,f3185,f3189,f3190,f3191,f3192,f3641,f3501,f3502,f3503,f3506,f3508,f3509,f3510,f3643,f3516,f3517,f3518,f3519,f3520,f3647,f3648,f3523,f3524,f3533,f3534,f3535,f3536,f3552,f3553,f3556,f3558,f3559,f3560,f3565,f3566,f3574,f3583,f3591,f3594,f3595,f3596,f3599,f3600,f3610,f3619,f3627,f3630,f3631,f3632,f3633,f3634,f3635,f3636,f3637,f3638,f3659,f3661,f3664,f3827,f3671,f3672,f3673,f3676,f3678,f3679,f3680,f3830,f3686,f3687,f3688,f3689,f3690,f3691,f3831,f3832,f3694,f3695,f3696,f3697,f3710,f3711,f3713,f3714,f3730,f3731,f3734,f3736,f3737,f3738,f3743,f3744,f3752,f3761,f3771,f3774,f3775,f3776,f3777,f3781,f3782,f3792,f3801,f3811,f3814,f3815,f3816,f3817,f3819,f3820,f3822,f3823,f3824,f3839,f3845,f3856,f3829,f40,f3867,f3873,f3883,f3857,f3644,f3646,f3895,f3901,f3858,f3955,f3972,f3975,f3962,f3983])).
fof(f3983,plain,(
  ( ! [X12,X11] : (s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) | s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10(sK6,sK1,sK0)))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) | ($spl40 | $spl42 | $spl51)),
  inference(forward_demodulation,[],[f3982,f3955])).
fof(f3982,plain,(
  ( ! [X12,X11] : (s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10(sK6,sK1,sK0)))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) | ($spl40 | $spl42 | $spl51)),
  inference(forward_demodulation,[],[f3965,f3644])).
fof(f3965,plain,(
  ( ! [X12,X13,X11] : (s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X11),X13),s(fun(sK1,X11),X12))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) | ($spl42 | $spl51)),
  inference(superposition,[],[f328,f3955])).
fof(f3962,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,bool),X3),s(fun(sK1,bool),X4))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),X4),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) | ($spl42 | $spl51)),
  inference(superposition,[],[f296,f3955])).
fof(f3975,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK6,sK1,sK0)))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK6,sK1,sK0)))) | ($spl42 | $spl48 | $spl51)),
  inference(forward_demodulation,[],[f3974,f3857])).
fof(f3974,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK6,sK1,sK0)))) = s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | ($spl42 | $spl48 | $spl51)),
  inference(forward_demodulation,[],[f3970,f3857])).
fof(f3970,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | ($spl42 | $spl51)),
  inference(trivial_inequality_removal,[],[f3959])).
fof(f3959,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),sK5) | s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | ($spl42 | $spl51)),
  inference(superposition,[],[f39,f3955])).
fof(f3972,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK6,sK1,sK0)))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK6,sK1,sK0)))) | ($spl42 | $spl48 | $spl51)),
  inference(forward_demodulation,[],[f3971,f3857])).
fof(f3971,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK6,sK1,sK0)))) = s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | ($spl42 | $spl48 | $spl51)),
  inference(forward_demodulation,[],[f3956,f3857])).
fof(f3956,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) = s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | ($spl42 | $spl51)),
  inference(unit_resulting_resolution,[],[f3955,f39])).
fof(f3955,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | ($spl42 | $spl51)),
  inference(forward_demodulation,[],[f3954,f3646])).
fof(f3954,plain,(
  s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | ($spl42 | $spl51)),
  inference(subsumption_resolution,[],[f3950,f3646])).
fof(f3950,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | $spl51),
  inference(superposition,[],[f3858,f34])).
fof(f3858,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | $spl51),
  inference(cnf_transformation,[],[f3858_D])).
fof(f3858_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f3901,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,bool),X3),s(fun(sK1,bool),X4))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK9(sK6,sK1,sK0))))) ) | $spl42),
  inference(superposition,[],[f259,f3646])).
fof(f3873,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,bool),X3),s(fun(sK1,bool),X4))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),X3),s(sK0,sK9(sK6,sK1,sK0))))) ) | $spl42),
  inference(superposition,[],[f259,f3646])).
fof(f3867,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) = s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,sK9(sK6,sK1,sK0)))) | $spl42),
  inference(unit_resulting_resolution,[],[f3646,f38])).
fof(f3981,plain,(
  $false | ($spl40 | $spl42 | $spl48 | $spl51 | $spl52)),
  inference(global_subsumption,[],[f3980,f30,f29,f28,f27,f38,f39,f31,f32,f33,f34,f35,f36,f37,f41,f43,f46,f78,f110,f212,f214,f216,f217,f219,f221,f255,f257,f259,f292,f294,f296,f317,f319,f321,f322,f324,f326,f327,f328,f331,f332,f350,f1511,f383,f384,f1371,f425,f426,f1369,f446,f447,f448,f474,f475,f562,f565,f566,f567,f606,f609,f610,f611,f981,f628,f631,f632,f633,f636,f656,f659,f660,f661,f855,f1573,f772,f1648,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1586,f1587,f1588,f1589,f1590,f1591,f1643,f1615,f1848,f1849,f1850,f1853,f1855,f1856,f1857,f1861,f1863,f1865,f1866,f1941,f1870,f1872,f1873,f1874,f1875,f1876,f1933,f2632,f2633,f2634,f2637,f2639,f2640,f2641,f2645,f2647,f2649,f2650,f2727,f2655,f2657,f2658,f2659,f2660,f2661,f2719,f3044,f3045,f3046,f3049,f3051,f3052,f3053,f3057,f3058,f3059,f3060,f3177,f3178,f3181,f3183,f3184,f3185,f3189,f3190,f3191,f3192,f3641,f3501,f3502,f3503,f3506,f3508,f3509,f3510,f3643,f3516,f3517,f3518,f3519,f3520,f3647,f3648,f3523,f3524,f3533,f3534,f3535,f3536,f3552,f3553,f3556,f3558,f3559,f3560,f3565,f3566,f3574,f3583,f3591,f3594,f3595,f3596,f3599,f3600,f3610,f3619,f3627,f3630,f3631,f3632,f3633,f3634,f3635,f3636,f3637,f3638,f3659,f3661,f3664,f3827,f3671,f3672,f3673,f3676,f3678,f3679,f3680,f3830,f3686,f3687,f3688,f3689,f3690,f3691,f3831,f3832,f3694,f3695,f3696,f3697,f3710,f3711,f3713,f3714,f3730,f3731,f3734,f3736,f3737,f3738,f3743,f3744,f3752,f3761,f3771,f3774,f3775,f3776,f3777,f3781,f3782,f3792,f3801,f3811,f3814,f3815,f3816,f3817,f3819,f3820,f3822,f3823,f3824,f3839,f3845,f3856,f3829,f40,f3867,f3873,f3883,f3857,f3644,f3646,f3895,f3901,f3858,f3955,f3972,f3975,f3962])).
fof(f3980,plain,(
  ( ! [X8,X9] : (s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) | s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9(sK6,sK1,sK0)))) = s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) | ($spl42 | $spl51 | $spl52)),
  inference(forward_demodulation,[],[f3979,f3955])).
fof(f3979,plain,(
  ( ! [X8,X9] : (s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9(sK6,sK1,sK0)))) = s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) | ($spl42 | $spl51 | $spl52)),
  inference(forward_demodulation,[],[f3964,f3862])).
fof(f3964,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X8),X9),s(fun(sK1,X8),X10))) = s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9('const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) | ($spl42 | $spl51)),
  inference(superposition,[],[f332,f3955])).
fof(f3977,plain,(
  $false | ($spl40 | $spl42 | $spl48 | $spl51)),
  inference(global_subsumption,[],[f30,f29,f28,f27,f38,f39,f31,f32,f33,f34,f35,f36,f37,f41,f43,f46,f78,f110,f212,f214,f216,f217,f219,f221,f255,f257,f259,f292,f294,f296,f317,f319,f321,f322,f324,f326,f327,f328,f331,f332,f350,f1511,f383,f384,f1371,f425,f426,f1369,f446,f447,f448,f474,f475,f562,f565,f566,f567,f606,f609,f610,f611,f981,f628,f631,f632,f633,f636,f656,f659,f660,f661,f855,f1573,f772,f1648,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1586,f1587,f1588,f1589,f1590,f1591,f1643,f1615,f1848,f1849,f1850,f1853,f1855,f1856,f1857,f1861,f1863,f1865,f1866,f1941,f1870,f1872,f1873,f1874,f1875,f1876,f1933,f2632,f2633,f2634,f2637,f2639,f2640,f2641,f2645,f2647,f2649,f2650,f2727,f2655,f2657,f2658,f2659,f2660,f2661,f2719,f3044,f3045,f3046,f3049,f3051,f3052,f3053,f3057,f3058,f3059,f3060,f3177,f3178,f3181,f3183,f3184,f3185,f3189,f3190,f3191,f3192,f3641,f3501,f3502,f3503,f3506,f3508,f3509,f3510,f3643,f3516,f3517,f3518,f3519,f3520,f3647,f3648,f3523,f3524,f3533,f3534,f3535,f3536,f3552,f3553,f3556,f3558,f3559,f3560,f3565,f3566,f3574,f3583,f3591,f3594,f3595,f3596,f3599,f3600,f3610,f3619,f3627,f3630,f3631,f3632,f3633,f3634,f3635,f3636,f3637,f3638,f3659,f3661,f3664,f3827,f3671,f3672,f3673,f3676,f3678,f3679,f3680,f3830,f3686,f3687,f3688,f3689,f3690,f3691,f3831,f3832,f3694,f3695,f3696,f3697,f3710,f3711,f3713,f3714,f3730,f3731,f3734,f3736,f3737,f3738,f3743,f3744,f3752,f3761,f3771,f3774,f3775,f3776,f3777,f3781,f3782,f3792,f3801,f3811,f3814,f3815,f3816,f3817,f3819,f3820,f3822,f3823,f3824,f3839,f3845,f3856,f3829,f40,f3867,f3873,f3883,f3857,f3644,f3646,f3895,f3901,f3858,f3955,f3972,f3975,f3962])).
fof(f3976,plain,(
  $false | ($spl40 | $spl42 | $spl48 | $spl51)),
  inference(global_subsumption,[],[f30,f29,f28,f27,f38,f39,f31,f32,f33,f34,f35,f36,f37,f41,f43,f46,f78,f110,f212,f214,f216,f217,f219,f221,f255,f257,f259,f292,f294,f296,f317,f319,f321,f322,f324,f326,f327,f328,f331,f332,f350,f1511,f383,f384,f1371,f425,f426,f1369,f446,f447,f448,f474,f475,f562,f565,f566,f567,f606,f609,f610,f611,f981,f628,f631,f632,f633,f636,f656,f659,f660,f661,f855,f1573,f772,f1648,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1586,f1587,f1588,f1589,f1590,f1591,f1643,f1615,f1848,f1849,f1850,f1853,f1855,f1856,f1857,f1861,f1863,f1865,f1866,f1941,f1870,f1872,f1873,f1874,f1875,f1876,f1933,f2632,f2633,f2634,f2637,f2639,f2640,f2641,f2645,f2647,f2649,f2650,f2727,f2655,f2657,f2658,f2659,f2660,f2661,f2719,f3044,f3045,f3046,f3049,f3051,f3052,f3053,f3057,f3058,f3059,f3060,f3177,f3178,f3181,f3183,f3184,f3185,f3189,f3190,f3191,f3192,f3641,f3501,f3502,f3503,f3506,f3508,f3509,f3510,f3643,f3516,f3517,f3518,f3519,f3520,f3647,f3648,f3523,f3524,f3533,f3534,f3535,f3536,f3552,f3553,f3556,f3558,f3559,f3560,f3565,f3566,f3574,f3583,f3591,f3594,f3595,f3596,f3599,f3600,f3610,f3619,f3627,f3630,f3631,f3632,f3633,f3634,f3635,f3636,f3637,f3638,f3659,f3661,f3664,f3827,f3671,f3672,f3673,f3676,f3678,f3679,f3680,f3830,f3686,f3687,f3688,f3689,f3690,f3691,f3831,f3832,f3694,f3695,f3696,f3697,f3710,f3711,f3713,f3714,f3730,f3731,f3734,f3736,f3737,f3738,f3743,f3744,f3752,f3761,f3771,f3774,f3775,f3776,f3777,f3781,f3782,f3792,f3801,f3811,f3814,f3815,f3816,f3817,f3819,f3820,f3822,f3823,f3824,f3839,f3845,f3856,f3829,f40,f3867,f3873,f3883,f3857,f3644,f3646,f3895,f3901,f3858,f3955,f3972,f3975])).
fof(f3973,plain,(
  $false | ($spl40 | $spl42 | $spl48 | $spl51)),
  inference(global_subsumption,[],[f30,f29,f28,f27,f38,f39,f31,f32,f33,f34,f35,f36,f37,f41,f43,f46,f78,f110,f212,f214,f216,f217,f219,f221,f255,f257,f259,f292,f294,f296,f317,f319,f321,f322,f324,f326,f327,f328,f331,f332,f350,f1511,f383,f384,f1371,f425,f426,f1369,f446,f447,f448,f474,f475,f562,f565,f566,f567,f606,f609,f610,f611,f981,f628,f631,f632,f633,f636,f656,f659,f660,f661,f855,f1573,f772,f1648,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1586,f1587,f1588,f1589,f1590,f1591,f1643,f1615,f1848,f1849,f1850,f1853,f1855,f1856,f1857,f1861,f1863,f1865,f1866,f1941,f1870,f1872,f1873,f1874,f1875,f1876,f1933,f2632,f2633,f2634,f2637,f2639,f2640,f2641,f2645,f2647,f2649,f2650,f2727,f2655,f2657,f2658,f2659,f2660,f2661,f2719,f3044,f3045,f3046,f3049,f3051,f3052,f3053,f3057,f3058,f3059,f3060,f3177,f3178,f3181,f3183,f3184,f3185,f3189,f3190,f3191,f3192,f3641,f3501,f3502,f3503,f3506,f3508,f3509,f3510,f3643,f3516,f3517,f3518,f3519,f3520,f3647,f3648,f3523,f3524,f3533,f3534,f3535,f3536,f3552,f3553,f3556,f3558,f3559,f3560,f3565,f3566,f3574,f3583,f3591,f3594,f3595,f3596,f3599,f3600,f3610,f3619,f3627,f3630,f3631,f3632,f3633,f3634,f3635,f3636,f3637,f3638,f3659,f3661,f3664,f3827,f3671,f3672,f3673,f3676,f3678,f3679,f3680,f3830,f3686,f3687,f3688,f3689,f3690,f3691,f3831,f3832,f3694,f3695,f3696,f3697,f3710,f3711,f3713,f3714,f3730,f3731,f3734,f3736,f3737,f3738,f3743,f3744,f3752,f3761,f3771,f3774,f3775,f3776,f3777,f3781,f3782,f3792,f3801,f3811,f3814,f3815,f3816,f3817,f3819,f3820,f3822,f3823,f3824,f3839,f3845,f3856,f3829,f40,f3867,f3873,f3883,f3857,f3644,f3646,f3895,f3901,f3858,f3955,f3972])).
fof(f3948,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK5,sK1,sK0)))) | $spl71),
  inference(cnf_transformation,[],[f3948_D])).
fof(f3948_D,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK5,sK1,sK0)))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f3946,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK5,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0)))) | ($spl47 | $spl52)),
  inference(trivial_inequality_removal,[],[f3945])).
fof(f3945,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK3),s(sK1,sK10(sK5,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),sK5) | s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0)))) | ($spl47 | $spl52)),
  inference(superposition,[],[f3866,f327])).
fof(f3866,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) != s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK5,sK1,sK0)))) | ($spl47 | $spl52)),
  inference(backward_demodulation,[],[f3862,f3656])).
fof(f3939,plain,(
  ( ! [X12,X11] : (s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10(sK6,sK1,sK0)))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) | $spl66),
  inference(cnf_transformation,[],[f3939_D])).
fof(f3939_D,plain,(
  ( ! [X12,X11] : (s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10(sK6,sK1,sK0)))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f3941,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) | $spl68),
  inference(cnf_transformation,[],[f3941_D])).
fof(f3941_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f3938,plain,(
  ( ! [X12,X11] : (s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) | s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10(sK6,sK1,sK0)))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) | ($spl40 | $spl50)),
  inference(forward_demodulation,[],[f3937,f3859])).
fof(f3937,plain,(
  ( ! [X12,X11] : (s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10(sK6,sK1,sK0)))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) | ($spl40 | $spl50)),
  inference(forward_demodulation,[],[f3923,f3644])).
fof(f3923,plain,(
  ( ! [X12,X13,X11] : (s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X11),X13),s(fun(sK1,X11),X12))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) | $spl50),
  inference(superposition,[],[f328,f3859])).
fof(f3934,plain,(
  ( ! [X8,X9] : (s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9(sK6,sK1,sK0)))) = s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) | $spl62),
  inference(cnf_transformation,[],[f3934_D])).
fof(f3934_D,plain,(
  ( ! [X8,X9] : (s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9(sK6,sK1,sK0)))) = s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f3936,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) | $spl64),
  inference(cnf_transformation,[],[f3936_D])).
fof(f3936_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f3933,plain,(
  ( ! [X8,X9] : (s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) | s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9(sK6,sK1,sK0)))) = s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) | ($spl50 | $spl52)),
  inference(forward_demodulation,[],[f3932,f3859])).
fof(f3932,plain,(
  ( ! [X8,X9] : (s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9(sK6,sK1,sK0)))) = s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) | ($spl50 | $spl52)),
  inference(forward_demodulation,[],[f3922,f3862])).
fof(f3922,plain,(
  ( ! [X10,X8,X9] : (s(X8,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X8),X9),s(fun(sK1,X8),X10))) = s(X8,happ(s(fun(sK0,X8),X9),s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))),sK1,sK0))))) ) | $spl50),
  inference(superposition,[],[f332,f3859])).
fof(f3894,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) != s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,sK9(sK5,sK1,sK0)))) | $spl61),
  inference(cnf_transformation,[],[f3894_D])).
fof(f3894_D,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) != s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,sK9(sK5,sK1,sK0)))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f3889,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) != s(sK2,happ(s(fun(sK0,sK2),sK7),s(sK0,sK9(sK5,sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK5,sK1,sK0)))) | $spl52),
  inference(superposition,[],[f3865,f332])).
fof(f3865,plain,(
  s(sK2,happ(s(fun(sK0,sK2),sK4),s(sK0,sK9(sK6,sK1,sK0)))) != s(sK2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK2),sK7),s(fun(sK1,sK2),sK8))) | $spl52),
  inference(backward_demodulation,[],[f3862,f40])).
fof(f3888,plain,(
  $false | ($spl40 | $spl43)),
  inference(global_subsumption,[],[f30,f29,f28,f27,f38,f39,f31,f32,f33,f34,f35,f36,f37,f41,f43,f46,f78,f110,f212,f214,f216,f217,f219,f221,f255,f257,f259,f292,f294,f296,f317,f319,f321,f322,f324,f326,f327,f328,f331,f332,f350,f1511,f383,f384,f1371,f425,f426,f1369,f446,f447,f448,f474,f475,f562,f565,f566,f567,f606,f609,f610,f611,f981,f628,f631,f632,f633,f636,f656,f659,f660,f661,f855,f1573,f772,f1648,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1586,f1587,f1588,f1589,f1590,f1591,f1643,f1615,f1848,f1849,f1850,f1853,f1855,f1856,f1857,f1861,f1863,f1865,f1866,f1941,f1870,f1872,f1873,f1874,f1875,f1876,f1933,f2632,f2633,f2634,f2637,f2639,f2640,f2641,f2645,f2647,f2649,f2650,f2727,f2655,f2657,f2658,f2659,f2660,f2661,f2719,f3044,f3045,f3046,f3049,f3051,f3052,f3053,f3057,f3058,f3059,f3060,f3177,f3178,f3181,f3183,f3184,f3185,f3189,f3190,f3191,f3192,f3641,f3501,f3502,f3503,f3506,f3508,f3509,f3510,f3643,f3516,f3517,f3518,f3519,f3520,f3647,f3648,f3523,f3524,f3533,f3534,f3535,f3536,f3552,f3553,f3556,f3558,f3559,f3560,f3565,f3566,f3574,f3583,f3591,f3594,f3595,f3596,f3599,f3600,f3610,f3619,f3627,f3630,f3631,f3632,f3633,f3634,f3635,f3636,f3637,f3638,f3659,f3661,f3664,f3827,f3671,f3672,f3673,f3676,f3678,f3679,f3680,f3830,f3686,f3687,f3688,f3689,f3690,f3691,f3831,f3832,f3694,f3695,f3696,f3697,f3710,f3711,f3713,f3714,f3730,f3731,f3734,f3736,f3737,f3738,f3743,f3744,f3752,f3761,f3771,f3774,f3775,f3776,f3777,f3781,f3782,f3792,f3801,f3811,f3814,f3815,f3816,f3817,f3819,f3820,f3822,f3823,f3824,f3839,f3845,f3856,f3829,f40,f3645,f3644])).
fof(f3645,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) | $spl43),
  inference(cnf_transformation,[],[f3645_D])).
fof(f3645_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) != s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f3887,plain,(
  $false | ($spl40 | $spl43)),
  inference(global_subsumption,[],[f30,f29,f28,f27,f38,f39,f31,f32,f33,f34,f35,f36,f37,f41,f43,f46,f78,f110,f212,f214,f216,f217,f219,f221,f255,f257,f259,f292,f294,f296,f317,f319,f321,f322,f324,f326,f327,f328,f331,f332,f350,f1511,f383,f384,f1371,f425,f426,f1369,f446,f447,f448,f474,f475,f562,f565,f566,f567,f606,f609,f610,f611,f981,f628,f631,f632,f633,f636,f656,f659,f660,f661,f855,f1573,f772,f1648,f1579,f1580,f1581,f1582,f1583,f1584,f1585,f1586,f1587,f1588,f1589,f1590,f1591,f1643,f1615,f1848,f1849,f1850,f1853,f1855,f1856,f1857,f1861,f1863,f1865,f1866,f1941,f1870,f1872,f1873,f1874,f1875,f1876,f1933,f2632,f2633,f2634,f2637,f2639,f2640,f2641,f2645,f2647,f2649,f2650,f2727,f2655,f2657,f2658,f2659,f2660,f2661,f2719,f3044,f3045,f3046,f3049,f3051,f3052,f3053,f3057,f3058,f3059,f3060,f3177,f3178,f3181,f3183,f3184,f3185,f3189,f3190,f3191,f3192,f3641,f3501,f3502,f3503,f3506,f3508,f3509,f3510,f3643,f3516,f3517,f3518,f3519,f3520,f3647,f3648,f3523,f3524,f3533,f3534,f3535,f3536,f3552,f3553,f3556,f3558,f3559,f3560,f3565,f3566,f3574,f3583,f3591,f3594,f3595,f3596,f3599,f3600,f3610,f3619,f3627,f3630,f3631,f3632,f3633,f3634,f3635,f3636,f3637,f3638,f3644,f3659,f3661,f3664,f3827,f3671,f3672,f3673,f3676,f3678,f3679,f3680,f3830,f3686,f3687,f3688,f3689,f3690,f3691,f3831,f3832,f3694,f3695,f3696,f3697,f3710,f3711,f3713,f3714,f3730,f3731,f3734,f3736,f3737,f3738,f3743,f3744,f3752,f3761,f3771,f3774,f3775,f3776,f3777,f3781,f3782,f3792,f3801,f3811,f3814,f3815,f3816,f3817,f3819,f3820,f3822,f3823,f3824,f3839,f3845,f3856,f3829,f40,f3645])).
fof(f3884,plain,(
  ( ! [X12,X13,X11] : (s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X11),X13),s(fun(sK1,X11),X12))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) | $spl56),
  inference(cnf_transformation,[],[f3884_D])).
fof(f3884_D,plain,(
  ( ! [X12,X13,X11] : (s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X11),X13),s(fun(sK1,X11),X12))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f3883,plain,(
  ( ! [X12,X13,X11] : (s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X11),X13),s(fun(sK1,X11),X12))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) | $spl42),
  inference(forward_demodulation,[],[f3876,f3646])).
fof(f3876,plain,(
  ( ! [X12,X13,X11] : (s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X11),X13),s(fun(sK1,X11),X12))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) | $spl42),
  inference(superposition,[],[f328,f3646])).
fof(f3862,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK0,X0),X1),s(sK0,sK9(sK6,sK1,sK0)))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X0),X1),s(fun(sK1,X0),X2)))) ) | $spl52),
  inference(cnf_transformation,[],[f3862_D])).
fof(f3862_D,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK0,X0),X1),s(sK0,sK9(sK6,sK1,sK0)))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X0),X1),s(fun(sK1,X0),X2)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f3864,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK6,sK1,sK0)))) | $spl54),
  inference(cnf_transformation,[],[f3864_D])).
fof(f3864_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK6,sK1,sK0)))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f3829,plain,(
  ( ! [X2,X0,X1] : (s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK6,sK1,sK0)))) | s(X0,happ(s(fun(sK0,X0),X1),s(sK0,sK9(sK6,sK1,sK0)))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X0),X1),s(fun(sK1,X0),X2)))) )),
  inference(forward_demodulation,[],[f3685,f27])).
fof(f3685,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK0,X0),X1),s(sK0,sK9(sK6,sK1,sK0)))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X0),X1),s(fun(sK1,X0),X2))) | s('type.sum.sum'(sK0,sK1),sK6) = s('type.sum.sum'(sK0,sK1),'const.sum.INR_1'(s(sK1,sK10(sK6,sK1,sK0))))) )),
  inference(superposition,[],[f332,f27])).
fof(f3857,plain,(
  ( ! [X12,X11] : (s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10(sK6,sK1,sK0)))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) | $spl48),
  inference(cnf_transformation,[],[f3857_D])).
fof(f3857_D,plain,(
  ( ! [X12,X11] : (s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10(sK6,sK1,sK0)))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f3856,plain,(
  ( ! [X12,X11] : (s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10(sK6,sK1,sK0)))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) | $spl40),
  inference(forward_demodulation,[],[f3855,f3664])).
fof(f3855,plain,(
  ( ! [X12,X11] : (s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10(sK6,sK1,sK0)))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) | $spl40),
  inference(forward_demodulation,[],[f3848,f3644])).
fof(f3848,plain,(
  ( ! [X12,X13,X11] : (s(X11,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X11),X13),s(fun(sK1,X11),X12))) = s(X11,happ(s(fun(sK1,X11),X12),s(sK1,sK10('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0)))) | s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9('const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))),sK1,sK0))))) ) | $spl40),
  inference(superposition,[],[f328,f3664])).
fof(f3656,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK5,sK1,sK0)))) != s(sK2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK2),sK4),s(fun(sK1,sK2),sK3))) | $spl47),
  inference(cnf_transformation,[],[f3656_D])).
fof(f3656_D,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK5,sK1,sK0)))) != s(sK2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK2),sK4),s(fun(sK1,sK2),sK3))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f3633,plain,(
  s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK10(sK5,sK1,sK0)))) != s(sK2,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,sK2),sK4),s(fun(sK1,sK2),sK3))) | s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK5,sK1,sK0))))),
  inference(superposition,[],[f40,f328])).
fof(f3646,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) | $spl42),
  inference(cnf_transformation,[],[f3646_D])).
fof(f3646_D,plain,(
  s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f3643,plain,(
  ( ! [X2,X0,X1] : (s('type.sum.sum'(sK0,sK1),sK5) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0)))) | s(X0,happ(s(fun(sK1,X0),X1),s(sK1,sK10(sK6,sK1,sK0)))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X0),X2),s(fun(sK1,X0),X1)))) )),
  inference(forward_demodulation,[],[f3515,f27])).
fof(f3515,plain,(
  ( ! [X2,X0,X1] : (s(X0,happ(s(fun(sK1,X0),X1),s(sK1,sK10(sK6,sK1,sK0)))) = s(X0,'const.sum.sum_CASE_3'(s('type.sum.sum'(sK0,sK1),sK5),s(fun(sK0,X0),X2),s(fun(sK1,X0),X1))) | s('type.sum.sum'(sK0,sK1),sK6) = s('type.sum.sum'(sK0,sK1),'const.sum.INL_1'(s(sK0,sK9(sK6,sK1,sK0))))) )),
  inference(superposition,[],[f328,f27])).
fof(f2601,plain,(
  ( ! [X64] : (s(bool,'T_0') = s(bool,X64)) ) | $spl26),
  inference(cnf_transformation,[],[f2601_D])).
fof(f2601_D,plain,(
  ( ! [X64] : (s(bool,'T_0') = s(bool,X64)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2623,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) = s(bool,g(s(bool,X15)))) ) | $spl38),
  inference(cnf_transformation,[],[f2623_D])).
fof(f2623_D,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) = s(bool,g(s(bool,X15)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2388,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,g(s(bool,X15)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2234])).
fof(f2234,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,g(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) ) | $spl8),
  inference(superposition,[],[f43,f1665])).
fof(f1665,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,g(s(bool,X8))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl8),
  inference(superposition,[],[f1653,f41])).
fof(f2613,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s(bool,X1)))) ) | $spl36),
  inference(cnf_transformation,[],[f2613_D])).
fof(f2613_D,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s(bool,X1)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2393,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,g(s(bool,X1)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2229])).
fof(f2229,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f33,f1665])).
fof(f2610,plain,(
  ( ! [X103,X101,X102] : (s(bool,sK11(X101,X102,X103,bool)) = s(bool,g(s(bool,sK11(X101,X102,X103,bool)))) | s(X103,f(s(bool,'F_0'))) != s(X103,g(s(bool,'F_0'))) | s(fun(bool,X103),X101) = s(fun(bool,X103),X102)) ) | $spl34),
  inference(cnf_transformation,[],[f2610_D])).
fof(f2610_D,plain,(
  ( ! [X103,X101,X102] : (s(bool,sK11(X101,X102,X103,bool)) = s(bool,g(s(bool,sK11(X101,X102,X103,bool)))) | s(X103,f(s(bool,'F_0'))) != s(X103,g(s(bool,'F_0'))) | s(fun(bool,X103),X101) = s(fun(bool,X103),X102)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2394,plain,(
  ( ! [X101,X99,X102,X100] : (s(bool,'T_0') = s(bool,X102) | s(fun(bool,X101),X99) = s(fun(bool,X101),X100) | s(X101,f(s(bool,'F_0'))) != s(X101,g(s(bool,'F_0'))) | s(bool,sK11(X99,X100,X101,bool)) = s(bool,g(s(bool,sK11(X99,X100,X101,bool))))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2218])).
fof(f2218,plain,(
  ( ! [X101,X99,X102,X100] : (s(bool,'T_0') = s(bool,X102) | s(fun(bool,X101),X99) = s(fun(bool,X101),X100) | s(X101,f(s(bool,'F_0'))) != s(X101,g(s(bool,'F_0'))) | s(bool,sK11(X99,X100,X101,bool)) = s(bool,g(s(bool,sK11(X99,X100,X101,bool)))) | s(bool,'T_0') = s(bool,X102)) ) | $spl8),
  inference(superposition,[],[f216,f1665])).
fof(f2609,plain,(
  ( ! [X97,X95,X98] : (s(bool,happ(s(fun(X95,bool),X98),s(X95,X97))) = s(bool,g(s(bool,happ(s(fun(X95,bool),X98),s(X95,X97))))) | s(bool,'F_0') = s(bool,happ(s(fun(X95,bool),X98),s(X95,X97)))) ) | $spl32),
  inference(cnf_transformation,[],[f2609_D])).
fof(f2609_D,plain,(
  ( ! [X97,X95,X98] : (s(bool,happ(s(fun(X95,bool),X98),s(X95,X97))) = s(bool,g(s(bool,happ(s(fun(X95,bool),X98),s(X95,X97))))) | s(bool,'F_0') = s(bool,happ(s(fun(X95,bool),X98),s(X95,X97)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2620,plain,(
  ( ! [X95,X93,X98,X96] : (s(bool,happ(s(fun(X93,bool),X96),s(X93,X95))) = s(bool,g(s(bool,happ(s(fun(X93,bool),X96),s(X93,X95))))) | s(bool,'T_0') = s(bool,X98) | s(bool,'F_0') = s(bool,happ(s(fun(X93,bool),X96),s(X93,X95)))) ) | $spl8),
  inference(forward_demodulation,[],[f2395,f36])).
fof(f2395,plain,(
  ( ! [X94,X97,X95,X93,X98,X96] : (s(bool,'T_0') = s(bool,X98) | s(bool,'F_0') = s(bool,happ(s(fun(X93,bool),X96),s(X93,X95))) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X93,X94),'const.sum.INL_1'(s(X93,X95))),s(fun(X93,bool),X96),s(fun(X94,bool),X97))) = s(bool,g(s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X93,X94),'const.sum.INL_1'(s(X93,X95))),s(fun(X93,bool),X96),s(fun(X94,bool),X97)))))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2217])).
fof(f2217,plain,(
  ( ! [X94,X97,X95,X93,X98,X96] : (s(bool,'T_0') = s(bool,X98) | s(bool,'F_0') = s(bool,happ(s(fun(X93,bool),X96),s(X93,X95))) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X93,X94),'const.sum.INL_1'(s(X93,X95))),s(fun(X93,bool),X96),s(fun(X94,bool),X97))) = s(bool,g(s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X93,X94),'const.sum.INL_1'(s(X93,X95))),s(fun(X93,bool),X96),s(fun(X94,bool),X97))))) | s(bool,'T_0') = s(bool,X98)) ) | $spl8),
  inference(superposition,[],[f259,f1665])).
fof(f2603,plain,(
  ( ! [X85,X86,X84] : (s(bool,g(s(X84,sK11(X85,X86,bool,X84)))) = s(bool,g(s(bool,g(s(X84,sK11(X85,X86,bool,X84)))))) | s(bool,'F_0') != s(bool,f(s(X84,sK11(X85,X86,bool,X84)))) | s(fun(X84,bool),X85) = s(fun(X84,bool),X86)) ) | $spl30),
  inference(cnf_transformation,[],[f2603_D])).
fof(f2603_D,plain,(
  ( ! [X85,X86,X84] : (s(bool,g(s(X84,sK11(X85,X86,bool,X84)))) = s(bool,g(s(bool,g(s(X84,sK11(X85,X86,bool,X84)))))) | s(bool,'F_0') != s(bool,f(s(X84,sK11(X85,X86,bool,X84)))) | s(fun(X84,bool),X85) = s(fun(X84,bool),X86)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2400,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') = s(bool,X88) | s(fun(X85,bool),X86) = s(fun(X85,bool),X87) | s(bool,'F_0') != s(bool,f(s(X85,sK11(X86,X87,bool,X85)))) | s(bool,g(s(X85,sK11(X86,X87,bool,X85)))) = s(bool,g(s(bool,g(s(X85,sK11(X86,X87,bool,X85))))))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2213])).
fof(f2213,plain,(
  ( ! [X88,X87,X85,X86] : (s(bool,'T_0') = s(bool,X88) | s(fun(X85,bool),X86) = s(fun(X85,bool),X87) | s(bool,'F_0') != s(bool,f(s(X85,sK11(X86,X87,bool,X85)))) | s(bool,g(s(X85,sK11(X86,X87,bool,X85)))) = s(bool,g(s(bool,g(s(X85,sK11(X86,X87,bool,X85)))))) | s(bool,'T_0') = s(bool,X88)) ) | $spl8),
  inference(superposition,[],[f221,f1665])).
fof(f2455,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') = s(bool,X54) | s(bool,X55) = s(bool,g(s(bool,X55))) | s(bool,'F_0') = s(bool,X55)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2160])).
fof(f2160,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') = s(bool,X54) | s(bool,X55) = s(bool,g(s(bool,X55))) | s(bool,'F_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X54)) ) | $spl8),
  inference(superposition,[],[f1665,f33])).
fof(f2602,plain,(
  ( ! [X65,X63] : (s(bool,X63) = s(bool,g(s(bool,X63))) | s(bool,X63) = s(bool,X65) | s(bool,X65) = s(bool,g(s(bool,X65)))) ) | $spl28),
  inference(cnf_transformation,[],[f2602_D])).
fof(f2602_D,plain,(
  ( ! [X65,X63] : (s(bool,X63) = s(bool,g(s(bool,X63))) | s(bool,X63) = s(bool,X65) | s(bool,X65) = s(bool,g(s(bool,X65)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2456,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') = s(bool,X44) | s(bool,X45) = s(bool,g(s(bool,X45))) | s(bool,X43) = s(bool,X45) | s(bool,X43) = s(bool,g(s(bool,X43)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2158])).
fof(f2158,plain,(
  ( ! [X45,X43,X44] : (s(bool,'T_0') = s(bool,X44) | s(bool,X45) = s(bool,g(s(bool,X45))) | s(bool,X43) = s(bool,X45) | s(bool,X43) = s(bool,g(s(bool,X43))) | s(bool,'T_0') = s(bool,X44)) ) | $spl8),
  inference(superposition,[],[f1665,f1665])).
fof(f2487,plain,(
  ( ! [X61,X62,X60] : (s(bool,X61) = s(bool,X62) | s(bool,X62) = s(bool,g(s(bool,X62))) | s(bool,'T_0') = s(bool,X60) | s(bool,X61) = s(bool,g(s(bool,X61)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2113])).
fof(f2113,plain,(
  ( ! [X61,X62,X60] : (s(bool,X61) = s(bool,X62) | s(bool,X62) = s(bool,g(s(bool,X62))) | s(bool,'T_0') = s(bool,X60) | s(bool,X61) = s(bool,g(s(bool,X61))) | s(bool,'T_0') = s(bool,X60)) ) | $spl8),
  inference(superposition,[],[f1665,f1665])).
fof(f2504,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2098])).
fof(f2098,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(superposition,[],[f1665,f33])).
fof(f2505,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'T_0') = s(bool,X104) | s(bool,sK11(X101,X102,X103,bool)) = s(bool,g(s(bool,sK11(X101,X102,X103,bool)))) | s(fun(bool,X103),X101) = s(fun(bool,X103),X102) | s(X103,f(s(bool,'F_0'))) != s(X103,g(s(bool,'F_0')))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2095])).
fof(f2095,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,'T_0') = s(bool,X104) | s(bool,sK11(X101,X102,X103,bool)) = s(bool,g(s(bool,sK11(X101,X102,X103,bool)))) | s(bool,'T_0') = s(bool,X104) | s(fun(bool,X103),X101) = s(fun(bool,X103),X102) | s(X103,f(s(bool,'F_0'))) != s(X103,g(s(bool,'F_0')))) ) | $spl8),
  inference(superposition,[],[f1665,f216])).
fof(f2608,plain,(
  ( ! [X97,X95,X100,X98] : (s(bool,happ(s(fun(X95,bool),X98),s(X95,X97))) = s(bool,g(s(bool,happ(s(fun(X95,bool),X98),s(X95,X97))))) | s(bool,'T_0') = s(bool,X100) | s(bool,'F_0') = s(bool,happ(s(fun(X95,bool),X98),s(X95,X97)))) ) | $spl8),
  inference(forward_demodulation,[],[f2506,f36])).
fof(f2506,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (s(bool,'T_0') = s(bool,X100) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X95,X96),'const.sum.INL_1'(s(X95,X97))),s(fun(X95,bool),X98),s(fun(X96,bool),X99))) = s(bool,g(s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X95,X96),'const.sum.INL_1'(s(X95,X97))),s(fun(X95,bool),X98),s(fun(X96,bool),X99))))) | s(bool,'F_0') = s(bool,happ(s(fun(X95,bool),X98),s(X95,X97)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2094])).
fof(f2094,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (s(bool,'T_0') = s(bool,X100) | s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X95,X96),'const.sum.INL_1'(s(X95,X97))),s(fun(X95,bool),X98),s(fun(X96,bool),X99))) = s(bool,g(s(bool,'const.sum.sum_CASE_3'(s('type.sum.sum'(X95,X96),'const.sum.INL_1'(s(X95,X97))),s(fun(X95,bool),X98),s(fun(X96,bool),X99))))) | s(bool,'T_0') = s(bool,X100) | s(bool,'F_0') = s(bool,happ(s(fun(X95,bool),X98),s(X95,X97)))) ) | $spl8),
  inference(superposition,[],[f1665,f259])).
fof(f2513,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') = s(bool,X87) | s(bool,g(s(X84,sK11(X85,X86,bool,X84)))) = s(bool,g(s(bool,g(s(X84,sK11(X85,X86,bool,X84)))))) | s(fun(X84,bool),X85) = s(fun(X84,bool),X86) | s(bool,'F_0') != s(bool,f(s(X84,sK11(X85,X86,bool,X84))))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2089])).
fof(f2089,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') = s(bool,X87) | s(bool,g(s(X84,sK11(X85,X86,bool,X84)))) = s(bool,g(s(bool,g(s(X84,sK11(X85,X86,bool,X84)))))) | s(bool,'T_0') = s(bool,X87) | s(fun(X84,bool),X85) = s(fun(X84,bool),X86) | s(bool,'F_0') != s(bool,f(s(X84,sK11(X85,X86,bool,X84))))) ) | $spl8),
  inference(superposition,[],[f1665,f221])).
fof(f2531,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') = s(bool,X55) | s(bool,X54) = s(bool,g(s(bool,X54))) | s(bool,X56) = s(bool,g(s(bool,X56))) | s(bool,X54) = s(bool,X56)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2080])).
fof(f2080,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') = s(bool,X55) | s(bool,X54) = s(bool,g(s(bool,X54))) | s(bool,'T_0') = s(bool,X55) | s(bool,X56) = s(bool,g(s(bool,X56))) | s(bool,X54) = s(bool,X56)) ) | $spl8),
  inference(superposition,[],[f1665,f1665])).
fof(f2575,plain,(
  ( ! [X64,X65,X63] : (s(bool,X63) = s(bool,g(s(bool,X63))) | s(bool,'T_0') = s(bool,X64) | s(bool,X65) = s(bool,g(s(bool,X65))) | s(bool,X63) = s(bool,X65)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f2033])).
fof(f2033,plain,(
  ( ! [X64,X65,X63] : (s(bool,X63) = s(bool,g(s(bool,X63))) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X64) | s(bool,X65) = s(bool,g(s(bool,X65))) | s(bool,X63) = s(bool,X65)) ) | $spl8),
  inference(superposition,[],[f1665,f1665])).
fof(f1654,plain,(
  ( ! [X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK11(X2,X3,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X2,X3,bool,bool))))) ) | $spl10),
  inference(cnf_transformation,[],[f1654_D])).
fof(f1654_D,plain,(
  ( ! [X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK11(X2,X3,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X2,X3,bool,bool))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1662,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') = s(bool,g(s(bool,X38))) | s(bool,'T_0') = s(bool,X40) | s(bool,X39) = s(bool,X40) | s(bool,X38) = s(bool,X39)) ) | $spl24),
  inference(cnf_transformation,[],[f1662_D])).
fof(f1662_D,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') = s(bool,g(s(bool,X38))) | s(bool,'T_0') = s(bool,X40) | s(bool,X39) = s(bool,X40) | s(bool,X38) = s(bool,X39)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1591,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,'T_0') = s(bool,g(s(bool,X53))) | s(fun(bool,bool),X51) = s(fun(bool,bool),X52) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X51,X52,bool,bool)))) | s(bool,X53) = s(bool,X54) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,sK11(X51,X52,bool,bool))) )),
  inference(superposition,[],[f221,f78])).
fof(f1661,plain,(
  ( ! [X35,X31,X34,X32] : (s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,X35) | s(bool,X34) = s(bool,sK11(X31,X32,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X31,X32,bool,bool))))) ) | $spl22),
  inference(cnf_transformation,[],[f1661_D])).
fof(f1661_D,plain,(
  ( ! [X35,X31,X34,X32] : (s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,X35) | s(bool,X34) = s(bool,sK11(X31,X32,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X31,X32,bool,bool))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1658,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,g(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl16),
  inference(cnf_transformation,[],[f1658_D])).
fof(f1658_D,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,g(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1590,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,g(s(bool,X48))) | s(fun(bool,bool),X46) = s(fun(bool,bool),X47) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X46,X47,bool,bool)))) | s(bool,X49) = s(bool,sK11(X46,X47,bool,bool)) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f221,f78])).
fof(f1660,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,'T_0') = s(bool,X30) | s(bool,X28) = s(bool,X30) | s(bool,X29) = s(bool,sK11(X26,X27,bool,bool)) | s(bool,X28) = s(bool,X29) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X26,X27,bool,bool))))) ) | $spl20),
  inference(cnf_transformation,[],[f1660_D])).
fof(f1660_D,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,'T_0') = s(bool,X30) | s(bool,X28) = s(bool,X30) | s(bool,X29) = s(bool,sK11(X26,X27,bool,bool)) | s(bool,X28) = s(bool,X29) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X26,X27,bool,bool))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1653,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl8),
  inference(cnf_transformation,[],[f1653_D])).
fof(f1653_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1589,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X41) = s(fun(bool,bool),X42) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X41,X42,bool,bool)))) | s(bool,X43) = s(bool,X44) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,X43) = s(bool,sK11(X41,X42,bool,bool))) )),
  inference(superposition,[],[f221,f78])).
fof(f1588,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'T_0') = s(bool,g(s(bool,X38))) | s(fun(bool,bool),X36) = s(fun(bool,bool),X37) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X36,X37,bool,bool)))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,sK11(X36,X37,bool,bool)) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f221,f78])).
fof(f1587,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,g(s(bool,X33))) | s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X31,X32,bool,bool)))) | s(bool,X34) = s(bool,sK11(X31,X32,bool,bool)) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f221,f78])).
fof(f1586,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X26,X27,bool,bool)))) | s(bool,X28) = s(bool,X29) | s(bool,X29) = s(bool,sK11(X26,X27,bool,bool)) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f221,f78])).
fof(f1657,plain,(
  ( ! [X6,X4,X5] : (s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X6) | s(bool,X6) = s(bool,sK11(X4,X5,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X4,X5,bool,bool))))) ) | $spl14),
  inference(cnf_transformation,[],[f1657_D])).
fof(f1657_D,plain,(
  ( ! [X6,X4,X5] : (s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X6) | s(bool,X6) = s(bool,sK11(X4,X5,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X4,X5,bool,bool))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1659,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') = s(bool,g(s(bool,X18))) | s(bool,'T_0') = s(bool,X19) | s(bool,X18) = s(bool,X19)) ) | $spl18),
  inference(cnf_transformation,[],[f1659_D])).
fof(f1659_D,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') = s(bool,g(s(bool,X18))) | s(bool,'T_0') = s(bool,X19) | s(bool,X18) = s(bool,X19)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1585,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,g(s(bool,X23))) | s(fun(bool,bool),X21) = s(fun(bool,bool),X22) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X21,X22,bool,bool)))) | s(bool,X24) = s(bool,sK11(X21,X22,bool,bool)) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f221,f78])).
fof(f1584,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,g(s(bool,X18))) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X16,X17,bool,bool)))) | s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,sK11(X16,X17,bool,bool)) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f221,f78])).
fof(f1583,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X13) = s(fun(bool,bool),X14) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X13,X14,bool,bool)))) | s(bool,X15) = s(bool,sK11(X13,X14,bool,bool)) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f221,f41])).
fof(f1582,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,g(s(bool,X12))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X10,X11,bool,bool)))) | s(bool,'T_0') = s(bool,sK11(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f221,f41])).
fof(f1581,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') = s(bool,g(s(bool,X9))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X7,X8,bool,bool)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK11(X7,X8,bool,bool))) )),
  inference(superposition,[],[f221,f41])).
fof(f1580,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X4,X5,bool,bool)))) | s(bool,X6) = s(bool,sK11(X4,X5,bool,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f221,f41])).
fof(f1656,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl12),
  inference(cnf_transformation,[],[f1656_D])).
fof(f1656_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1579,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X2,X3,bool,bool)))) | s(bool,'T_0') = s(bool,sK11(X2,X3,bool,bool))) )),
  inference(superposition,[],[f221,f33])).
fof(f1650,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0'))) | $spl5),
  inference(cnf_transformation,[],[f1650_D])).
fof(f1650_D,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0'))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f1651,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X0,X1,bool,bool))))) ) | $spl6),
  inference(cnf_transformation,[],[f1651_D])).
fof(f1651_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X0,X1,bool,bool))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1648,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X0,X1,bool,bool)))) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f1578])).
fof(f1578,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK11(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f221,f216])).
fof(f47,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(cnf_transformation,[],[f47_D])).
fof(f47_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f49,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f49_D])).
fof(f49_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f46,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f33])).
% SZS output end Proof for thm.sum
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 66
% Passive clauses: 493
% Generated clauses: 4515
% Final active clauses: 36
% Final passive clauses: 259
% Input formulas: 8
% Initial clauses: 11
% 
% Duplicate literals: 636
% Trivial inequalities: 48
% Fw subsumption resolutions: 40
% Fw demodulations: 243
% Bw demodulations: 51
% Condensations: 875
% Global subsumptions: 8
% 
% Simple tautologies: 57
% Equational tautologies: 164
% Forward subsumptions: 1858
% 
% Unit resulting resolution: 234
% Forward superposition: 966
% Backward superposition: 1068
% Self superposition: 169
% Equality factoring: 122
% 
% Split clauses: 124
% Split components: 249
% SAT solver clauses: 767
% SAT solver unit clauses: 11
% SAT solver binary clauses: 113
% 
% Sat splits: 124
% Sat splitting refutations: 16
% 
% Memory used [KB]: 4861
% Time elapsed: 0.281 s
% ------------------------------
% ------------------------------
% Success in time 1.64 s
