Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.extreal
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.extreal
% SZS output start Proof for thm.extreal
fof(f9968,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f9967,f6084,f6081,f248,f365,f367,f229,f368])).
fof(f368,plain,(
  ~p(s(bool,'F_0')) | $spl3),
  inference(cnf_transformation,[],[f368_D])).
fof(f368_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f229,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_THM|split|0',aHL_FALSITY)).
fof(f367,plain,(
  p(s(bool,'F_0')) | $spl2),
  inference(cnf_transformation,[],[f367_D])).
fof(f367_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f365,plain,(
  p(s(bool,'T_0')) | $spl0),
  inference(cnf_transformation,[],[f365_D])).
fof(f365_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f248,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f173])).
fof(f173,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X8] : ((p(s(bool,'F_0')) => p(s(bool,X8))) <=> p(s(bool,'T_0')))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_THM|split|0','thm.bool.IMP_CLAUSES|split|2')).
fof(f6081,plain,(
  ( ! [X4,X2,X0,X3] : (s(X2,X4) = s(X2,'const.pred_set.ITSET_3'(s(fun(X0,fun(X2,X2)),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X2,X4)))) ) | $spl170),
  inference(cnf_transformation,[],[f6081_D])).
fof(f6081_D,plain,(
  ( ! [X4,X2,X0,X3] : (s(X2,X4) = s(X2,'const.pred_set.ITSET_3'(s(fun(X0,fun(X2,X2)),X3),s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(X2,X4)))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f6084,plain,(
  ( ! [X6,X8,X7,X5] : (s(X6,X8) = s(X6,'const.pred_set.ITSET_3'(s(fun(X5,fun(X6,X6)),X7),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(X6,X8)))) ) | $spl0),
  inference(forward_demodulation,[],[f6083,f271])).
fof(f271,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X11,X12] : s(X6,X11) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X11),s(X6,X12)))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_THM|split|0','thm.bool.COND_CLAUSES|split|0')).
fof(f6083,plain,(
  ( ! [X6,X8,X7,X5] : (s(X6,'const.pred_set.ITSET_3'(s(fun(X5,fun(X6,X6)),X7),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(X6,X8))) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X8),s(X6,'const.pred_set.ITSET_3'(s(fun(X5,fun(X6,X6)),X7),s(fun(X5,bool),'const.pred_set.REST_1'(s(fun(X5,bool),'const.pred_set.EMPTY_0'))),s(X6,happ(s(fun(X6,X6),happ(s(fun(X5,fun(X6,X6)),X7),s(X5,'const.pred_set.CHOICE_1'(s(fun(X5,bool),'const.pred_set.EMPTY_0'))))),s(X6,X8)))))))) ) | $spl0),
  inference(forward_demodulation,[],[f6082,f1228])).
fof(f1228,plain,(
  ( ! [X4,X2,X3,X1] : (s(bool,'T_0') = s(bool,sK10(X1,X2,'const.pred_set.EMPTY_0',X3,X4))) ) | $spl0),
  inference(resolution,[],[f1226,f241])).
fof(f241,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_THM|split|0','thm.bool.EQ_CLAUSES|split|1')).
fof(f1226,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,sK10(X0,X1,'const.pred_set.EMPTY_0',X2,X3)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1225,f365])).
fof(f1225,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,sK10(X0,X1,'const.pred_set.EMPTY_0',X2,X3)))) )),
  inference(forward_demodulation,[],[f1222,f391])).
fof(f391,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(resolution,[],[f241,f230])).
fof(f230,plain,(
  ( ! [X0] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_THM|split|0','thm.pred_set.FINITE_EMPTY')).
fof(f1222,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,sK10(X0,X1,'const.pred_set.EMPTY_0',X2,X3))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X3,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(equality_resolution,[],[f302])).
fof(f302,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X2) | p(s(bool,sK10(X4,X3,X2,X1,X0))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ((~p(s(bool,sK10(X4,X3,X2,X1,X0))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X2)) & (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X2) | p(s(bool,sK10(X4,X3,X2,X1,X0)))) & s(X1,'const.pred_set.ITSET_3'(s(fun(X0,fun(X1,X1)),X3),s(fun(X0,bool),X2),s(X1,X4))) = s(X1,'const.bool.COND_3'(s(bool,sK10(X4,X3,X2,X1,X0)),s(X1,X4),s(X1,'const.pred_set.ITSET_3'(s(fun(X0,fun(X1,X1)),X3),s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X2))),s(X1,happ(s(fun(X1,X1),happ(s(fun(X0,fun(X1,X1)),X3),s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X2))))),s(X1,X4)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10])],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ? [X5] : ((~p(s(bool,X5)) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X2)) & (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X2) | p(s(bool,X5))) & s(X1,'const.pred_set.ITSET_3'(s(fun(X0,fun(X1,X1)),X3),s(fun(X0,bool),X2),s(X1,X4))) = s(X1,'const.bool.COND_3'(s(bool,X5),s(X1,X4),s(X1,'const.pred_set.ITSET_3'(s(fun(X0,fun(X1,X1)),X3),s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X2))),s(X1,happ(s(fun(X1,X1),happ(s(fun(X0,fun(X1,X1)),X3),s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X2))))),s(X1,X4)))))))))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ? [X5] : (((~p(s(bool,X5)) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X2)) & (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X2) | p(s(bool,X5)))) & s(X1,'const.pred_set.ITSET_3'(s(fun(X0,fun(X1,X1)),X3),s(fun(X0,bool),X2),s(X1,X4))) = s(X1,'const.bool.COND_3'(s(bool,X5),s(X1,X4),s(X1,'const.pred_set.ITSET_3'(s(fun(X0,fun(X1,X1)),X3),s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X2))),s(X1,happ(s(fun(X1,X1),happ(s(fun(X0,fun(X1,X1)),X3),s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X2))))),s(X1,X4)))))))))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | ? [X5] : ((p(s(bool,X5)) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X2)) & s(X1,'const.pred_set.ITSET_3'(s(fun(X0,fun(X1,X1)),X3),s(fun(X0,bool),X2),s(X1,X4))) = s(X1,'const.bool.COND_3'(s(bool,X5),s(X1,X4),s(X1,'const.pred_set.ITSET_3'(s(fun(X0,fun(X1,X1)),X3),s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X2))),s(X1,happ(s(fun(X1,X1),happ(s(fun(X0,fun(X1,X1)),X3),s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X2))))),s(X1,X4)))))))))),
  inference(ennf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => ? [X5] : ((p(s(bool,X5)) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X2)) & s(X1,'const.pred_set.ITSET_3'(s(fun(X0,fun(X1,X1)),X3),s(fun(X0,bool),X2),s(X1,X4))) = s(X1,'const.bool.COND_3'(s(bool,X5),s(X1,X4),s(X1,'const.pred_set.ITSET_3'(s(fun(X0,fun(X1,X1)),X3),s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X2))),s(X1,happ(s(fun(X1,X1),happ(s(fun(X0,fun(X1,X1)),X3),s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X2))))),s(X1,X4)))))))))),
  inference(flattening,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => ? [X5] : ((p(s(bool,X5)) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X2)) & s(X1,'const.pred_set.ITSET_3'(s(fun(X0,fun(X1,X1)),X3),s(fun(X0,bool),X2),s(X1,X4))) = s(X1,'const.bool.COND_3'(s(bool,X5),s(X1,X4),s(X1,'const.pred_set.ITSET_3'(s(fun(X0,fun(X1,X1)),X3),s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X2))),s(X1,happ(s(fun(X1,X1),happ(s(fun(X0,fun(X1,X1)),X3),s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X2))))),s(X1,X4)))))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6,X20] : ! [X19,X7,X21] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X19)))) => ? [X22] : ((p(s(bool,X22)) <=> s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X19)) & s(X20,'const.pred_set.ITSET_3'(s(fun(X6,fun(X20,X20)),X7),s(fun(X6,bool),X19),s(X20,X21))) = s(X20,'const.bool.COND_3'(s(bool,X22),s(X20,X21),s(X20,'const.pred_set.ITSET_3'(s(fun(X6,fun(X20,X20)),X7),s(fun(X6,bool),'const.pred_set.REST_1'(s(fun(X6,bool),X19))),s(X20,happ(s(fun(X20,X20),happ(s(fun(X6,fun(X20,X20)),X7),s(X6,'const.pred_set.CHOICE_1'(s(fun(X6,bool),X19))))),s(X20,X21)))))))))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_THM|split|0','thm.pred_set.ITSET_THM')).
fof(f6082,plain,(
  ( ! [X6,X8,X7,X5] : (s(X6,'const.pred_set.ITSET_3'(s(fun(X5,fun(X6,X6)),X7),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(X6,X8))) = s(X6,'const.bool.COND_3'(s(bool,sK10(X8,X7,'const.pred_set.EMPTY_0',X6,X5)),s(X6,X8),s(X6,'const.pred_set.ITSET_3'(s(fun(X5,fun(X6,X6)),X7),s(fun(X5,bool),'const.pred_set.REST_1'(s(fun(X5,bool),'const.pred_set.EMPTY_0'))),s(X6,happ(s(fun(X6,X6),happ(s(fun(X5,fun(X6,X6)),X7),s(X5,'const.pred_set.CHOICE_1'(s(fun(X5,bool),'const.pred_set.EMPTY_0'))))),s(X6,X8)))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f6064,f365])).
fof(f6064,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | s(X6,'const.pred_set.ITSET_3'(s(fun(X5,fun(X6,X6)),X7),s(fun(X5,bool),'const.pred_set.EMPTY_0'),s(X6,X8))) = s(X6,'const.bool.COND_3'(s(bool,sK10(X8,X7,'const.pred_set.EMPTY_0',X6,X5)),s(X6,X8),s(X6,'const.pred_set.ITSET_3'(s(fun(X5,fun(X6,X6)),X7),s(fun(X5,bool),'const.pred_set.REST_1'(s(fun(X5,bool),'const.pred_set.EMPTY_0'))),s(X6,happ(s(fun(X6,X6),happ(s(fun(X5,fun(X6,X6)),X7),s(X5,'const.pred_set.CHOICE_1'(s(fun(X5,bool),'const.pred_set.EMPTY_0'))))),s(X6,X8)))))))) )),
  inference(superposition,[],[f303,f391])).
fof(f303,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | s(X1,'const.pred_set.ITSET_3'(s(fun(X0,fun(X1,X1)),X3),s(fun(X0,bool),X2),s(X1,X4))) = s(X1,'const.bool.COND_3'(s(bool,sK10(X4,X3,X2,X1,X0)),s(X1,X4),s(X1,'const.pred_set.ITSET_3'(s(fun(X0,fun(X1,X1)),X3),s(fun(X0,bool),'const.pred_set.REST_1'(s(fun(X0,bool),X2))),s(X1,happ(s(fun(X1,X1),happ(s(fun(X0,fun(X1,X1)),X3),s(X0,'const.pred_set.CHOICE_1'(s(fun(X0,bool),X2))))),s(X1,X4)))))))) )),
  inference(cnf_transformation,[],[f196])).
fof(f9967,plain,(
  $false | $spl170),
  inference(trivial_inequality_removal,[],[f9951])).
fof(f9951,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl170),
  inference(superposition,[],[f222,f9732])).
fof(f9732,plain,(
  ( ! [X61,X62] : (s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X61,'type.extreal.extreal'),X62),s(fun(X61,bool),'const.pred_set.EMPTY_0')))) ) | $spl170),
  inference(superposition,[],[f6081,f272])).
fof(f272,plain,(
  ( ! [X2,X0,X1] : (s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X1),s(fun(X0,bool),X2))) = s('type.extreal.extreal','const.pred_set.ITSET_3'(s(fun(X0,fun('type.extreal.extreal','type.extreal.extreal')),'_dst_x0x1_2'(s(fun('type.extreal.extreal',fun('type.extreal.extreal','type.extreal.extreal')),'const.extreal.extreal_add_0'),s(fun(X0,'type.extreal.extreal'),X1))),s(fun(X0,bool),X2),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X1),s(fun(X0,bool),X2))) = s('type.extreal.extreal','const.pred_set.ITSET_3'(s(fun(X0,fun('type.extreal.extreal','type.extreal.extreal')),'_dst_x0x1_2'(s(fun('type.extreal.extreal',fun('type.extreal.extreal','type.extreal.extreal')),'const.extreal.extreal_add_0'),s(fun(X0,'type.extreal.extreal'),X1))),s(fun(X0,bool),X2),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : ! [X1,X2] : s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X1),s(fun(X0,bool),X2))) = s('type.extreal.extreal','const.pred_set.ITSET_3'(s(fun(X0,fun('type.extreal.extreal','type.extreal.extreal')),'_dst_x0x1_2'(s(fun('type.extreal.extreal',fun('type.extreal.extreal','type.extreal.extreal')),'const.extreal.extreal_add_0'),s(fun(X0,'type.extreal.extreal'),X1))),s(fun(X0,bool),X2),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X7,X19] : s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X6,'type.extreal.extreal'),X7),s(fun(X6,bool),X19))) = s('type.extreal.extreal','const.pred_set.ITSET_3'(s(fun(X6,fun('type.extreal.extreal','type.extreal.extreal')),'_dst_x0x1_2'(s(fun('type.extreal.extreal',fun('type.extreal.extreal','type.extreal.extreal')),'const.extreal.extreal_add_0'),s(fun(X6,'type.extreal.extreal'),X7))),s(fun(X6,bool),X19),s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_THM|split|0','thm.extreal.EXTREAL_SUM_IMAGE_DEF')).
fof(f222,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK7,'type.extreal.extreal'),sK8),s(fun(sK7,bool),'const.pred_set.EMPTY_0')))),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(sK7,'type.extreal.extreal'),sK8),s(fun(sK7,bool),'const.pred_set.EMPTY_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8])],[f122])).
fof(f122,plain,(
  ? [X0,X1] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(ennf_transformation,[],[f51])).
fof(f51,plain,(
  ~! [X0,X1] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ~! [X0] : ! [X1] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X0,'type.extreal.extreal'),X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X6,'type.extreal.extreal'),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s('type.extreal.extreal','const.extreal.extreal_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.extreal.extreal','const.extreal.EXTREAL_SUM_IMAGE_2'(s(fun(X6,'type.extreal.extreal'),X7),s(fun(X6,bool),'const.pred_set.EMPTY_0')))),
  file('thm.extreal.EXTREAL_SUM_IMAGE_THM|split|0','thm.extreal.EXTREAL_SUM_IMAGE_THM|split|0')).
% SZS output end Proof for thm.extreal
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 186
% Passive clauses: 1545
% Generated clauses: 9487
% Final active clauses: 175
% Final passive clauses: 1169
% Input formulas: 48
% Initial clauses: 139
% 
% Duplicate literals: 382
% Trivial inequalities: 90
% Fw subsumption resolutions: 936
% Fw demodulations: 1351
% Bw demodulations: 5
% 
% Simple tautologies: 204
% Equational tautologies: 37
% Forward subsumptions: 4686
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 407
% Factoring: 8
% Forward superposition: 2451
% Backward superposition: 3469
% Self superposition: 117
% Equality factoring: 62
% Equality resolution: 18
% 
% Split clauses: 218
% Split components: 440
% SAT solver clauses: 347
% SAT solver unit clauses: 6
% SAT solver binary clauses: 57
% SAT solver learnt clauses: 1
% 
% Sat splits: 218
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 345
% TWLsolver calls for satisfiability: 61
% 
% Memory used [KB]: 6652
% Time elapsed: 0.254 s
% ------------------------------
% ------------------------------
% Success in time 0.263 s
