Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.rat
% (94663)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 125
% Passive clauses: 1729
% Generated clauses: 24387
% Final active clauses: 107
% Final passive clauses: 1348
% Input formulas: 20
% Initial clauses: 31
% 
% Duplicate literals: 1203
% Trivial inequalities: 379
% Fw subsumption resolutions: 3987
% Fw demodulations: 3059
% Bw demodulations: 9
% 
% Simple tautologies: 411
% Equational tautologies: 77
% Forward subsumptions: 10718
% Fw demodulations to eq. taut.: 3
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 126
% Factoring: 2
% Forward superposition: 7318
% Backward superposition: 7517
% Self superposition: 338
% Equality factoring: 233
% Equality resolution: 33
% 
% Split clauses: 929
% Split components: 1906
% SAT solver clauses: 1843
% SAT solver unit clauses: 8
% SAT solver binary clauses: 85
% 
% Sat splits: 929
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 1838
% TWLsolver calls for satisfiability: 144
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (94677)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 134
% Passive clauses: 444
% Generated clauses: 12724
% Final active clauses: 102
% Final passive clauses: 240
% Input formulas: 20
% Initial clauses: 31
% 
% Duplicate literals: 1242
% Trivial inequalities: 75
% Fw subsumption resolutions: 889
% Fw demodulations: 1425
% Bw demodulations: 84
% Condensations: 1589
% Global subsumptions: 5
% 
% Simple tautologies: 178
% Equational tautologies: 51
% Forward subsumptions: 5034
% Fw demodulations to eq. taut.: 19
% Bw demodulations to eq. taut.: 11
% 
% Binary resolution: 122
% Factoring: 2
% Forward superposition: 3332
% Backward superposition: 3186
% Self superposition: 161
% Equality factoring: 128
% Equality resolution: 13
% 
% Split clauses: 103
% Split components: 216
% SAT solver clauses: 801
% SAT solver unit clauses: 46
% SAT solver binary clauses: 55
% SAT solver learnt clauses: 10
% 
% Sat splits: 103
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 781
% TWLsolver calls for satisfiability: 207
% 
% Memory used [KB]: 7419
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f11778,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f62,f96,f94,f103,f111,f105,f131,f135,f133,f142,f143,f146,f1167,f1201,f1187,f1189,f1321,f1336,f1335,f1582,f1584,f1569,f1570,f1571,f1572,f1592,f1573,f1593,f1588,f2650,f2641,f2636,f2634,f2679,f2622,f2672,f2620,f2678,f2526,f2790,f2969,f2972,f2973,f2974,f4450,f4466,f4444,f4443,f4473,f4442,f4441,f4476,f4440,f4477,f4484,f4382,f4487,f113,f4488,f4491,f4493,f4495,f4498,f4501,f4503,f4505,f4486,f3673,f4510,f3423,f2789,f4514,f109,f4516,f3277,f114,f2545,f4533,f4572,f4604,f3278,f4618,f4619,f2677,f4655,f4661,f5050,f5064,f5051,f5065,f5063,f2524,f5072,f2711,f1587,f1591,f5146,f5165,f5164,f5291,f5300,f5341,f5350,f5354,f5357,f5362,f5458,f5464,f5547,f5518,f5565,f5791,f5814,f5812,f7083,f7149,f1590,f7379,f8068,f8283,f8284,f8285,f8305,f8307,f8310,f8314,f8319,f8323,f8324,f8393,f8397,f8403,f8406,f8767,f4606,f8768,f8769,f4605,f9984,f9997,f9978,f4470,f9977,f8131,f10380,f10409,f10414,f10415,f112,f10418,f10420,f1585,f11057,f11059,f11061,f5550,f11040,f11062,f5549,f11103,f11105,f11120,f11121,f11122,f11123,f11124,f11125,f11126,f11128,f11348,f11387,f11393,f11399,f11405,f11661,f4475,f132,f11762,f11764,f11744,f11768,f11767])).
fof(f11767,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))) | $spl116),
  inference(cnf_transformation,[],[f11767_D])).
fof(f11767_D,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f11768,plain,(
  ( ! [X1] : (s('type.integer.int',X1) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))))) ) | $spl118),
  inference(cnf_transformation,[],[f11768_D])).
fof(f11768_D,plain,(
  ( ! [X1] : (s('type.integer.int',X1) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f11744,plain,(
  ( ! [X1] : (s('type.integer.int',X1) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))) ) | ($spl4 | $spl83)),
  inference(superposition,[],[f404,f11524])).
fof(f11524,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))) | ($spl4 | $spl83)),
  inference(unit_resulting_resolution,[],[f11498,f327])).
fof(f327,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) | s('type.integer.int',X1) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0)))))) )),
  inference(superposition,[],[f89,f72])).
fof(f72,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0)))))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0] : s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0)))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X6))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X6)))))))),
  file('thm.rat.RAT_MINV_LES','thm.rat.RAT_SGN_AINV')).
fof(f89,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : ((s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) <=> s('type.integer.int',X0) = s('type.integer.int',X1))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X11,X12] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X11))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X12))) <=> s('type.integer.int',X11) = s('type.integer.int',X12))),
  file('thm.rat.RAT_MINV_LES','thm.integer.INT_EQ_NEG')).
fof(f11498,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | ($spl4 | $spl83)),
  inference(forward_demodulation,[],[f11497,f9530])).
fof(f9530,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl83),
  inference(unit_resulting_resolution,[],[f5163,f75])).
fof(f75,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))))))),
  inference(ennf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : (s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))))))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0] : (~s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : (~s('type.rat.rat',X6) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X6))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X6))))))),
  file('thm.rat.RAT_MINV_LES','thm.rat.RAT_SGN_MINV')).
fof(f5163,plain,(
  s('type.rat.rat',sK0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl83),
  inference(cnf_transformation,[],[f5163_D])).
fof(f5163_D,plain,(
  s('type.rat.rat',sK0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f11497,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))) | ($spl4 | $spl83)),
  inference(forward_demodulation,[],[f11478,f72])).
fof(f11478,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))))) | ($spl4 | $spl83)),
  inference(unit_resulting_resolution,[],[f5163,f11466,f404])).
fof(f11466,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))) | $spl4),
  inference(forward_demodulation,[],[f11465,f71])).
fof(f71,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X0)))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0] : s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X6))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X6)))),
  file('thm.rat.RAT_MINV_LES','thm.rat.RAT_SUB_LID')).
fof(f11465,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))) | $spl4),
  inference(subsumption_resolution,[],[f11430,f63])).
fof(f63,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rat.RAT_MINV_LES','thm.bool.TRUTH')).
fof(f11430,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))) | $spl4),
  inference(superposition,[],[f884,f109])).
fof(f884,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))) )),
  inference(backward_demodulation,[],[f883,f87])).
fof(f87,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))) & (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) <=> s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6,X13] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat',X13)))) <=> s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X13),s('type.rat.rat',X6))))))),
  file('thm.rat.RAT_MINV_LES','thm.rat.rat_les_def')).
fof(f883,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(forward_demodulation,[],[f882,f70])).
fof(f70,plain,(
  ( ! [X0] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0] : s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : s('type.rat.rat',X6) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.rat.RAT_MINV_LES','thm.rat.RAT_SUB_RID')).
fof(f882,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(subsumption_resolution,[],[f865,f63])).
fof(f865,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f87,f106])).
fof(f106,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))),
  inference(subsumption_resolution,[],[f102,f97])).
fof(f97,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f66,f63])).
fof(f66,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f23])).
fof(f23,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f12])).
fof(f12,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.rat.RAT_MINV_LES','thm.bool.NOT_CLAUSES|split|2')).
fof(f102,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))),
  inference(superposition,[],[f61,f73])).
fof(f73,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.rat.RAT_MINV_LES',aHL_BOOL_CASES)).
fof(f61,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))))),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) & (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f45])).
fof(f45,plain,(
  ? [X0] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X0)))) & (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X0))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0)))))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X0)))) => (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) & s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X0))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X6)))) => (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X6))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) & s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X6))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X6)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X6)))) => (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X6))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) & s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X6))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X6)))))))),
  file('thm.rat.RAT_MINV_LES','thm.rat.RAT_MINV_LES')).
fof(f404,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))))))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f90,f75])).
fof(f90,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f59])).
fof(f11764,plain,(
  $false | ($spl4 | $spl6 | $spl83)),
  inference(subsumption_resolution,[],[f11763,f97])).
fof(f11763,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl6 | $spl83)),
  inference(forward_demodulation,[],[f11756,f110])).
fof(f110,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl6),
  inference(cnf_transformation,[],[f110_D])).
fof(f110_D,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f11756,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl4 | $spl83)),
  inference(trivial_inequality_removal,[],[f11740])).
fof(f11740,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl4 | $spl83)),
  inference(superposition,[],[f982,f11524])).
fof(f982,plain,(
  ( ! [X4] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X4))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f981,f233])).
fof(f233,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f71,f70])).
fof(f981,plain,(
  ( ! [X4] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X4))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(forward_demodulation,[],[f973,f71])).
fof(f973,plain,(
  ( ! [X4] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X4))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(superposition,[],[f885,f233])).
fof(f885,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))) )),
  inference(backward_demodulation,[],[f883,f88])).
fof(f88,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f11762,plain,(
  $false | ($spl4 | $spl6 | $spl83)),
  inference(subsumption_resolution,[],[f11761,f97])).
fof(f11761,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl6 | $spl83)),
  inference(forward_demodulation,[],[f11735,f110])).
fof(f11735,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl4 | $spl83)),
  inference(unit_resulting_resolution,[],[f11524,f982])).
fof(f132,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f132_D])).
fof(f132_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f4475,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,X30)) ) | $spl54),
  inference(cnf_transformation,[],[f4475_D])).
fof(f4475_D,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,X30)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f11661,plain,(
  ( ! [X39,X38] : (s(bool,X39) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X39)) ) | ($spl1 | $spl4)),
  inference(duplicate_literal_removal,[],[f11610])).
fof(f11610,plain,(
  ( ! [X39,X38] : (s(bool,X39) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X38)) ) | ($spl1 | $spl4)),
  inference(superposition,[],[f6453,f11195])).
fof(f11195,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl4)),
  inference(trivial_inequality_removal,[],[f11185])).
fof(f11185,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl4)),
  inference(superposition,[],[f11133,f98])).
fof(f98,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f73,f73])).
fof(f11133,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f94,f109])).
fof(f6453,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f2554,f4620])).
fof(f4620,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'F_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f2995,f63])).
fof(f2995,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | s(bool,X8) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f97,f855])).
fof(f855,plain,(
  ( ! [X45,X44] : (s(bool,X44) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f823,f63])).
fof(f823,plain,(
  ( ! [X45,X44] : (~p(s(bool,'T_0')) | s(bool,X44) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f756,f106])).
fof(f756,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(subsumption_resolution,[],[f563,f63])).
fof(f563,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X9)) | s(bool,X8) = s(bool,X9) | s(bool,X8) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f86,f98])).
fof(f86,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7,X8] : ((p(s(bool,X7)) => p(s(bool,X8))) => ((p(s(bool,X8)) => p(s(bool,X7))) => s(bool,X7) = s(bool,X8)))),
  file('thm.rat.RAT_MINV_LES','thm.bool.IMP_ANTISYM_AX')).
fof(f2554,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f2532])).
fof(f2532,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) )),
  inference(equality_factoring,[],[f855])).
fof(f11405,plain,(
  $false | ($spl83 | $spl88)),
  inference(subsumption_resolution,[],[f11404,f97])).
fof(f11404,plain,(
  p(s(bool,'F_0')) | ($spl83 | $spl88)),
  inference(forward_demodulation,[],[f11403,f221])).
fof(f221,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f97,f214,f83])).
fof(f83,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f50])).
fof(f214,plain,(
  ( ! [X0] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0))))) )),
  inference(forward_demodulation,[],[f210,f70])).
fof(f210,plain,(
  ( ! [X0] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat',X0))))) )),
  inference(unit_resulting_resolution,[],[f70,f82])).
fof(f82,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',X0) != s('type.rat.rat',X1) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))) )),
  inference(cnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) | s('type.rat.rat',X0) != s('type.rat.rat',X1))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) => s('type.rat.rat',X0) != s('type.rat.rat',X1))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) => ~s('type.rat.rat',X0) = s('type.rat.rat',X1))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X13] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat',X13)))) => ~s('type.rat.rat',X6) = s('type.rat.rat',X13))),
  file('thm.rat.RAT_MINV_LES','thm.rat.RAT_LES_IMP_NEQ')).
fof(f11403,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl83 | $spl88)),
  inference(forward_demodulation,[],[f11402,f233])).
fof(f11402,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl83 | $spl88)),
  inference(subsumption_resolution,[],[f2215,f11347])).
fof(f11347,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl83 | $spl88)),
  inference(forward_demodulation,[],[f11346,f9530])).
fof(f11346,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl88),
  inference(forward_demodulation,[],[f11320,f271])).
fof(f271,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f72,f233])).
fof(f11320,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl88),
  inference(superposition,[],[f72,f5549])).
fof(f2215,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(forward_demodulation,[],[f2210,f233])).
fof(f2210,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(superposition,[],[f982,f233])).
fof(f11399,plain,(
  $false | ($spl83 | $spl88)),
  inference(subsumption_resolution,[],[f11398,f97])).
fof(f11398,plain,(
  p(s(bool,'F_0')) | ($spl83 | $spl88)),
  inference(forward_demodulation,[],[f11397,f221])).
fof(f11397,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl83 | $spl88)),
  inference(forward_demodulation,[],[f11396,f233])).
fof(f11396,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl83 | $spl88)),
  inference(subsumption_resolution,[],[f6380,f11347])).
fof(f6380,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(superposition,[],[f986,f233])).
fof(f986,plain,(
  ( ! [X4] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X4))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X4))))) )),
  inference(forward_demodulation,[],[f985,f233])).
fof(f985,plain,(
  ( ! [X4] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X4))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat',X4))))) )),
  inference(forward_demodulation,[],[f976,f70])).
fof(f976,plain,(
  ( ! [X4] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X4),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat',X4))))) )),
  inference(superposition,[],[f885,f233])).
fof(f11393,plain,(
  $false | ($spl83 | $spl88)),
  inference(subsumption_resolution,[],[f11392,f97])).
fof(f11392,plain,(
  p(s(bool,'F_0')) | ($spl83 | $spl88)),
  inference(forward_demodulation,[],[f11391,f221])).
fof(f11391,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl83 | $spl88)),
  inference(forward_demodulation,[],[f11390,f233])).
fof(f11390,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl83 | $spl88)),
  inference(subsumption_resolution,[],[f10764,f11347])).
fof(f10764,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(superposition,[],[f986,f233])).
fof(f11387,plain,(
  $false | ($spl83 | $spl88)),
  inference(subsumption_resolution,[],[f11386,f97])).
fof(f11386,plain,(
  p(s(bool,'F_0')) | ($spl83 | $spl88)),
  inference(forward_demodulation,[],[f11385,f221])).
fof(f11385,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl83 | $spl88)),
  inference(forward_demodulation,[],[f11384,f233])).
fof(f11384,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl83 | $spl88)),
  inference(subsumption_resolution,[],[f10966,f11347])).
fof(f10966,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(forward_demodulation,[],[f10958,f233])).
fof(f10958,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(superposition,[],[f982,f233])).
fof(f11348,plain,(
  $false | ($spl19 | $spl83 | $spl88)),
  inference(subsumption_resolution,[],[f11347,f1189])).
fof(f11128,plain,(
  $false | ($spl1 | $spl2 | $spl4 | $spl7)),
  inference(global_subsumption,[],[f11127,f62,f61,f63,f72,f73,f86,f83,f90,f89,f91,f98,f100,f103,f105,f131,f134,f135,f88,f87,f155,f153,f160,f161,f214,f222,f283,f284,f326,f327,f396,f479,f756,f758,f759,f856,f883,f888,f902,f910,f914,f958,f962,f965,f968,f987,f1007,f1000,f1024,f1030,f1037,f1050,f1079,f1080,f1107,f1111,f1112,f1128,f1118,f1129,f1145,f1138,f1167,f1201,f1208,f1277,f1306,f1316,f1321,f1337,f1323,f1338,f380,f1582,f1569,f1570,f1571,f1572,f1573,f1593,f2650,f2641,f2636,f2634,f2622,f2620,f2526,f3239,f3277,f2545,f1576,f3278,f4780,f4789,f4781,f4790,f2524,f2711,f5467,f5510,f5520,f5522,f5519,f5518,f5517,f6401,f6402,f6835,f7083,f7149,f6453,f94,f111,f7361,f6960,f6591,f6522,f6506,f1237,f1227,f757,f221,f97,f955,f6580,f6324,f6323,f4620,f2554,f2549,f2523,f2472,f2312,f2311,f1535,f1532,f1530,f1525,f1517,f1296,f1293,f1285,f1194,f1193,f1110,f1099,f1092,f1083,f1039,f1036,f1027,f855,f781,f777,f406,f404,f401,f271,f233,f106,f75,f71,f70,f8721,f1236,f1195,f986,f991,f1196,f982,f6521,f1200,f1198,f1191,f1190,f890,f885,f884,f82,f7195,f7084,f6962,f6525,f2525,f2474,f1581,f766,f765,f764,f9978,f109])).
fof(f764,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK1(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK1(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f91,f73])).
fof(f765,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK1(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK1(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f91,f98])).
fof(f766,plain,(
  ( ! [X10,X8,X7,X9] : (s(bool,X10) != s(bool,f(s(X7,sK1(X8,X9,bool,X7)))) | s(fun(X7,bool),X8) = s(fun(X7,bool),X9) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,g(s(X7,sK1(X8,X9,bool,X7))))) )),
  inference(superposition,[],[f91,f98])).
fof(f1581,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK1(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,f(s(X0,sK1(X1,X2,bool,X0))))) )),
  inference(trivial_inequality_removal,[],[f1575])).
fof(f1575,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK1(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,f(s(X0,sK1(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f764,f73])).
fof(f2474,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,X111) != s(bool,f(s(X108,sK1(X109,X110,bool,X108)))) | s(fun(X108,bool),X109) = s(fun(X108,bool),X110) | s(bool,X111) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,g(s(X108,sK1(X109,X110,bool,X108))))) )),
  inference(superposition,[],[f91,f855])).
fof(f2525,plain,(
  ( ! [X94,X97,X95,X96] : (s(bool,'T_0') != s(bool,f(s(X94,sK1(X95,X96,bool,X94)))) | s(fun(X94,bool),X95) = s(fun(X94,bool),X96) | s(bool,X97) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X97) = s(bool,g(s(X94,sK1(X95,X96,bool,X94))))) )),
  inference(superposition,[],[f91,f855])).
fof(f6525,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') != s(bool,f(s(X64,sK1(X65,X66,bool,X64)))) | s(fun(X64,bool),X65) = s(fun(X64,bool),X66) | s(bool,'F_0') = s(bool,g(s(X64,sK1(X65,X66,bool,X64))))) )),
  inference(forward_demodulation,[],[f6452,f106])).
fof(f6452,plain,(
  ( ! [X66,X64,X65] : (s(bool,f(s(X64,sK1(X65,X66,bool,X64)))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(fun(X64,bool),X65) = s(fun(X64,bool),X66) | s(bool,'F_0') = s(bool,g(s(X64,sK1(X65,X66,bool,X64))))) )),
  inference(superposition,[],[f91,f4620])).
fof(f6962,plain,(
  ( ! [X146,X147,X145,X148] : (s(bool,X148) != s(bool,f(s(X145,sK1(X146,X147,bool,X145)))) | s(fun(X145,bool),X146) = s(fun(X145,bool),X147) | s(bool,'F_0') = s(bool,X148) | s(bool,'F_0') = s(bool,g(s(X145,sK1(X146,X147,bool,X145))))) )),
  inference(superposition,[],[f91,f6401])).
fof(f7084,plain,(
  ( ! [X123,X121,X122,X120] : (s(bool,'F_0') != s(bool,f(s(X120,sK1(X121,X122,bool,X120)))) | s(fun(X120,bool),X121) = s(fun(X120,bool),X122) | s(bool,X123) = s(bool,g(s(X120,sK1(X121,X122,bool,X120)))) | s(bool,'F_0') = s(bool,X123)) )),
  inference(superposition,[],[f91,f6401])).
fof(f7195,plain,(
  ( ! [X80,X78,X79,X77] : (s(fun(X77,bool),X78) = s(fun(X77,bool),X79) | s(bool,'T_0') = s(bool,g(s(X77,sK1(X78,X79,bool,X77)))) | s(bool,X80) = s(bool,f(s(X77,sK1(X78,X79,bool,X77)))) | s(bool,'F_0') = s(bool,X80)) )),
  inference(trivial_inequality_removal,[],[f7068])).
fof(f7068,plain,(
  ( ! [X80,X78,X79,X77] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X77,bool),X78) = s(fun(X77,bool),X79) | s(bool,'T_0') = s(bool,g(s(X77,sK1(X78,X79,bool,X77)))) | s(bool,X80) = s(bool,f(s(X77,sK1(X78,X79,bool,X77)))) | s(bool,'F_0') = s(bool,X80)) )),
  inference(superposition,[],[f764,f6401])).
fof(f890,plain,(
  ( ! [X6,X8,X7] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X7),s('type.rat.rat',X6))))) | s(bool,X8) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(forward_demodulation,[],[f889,f883])).
fof(f889,plain,(
  ( ! [X6,X8,X7] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X7),s('type.rat.rat',X6))))) | s(bool,X8) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f870,f63])).
fof(f870,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X7),s('type.rat.rat',X6))))) | s(bool,X8) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f87,f98])).
fof(f1190,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat',X7)))) | s('type.rat.rat',X7) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X6))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X6)))))) )),
  inference(forward_demodulation,[],[f1173,f70])).
fof(f1173,plain,(
  ( ! [X6,X7] : (s('type.rat.rat',X7) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat',X7)))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X6))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X6)))))) )),
  inference(superposition,[],[f82,f401])).
fof(f1191,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X9),s('type.rat.rat',X8)))) | s('type.rat.rat',X9) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X8))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X8)))))) )),
  inference(forward_demodulation,[],[f1174,f70])).
fof(f1174,plain,(
  ( ! [X8,X9] : (s('type.rat.rat',X9) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X9),s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X8),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X8))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X8)))))) )),
  inference(superposition,[],[f82,f401])).
fof(f1198,plain,(
  ( ! [X23,X22] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X22),s('type.rat.rat',X23)))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X23))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X22))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X22)))))) )),
  inference(forward_demodulation,[],[f1197,f70])).
fof(f1197,plain,(
  ( ! [X23,X22] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X23))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat',X23)))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X22))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X22)))))) )),
  inference(forward_demodulation,[],[f1182,f70])).
fof(f1182,plain,(
  ( ! [X23,X22] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat',X23)))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X22))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X22)))))) )),
  inference(superposition,[],[f885,f401])).
fof(f1200,plain,(
  ( ! [X24,X25] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X25),s('type.rat.rat',X24)))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X25))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X24))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X24)))))) )),
  inference(forward_demodulation,[],[f1199,f70])).
fof(f1199,plain,(
  ( ! [X24,X25] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X25))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X25),s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X24))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X24)))))) )),
  inference(forward_demodulation,[],[f1183,f71])).
fof(f1183,plain,(
  ( ! [X24,X25] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X25))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X25),s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X24))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X24)))))) )),
  inference(superposition,[],[f885,f401])).
fof(f6521,plain,(
  ( ! [X48,X49] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X49),s('type.rat.rat',X48))))) | s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49)))) )),
  inference(subsumption_resolution,[],[f6520,f63])).
fof(f6520,plain,(
  ( ! [X48,X49] : (~p(s(bool,'T_0')) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X49),s('type.rat.rat',X48))))) | s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49)))) )),
  inference(forward_demodulation,[],[f6442,f106])).
fof(f6442,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X49),s('type.rat.rat',X48))))) | s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49)))) )),
  inference(superposition,[],[f884,f4620])).
fof(f1196,plain,(
  ( ! [X21,X20] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X20))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X20)))))) )),
  inference(forward_demodulation,[],[f1181,f70])).
fof(f1181,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat',X21))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X20))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X20)))))) )),
  inference(superposition,[],[f884,f401])).
fof(f991,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0)))))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f986,f75])).
fof(f1195,plain,(
  ( ! [X19,X18] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X19),s('type.rat.rat',X18))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X19)))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X18))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X18)))))) )),
  inference(forward_demodulation,[],[f1180,f70])).
fof(f1180,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X19)))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X18))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X18)))))) )),
  inference(superposition,[],[f884,f401])).
fof(f1236,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,X14)) | s(bool,X14) = s(bool,sK1(X11,X12,X13,bool)) | s(fun(bool,X13),X11) = s(fun(bool,X13),X12) | s(X13,f(s(bool,'F_0'))) != s(X13,g(s(bool,'F_0')))) )),
  inference(subsumption_resolution,[],[f1217,f63])).
fof(f1217,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | s(bool,X14) = s(bool,sK1(X11,X12,X13,bool)) | s(fun(bool,X13),X11) = s(fun(bool,X13),X12) | s(X13,f(s(bool,'F_0'))) != s(X13,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f86,f757])).
fof(f8721,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,X8)) ) | $spl7),
  inference(trivial_inequality_removal,[],[f8720])).
fof(f8720,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X8) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,X8)) ) | $spl7),
  inference(superposition,[],[f111,f6401])).
fof(f401,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0)))))) )),
  inference(superposition,[],[f75,f70])).
fof(f406,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X4))))) | s('type.integer.int',X5) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X4))))) | s('type.rat.rat',X4) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f89,f75])).
fof(f777,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f89,f271])).
fof(f781,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f89,f271])).
fof(f1027,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X5))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X4))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X4))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X5))))) | s('type.rat.rat',X4) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f327,f75])).
fof(f1036,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) | s('type.integer.int',X1) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))))))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f327,f75])).
fof(f1039,plain,(
  ( ! [X2] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X2))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X2))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f1038,f271])).
fof(f1038,plain,(
  ( ! [X2] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X2))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X2))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(forward_demodulation,[],[f1025,f271])).
fof(f1025,plain,(
  ( ! [X2] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X2))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X2))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f327,f271])).
fof(f1083,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X4))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X5))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X4))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X5))))))) | s('type.rat.rat',X4) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f326,f75])).
fof(f1092,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))))))))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f326,f75])).
fof(f1099,plain,(
  ( ! [X2] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X2))))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X2)))))) )),
  inference(forward_demodulation,[],[f1098,f271])).
fof(f1098,plain,(
  ( ! [X2] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X2))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X2)))))))) )),
  inference(forward_demodulation,[],[f1081,f271])).
fof(f1081,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X2))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X2)))))))) )),
  inference(superposition,[],[f326,f271])).
fof(f1110,plain,(
  ( ! [X2] : (s('type.integer.int',X2) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X2))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f1095,f233])).
fof(f1095,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X2))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) )),
  inference(superposition,[],[f326,f271])).
fof(f1193,plain,(
  ( ! [X14,X13] : (s('type.integer.int',X14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X13))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X14))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X13))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X13)))))) )),
  inference(forward_demodulation,[],[f1192,f70])).
fof(f1192,plain,(
  ( ! [X14,X13] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X14))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X13),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X13))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X13)))))) )),
  inference(forward_demodulation,[],[f1177,f271])).
fof(f1177,plain,(
  ( ! [X14,X13] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X14))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int',X14) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X13),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X13))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X13)))))) )),
  inference(superposition,[],[f326,f401])).
fof(f1194,plain,(
  ( ! [X15,X16] : (s('type.integer.int',X16) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X15))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X16))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X15))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X15)))))) )),
  inference(forward_demodulation,[],[f1178,f70])).
fof(f1178,plain,(
  ( ! [X15,X16] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X16))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X16) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X15))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X15)))))) )),
  inference(superposition,[],[f327,f401])).
fof(f1285,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X4))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X5))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X4))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X5))))))) | s('type.rat.rat',X5) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.rat.rat',X4) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f404,f75])).
fof(f1293,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f404,f75])).
fof(f1296,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))))))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(equality_resolution,[],[f404])).
fof(f1517,plain,(
  ( ! [X4,X3] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X3))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X4))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X3))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X4))))) | s('type.rat.rat',X4) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.rat.rat',X3) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f406,f75])).
fof(f1525,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))))) | s('type.integer.int',X1) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f406,f75])).
fof(f1530,plain,(
  ( ! [X2] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X2))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X2))))) | s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1529,f271])).
fof(f1529,plain,(
  ( ! [X2] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X2))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X2))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1516,f271])).
fof(f1516,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X2))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X2))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f406,f271])).
fof(f1532,plain,(
  ( ! [X2,X3] : (s('type.integer.int',X3) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X2))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X2))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f1522,f70])).
fof(f1522,plain,(
  ( ! [X2,X3] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X2))))) | s('type.integer.int',X3) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f406,f70])).
fof(f1535,plain,(
  ( ! [X4,X5] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X4))) | s('type.integer.int',X5) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X4))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X4)))) )),
  inference(forward_demodulation,[],[f1534,f71])).
fof(f1534,plain,(
  ( ! [X4,X5] : (s('type.integer.int',X5) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X4))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X4))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X4)))) )),
  inference(forward_demodulation,[],[f1533,f71])).
fof(f1533,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X4))) | s('type.integer.int',X5) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X4))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X4)))) )),
  inference(forward_demodulation,[],[f1523,f72])).
fof(f1523,plain,(
  ( ! [X4,X5] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X5))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X4))))))) | s('type.integer.int',X5) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X4))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X4)))) )),
  inference(superposition,[],[f406,f71])).
fof(f2311,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f855,f855])).
fof(f2312,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f855,f98])).
fof(f2472,plain,(
  ( ! [X103,X101,X102,X104] : (s(X103,f(s(bool,X104))) != s(X103,g(s(bool,X104))) | s(fun(bool,X103),X101) = s(fun(bool,X103),X102) | s(bool,X104) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,sK1(X101,X102,X103,bool))) )),
  inference(superposition,[],[f91,f855])).
fof(f2523,plain,(
  ( ! [X90,X88,X87,X89] : (s(X89,f(s(bool,'T_0'))) != s(X89,g(s(bool,'T_0'))) | s(fun(bool,X89),X87) = s(fun(bool,X89),X88) | s(bool,X90) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X90) = s(bool,sK1(X87,X88,X89,bool))) )),
  inference(superposition,[],[f91,f855])).
fof(f2549,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) )),
  inference(condensation,[],[f2539])).
fof(f2539,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X0) = s(bool,X1)) )),
  inference(equality_factoring,[],[f855])).
fof(f6323,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X2) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X0) = s(bool,X2)) )),
  inference(superposition,[],[f2549,f855])).
fof(f6324,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,X4) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X3) = s(bool,X5) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2549,f98])).
fof(f6580,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) = s(bool,X27) | s(bool,'F_0') = s(bool,X27) | s(bool,X28) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X26) = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f6553,f63])).
fof(f6553,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | s(bool,X26) = s(bool,X27) | s(bool,'F_0') = s(bool,X27) | s(bool,X28) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X26) = s(bool,X28)) )),
  inference(superposition,[],[f6506,f855])).
fof(f955,plain,(
  ( ! [X3] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f89,f883])).
fof(f757,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,sK1(X0,X1,X2,bool)) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f91,f73])).
fof(f1227,plain,(
  ( ! [X52,X53,X51] : (s(X53,f(s(bool,'T_0'))) != s(X53,g(s(bool,'T_0'))) | s(fun(bool,X53),X51) = s(fun(bool,X53),X52) | s(X53,f(s(bool,'F_0'))) != s(X53,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f1226])).
fof(f1226,plain,(
  ( ! [X52,X53,X51] : (s(X53,f(s(bool,'T_0'))) != s(X53,g(s(bool,'T_0'))) | s(fun(bool,X53),X51) = s(fun(bool,X53),X52) | s(fun(bool,X53),X51) = s(fun(bool,X53),X52) | s(X53,f(s(bool,'F_0'))) != s(X53,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f91,f757])).
fof(f1237,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,X49) = s(bool,sK1(X46,X47,X48,bool)) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(fun(bool,X48),X46) = s(fun(bool,X48),X47) | s(X48,f(s(bool,'F_0'))) != s(X48,g(s(bool,'F_0')))) )),
  inference(subsumption_resolution,[],[f1225,f63])).
fof(f1225,plain,(
  ( ! [X47,X50,X48,X46,X49] : (~p(s(bool,'T_0')) | s(bool,X49) = s(bool,sK1(X46,X47,X48,bool)) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(fun(bool,X48),X46) = s(fun(bool,X48),X47) | s(X48,f(s(bool,'F_0'))) != s(X48,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f756,f757])).
fof(f6506,plain,(
  ( ! [X6,X5] : (~p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6505,f63])).
fof(f6505,plain,(
  ( ! [X6,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X5)) )),
  inference(forward_demodulation,[],[f6419,f106])).
fof(f6419,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) | ~p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X5)) )),
  inference(superposition,[],[f86,f4620])).
fof(f6522,plain,(
  ( ! [X61,X59,X60] : (s(X61,f(s(bool,'T_0'))) != s(X61,g(s(bool,'T_0'))) | s(fun(bool,X61),X59) = s(fun(bool,X61),X60) | s(bool,'F_0') = s(bool,sK1(X59,X60,X61,bool))) )),
  inference(forward_demodulation,[],[f6450,f106])).
fof(f6450,plain,(
  ( ! [X61,X59,X60] : (s(X61,f(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))))) != s(X61,g(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))))) | s(fun(bool,X61),X59) = s(fun(bool,X61),X60) | s(bool,'F_0') = s(bool,sK1(X59,X60,X61,bool))) )),
  inference(superposition,[],[f91,f4620])).
fof(f6591,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,X62) = s(bool,sK1(X59,X60,X61,bool)) | s(bool,'F_0') = s(bool,X62) | s(fun(bool,X61),X59) = s(fun(bool,X61),X60) | s(X61,f(s(bool,'F_0'))) != s(X61,g(s(bool,'F_0')))) )),
  inference(subsumption_resolution,[],[f6569,f63])).
fof(f6569,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | s(bool,X62) = s(bool,sK1(X59,X60,X61,bool)) | s(bool,'F_0') = s(bool,X62) | s(fun(bool,X61),X59) = s(fun(bool,X61),X60) | s(X61,f(s(bool,'F_0'))) != s(X61,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f6506,f757])).
fof(f6960,plain,(
  ( ! [X140,X138,X141,X139] : (s(X140,f(s(bool,X141))) != s(X140,g(s(bool,X141))) | s(fun(bool,X140),X138) = s(fun(bool,X140),X139) | s(bool,'F_0') = s(bool,X141) | s(bool,'F_0') = s(bool,sK1(X138,X139,X140,bool))) )),
  inference(superposition,[],[f91,f6401])).
fof(f7361,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X2)) | s(bool,X1) = s(bool,X2) | s(bool,X1) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f7040,f97])).
fof(f7040,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,X2)) | s(bool,X1) = s(bool,X2) | s(bool,X1) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f83,f6401])).
fof(f6835,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'F_0') = s(bool,X3)) )),
  inference(superposition,[],[f6401,f6401])).
fof(f6402,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f4620,f855])).
fof(f6401,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f4620,f4620])).
fof(f5517,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl4),
  inference(superposition,[],[f72,f5510])).
fof(f5519,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))))))))) ) | $spl4),
  inference(superposition,[],[f326,f5510])).
fof(f5522,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X4) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))))) ) | $spl4),
  inference(superposition,[],[f89,f5510])).
fof(f5520,plain,(
  ( ! [X2] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))))))) ) | $spl4),
  inference(superposition,[],[f327,f5510])).
fof(f5510,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))) | $spl4),
  inference(forward_demodulation,[],[f5509,f71])).
fof(f5509,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))) | $spl4),
  inference(subsumption_resolution,[],[f5486,f63])).
fof(f5486,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))) | $spl4),
  inference(superposition,[],[f884,f109])).
fof(f5467,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl4),
  inference(subsumption_resolution,[],[f5442,f63])).
fof(f5442,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | s(bool,X7) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl4),
  inference(superposition,[],[f756,f109])).
fof(f4790,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl1 | $spl7)),
  inference(global_subsumption,[],[f62,f61,f63,f70,f71,f72,f73,f75,f86,f83,f90,f89,f91,f98,f100,f106,f103,f105,f131,f134,f135,f88,f87,f155,f153,f160,f161,f214,f222,f233,f271,f283,f284,f326,f327,f396,f401,f404,f406,f479,f756,f758,f759,f777,f781,f855,f856,f883,f888,f902,f910,f914,f958,f962,f965,f968,f987,f1007,f1000,f1024,f1039,f1027,f1030,f1036,f1037,f1050,f1079,f1080,f1099,f1083,f1107,f1092,f1110,f1111,f1112,f1128,f1118,f1129,f1145,f1138,f1167,f1193,f1194,f1201,f1208,f1277,f1285,f1306,f1293,f1296,f1316,f1321,f1337,f1323,f1338,f380,f955,f1530,f1517,f1532,f1535,f1525,f1582,f1569,f1570,f1571,f1572,f1573,f1593,f2650,f2641,f2636,f2634,f2311,f2312,f2622,f2620,f2472,f2523,f2526,f2554,f1237,f1227,f757,f221,f97,f3239,f3277,f2545,f2549,f1576,f3278,f4620,f1236,f991,f986,f1195,f1200,f1198,f1191,f1190,f890,f885,f884,f82,f982,f1196,f2525,f2474,f1581,f766,f765,f764,f111,f94,f4780,f4789,f4781])).
fof(f4781,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl1),
  inference(superposition,[],[f94,f98])).
fof(f4789,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl7)),
  inference(global_subsumption,[],[f62,f61,f63,f70,f71,f72,f73,f75,f86,f83,f90,f89,f91,f98,f100,f106,f103,f105,f131,f134,f135,f88,f87,f155,f153,f160,f161,f214,f222,f233,f271,f283,f284,f326,f327,f396,f401,f404,f406,f479,f756,f758,f759,f777,f781,f855,f856,f883,f888,f902,f910,f914,f958,f962,f965,f968,f987,f1007,f1000,f1024,f1039,f1027,f1030,f1036,f1037,f1050,f1079,f1080,f1099,f1083,f1107,f1092,f1110,f1111,f1112,f1128,f1118,f1129,f1145,f1138,f1167,f1193,f1194,f1201,f1208,f1277,f1285,f1306,f1293,f1296,f1316,f1321,f1337,f1323,f1338,f380,f955,f1530,f1517,f1532,f1535,f1525,f1582,f1569,f1570,f1571,f1572,f1573,f1593,f2650,f2641,f2636,f2634,f2311,f2312,f2622,f2620,f2472,f2523,f2526,f2554,f1237,f1227,f757,f221,f97,f3239,f3277,f2545,f2549,f1576,f3278,f4620,f1236,f991,f986,f1195,f1200,f1198,f1191,f1190,f890,f885,f884,f82,f982,f1196,f2525,f2474,f1581,f766,f765,f764,f111,f94,f4780])).
fof(f4780,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl1),
  inference(superposition,[],[f94,f98])).
fof(f1576,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,'F_0') | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,'T_0') = s(bool,g(s(X3,sK1(X4,X5,bool,X3)))) | s(bool,X6) = s(bool,f(s(X3,sK1(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f764,f98])).
fof(f3239,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2554,f855])).
fof(f380,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X1))) | s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X1)))) )),
  inference(equality_factoring,[],[f283])).
fof(f1338,plain,(
  ( ! [X3] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))) | s('type.integer.int',X3) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))))) ) | $spl7),
  inference(forward_demodulation,[],[f1325,f902])).
fof(f1325,plain,(
  ( ! [X3] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X3) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))))) ) | $spl7),
  inference(superposition,[],[f89,f1208])).
fof(f1323,plain,(
  ( ! [X2] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))))))) ) | $spl7),
  inference(superposition,[],[f327,f1208])).
fof(f1337,plain,(
  ( ! [X1] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))))))))) ) | $spl7),
  inference(forward_demodulation,[],[f1322,f902])).
fof(f1322,plain,(
  ( ! [X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))))))))) ) | $spl7),
  inference(superposition,[],[f326,f1208])).
fof(f1316,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f1208,f986])).
fof(f1306,plain,(
  ( ! [X7] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X7))))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X7))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl7),
  inference(forward_demodulation,[],[f1287,f72])).
fof(f1287,plain,(
  ( ! [X7] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X7))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X7))))))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl7),
  inference(superposition,[],[f404,f888])).
fof(f1277,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl7),
  inference(subsumption_resolution,[],[f1264,f63])).
fof(f1264,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | s(bool,X7) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl7),
  inference(superposition,[],[f756,f1145])).
fof(f1208,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))))) | $spl7),
  inference(forward_demodulation,[],[f1207,f1050])).
fof(f1207,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f1138,f75])).
fof(f1138,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))) | $spl7),
  inference(unit_resulting_resolution,[],[f1112,f82])).
fof(f1145,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f63,f1112,f86])).
fof(f1129,plain,(
  ( ! [X2] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))) ) | $spl7),
  inference(forward_demodulation,[],[f1120,f902])).
fof(f1120,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X2) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))) ) | $spl7),
  inference(superposition,[],[f89,f1050])).
fof(f1118,plain,(
  ( ! [X1] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X1) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))))) ) | $spl7),
  inference(superposition,[],[f327,f1050])).
fof(f1128,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))))))) ) | $spl7),
  inference(forward_demodulation,[],[f1117,f902])).
fof(f1117,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))))))) ) | $spl7),
  inference(superposition,[],[f326,f1050])).
fof(f1112,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f1050,f986])).
fof(f1111,plain,(
  ( ! [X3] : (s('type.integer.int',X3) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3)))) ) | $spl7),
  inference(forward_demodulation,[],[f1096,f888])).
fof(f1096,plain,(
  ( ! [X3] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))) ) | $spl7),
  inference(superposition,[],[f326,f902])).
fof(f1107,plain,(
  ( ! [X2,X3] : (s('type.integer.int',X3) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X2))))))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X2)))) )),
  inference(forward_demodulation,[],[f1106,f71])).
fof(f1106,plain,(
  ( ! [X2,X3] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X2))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X2)))))))))) )),
  inference(forward_demodulation,[],[f1090,f72])).
fof(f1090,plain,(
  ( ! [X2,X3] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X3))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X2))))))) | s('type.integer.int',X3) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X2)))))))))) )),
  inference(superposition,[],[f326,f71])).
fof(f1080,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X1))))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0))))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X1)))))))) )),
  inference(superposition,[],[f326,f72])).
fof(f1079,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl7),
  inference(subsumption_resolution,[],[f1066,f63])).
fof(f1066,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | s(bool,X7) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl7),
  inference(superposition,[],[f756,f1007])).
fof(f1050,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))) | $spl7),
  inference(forward_demodulation,[],[f1049,f888])).
fof(f1049,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f1000,f75])).
fof(f1037,plain,(
  ( ! [X2] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))) ) | $spl7),
  inference(superposition,[],[f327,f888])).
fof(f1030,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X1))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X1)))))) )),
  inference(superposition,[],[f327,f72])).
fof(f1024,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X1))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0)))))))) )),
  inference(superposition,[],[f327,f72])).
fof(f1000,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))) | $spl7),
  inference(unit_resulting_resolution,[],[f987,f82])).
fof(f1007,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))) | $spl7),
  inference(unit_resulting_resolution,[],[f63,f987,f86])).
fof(f987,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))) | $spl7),
  inference(unit_resulting_resolution,[],[f888,f986])).
fof(f968,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0)))))) ) | $spl7),
  inference(superposition,[],[f910,f72])).
fof(f965,plain,(
  ( ! [X1] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X1))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X1))))) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl7),
  inference(superposition,[],[f910,f75])).
fof(f962,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0)))))))) ) | $spl7),
  inference(superposition,[],[f910,f72])).
fof(f958,plain,(
  ( ! [X2] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl7),
  inference(forward_demodulation,[],[f954,f902])).
fof(f954,plain,(
  ( ! [X2] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f89,f883])).
fof(f914,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X4))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X4) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0)))))) ) | $spl7),
  inference(superposition,[],[f89,f902])).
fof(f910,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0)))) ) | $spl7),
  inference(superposition,[],[f89,f902])).
fof(f902,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | $spl7),
  inference(forward_demodulation,[],[f896,f72])).
fof(f896,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f888,f90])).
fof(f888,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))) | $spl7),
  inference(forward_demodulation,[],[f887,f883])).
fof(f887,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl7),
  inference(forward_demodulation,[],[f886,f71])).
fof(f886,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))))) | $spl7),
  inference(subsumption_resolution,[],[f868,f63])).
fof(f868,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))))) | $spl7),
  inference(superposition,[],[f87,f153])).
fof(f856,plain,(
  ( ! [X52,X51] : (s(bool,X51) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) ) | $spl7),
  inference(subsumption_resolution,[],[f826,f63])).
fof(f826,plain,(
  ( ! [X52,X51] : (~p(s(bool,'T_0')) | s(bool,X51) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) ) | $spl7),
  inference(superposition,[],[f756,f153])).
fof(f759,plain,(
  ( ! [X10,X8,X7,X9] : (s(X9,f(s(bool,X10))) != s(X9,g(s(bool,X10))) | s(fun(bool,X9),X7) = s(fun(bool,X9),X8) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,sK1(X7,X8,X9,bool))) )),
  inference(superposition,[],[f91,f98])).
fof(f758,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK1(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f91,f98])).
fof(f479,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f98,f98])).
fof(f396,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X32),s('type.rat.rat',X32)))) ) | ($spl1 | $spl7)),
  inference(subsumption_resolution,[],[f369,f160])).
fof(f369,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X32),s('type.rat.rat',X32))) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl1 | $spl7)),
  inference(superposition,[],[f160,f283])).
fof(f326,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X0))))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0)))))))) )),
  inference(superposition,[],[f89,f72])).
fof(f284,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X7)) ) | ($spl1 | $spl7)),
  inference(subsumption_resolution,[],[f281,f97])).
fof(f281,plain,(
  ( ! [X7] : (p(s(bool,'F_0')) | s(bool,X7) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X7)) ) | ($spl1 | $spl7)),
  inference(superposition,[],[f134,f161])).
fof(f283,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat',X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f279,f97])).
fof(f279,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat',X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f134,f221])).
fof(f222,plain,(
  ( ! [X0,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X1)))) )),
  inference(unit_resulting_resolution,[],[f214,f214,f83])).
fof(f161,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl7)),
  inference(unit_resulting_resolution,[],[f160,f73])).
fof(f160,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl7)),
  inference(backward_demodulation,[],[f153,f94])).
fof(f153,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl7),
  inference(unit_resulting_resolution,[],[f111,f73])).
fof(f155,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl7),
  inference(unit_resulting_resolution,[],[f97,f111,f83])).
fof(f134,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f126,f97])).
fof(f126,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f83,f73])).
fof(f100,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f97,f73])).
fof(f91,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK1(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK1(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK1(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK1(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.rat.RAT_MINV_LES',aHL_EXT)).
fof(f11127,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl2),
  inference(forward_demodulation,[],[f95,f106])).
fof(f95,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl2),
  inference(cnf_transformation,[],[f95_D])).
fof(f95_D,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f11126,plain,(
  $false | ($spl1 | $spl4 | $spl7)),
  inference(global_subsumption,[],[f62,f61,f63,f72,f73,f86,f83,f90,f89,f91,f98,f100,f103,f105,f131,f134,f135,f88,f87,f155,f153,f160,f161,f214,f222,f283,f284,f326,f327,f396,f479,f756,f758,f759,f856,f883,f888,f902,f910,f914,f958,f962,f965,f968,f987,f1007,f1000,f1024,f1030,f1037,f1050,f1079,f1080,f1107,f1111,f1112,f1128,f1118,f1129,f1145,f1138,f1167,f1201,f1208,f1277,f1306,f1316,f1321,f1337,f1323,f1338,f380,f1582,f1569,f1570,f1571,f1572,f1573,f1593,f2650,f2641,f2636,f2634,f2622,f2620,f2526,f3239,f3277,f2545,f1576,f3278,f4780,f4789,f4781,f4790,f2524,f2711,f5467,f5510,f5520,f5522,f5519,f5518,f5517,f6401,f6402,f6835,f7083,f7149,f6453,f94,f111,f7361,f6960,f6591,f6522,f6506,f1237,f1227,f757,f221,f97,f955,f6580,f6324,f6323,f4620,f2554,f2549,f2523,f2472,f2312,f2311,f1535,f1532,f1530,f1525,f1517,f1296,f1293,f1285,f1194,f1193,f1110,f1099,f1092,f1083,f1039,f1036,f1027,f855,f781,f777,f406,f404,f401,f271,f233,f106,f75,f71,f70,f8721,f1236,f1195,f986,f991,f1196,f982,f6521,f1200,f1198,f1191,f1190,f890,f885,f884,f82,f7195,f7084,f6962,f6525,f2525,f2474,f1581,f766,f765,f764,f9978,f109])).
fof(f11125,plain,(
  $false | ($spl1 | $spl4 | $spl7 | $spl61)),
  inference(global_subsumption,[],[f4485,f62,f61,f63,f72,f73,f86,f83,f90,f89,f91,f98,f100,f103,f105,f131,f134,f135,f88,f87,f155,f153,f160,f161,f214,f222,f283,f284,f326,f327,f396,f479,f756,f758,f759,f856,f883,f888,f902,f910,f914,f958,f962,f965,f968,f987,f1007,f1000,f1024,f1030,f1037,f1050,f1079,f1080,f1107,f1111,f1112,f1128,f1118,f1129,f1145,f1138,f1167,f1201,f1208,f1277,f1306,f1316,f1321,f1337,f1323,f1338,f380,f1582,f1569,f1570,f1571,f1572,f1573,f1593,f2650,f2641,f2636,f2634,f2622,f2620,f2526,f3239,f3277,f2545,f1576,f3278,f4780,f4789,f4781,f4790,f2524,f2711,f109,f5467,f5510,f5520,f5522,f5519,f5518,f5517,f6401,f6402,f6835,f7083,f7149,f6453,f94,f111,f7361,f6960,f6591,f6522,f6506,f1237,f1227,f757,f221,f97,f955,f6580,f6324,f6323,f4620,f2554,f2549,f2523,f2472,f2312,f2311,f1535,f1532,f1530,f1525,f1517,f1296,f1293,f1285,f1194,f1193,f1110,f1099,f1092,f1083,f1039,f1036,f1027,f855,f781,f777,f406,f404,f401,f271,f233,f106,f75,f71,f70,f8721,f1236,f1195,f986,f991,f1196,f982,f6521,f1200,f1198,f1191,f1190,f890,f885,f884,f82,f7195,f7084,f6962,f6525,f2525,f2474,f1581,f766,f765,f764,f9978])).
fof(f4485,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl61),
  inference(cnf_transformation,[],[f4485_D])).
fof(f4485_D,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f11124,plain,(
  $false | ($spl1 | $spl4 | $spl7 | $spl75)),
  inference(global_subsumption,[],[f5062,f62,f61,f63,f72,f73,f86,f83,f90,f89,f91,f98,f100,f103,f105,f131,f134,f135,f88,f87,f155,f153,f160,f161,f214,f222,f283,f284,f326,f327,f396,f479,f756,f758,f759,f856,f883,f888,f902,f910,f914,f958,f962,f965,f968,f987,f1007,f1000,f1024,f1030,f1037,f1050,f1079,f1080,f1107,f1111,f1112,f1128,f1118,f1129,f1145,f1138,f1167,f1201,f1208,f1277,f1306,f1316,f1321,f1337,f1323,f1338,f380,f1582,f1569,f1570,f1571,f1572,f1573,f1593,f2650,f2641,f2636,f2634,f2622,f2620,f2526,f3239,f3277,f2545,f1576,f3278,f4780,f4789,f4781,f4790,f2524,f2711,f109,f5467,f5510,f5520,f5522,f5519,f5518,f5517,f6401,f6402,f6835,f7083,f7149,f6453,f94,f111,f7361,f6960,f6591,f6522,f6506,f1237,f1227,f757,f221,f97,f955,f6580,f6324,f6323,f4620,f2554,f2549,f2523,f2472,f2312,f2311,f1535,f1532,f1530,f1525,f1517,f1296,f1293,f1285,f1194,f1193,f1110,f1099,f1092,f1083,f1039,f1036,f1027,f855,f781,f777,f406,f404,f401,f271,f233,f106,f75,f71,f70,f8721,f1236,f1195,f986,f991,f1196,f982,f6521,f1200,f1198,f1191,f1190,f890,f885,f884,f82,f7195,f7084,f6962,f6525,f2525,f2474,f1581,f766,f765,f764,f9978])).
fof(f5062,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))) | $spl75),
  inference(cnf_transformation,[],[f5062_D])).
fof(f5062_D,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f11123,plain,(
  $false | ($spl1 | $spl4 | $spl7 | $spl88)),
  inference(global_subsumption,[],[f5549,f62,f61,f63,f72,f73,f86,f83,f90,f89,f91,f98,f100,f103,f105,f131,f134,f135,f88,f87,f155,f153,f160,f161,f214,f222,f283,f284,f326,f327,f396,f479,f756,f758,f759,f856,f883,f888,f902,f910,f914,f958,f962,f965,f968,f987,f1007,f1000,f1024,f1030,f1037,f1050,f1079,f1080,f1107,f1111,f1112,f1128,f1118,f1129,f1145,f1138,f1167,f1201,f1208,f1277,f1306,f1316,f1321,f1337,f1323,f1338,f380,f1582,f1569,f1570,f1571,f1572,f1573,f1593,f2650,f2641,f2636,f2634,f2622,f2620,f2526,f3239,f3277,f2545,f1576,f3278,f4780,f4789,f4781,f4790,f2524,f2711,f109,f5467,f5510,f5520,f5522,f5519,f5518,f5517,f6401,f6402,f6835,f7083,f7149,f6453,f94,f111,f7361,f6960,f6591,f6522,f6506,f1237,f1227,f757,f221,f97,f955,f6580,f6324,f6323,f4620,f2554,f2549,f2523,f2472,f2312,f2311,f1535,f1532,f1530,f1525,f1517,f1296,f1293,f1285,f1194,f1193,f1110,f1099,f1092,f1083,f1039,f1036,f1027,f855,f781,f777,f406,f404,f401,f271,f233,f106,f75,f71,f70,f8721,f1236,f1195,f986,f991,f1196,f982,f6521,f1200,f1198,f1191,f1190,f890,f885,f884,f82,f7195,f7084,f6962,f6525,f2525,f2474,f1581,f766,f765,f764,f9978])).
fof(f11122,plain,(
  $false | ($spl1 | $spl4 | $spl7 | $spl83)),
  inference(global_subsumption,[],[f11084,f62,f61,f63,f72,f73,f86,f83,f90,f89,f91,f98,f100,f103,f105,f131,f134,f135,f88,f87,f155,f153,f160,f161,f214,f222,f283,f284,f326,f327,f396,f479,f756,f758,f759,f856,f883,f888,f902,f910,f914,f958,f962,f965,f968,f987,f1007,f1000,f1024,f1030,f1037,f1050,f1079,f1080,f1107,f1111,f1112,f1128,f1118,f1129,f1145,f1138,f1167,f1201,f1208,f1277,f1306,f1316,f1321,f1337,f1323,f1338,f380,f1582,f1569,f1570,f1571,f1572,f1573,f1593,f2650,f2641,f2636,f2634,f2622,f2620,f2526,f3239,f3277,f2545,f1576,f3278,f4780,f4789,f4781,f4790,f2524,f2711,f109,f5467,f5510,f5520,f5522,f5519,f5518,f5517,f6401,f6402,f6835,f7083,f7149,f6453,f94,f111,f7361,f6960,f6591,f6522,f6506,f1237,f1227,f757,f221,f97,f955,f6580,f6324,f6323,f4620,f2554,f2549,f2523,f2472,f2312,f2311,f1535,f1532,f1530,f1525,f1517,f1296,f1293,f1285,f1194,f1193,f1110,f1099,f1092,f1083,f1039,f1036,f1027,f855,f781,f777,f406,f404,f401,f271,f233,f106,f75,f71,f70,f8721,f1236,f1195,f986,f991,f1196,f982,f6521,f1200,f1198,f1191,f1190,f890,f885,f884,f82,f7195,f7084,f6962,f6525,f2525,f2474,f1581,f766,f765,f764,f9978])).
fof(f11084,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))) | $spl83),
  inference(unit_resulting_resolution,[],[f9530,f986])).
fof(f11121,plain,(
  $false | ($spl1 | $spl4 | $spl7 | $spl83)),
  inference(global_subsumption,[],[f11101,f62,f61,f63,f72,f73,f86,f83,f90,f89,f91,f98,f100,f103,f105,f131,f134,f135,f88,f87,f155,f153,f160,f161,f214,f222,f283,f284,f326,f327,f396,f479,f756,f758,f759,f856,f883,f888,f902,f910,f914,f958,f962,f965,f968,f987,f1007,f1000,f1024,f1030,f1037,f1050,f1079,f1080,f1107,f1111,f1112,f1128,f1118,f1129,f1145,f1138,f1167,f1201,f1208,f1277,f1306,f1316,f1321,f1337,f1323,f1338,f380,f1582,f1569,f1570,f1571,f1572,f1573,f1593,f2650,f2641,f2636,f2634,f2622,f2620,f2526,f3239,f3277,f2545,f1576,f3278,f4780,f4789,f4781,f4790,f2524,f2711,f109,f5467,f5510,f5520,f5522,f5519,f5518,f5517,f6401,f6402,f6835,f7083,f7149,f6453,f94,f111,f7361,f6960,f6591,f6522,f6506,f1237,f1227,f757,f221,f97,f955,f6580,f6324,f6323,f4620,f2554,f2549,f2523,f2472,f2312,f2311,f1535,f1532,f1530,f1525,f1517,f1296,f1293,f1285,f1194,f1193,f1110,f1099,f1092,f1083,f1039,f1036,f1027,f855,f781,f777,f406,f404,f401,f271,f233,f106,f75,f71,f70,f8721,f1236,f1195,f986,f991,f1196,f982,f6521,f1200,f1198,f1191,f1190,f890,f885,f884,f82,f7195,f7084,f6962,f6525,f2525,f2474,f1581,f766,f765,f764,f9978])).
fof(f11101,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))) | $spl83),
  inference(trivial_inequality_removal,[],[f11088])).
fof(f11088,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))) | $spl83),
  inference(superposition,[],[f986,f9530])).
fof(f11120,plain,(
  $false | ($spl1 | $spl4 | $spl7 | $spl11)),
  inference(global_subsumption,[],[f11119,f62,f61,f63,f72,f73,f86,f83,f90,f89,f91,f98,f100,f103,f105,f131,f134,f135,f88,f87,f155,f153,f160,f161,f214,f222,f283,f284,f326,f327,f396,f479,f756,f758,f759,f856,f883,f888,f902,f910,f914,f958,f962,f965,f968,f987,f1007,f1000,f1024,f1030,f1037,f1050,f1079,f1080,f1107,f1111,f1112,f1128,f1118,f1129,f1145,f1138,f1167,f1201,f1208,f1277,f1306,f1316,f1321,f1337,f1323,f1338,f380,f1582,f1569,f1570,f1571,f1572,f1573,f1593,f2650,f2641,f2636,f2634,f2622,f2620,f2526,f3239,f3277,f2545,f1576,f3278,f4780,f4789,f4781,f4790,f2524,f2711,f109,f5467,f5510,f5520,f5522,f5519,f5518,f5517,f6401,f6402,f6835,f7083,f7149,f6453,f94,f111,f7361,f6960,f6591,f6522,f6506,f1237,f1227,f757,f221,f97,f955,f6580,f6324,f6323,f4620,f2554,f2549,f2523,f2472,f2312,f2311,f1535,f1532,f1530,f1525,f1517,f1296,f1293,f1285,f1194,f1193,f1110,f1099,f1092,f1083,f1039,f1036,f1027,f855,f781,f777,f406,f404,f401,f271,f233,f106,f75,f71,f70,f8721,f1236,f1195,f986,f991,f1196,f982,f6521,f1200,f1198,f1191,f1190,f890,f885,f884,f82,f7195,f7084,f6962,f6525,f2525,f2474,f1581,f766,f765,f764,f9978])).
fof(f11119,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl7 | $spl11)),
  inference(forward_demodulation,[],[f94,f8707])).
fof(f8707,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl7 | $spl11)),
  inference(unit_resulting_resolution,[],[f63,f4609,f111,f756])).
fof(f4609,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl11),
  inference(global_subsumption,[],[f62,f61,f63,f70,f71,f72,f73,f75,f86,f83,f90,f89,f91,f98,f100,f106,f105,f114,f131,f134,f135,f88,f87,f214,f222,f233,f271,f283,f326,f327,f401,f404,f406,f479,f756,f758,f759,f777,f781,f855,f883,f1024,f1039,f1027,f1030,f1036,f1080,f1099,f1083,f1107,f1092,f1110,f1167,f1193,f1194,f1201,f1285,f1293,f1296,f380,f955,f1530,f1517,f1532,f1535,f1525,f1582,f1569,f1570,f1571,f1572,f1573,f1593,f2650,f2641,f2636,f2634,f2311,f2312,f2622,f2620,f2685,f2756,f2472,f2523,f2526,f2554,f986,f1195,f982,f1196,f1200,f1198,f1191,f1190,f890,f885,f884,f82,f765,f766,f1581,f2474,f2525,f1237,f1236,f1227,f764,f757,f221,f97,f991,f3239,f3277,f2545,f2549,f1576])).
fof(f2756,plain,(
  ( ! [X62] : (s(bool,'T_0') != s(bool,X62) | s(bool,X62) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) | $spl11),
  inference(subsumption_resolution,[],[f2453,f114])).
fof(f2453,plain,(
  ( ! [X62] : (s(bool,'T_0') != s(bool,X62) | s(bool,X62) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl11),
  inference(superposition,[],[f114,f855])).
fof(f2685,plain,(
  ( ! [X37,X38] : (s(bool,X38) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X38) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X37),s('type.rat.rat',X37)))) ) | $spl11),
  inference(subsumption_resolution,[],[f2333,f114])).
fof(f2333,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X38) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X38) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X37),s('type.rat.rat',X37)))) )),
  inference(superposition,[],[f855,f221])).
fof(f11105,plain,(
  $false | ($spl60 | $spl83)),
  inference(subsumption_resolution,[],[f11104,f97])).
fof(f11104,plain,(
  p(s(bool,'F_0')) | ($spl60 | $spl83)),
  inference(forward_demodulation,[],[f11101,f4486])).
fof(f11103,plain,(
  $false | ($spl60 | $spl83)),
  inference(subsumption_resolution,[],[f11102,f97])).
fof(f11102,plain,(
  p(s(bool,'F_0')) | ($spl60 | $spl83)),
  inference(forward_demodulation,[],[f11084,f4486])).
fof(f5549,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl88),
  inference(cnf_transformation,[],[f5549_D])).
fof(f5549_D,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f11062,plain,(
  ( ! [X1] : (s('type.integer.int',X1) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))))))))) ) | $spl114),
  inference(cnf_transformation,[],[f11062_D])).
fof(f11062_D,plain,(
  ( ! [X1] : (s('type.integer.int',X1) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))))))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f11040,plain,(
  ( ! [X1] : (s('type.integer.int',X1) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))) ) | ($spl2 | $spl7 | $spl11)),
  inference(superposition,[],[f404,f11005])).
fof(f11005,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))) | ($spl2 | $spl7 | $spl11)),
  inference(unit_resulting_resolution,[],[f8856,f8795])).
fof(f8795,plain,(
  ( ! [X1] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) | s('type.integer.int',X1) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))))) ) | $spl2),
  inference(superposition,[],[f327,f8655])).
fof(f8655,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl2),
  inference(forward_demodulation,[],[f8654,f70])).
fof(f8654,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl2),
  inference(subsumption_resolution,[],[f8643,f63])).
fof(f8643,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl2),
  inference(superposition,[],[f884,f8410])).
fof(f8410,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl2),
  inference(backward_demodulation,[],[f106,f95])).
fof(f8856,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | ($spl7 | $spl11)),
  inference(forward_demodulation,[],[f8846,f72])).
fof(f8846,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))) | ($spl7 | $spl11)),
  inference(unit_resulting_resolution,[],[f8841,f90])).
fof(f8841,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))) | ($spl7 | $spl11)),
  inference(forward_demodulation,[],[f8840,f71])).
fof(f8840,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))))) | ($spl7 | $spl11)),
  inference(subsumption_resolution,[],[f8816,f63])).
fof(f8816,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))))) | ($spl7 | $spl11)),
  inference(superposition,[],[f884,f8707])).
fof(f5550,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))))))) ) | $spl90),
  inference(cnf_transformation,[],[f5550_D])).
fof(f5550_D,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f11061,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))) ) | ($spl2 | $spl7 | $spl11)),
  inference(forward_demodulation,[],[f11039,f8856])).
fof(f11039,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))) ) | ($spl2 | $spl7 | $spl11)),
  inference(superposition,[],[f406,f11005])).
fof(f11059,plain,(
  $false | ($spl2 | $spl5 | $spl7 | $spl11)),
  inference(subsumption_resolution,[],[f11058,f97])).
fof(f11058,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl5 | $spl7 | $spl11)),
  inference(forward_demodulation,[],[f11053,f10011])).
fof(f10011,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl5),
  inference(condensation,[],[f10010])).
fof(f10010,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,X7)) ) | $spl5),
  inference(subsumption_resolution,[],[f9667,f73])).
fof(f9667,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,X7)) ) | $spl5),
  inference(superposition,[],[f108,f6401])).
fof(f108,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl5),
  inference(cnf_transformation,[],[f108_D])).
fof(f108_D,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f11053,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl2 | $spl7 | $spl11)),
  inference(trivial_inequality_removal,[],[f11036])).
fof(f11036,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl2 | $spl7 | $spl11)),
  inference(superposition,[],[f982,f11005])).
fof(f11057,plain,(
  $false | ($spl2 | $spl5 | $spl7 | $spl11)),
  inference(subsumption_resolution,[],[f11056,f97])).
fof(f11056,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl5 | $spl7 | $spl11)),
  inference(forward_demodulation,[],[f11031,f10011])).
fof(f11031,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl2 | $spl7 | $spl11)),
  inference(unit_resulting_resolution,[],[f11005,f982])).
fof(f1585,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X0,X1,bool,bool))))) ) | $spl26),
  inference(cnf_transformation,[],[f1585_D])).
fof(f1585_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X0,X1,bool,bool))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f10420,plain,(
  s(bool,'F_0') != s(bool,g(s(bool,'F_0'))) | $spl113),
  inference(cnf_transformation,[],[f10420_D])).
fof(f10420_D,plain,(
  s(bool,'F_0') != s(bool,g(s(bool,'F_0'))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f10418,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,g(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X0,X1,bool,bool))))) ) | ($spl28 | $spl32)),
  inference(forward_demodulation,[],[f10417,f1589])).
fof(f1589,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl32),
  inference(cnf_transformation,[],[f1589_D])).
fof(f1589_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f10417,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X0,X1,bool,bool)))) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) ) | $spl28),
  inference(subsumption_resolution,[],[f6059,f1586])).
fof(f1586,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) | $spl28),
  inference(cnf_transformation,[],[f1586_D])).
fof(f1586_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'T_0'))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f6059,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X0,X1,bool,bool)))) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f6037])).
fof(f6037,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f764,f757])).
fof(f112,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f112_D])).
fof(f112_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f10415,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X23)) ) | $spl110),
  inference(cnf_transformation,[],[f10415_D])).
fof(f10415_D,plain,(
  ( ! [X23,X22] : (s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X23) | s(bool,X22) = s(bool,X23)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f10414,plain,(
  ( ! [X23,X21,X22] : (s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X21)) ) | ($spl14 | $spl25 | $spl32)),
  inference(subsumption_resolution,[],[f10365,f133])).
fof(f10365,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X21)) ) | ($spl25 | $spl32)),
  inference(superposition,[],[f756,f10223])).
fof(f10223,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X5)) ) | ($spl25 | $spl32)),
  inference(trivial_inequality_removal,[],[f10218])).
fof(f10218,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X5)) ) | ($spl25 | $spl32)),
  inference(superposition,[],[f10044,f6401])).
fof(f10044,plain,(
  s(bool,'F_0') != s(bool,g(s(bool,'F_0'))) | ($spl25 | $spl32)),
  inference(backward_demodulation,[],[f1589,f1584])).
fof(f10409,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X33)) ) | $spl108),
  inference(cnf_transformation,[],[f10409_D])).
fof(f10409_D,plain,(
  ( ! [X33] : (s(bool,X33) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X33)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f10380,plain,(
  ( ! [X33,X32] : (s(bool,X33) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X33) | s(bool,'F_0') = s(bool,X32)) ) | ($spl25 | $spl32)),
  inference(duplicate_literal_removal,[],[f10340])).
fof(f10340,plain,(
  ( ! [X33,X32] : (s(bool,X33) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X33) | s(bool,'F_0') = s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | ($spl25 | $spl32)),
  inference(superposition,[],[f6453,f10223])).
fof(f8131,plain,(
  ( ! [X40] : (s(bool,X40) != s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X40)) ) | $spl100),
  inference(cnf_transformation,[],[f8131_D])).
fof(f8131_D,plain,(
  ( ! [X40] : (s(bool,X40) != s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X40)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f9977,plain,(
  ( ! [X31,X32] : (s(bool,X32) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32)) ) | ($spl11 | $spl25 | $spl33)),
  inference(duplicate_literal_removal,[],[f9957])).
fof(f9957,plain,(
  ( ! [X31,X32] : (s(bool,X32) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X31)) ) | ($spl11 | $spl25 | $spl33)),
  inference(superposition,[],[f6453,f7579])).
fof(f7579,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl11 | $spl25 | $spl33)),
  inference(trivial_inequality_removal,[],[f7566])).
fof(f7566,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl11 | $spl25 | $spl33)),
  inference(superposition,[],[f7464,f98])).
fof(f7464,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | ($spl11 | $spl25 | $spl33)),
  inference(backward_demodulation,[],[f7462,f1584])).
fof(f7462,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl11 | $spl33)),
  inference(forward_demodulation,[],[f7427,f106])).
fof(f7427,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | ($spl11 | $spl33)),
  inference(unit_resulting_resolution,[],[f4609,f1590,f855])).
fof(f4470,plain,(
  ( ! [X15,X16] : (s(bool,X16) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X15),s('type.rat.rat',X15))) | s(bool,'F_0') = s(bool,X16)) ) | $spl50),
  inference(cnf_transformation,[],[f4470_D])).
fof(f4470_D,plain,(
  ( ! [X15,X16] : (s(bool,X16) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X15),s('type.rat.rat',X15))) | s(bool,'F_0') = s(bool,X16)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f9978,plain,(
  ( ! [X14,X15,X16] : (s(bool,X16) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X15),s('type.rat.rat',X15))) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X16)) )),
  inference(duplicate_literal_removal,[],[f9951])).
fof(f9951,plain,(
  ( ! [X14,X15,X16] : (s(bool,X16) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X15),s('type.rat.rat',X15))) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f6453,f283])).
fof(f9997,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X30)) ) | $spl106),
  inference(cnf_transformation,[],[f9997_D])).
fof(f9997_D,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X30)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f9984,plain,(
  ( ! [X30,X29] : (s(bool,X30) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X29)) ) | $spl33),
  inference(duplicate_literal_removal,[],[f9926])).
fof(f9926,plain,(
  ( ! [X30,X29] : (s(bool,X30) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X30) | s(bool,'F_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X29)) ) | $spl33),
  inference(superposition,[],[f6453,f7458])).
fof(f7458,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X8)) ) | $spl33),
  inference(trivial_inequality_removal,[],[f7453])).
fof(f7453,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X8) = s(bool,f(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X8)) ) | $spl33),
  inference(superposition,[],[f1590,f6401])).
fof(f4605,plain,(
  ( ! [X22] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X22))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X22)))))) ) | $spl68),
  inference(cnf_transformation,[],[f4605_D])).
fof(f4605_D,plain,(
  ( ! [X22] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X22))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X22)))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f8769,plain,(
  ( ! [X25] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X25)))))) ) | $spl104),
  inference(cnf_transformation,[],[f8769_D])).
fof(f8769_D,plain,(
  ( ! [X25] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X25)))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f8768,plain,(
  ( ! [X24,X25] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X25))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X24))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X24)))))) ) | ($spl7 | $spl54)),
  inference(subsumption_resolution,[],[f8750,f97])).
fof(f8750,plain,(
  ( ! [X24,X25] : (p(s(bool,'F_0')) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X25))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X24))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X24)))))) ) | ($spl7 | $spl54)),
  inference(backward_demodulation,[],[f8732,f1200])).
fof(f8732,plain,(
  ( ! [X5] : (s(bool,'F_0') = s(bool,X5)) ) | ($spl7 | $spl54)),
  inference(subsumption_resolution,[],[f8722,f4475])).
fof(f8722,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,X5)) ) | $spl7),
  inference(trivial_inequality_removal,[],[f8717])).
fof(f8717,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X5) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,X5)) ) | $spl7),
  inference(superposition,[],[f111,f6401])).
fof(f4606,plain,(
  ( ! [X23] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X23)))) ) | $spl70),
  inference(cnf_transformation,[],[f4606_D])).
fof(f4606_D,plain,(
  ( ! [X23] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X23)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f8767,plain,(
  ( ! [X23,X22] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X23))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X22))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X22)))))) ) | ($spl7 | $spl54)),
  inference(subsumption_resolution,[],[f8749,f97])).
fof(f8749,plain,(
  ( ! [X23,X22] : (p(s(bool,'F_0')) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X23))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X22))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X22)))))) ) | ($spl7 | $spl54)),
  inference(backward_demodulation,[],[f8732,f1198])).
fof(f8406,plain,(
  $false | ($spl1 | $spl82)),
  inference(subsumption_resolution,[],[f8405,f8326])).
fof(f8326,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0)))) ) | $spl82),
  inference(subsumption_resolution,[],[f8260,f283])).
fof(f8260,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) | s(bool,'T_0') != s(bool,X0)) ) | $spl82),
  inference(backward_demodulation,[],[f5164,f2549])).
fof(f8405,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl82)),
  inference(forward_demodulation,[],[f94,f8326])).
fof(f8403,plain,(
  $false | ($spl7 | $spl54 | $spl82)),
  inference(subsumption_resolution,[],[f111,f8400])).
fof(f8400,plain,(
  ( ! [X30] : (s(bool,'F_0') = s(bool,X30)) ) | ($spl54 | $spl82)),
  inference(subsumption_resolution,[],[f8265,f8326])).
fof(f8265,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) | s(bool,'F_0') = s(bool,X30)) ) | ($spl54 | $spl82)),
  inference(backward_demodulation,[],[f5164,f4475])).
fof(f8397,plain,(
  $false | ($spl11 | $spl25 | $spl33 | $spl82)),
  inference(subsumption_resolution,[],[f8365,f8345])).
fof(f8345,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl82),
  inference(subsumption_resolution,[],[f8261,f283])).
fof(f8261,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl82),
  inference(backward_demodulation,[],[f5164,f2554])).
fof(f8365,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | ($spl11 | $spl25 | $spl33 | $spl82)),
  inference(backward_demodulation,[],[f8345,f7464])).
fof(f8393,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f8346,f63])).
fof(f8346,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(backward_demodulation,[],[f8345,f97])).
fof(f8324,plain,(
  ( ! [X109,X110,X108] : (s(fun(X108,bool),X109) = s(fun(X108,bool),X110) | s(bool,'T_0') = s(bool,g(s(X108,sK1(X109,X110,bool,X108))))) ) | $spl102),
  inference(cnf_transformation,[],[f8324_D])).
fof(f8324_D,plain,(
  ( ! [X109,X110,X108] : (s(fun(X108,bool),X109) = s(fun(X108,bool),X110) | s(bool,'T_0') = s(bool,g(s(X108,sK1(X109,X110,bool,X108))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f8323,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'F_0') = s(bool,X111) | s(fun(X108,bool),X109) = s(fun(X108,bool),X110) | s(bool,'T_0') = s(bool,g(s(X108,sK1(X109,X110,bool,X108))))) ) | $spl82),
  inference(subsumption_resolution,[],[f8322,f7195])).
fof(f8322,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'F_0') = s(bool,X111) | s(bool,X111) != s(bool,f(s(X108,sK1(X109,X110,bool,X108)))) | s(fun(X108,bool),X109) = s(fun(X108,bool),X110) | s(bool,'T_0') = s(bool,g(s(X108,sK1(X109,X110,bool,X108))))) ) | $spl82),
  inference(forward_demodulation,[],[f8257,f221])).
fof(f8257,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,X111) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) | s(bool,X111) != s(bool,f(s(X108,sK1(X109,X110,bool,X108)))) | s(fun(X108,bool),X109) = s(fun(X108,bool),X110) | s(bool,'T_0') = s(bool,g(s(X108,sK1(X109,X110,bool,X108))))) ) | $spl82),
  inference(backward_demodulation,[],[f5164,f2474])).
fof(f8319,plain,(
  $false | ($spl19 | $spl82)),
  inference(subsumption_resolution,[],[f8318,f8215])).
fof(f8215,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | $spl82),
  inference(backward_demodulation,[],[f5164,f271])).
fof(f8318,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | ($spl19 | $spl82)),
  inference(forward_demodulation,[],[f8317,f8215])).
fof(f8317,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))) | ($spl19 | $spl82)),
  inference(forward_demodulation,[],[f8316,f8215])).
fof(f8316,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))))) | ($spl19 | $spl82)),
  inference(forward_demodulation,[],[f8315,f8215])).
fof(f8315,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))))))) | ($spl19 | $spl82)),
  inference(forward_demodulation,[],[f8253,f8215])).
fof(f8253,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))))))))) | ($spl19 | $spl82)),
  inference(backward_demodulation,[],[f5164,f2205])).
fof(f2205,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))))))))) | $spl19),
  inference(forward_demodulation,[],[f2204,f271])).
fof(f2204,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f1982,f89])).
fof(f1982,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))))))) | $spl19),
  inference(forward_demodulation,[],[f1981,f271])).
fof(f1981,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f1978,f89])).
fof(f1978,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))))) | $spl19),
  inference(forward_demodulation,[],[f1977,f271])).
fof(f1977,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f1565,f89])).
fof(f1565,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))) | $spl19),
  inference(forward_demodulation,[],[f1564,f271])).
fof(f1564,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl19),
  inference(forward_demodulation,[],[f1562,f271])).
fof(f1562,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f1202,f89])).
fof(f1202,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f1189,f89])).
fof(f8314,plain,(
  $false | ($spl19 | $spl82)),
  inference(subsumption_resolution,[],[f8313,f8215])).
fof(f8313,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | ($spl19 | $spl82)),
  inference(forward_demodulation,[],[f8312,f8215])).
fof(f8312,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))) | ($spl19 | $spl82)),
  inference(forward_demodulation,[],[f8311,f8215])).
fof(f8311,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))))) | ($spl19 | $spl82)),
  inference(forward_demodulation,[],[f8252,f8215])).
fof(f8252,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))))))) | ($spl19 | $spl82)),
  inference(backward_demodulation,[],[f5164,f1982])).
fof(f8310,plain,(
  $false | ($spl19 | $spl82)),
  inference(subsumption_resolution,[],[f8309,f8215])).
fof(f8309,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | ($spl19 | $spl82)),
  inference(forward_demodulation,[],[f8308,f8215])).
fof(f8308,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))) | ($spl19 | $spl82)),
  inference(forward_demodulation,[],[f8251,f8215])).
fof(f8251,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))))) | ($spl19 | $spl82)),
  inference(backward_demodulation,[],[f5164,f1978])).
fof(f8307,plain,(
  $false | ($spl19 | $spl82)),
  inference(subsumption_resolution,[],[f8306,f8215])).
fof(f8306,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | ($spl19 | $spl82)),
  inference(forward_demodulation,[],[f8250,f8215])).
fof(f8250,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))))) | ($spl19 | $spl82)),
  inference(backward_demodulation,[],[f5164,f1565])).
fof(f8305,plain,(
  $false | ($spl19 | $spl82)),
  inference(subsumption_resolution,[],[f8249,f8215])).
fof(f8249,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | ($spl19 | $spl82)),
  inference(backward_demodulation,[],[f5164,f1563])).
fof(f1563,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl19),
  inference(superposition,[],[f1202,f271])).
fof(f8285,plain,(
  $false | ($spl72 | $spl82)),
  inference(subsumption_resolution,[],[f8212,f4618])).
fof(f8212,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) | $spl82),
  inference(backward_demodulation,[],[f5164,f106])).
fof(f8284,plain,(
  $false | ($spl19 | $spl82)),
  inference(trivial_inequality_removal,[],[f8232])).
fof(f8232,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | ($spl19 | $spl82)),
  inference(backward_demodulation,[],[f5164,f1189])).
fof(f8283,plain,(
  $false | ($spl19 | $spl82)),
  inference(trivial_inequality_removal,[],[f8239])).
fof(f8239,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | ($spl19 | $spl82)),
  inference(backward_demodulation,[],[f5164,f1202])).
fof(f8068,plain,(
  ( ! [X39,X40] : (s(bool,X40) != s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X40) | s(bool,'T_0') = s(bool,X39)) ) | ($spl11 | $spl25 | $spl33 | $spl50)),
  inference(duplicate_literal_removal,[],[f8001])).
fof(f8001,plain,(
  ( ! [X39,X40] : (s(bool,X40) != s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X40) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | ($spl11 | $spl25 | $spl33 | $spl50)),
  inference(superposition,[],[f5948,f7579])).
fof(f5948,plain,(
  ( ! [X2,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X1)) ) | $spl50),
  inference(superposition,[],[f4470,f283])).
fof(f7379,plain,(
  ( ! [X118,X119,X117] : (s(fun(bool,bool),X117) = s(fun(bool,bool),X118) | s(bool,'F_0') = s(bool,X119) | s(bool,X119) = s(bool,sK1(X117,X118,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X117,X118,bool,bool))))) ) | $spl98),
  inference(cnf_transformation,[],[f7379_D])).
fof(f7379_D,plain,(
  ( ! [X118,X119,X117] : (s(fun(bool,bool),X117) = s(fun(bool,bool),X118) | s(bool,'F_0') = s(bool,X119) | s(bool,X119) = s(bool,sK1(X117,X118,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X117,X118,bool,bool))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f1590,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl33),
  inference(cnf_transformation,[],[f1590_D])).
fof(f1590_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f7149,plain,(
  ( ! [X118,X119,X117] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X117) = s(fun(bool,bool),X118) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X117,X118,bool,bool)))) | s(bool,X119) = s(bool,sK1(X117,X118,bool,bool)) | s(bool,'F_0') = s(bool,X119)) )),
  inference(superposition,[],[f764,f6401])).
fof(f7083,plain,(
  ( ! [X118,X119,X117] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X117) = s(fun(bool,bool),X118) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X117,X118,bool,bool)))) | s(bool,X119) = s(bool,sK1(X117,X118,bool,bool)) | s(bool,'F_0') = s(bool,X119)) )),
  inference(superposition,[],[f764,f6401])).
fof(f5812,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))) | $spl94),
  inference(cnf_transformation,[],[f5812_D])).
fof(f5812_D,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f5814,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))) | $spl96),
  inference(cnf_transformation,[],[f5814_D])).
fof(f5814_D,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f5791,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))) | $spl20),
  inference(superposition,[],[f75,f1335])).
fof(f5565,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))))))) ) | $spl92),
  inference(cnf_transformation,[],[f5565_D])).
fof(f5565_D,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f5518,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))) ) | $spl4),
  inference(superposition,[],[f406,f5510])).
fof(f5547,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))) ) | ($spl4 | $spl68)),
  inference(forward_demodulation,[],[f5518,f5528])).
fof(f5528,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | ($spl4 | $spl68)),
  inference(forward_demodulation,[],[f5517,f4605])).
fof(f5464,plain,(
  $false | ($spl4 | $spl11 | $spl58)),
  inference(subsumption_resolution,[],[f5463,f4609])).
fof(f5463,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl58)),
  inference(subsumption_resolution,[],[f5432,f4477])).
fof(f5432,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl4),
  inference(superposition,[],[f73,f109])).
fof(f5458,plain,(
  $false | ($spl4 | $spl11 | $spl58)),
  inference(subsumption_resolution,[],[f5457,f4609])).
fof(f5457,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl4 | $spl58)),
  inference(subsumption_resolution,[],[f5425,f4477])).
fof(f5425,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl4),
  inference(superposition,[],[f109,f73])).
fof(f5362,plain,(
  $false | ($spl2 | $spl11 | $spl60)),
  inference(subsumption_resolution,[],[f5361,f4609])).
fof(f5361,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl60)),
  inference(forward_demodulation,[],[f4486,f5359])).
fof(f5359,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl2),
  inference(forward_demodulation,[],[f95,f106])).
fof(f5357,plain,(
  $false | ($spl1 | $spl6 | $spl74)),
  inference(subsumption_resolution,[],[f5356,f110])).
fof(f5356,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl74)),
  inference(forward_demodulation,[],[f5355,f221])).
fof(f5355,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl74)),
  inference(forward_demodulation,[],[f94,f5063])).
fof(f5354,plain,(
  $false | ($spl2 | $spl11 | $spl74)),
  inference(subsumption_resolution,[],[f5353,f4609])).
fof(f5353,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl74)),
  inference(forward_demodulation,[],[f5352,f106])).
fof(f5352,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | ($spl2 | $spl74)),
  inference(forward_demodulation,[],[f5351,f221])).
fof(f5351,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl74)),
  inference(forward_demodulation,[],[f95,f5063])).
fof(f5350,plain,(
  $false | ($spl4 | $spl72 | $spl74)),
  inference(subsumption_resolution,[],[f5349,f4618])).
fof(f5349,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl4 | $spl74)),
  inference(forward_demodulation,[],[f109,f5063])).
fof(f5341,plain,(
  $false | ($spl8 | $spl25)),
  inference(subsumption_resolution,[],[f5326,f112])).
fof(f5326,plain,(
  s(bool,'F_0') != s(bool,g(s(bool,'F_0'))) | ($spl8 | $spl25)),
  inference(backward_demodulation,[],[f112,f1584])).
fof(f5300,plain,(
  ( ! [X16] : (s(bool,X16) != s(bool,f(s(bool,'T_0'))) | s(bool,X16) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) | $spl86),
  inference(cnf_transformation,[],[f5300_D])).
fof(f5300_D,plain,(
  ( ! [X16] : (s(bool,X16) != s(bool,f(s(bool,'T_0'))) | s(bool,X16) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f5291,plain,(
  ( ! [X15,X16] : (s(bool,X16) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) | $spl48),
  inference(duplicate_literal_removal,[],[f5257])).
fof(f5257,plain,(
  ( ! [X15,X16] : (s(bool,X16) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X15)) ) | $spl48),
  inference(superposition,[],[f4466,f855])).
fof(f5164,plain,(
  s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl82),
  inference(cnf_transformation,[],[f5164_D])).
fof(f5164_D,plain,(
  s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f5165,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))) ) | $spl84),
  inference(cnf_transformation,[],[f5165_D])).
fof(f5165_D,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f5146,plain,(
  ( ! [X0] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl7 | $spl11)),
  inference(superposition,[],[f406,f5090])).
fof(f5090,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | ($spl7 | $spl11)),
  inference(forward_demodulation,[],[f5076,f72])).
fof(f5076,plain,(
  s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))) | ($spl7 | $spl11)),
  inference(unit_resulting_resolution,[],[f5043,f90])).
fof(f5043,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))) | ($spl7 | $spl11)),
  inference(forward_demodulation,[],[f5042,f71])).
fof(f5042,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))))) | ($spl7 | $spl11)),
  inference(subsumption_resolution,[],[f5018,f63])).
fof(f5018,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))))) | ($spl7 | $spl11)),
  inference(superposition,[],[f884,f4818])).
fof(f4818,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl7 | $spl11)),
  inference(forward_demodulation,[],[f4807,f106])).
fof(f4807,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl7 | $spl11)),
  inference(unit_resulting_resolution,[],[f4609,f111,f855])).
fof(f1591,plain,(
  ( ! [X6,X4,X5] : (s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X6) | s(bool,X6) = s(bool,sK1(X4,X5,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X4,X5,bool,bool))))) ) | $spl34),
  inference(cnf_transformation,[],[f1591_D])).
fof(f1591_D,plain,(
  ( ! [X6,X4,X5] : (s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X6) | s(bool,X6) = s(bool,sK1(X4,X5,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X4,X5,bool,bool))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1587,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl29),
  inference(cnf_transformation,[],[f1587_D])).
fof(f1587_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f2711,plain,(
  ( ! [X78,X79,X77] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X77) = s(fun(bool,bool),X78) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X77,X78,bool,bool)))) | s(bool,X79) = s(bool,sK1(X77,X78,bool,bool)) | s(bool,'T_0') = s(bool,X79)) )),
  inference(forward_demodulation,[],[f2371,f106])).
fof(f2371,plain,(
  ( ! [X78,X79,X77] : (s(bool,'F_0') != s(bool,f(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))))) | s(fun(bool,bool),X77) = s(fun(bool,bool),X78) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X77,X78,bool,bool)))) | s(bool,X79) = s(bool,sK1(X77,X78,bool,bool)) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f764,f855])).
fof(f5072,plain,(
  ( ! [X92,X93,X91] : (s(fun(bool,bool),X91) = s(fun(bool,bool),X92) | s(bool,X93) = s(bool,sK1(X91,X92,bool,bool)) | s(bool,X93) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X91,X92,bool,bool))))) ) | $spl80),
  inference(cnf_transformation,[],[f5072_D])).
fof(f5072_D,plain,(
  ( ! [X92,X93,X91] : (s(fun(bool,bool),X91) = s(fun(bool,bool),X92) | s(bool,X93) = s(bool,sK1(X91,X92,bool,bool)) | s(bool,X93) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X91,X92,bool,bool))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f2524,plain,(
  ( ! [X92,X93,X91] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X91) = s(fun(bool,bool),X92) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X91,X92,bool,bool)))) | s(bool,X93) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X93) = s(bool,sK1(X91,X92,bool,bool))) )),
  inference(superposition,[],[f764,f855])).
fof(f5063,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))) | $spl74),
  inference(cnf_transformation,[],[f5063_D])).
fof(f5063_D,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f5065,plain,(
  ( ! [X1] : (s('type.integer.int',X1) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))))))) ) | $spl78),
  inference(cnf_transformation,[],[f5065_D])).
fof(f5065_D,plain,(
  ( ! [X1] : (s('type.integer.int',X1) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f5051,plain,(
  ( ! [X1] : (s('type.integer.int',X1) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))) ) | $spl2),
  inference(superposition,[],[f404,f5041])).
fof(f5041,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl2),
  inference(forward_demodulation,[],[f5040,f70])).
fof(f5040,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl2),
  inference(subsumption_resolution,[],[f5016,f63])).
fof(f5016,plain,(
  ~p(s(bool,'T_0')) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl2),
  inference(superposition,[],[f884,f4621])).
fof(f4621,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl2),
  inference(forward_demodulation,[],[f95,f106])).
fof(f5064,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))))) ) | $spl76),
  inference(cnf_transformation,[],[f5064_D])).
fof(f5064_D,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f5050,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))))) | s('type.integer.int',X0) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))) ) | $spl2),
  inference(superposition,[],[f406,f5041])).
fof(f4661,plain,(
  $false | ($spl2 | $spl11 | $spl56)),
  inference(subsumption_resolution,[],[f4660,f4609])).
fof(f4660,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl56)),
  inference(subsumption_resolution,[],[f4629,f4476])).
fof(f4629,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl2),
  inference(superposition,[],[f73,f4621])).
fof(f4655,plain,(
  $false | ($spl2 | $spl11 | $spl56)),
  inference(subsumption_resolution,[],[f4654,f4609])).
fof(f4654,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl2 | $spl56)),
  inference(subsumption_resolution,[],[f4622,f4476])).
fof(f4622,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl2),
  inference(superposition,[],[f4621,f73])).
fof(f2677,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) | $spl40),
  inference(cnf_transformation,[],[f2677_D])).
fof(f2677_D,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4619,plain,(
  ( ! [X14,X13] : (s(bool,'F_0') = s(bool,X13) | s(bool,X13) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X14)) ) | $spl14),
  inference(subsumption_resolution,[],[f3167,f100])).
fof(f3167,plain,(
  ( ! [X14,X13] : (p(s(bool,X14)) | s(bool,'F_0') = s(bool,X13) | s(bool,X13) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X14)) ) | $spl14),
  inference(superposition,[],[f133,f855])).
fof(f4618,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X10),s('type.rat.rat',X10)))) ) | $spl72),
  inference(cnf_transformation,[],[f4618_D])).
fof(f4618_D,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X10),s('type.rat.rat',X10)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3278,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X10),s('type.rat.rat',X10))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(forward_demodulation,[],[f3263,f106])).
fof(f3263,plain,(
  ( ! [X10,X9] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X10),s('type.rat.rat',X10))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(duplicate_literal_removal,[],[f3224])).
fof(f3224,plain,(
  ( ! [X10,X9] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X10),s('type.rat.rat',X10))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f2554,f283])).
fof(f4604,plain,(
  ( ! [X23,X22] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X23))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X22))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X22)))))) ) | $spl8),
  inference(subsumption_resolution,[],[f4590,f97])).
fof(f4590,plain,(
  ( ! [X23,X22] : (p(s(bool,'F_0')) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X23))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X22))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X22)))))) ) | $spl8),
  inference(backward_demodulation,[],[f112,f1198])).
fof(f4572,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f113,f4518])).
fof(f4518,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(forward_demodulation,[],[f4517,f106])).
fof(f4517,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | $spl3),
  inference(forward_demodulation,[],[f3215,f4509])).
fof(f4509,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl3),
  inference(condensation,[],[f4508])).
fof(f4508,plain,(
  ( ! [X109] : (s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'F_0') = s(bool,X109)) ) | $spl3),
  inference(subsumption_resolution,[],[f3995,f73])).
fof(f3995,plain,(
  ( ! [X109] : (s(bool,'T_0') != s(bool,X109) | s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'F_0') = s(bool,X109)) ) | $spl3),
  inference(superposition,[],[f1349,f3321])).
fof(f3321,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f2684,f2684])).
fof(f2684,plain,(
  ( ! [X48] : (s(bool,X48) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'F_0') = s(bool,X48)) ) | $spl3),
  inference(subsumption_resolution,[],[f2301,f1420])).
fof(f1420,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat',X4)))) ) | $spl3),
  inference(subsumption_resolution,[],[f1409,f1349])).
fof(f1409,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat',X4))) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))) ) | $spl3),
  inference(superposition,[],[f1349,f283])).
fof(f2301,plain,(
  ( ! [X47,X48] : (s(bool,'F_0') = s(bool,X48) | s(bool,X48) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X47),s('type.rat.rat',X47)))) )),
  inference(superposition,[],[f855,f221])).
fof(f1349,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl3),
  inference(forward_demodulation,[],[f96,f106])).
fof(f3215,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl3),
  inference(unit_resulting_resolution,[],[f1349,f2554])).
fof(f4533,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) | $spl66),
  inference(cnf_transformation,[],[f4533_D])).
fof(f4533_D,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2545,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f855])).
fof(f114,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl11),
  inference(cnf_transformation,[],[f114_D])).
fof(f114_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f3277,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(forward_demodulation,[],[f3268,f106])).
fof(f3268,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f3219])).
fof(f3219,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2554,f73])).
fof(f4516,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X20)) ) | $spl3),
  inference(forward_demodulation,[],[f4515,f106])).
fof(f4515,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X20)) ) | $spl3),
  inference(forward_demodulation,[],[f3259,f4509])).
fof(f3259,plain,(
  ( ! [X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X20)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f3230])).
fof(f3230,plain,(
  ( ! [X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) ) | $spl3),
  inference(superposition,[],[f2554,f1411])).
fof(f1411,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(trivial_inequality_removal,[],[f1405])).
fof(f1405,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f1349,f98])).
fof(f109,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl4),
  inference(cnf_transformation,[],[f109_D])).
fof(f109_D,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f4514,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f3237,f106])).
fof(f3237,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl6)),
  inference(superposition,[],[f2554,f1350])).
fof(f1350,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f93,f110])).
fof(f93,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl0),
  inference(cnf_transformation,[],[f93_D])).
fof(f93_D,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f2789,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) | $spl44),
  inference(cnf_transformation,[],[f2789_D])).
fof(f2789_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f3423,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) | $spl3),
  inference(forward_demodulation,[],[f3392,f106])).
fof(f3392,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) | $spl3),
  inference(equality_factoring,[],[f2684])).
fof(f4510,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'T_0')))) ) | $spl64),
  inference(cnf_transformation,[],[f4510_D])).
fof(f4510_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'T_0')))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f3673,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X0) = s(bool,f(s(bool,'T_0')))) ) | ($spl3 | $spl29)),
  inference(forward_demodulation,[],[f3636,f1612])).
fof(f1612,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl29),
  inference(unit_resulting_resolution,[],[f97,f1587,f134])).
fof(f3636,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,f(s(bool,'T_0')))) ) | ($spl3 | $spl29)),
  inference(equality_factoring,[],[f3527])).
fof(f3527,plain,(
  ( ! [X48] : (s(bool,X48) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X48)) ) | ($spl3 | $spl29)),
  inference(subsumption_resolution,[],[f3489,f63])).
fof(f3489,plain,(
  ( ! [X48] : (~p(s(bool,'T_0')) | s(bool,X48) = s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X48)) ) | ($spl3 | $spl29)),
  inference(superposition,[],[f3409,f1612])).
fof(f3409,plain,(
  ( ! [X6,X5] : (~p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X5)) ) | $spl3),
  inference(subsumption_resolution,[],[f3408,f63])).
fof(f3408,plain,(
  ( ! [X6,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X5)) ) | $spl3),
  inference(forward_demodulation,[],[f3339,f106])).
fof(f3339,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) | ~p(s(bool,X6)) | s(bool,X5) = s(bool,X6) | s(bool,'F_0') = s(bool,X5)) ) | $spl3),
  inference(superposition,[],[f86,f2684])).
fof(f4486,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl60),
  inference(cnf_transformation,[],[f4486_D])).
fof(f4486_D,plain,(
  s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f4505,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(condensation,[],[f4504])).
fof(f4504,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X51)) ) | $spl3),
  inference(subsumption_resolution,[],[f4040,f3370])).
fof(f3370,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f2554,f2684])).
fof(f4040,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X50) | s(bool,X50) = s(bool,X51) | s(bool,'F_0') = s(bool,X51)) ) | $spl3),
  inference(superposition,[],[f1411,f3321])).
fof(f4503,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X55)) ) | $spl3),
  inference(forward_demodulation,[],[f4502,f106])).
fof(f4502,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X55)) ) | $spl3),
  inference(subsumption_resolution,[],[f4042,f3370])).
fof(f4042,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X54) | s(bool,X54) = s(bool,X55) | s(bool,'F_0') = s(bool,X55)) ) | $spl3),
  inference(superposition,[],[f2554,f3321])).
fof(f4501,plain,(
  ( ! [X71] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X71)) ) | $spl3),
  inference(global_subsumption,[],[f4051,f62,f61,f63,f70,f71,f72,f73,f75,f86,f83,f90,f89,f91,f98,f100,f106,f105,f131,f134,f135,f88,f87,f214,f222,f233,f271,f283,f326,f327,f401,f404,f406,f479,f756,f758,f759,f777,f781,f855,f883,f1024,f1039,f1027,f1030,f1036,f1080,f1099,f1083,f1107,f1092,f1110,f1167,f1193,f1194,f1201,f1285,f1293,f1296,f380,f955,f1530,f1517,f1532,f1535,f1525,f1582,f1569,f1570,f1571,f1572,f1573,f1593,f2650,f2641,f2636,f2634,f2311,f2312,f2622,f2620,f2472,f2523,f2526,f2554,f986,f1195,f982,f1196,f1200,f1198,f1191,f1190,f890,f885,f884,f82,f765,f766,f1581,f2474,f2525,f1237,f1236,f1227,f764,f757,f221,f97,f991,f3239])).
fof(f4051,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X71) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X70),s('type.rat.rat',X70))) | s(bool,'F_0') = s(bool,X71)) ) | $spl3),
  inference(superposition,[],[f1420,f3321])).
fof(f4498,plain,(
  ( ! [X88] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X88)) ) | $spl3),
  inference(global_subsumption,[],[f4059,f62,f61,f63,f70,f71,f72,f73,f75,f86,f83,f90,f89,f91,f98,f100,f106,f105,f131,f134,f135,f88,f87,f214,f222,f233,f271,f283,f326,f327,f401,f404,f406,f479,f756,f758,f759,f777,f781,f855,f883,f1024,f1039,f1027,f1030,f1036,f1080,f1099,f1083,f1107,f1092,f1110,f1167,f1193,f1194,f1201,f1285,f1293,f1296,f380,f955,f1530,f1517,f1532,f1535,f1525,f1582,f1569,f1570,f1571,f1572,f1573,f1593,f2650,f2641,f2636,f2634,f2311,f2312,f2622,f2620,f2472,f2523,f2526,f2554,f986,f1195,f982,f1196,f1200,f1198,f1191,f1190,f890,f885,f884,f82,f765,f766,f1581,f2474,f2525,f1237,f1236,f1227,f764,f757,f221,f97,f991,f3239])).
fof(f4059,plain,(
  ( ! [X88] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X88) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'F_0') = s(bool,X88)) ) | $spl3),
  inference(superposition,[],[f1349,f3321])).
fof(f4495,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(condensation,[],[f4494])).
fof(f4494,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X51)) ) | $spl3),
  inference(subsumption_resolution,[],[f4102,f3370])).
fof(f4102,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X50) | s(bool,X50) = s(bool,X51) | s(bool,'F_0') = s(bool,X51)) ) | $spl3),
  inference(superposition,[],[f1411,f3321])).
fof(f4493,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X55)) ) | $spl3),
  inference(forward_demodulation,[],[f4492,f106])).
fof(f4492,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,X55)) ) | $spl3),
  inference(subsumption_resolution,[],[f4104,f3370])).
fof(f4104,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X54) | s(bool,X54) = s(bool,X55) | s(bool,'F_0') = s(bool,X55)) ) | $spl3),
  inference(superposition,[],[f2554,f3321])).
fof(f4491,plain,(
  ( ! [X71] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X71)) ) | $spl3),
  inference(global_subsumption,[],[f4113,f62,f61,f63,f70,f71,f72,f73,f75,f86,f83,f90,f89,f91,f98,f100,f106,f105,f131,f134,f135,f88,f87,f214,f222,f233,f271,f283,f326,f327,f401,f404,f406,f479,f756,f758,f759,f777,f781,f855,f883,f1024,f1039,f1027,f1030,f1036,f1080,f1099,f1083,f1107,f1092,f1110,f1167,f1193,f1194,f1201,f1285,f1293,f1296,f380,f955,f1530,f1517,f1532,f1535,f1525,f1582,f1569,f1570,f1571,f1572,f1573,f1593,f2650,f2641,f2636,f2634,f2311,f2312,f2622,f2620,f2472,f2523,f2526,f2554,f986,f1195,f982,f1196,f1200,f1198,f1191,f1190,f890,f885,f884,f82,f765,f766,f1581,f2474,f2525,f1237,f1236,f1227,f764,f757,f221,f97,f991,f3239])).
fof(f4113,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X71) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X70),s('type.rat.rat',X70))) | s(bool,'F_0') = s(bool,X71)) ) | $spl3),
  inference(superposition,[],[f1420,f3321])).
fof(f4488,plain,(
  ( ! [X88] : (s(bool,'F_0') = s(bool,X88) | s(bool,'T_0') != s(bool,'F_0')) ) | $spl3),
  inference(global_subsumption,[],[f4121,f62,f61,f63,f70,f71,f72,f73,f75,f86,f83,f90,f89,f91,f98,f100,f106,f105,f131,f134,f135,f88,f87,f214,f222,f233,f271,f283,f326,f327,f401,f404,f406,f479,f756,f758,f759,f777,f781,f855,f883,f1024,f1039,f1027,f1030,f1036,f1080,f1099,f1083,f1107,f1092,f1110,f1167,f1193,f1194,f1201,f1285,f1293,f1296,f380,f955,f1530,f1517,f1532,f1535,f1525,f1582,f1569,f1570,f1571,f1572,f1573,f1593,f2650,f2641,f2636,f2634,f2311,f2312,f2622,f2620,f2472,f2523,f2526,f2554,f986,f1195,f982,f1196,f1200,f1198,f1191,f1190,f890,f885,f884,f82,f765,f766,f1581,f2474,f2525,f1237,f1236,f1227,f764,f757,f221,f97,f991,f3239])).
fof(f4121,plain,(
  ( ! [X88] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X88) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'F_0') = s(bool,X88)) ) | $spl3),
  inference(superposition,[],[f1349,f3321])).
fof(f113,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl10),
  inference(cnf_transformation,[],[f113_D])).
fof(f113_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f4487,plain,(
  ( ! [X52,X53,X51] : (s(bool,X51) != s(bool,X52) | s(bool,X51) = s(bool,X53) | s(bool,X53) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X52)) ) | $spl62),
  inference(cnf_transformation,[],[f4487_D])).
fof(f4487_D,plain,(
  ( ! [X52,X53,X51] : (s(bool,X51) != s(bool,X52) | s(bool,X51) = s(bool,X53) | s(bool,X53) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X52)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f4382,plain,(
  ( ! [X52,X53,X51] : (s(bool,X51) != s(bool,X52) | s(bool,'T_0') = s(bool,X52) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X53) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X51) = s(bool,X53)) ) | $spl3),
  inference(superposition,[],[f3370,f855])).
fof(f4484,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0)))))) ) | $spl3),
  inference(condensation,[],[f4483])).
fof(f4483,plain,(
  ( ! [X80,X79] : (s(bool,'T_0') = s(bool,X80) | s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X79)) ) | $spl3),
  inference(subsumption_resolution,[],[f4395,f98])).
fof(f4395,plain,(
  ( ! [X80,X79] : (s(bool,X79) != s(bool,X80) | s(bool,'T_0') = s(bool,X80) | s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X79)) ) | $spl3),
  inference(superposition,[],[f3370,f1411])).
fof(f4477,plain,(
  ( ! [X34] : (s(bool,X34) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,X34)) ) | $spl58),
  inference(cnf_transformation,[],[f4477_D])).
fof(f4477_D,plain,(
  ( ! [X34] : (s(bool,X34) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'F_0') = s(bool,X34)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f4440,plain,(
  ( ! [X33,X34] : (s(bool,X34) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X33) | s(bool,'F_0') = s(bool,X34)) ) | ($spl0 | $spl3 | $spl6 | $spl11)),
  inference(duplicate_literal_removal,[],[f4413])).
fof(f4413,plain,(
  ( ! [X33,X34] : (s(bool,X34) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X33) | s(bool,'F_0') = s(bool,X34) | s(bool,'T_0') = s(bool,X33)) ) | ($spl0 | $spl3 | $spl6 | $spl11)),
  inference(superposition,[],[f3370,f1510])).
fof(f1510,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl6 | $spl11)),
  inference(subsumption_resolution,[],[f1492,f114])).
fof(f1492,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f1350,f98])).
fof(f4476,plain,(
  ( ! [X32] : (s(bool,X32) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'F_0') = s(bool,X32)) ) | $spl56),
  inference(cnf_transformation,[],[f4476_D])).
fof(f4476_D,plain,(
  ( ! [X32] : (s(bool,X32) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'F_0') = s(bool,X32)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4441,plain,(
  ( ! [X31,X32] : (s(bool,X32) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f4412])).
fof(f4412,plain,(
  ( ! [X31,X32] : (s(bool,X32) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X31)) ) | $spl3),
  inference(superposition,[],[f3370,f1411])).
fof(f4442,plain,(
  ( ! [X30,X29] : (s(bool,X30) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X30)) ) | ($spl3 | $spl6 | $spl11)),
  inference(duplicate_literal_removal,[],[f4411])).
fof(f4411,plain,(
  ( ! [X30,X29] : (s(bool,X30) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X29) | s(bool,'F_0') = s(bool,X30) | s(bool,'T_0') = s(bool,X29)) ) | ($spl3 | $spl6 | $spl11)),
  inference(superposition,[],[f3370,f1369])).
fof(f1369,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f1351,f114])).
fof(f1351,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(superposition,[],[f110,f98])).
fof(f4473,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'F_0') = s(bool,X25)) ) | $spl52),
  inference(cnf_transformation,[],[f4473_D])).
fof(f4473_D,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | s(bool,X24) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'F_0') = s(bool,X25)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f4443,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) != s(bool,X25) | s(bool,'T_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X25) | s(bool,X24) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f4409])).
fof(f4409,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) != s(bool,X25) | s(bool,'T_0') = s(bool,X23) | s(bool,'F_0') = s(bool,X25) | s(bool,X24) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X23)) ) | $spl3),
  inference(superposition,[],[f3370,f855])).
fof(f4444,plain,(
  ( ! [X14,X15,X16] : (s(bool,X16) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X15),s('type.rat.rat',X15))) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X16)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f4406])).
fof(f4406,plain,(
  ( ! [X14,X15,X16] : (s(bool,X16) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X15),s('type.rat.rat',X15))) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X14)) ) | $spl3),
  inference(superposition,[],[f3370,f283])).
fof(f4466,plain,(
  ( ! [X50] : (s(bool,X50) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X50)) ) | $spl48),
  inference(cnf_transformation,[],[f4466_D])).
fof(f4466_D,plain,(
  ( ! [X50] : (s(bool,X50) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X50)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f4450,plain,(
  ( ! [X50,X49] : (s(bool,X50) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X49)) ) | ($spl3 | $spl29)),
  inference(duplicate_literal_removal,[],[f4381])).
fof(f4381,plain,(
  ( ! [X50,X49] : (s(bool,X50) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X49)) ) | ($spl3 | $spl29)),
  inference(superposition,[],[f3370,f3527])).
fof(f2974,plain,(
  $false | ($spl0 | $spl3 | $spl6 | $spl11 | $spl44)),
  inference(subsumption_resolution,[],[f2858,f2934])).
fof(f2934,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl44),
  inference(subsumption_resolution,[],[f2931,f2554])).
fof(f2931,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) | $spl44),
  inference(condensation,[],[f2797])).
fof(f2797,plain,(
  ( ! [X4,X3] : (s(bool,X4) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | $spl44),
  inference(superposition,[],[f2789,f98])).
fof(f2858,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | ($spl0 | $spl3 | $spl6 | $spl11 | $spl44)),
  inference(superposition,[],[f2096,f2789])).
fof(f2096,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl3 | $spl6 | $spl11)),
  inference(subsumption_resolution,[],[f2026,f1420])).
fof(f2026,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X29),s('type.rat.rat',X29)))) ) | ($spl0 | $spl3 | $spl6 | $spl11)),
  inference(superposition,[],[f1420,f1510])).
fof(f2973,plain,(
  $false | ($spl3 | $spl44)),
  inference(subsumption_resolution,[],[f2857,f2934])).
fof(f2857,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | ($spl3 | $spl44)),
  inference(superposition,[],[f1349,f2789])).
fof(f2972,plain,(
  $false | ($spl3 | $spl44)),
  inference(subsumption_resolution,[],[f2849,f2934])).
fof(f2849,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | ($spl3 | $spl44)),
  inference(superposition,[],[f1420,f2789])).
fof(f2969,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f2935,f63])).
fof(f2935,plain,(
  ~p(s(bool,'T_0')) | $spl44),
  inference(backward_demodulation,[],[f2934,f97])).
fof(f2790,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X1)) ) | $spl46),
  inference(cnf_transformation,[],[f2790_D])).
fof(f2790_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2526,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X0) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(equality_factoring,[],[f855])).
fof(f2678,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X15),s('type.rat.rat',X15))) | s(bool,X16) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) | $spl42),
  inference(cnf_transformation,[],[f2678_D])).
fof(f2678_D,plain,(
  ( ! [X15,X16] : (s(bool,X16) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X15),s('type.rat.rat',X15))) | s(bool,X16) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2620,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X36) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X34),s('type.rat.rat',X34)))) )),
  inference(duplicate_literal_removal,[],[f2332])).
fof(f2332,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X36) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X34),s('type.rat.rat',X34))) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f855,f283])).
fof(f2672,plain,(
  ( ! [X28,X26] : (s(bool,X26) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X26) = s(bool,X28) | s(bool,X28) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) | $spl38),
  inference(cnf_transformation,[],[f2672_D])).
fof(f2672_D,plain,(
  ( ! [X28,X26] : (s(bool,X26) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X26) = s(bool,X28) | s(bool,X28) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2622,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X16) = s(bool,X18) | s(bool,X16) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) )),
  inference(duplicate_literal_removal,[],[f2326])).
fof(f2326,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X16) = s(bool,X18) | s(bool,X16) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f855,f855])).
fof(f2679,plain,(
  ( ! [X30,X29] : (s(bool,'F_0') = s(bool,X30) | s(bool,X30) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl6 | $spl11)),
  inference(forward_demodulation,[],[f2633,f110])).
fof(f2633,plain,(
  ( ! [X30,X29] : (s(bool,X30) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X30) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X29)) ) | ($spl6 | $spl11)),
  inference(duplicate_literal_removal,[],[f2292])).
fof(f2292,plain,(
  ( ! [X30,X29] : (s(bool,X30) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,X30) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) ) | ($spl6 | $spl11)),
  inference(superposition,[],[f855,f1369])).
fof(f2634,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) = s(bool,X25) | s(bool,X25) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0)))) )),
  inference(duplicate_literal_removal,[],[f2290])).
fof(f2290,plain,(
  ( ! [X24,X23,X25] : (s(bool,X24) = s(bool,X25) | s(bool,X25) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X23) | s(bool,X24) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f855,f855])).
fof(f2636,plain,(
  ( ! [X14,X15,X16] : (s(bool,X16) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X15),s('type.rat.rat',X15))) | s(bool,X16) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(duplicate_literal_removal,[],[f2287])).
fof(f2287,plain,(
  ( ! [X14,X15,X16] : (s(bool,X16) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X15),s('type.rat.rat',X15))) | s(bool,X16) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f855,f283])).
fof(f2641,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f2282])).
fof(f2282,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f855,f73])).
fof(f2650,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,X27) | s(bool,X26) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X28) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X26) = s(bool,X28)) )),
  inference(duplicate_literal_removal,[],[f2262])).
fof(f2262,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,X27) | s(bool,X26) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X27) | s(bool,X28) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) | s(bool,X26) = s(bool,X28)) )),
  inference(superposition,[],[f855,f855])).
fof(f1588,plain,(
  ( ! [X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK1(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X2,X3,bool,bool))))) ) | $spl30),
  inference(cnf_transformation,[],[f1588_D])).
fof(f1588_D,plain,(
  ( ! [X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK1(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X2,X3,bool,bool))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1593,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X16,X17,bool,bool)))) | s(bool,'T_0') = s(bool,sK1(X16,X17,bool,bool))) )),
  inference(forward_demodulation,[],[f1574,f221])).
fof(f1574,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,f(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X18),s('type.rat.rat',X18))))) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X16,X17,bool,bool)))) | s(bool,'T_0') = s(bool,sK1(X16,X17,bool,bool))) )),
  inference(superposition,[],[f764,f283])).
fof(f1573,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X13) = s(fun(bool,bool),X14) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X13,X14,bool,bool)))) | s(bool,X15) = s(bool,sK1(X13,X14,bool,bool)) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f764,f98])).
fof(f1592,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl36),
  inference(cnf_transformation,[],[f1592_D])).
fof(f1592_D,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f1572,plain,(
  ( ! [X12,X10,X11] : (s(bool,'F_0') != s(bool,f(s(bool,X12))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X10,X11,bool,bool)))) | s(bool,'T_0') = s(bool,sK1(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f764,f98])).
fof(f1571,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X9))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X7,X8,bool,bool)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK1(X7,X8,bool,bool))) )),
  inference(superposition,[],[f764,f98])).
fof(f1570,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X4,X5,bool,bool)))) | s(bool,X6) = s(bool,sK1(X4,X5,bool,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f764,f98])).
fof(f1569,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X2,X3,bool,bool)))) | s(bool,'T_0') = s(bool,sK1(X2,X3,bool,bool))) )),
  inference(superposition,[],[f764,f73])).
fof(f1584,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0'))) | $spl25),
  inference(cnf_transformation,[],[f1584_D])).
fof(f1584_D,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0'))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f1582,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X0,X1,bool,bool)))) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f1568])).
fof(f1568,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK1(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f764,f757])).
fof(f1335,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))) | $spl20),
  inference(cnf_transformation,[],[f1335_D])).
fof(f1335_D,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1336,plain,(
  ( ! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))))))))) ) | $spl22),
  inference(cnf_transformation,[],[f1336_D])).
fof(f1336_D,plain,(
  ( ! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))))))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1321,plain,(
  ( ! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0)))))))) ) | $spl7),
  inference(superposition,[],[f404,f1208])).
fof(f1189,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl19),
  inference(cnf_transformation,[],[f1189_D])).
fof(f1189_D,plain,(
  s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f1187,plain,(
  ( ! [X1] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X1)))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X1))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X1)))))) ) | $spl16),
  inference(cnf_transformation,[],[f1187_D])).
fof(f1187_D,plain,(
  ( ! [X1] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X1)))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X1))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X1)))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1201,plain,(
  ( ! [X26] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X26)))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X26))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X26)))))) )),
  inference(forward_demodulation,[],[f1184,f70])).
fof(f1184,plain,(
  ( ! [X26] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_sub_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X26))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X26)))))) )),
  inference(superposition,[],[f986,f401])).
fof(f1167,plain,(
  ( ! [X1] : (s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',sK0))) != s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X1)))) | s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat',X1))) = s('type.integer.int','const.rat.rat_sgn_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X1)))))) )),
  inference(superposition,[],[f885,f401])).
fof(f146,plain,(
  $false | ($spl1 | $spl12)),
  inference(subsumption_resolution,[],[f141,f132])).
fof(f141,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl1 | $spl12)),
  inference(backward_demodulation,[],[f132,f94])).
fof(f143,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f137,f63])).
fof(f137,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(backward_demodulation,[],[f132,f97])).
fof(f142,plain,(
  $false | ($spl7 | $spl12)),
  inference(subsumption_resolution,[],[f136,f132])).
fof(f136,plain,(
  s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl7 | $spl12)),
  inference(backward_demodulation,[],[f132,f111])).
fof(f133,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl14),
  inference(cnf_transformation,[],[f133_D])).
fof(f133_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f135,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f128,f100])).
fof(f128,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f83,f73])).
fof(f131,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f125,f100])).
fof(f125,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f83,f73])).
fof(f105,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f73])).
fof(f111,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl7),
  inference(cnf_transformation,[],[f111_D])).
fof(f111_D,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f103,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl1),
  inference(superposition,[],[f94,f73])).
fof(f94,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl1),
  inference(cnf_transformation,[],[f94_D])).
fof(f94_D,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f96,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | $spl3),
  inference(cnf_transformation,[],[f96_D])).
fof(f96_D,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f62,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK0))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f52])).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 219
% Passive clauses: 844
% Generated clauses: 12228
% Final active clauses: 59
% Final passive clauses: 113
% Input formulas: 20
% Initial clauses: 31
% 
% Duplicate literals: 813
% Trivial inequalities: 131
% Fw subsumption resolutions: 694
% Fw demodulations: 1702
% Bw demodulations: 331
% Condensations: 885
% Global subsumptions: 20
% 
% Simple tautologies: 231
% Equational tautologies: 332
% Forward subsumptions: 6408
% Fw demodulations to eq. taut.: 54
% Bw demodulations to eq. taut.: 13
% 
% Unit resulting resolution: 1170
% Factoring: 1
% Forward superposition: 2718
% Backward superposition: 2697
% Self superposition: 130
% Equality factoring: 124
% Equality resolution: 26
% 
% Split clauses: 73
% Split components: 149
% SAT solver clauses: 1060
% SAT solver unit clauses: 27
% SAT solver binary clauses: 290
% 
% Sat splits: 73
% Sat splitting refutations: 48
% 
% Memory used [KB]: 3709
% Time elapsed: 0.374 s
% ------------------------------
% ------------------------------
% Success in time 1.73 s
