Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (104480)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 143
% Passive clauses: 2642
% Generated clauses: 20774
% Final active clauses: 129
% Final passive clauses: 2353
% Input formulas: 37
% Initial clauses: 72
% 
% Duplicate literals: 734
% Trivial inequalities: 295
% Fw subsumption resolutions: 2726
% Fw demodulations: 3164
% Bw demodulations: 12
% 
% Simple tautologies: 298
% Equational tautologies: 68
% Forward subsumptions: 9762
% 
% Binary resolution: 569
% Factoring: 1
% Forward superposition: 8013
% Backward superposition: 4749
% Self superposition: 191
% Equality factoring: 126
% Equality resolution: 39
% 
% Split clauses: 394
% Split components: 796
% SAT solver clauses: 814
% SAT solver unit clauses: 21
% SAT solver binary clauses: 136
% 
% Sat splits: 394
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 811
% TWLsolver calls for satisfiability: 109
% 
% Memory used [KB]: 11513
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (104508)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 154
% Passive clauses: 1023
% Generated clauses: 18211
% Final active clauses: 111
% Final passive clauses: 576
% Input formulas: 37
% Initial clauses: 72
% 
% Duplicate literals: 1178
% Trivial inequalities: 187
% Fw subsumption resolutions: 854
% Fw demodulations: 1277
% Bw demodulations: 326
% Condensations: 2255
% Global subsumptions: 9
% 
% Simple tautologies: 283
% Equational tautologies: 31
% Forward subsumptions: 9979
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 236
% Factoring: 1
% Forward superposition: 5577
% Backward superposition: 5007
% Self superposition: 140
% Equality factoring: 49
% Equality resolution: 28
% 
% Split clauses: 123
% Split components: 247
% SAT solver clauses: 1190
% SAT solver unit clauses: 92
% SAT solver binary clauses: 40
% SAT solver learnt clauses: 18
% 
% Sat splits: 123
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 1154
% TWLsolver calls for satisfiability: 339
% 
% Memory used [KB]: 12409
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 287 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (104525)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 698
% Passive clauses: 6181
% Generated clauses: 168046
% Final active clauses: 53
% Final passive clauses: 129
% Input formulas: 37
% Initial clauses: 72
% 
% Duplicate literals: 6095
% Trivial inequalities: 1365
% Fw subsumption resolutions: 12008
% Fw demodulations: 24941
% Bw demodulations: 4220
% Condensations: 9970
% Global subsumptions: 3549
% 
% Simple tautologies: 1390
% Equational tautologies: 903
% Forward subsumptions: 96945
% Fw demodulations to eq. taut.: 59
% Bw demodulations to eq. taut.: 183
% 
% Unit resulting resolution: 7422
% Factoring: 6
% Forward superposition: 37900
% Backward superposition: 33578
% Self superposition: 649
% Equality factoring: 340
% Equality resolution: 110
% 
% Split clauses: 348
% Split components: 711
% SAT solver clauses: 14455
% SAT solver unit clauses: 191
% SAT solver binary clauses: 1522
% 
% Sat splits: 348
% Sat splitting refutations: 4026
% 
% Memory used [KB]: 85201
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 183 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% (104793)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1346
% Passive clauses: 26281
% Generated clauses: 171662
% Final active clauses: 361
% Final passive clauses: 13541
% Input formulas: 37
% Initial clauses: 69
% 
% Selected by SInE selection: 34
% SInE iterations: 1
% 
% Duplicate literals: 4490
% Trivial inequalities: 4752
% Fw subsumption resolutions: 23783
% Fw demodulations: 28119
% Bw demodulations: 486
% 
% Simple tautologies: 4164
% Equational tautologies: 111
% Forward subsumptions: 79696
% Fw demodulations to eq. taut.: 29
% Bw demodulations to eq. taut.: 10
% 
% Binary resolution: 6907
% Factoring: 4
% Forward superposition: 55239
% Backward superposition: 41227
% Self superposition: 433
% Equality factoring: 248
% Equality resolution: 304
% 
% Split clauses: 3491
% Split components: 7056
% SAT solver clauses: 3721
% SAT solver unit clauses: 36
% SAT solver binary clauses: 363
% 
% Sat splits: 3491
% Sat splitting refutations: 221
% 
% Memory used [KB]: 52323
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.pred_set
% (104926)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 156
% Passive clauses: 1694
% Generated clauses: 21455
% Final active clauses: 131
% Final passive clauses: 1444
% Input formulas: 37
% Initial clauses: 72
% 
% Duplicate literals: 519
% Trivial inequalities: 333
% Fw subsumption resolutions: 1864
% Fw demodulations: 3182
% Bw demodulations: 99
% Global subsumptions: 1
% 
% Simple tautologies: 572
% Equational tautologies: 104
% Forward subsumptions: 12992
% Fw demodulations to eq. taut.: 87
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 1560
% Forward superposition: 7359
% Backward superposition: 6177
% Self superposition: 166
% Equality factoring: 30
% Equality resolution: 26
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 1752
% SAT solver unit clauses: 108
% SAT solver binary clauses: 187
% 
% Sat splits: 4
% Sat splitting refutations: 51
% 
% Memory used [KB]: 12153
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f14680,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f137,f220,f223,f371,f369,f216,f766,f218,f1342,f1347,f1918,f1920,f1923,f1926,f1927,f1930,f1932,f1933,f1934,f1941,f1944,f1981,f1985,f2003,f2005,f2007,f2009,f2011,f2013,f2015,f2017,f2022,f2024,f2026,f2028,f2030,f2032,f2034,f2036,f2038,f2040,f2045,f2047,f2048,f2049,f2050,f2061,f2065,f2067,f2069,f2078,f2080,f2983,f2985,f2996,f2997,f2998,f2999,f3096,f3097,f3099,f3102,f3105,f3111,f3115,f3119,f3121,f3123,f3127,f3129,f3131,f3133,f3135,f3140,f3153,f3159,f3163,f3167,f3171,f3175,f3178,f3181,f3187,f3192,f3195,f3197,f3199,f3202,f3206,f3207,f3209,f4507,f4984,f4982,f8614,f8619,f8618,f8623,f8695,f8698,f8699,f8702,f9195,f9197,f9201,f9203,f9205,f9207,f9209,f9257,f9259,f9261,f9262,f9263,f9272,f9435,f9437,f9449,f9456,f9460,f9477,f9481,f9487,f9489,f9580,f9689,f9683,f9740,f9789,f9796,f9801,f9802,f9803,f9804,f9805,f9809,f9810,f9813,f9825,f9826,f9827,f14671])).
fof(f14671,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f14646,f3222])).
fof(f3222,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f3220,f139])).
fof(f139,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.SING_IFF_CARD|49|',aHL_TRUTH)).
fof(f3220,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f1353,f221])).
fof(f221,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | $spl2),
  inference(subsumption_resolution,[],[f215,f220])).
fof(f215,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))))),
  inference(forward_demodulation,[],[f138,f141])).
fof(f141,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.arithmetic.ONE')).
fof(f138,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))))),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  (p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))))) & (~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f102])).
fof(f102,plain,(
  ? [X0,X1] : ((p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1)))) | (s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))))) & (~p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ? [X0,X1] : ((p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1)))) | (s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))))) & (~p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1)))) | (s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))))))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1)))) <~> (s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1] : (p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1)))) <=> (s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1)))) <=> (s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.pred_set.SING_1'(s(fun(X6,bool),X7)))) <=> (s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.SING_1'(s(fun(X6,bool),X7)))) <=> (s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) & p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7))))))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.pred_set.SING_IFF_CARD|49|')).
fof(f1353,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))) | $spl0),
  inference(resolution,[],[f218,f150])).
fof(f150,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.bool.EQ_CLAUSES|split|1')).
fof(f14646,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ($spl0 | $spl2)),
  inference(superposition,[],[f10034,f1352])).
fof(f1352,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl0),
  inference(resolution,[],[f218,f195])).
fof(f195,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,sK4(X1,X0)),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & (! [X2] : s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f132])).
fof(f132,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1)))) | ? [X3] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & (! [X2] : s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1)))) | ? [X2] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) & (! [X2] : s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.SING_1'(s(fun(X6,bool),X7)))) <=> ? [X11] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.pred_set.SING_DEF')).
fof(f10034,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f10027,f199])).
fof(f199,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X8,X9] : s(X6,X9) = s(X6,'const.bool.COND_3'(s(bool,'F_0'),s(X6,X8),s(X6,X9)))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.bool.COND_CLAUSES|split|1')).
fof(f10027,plain,(
  ( ! [X0] : (s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f10006,f3489])).
fof(f3489,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f3488,f144])).
fof(f144,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.EMPTY_0')))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.pred_set.CARD_DEF|split|0')).
fof(f3488,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f3443,f139])).
fof(f3443,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f192,f3219])).
fof(f3219,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2)),
  inference(backward_demodulation,[],[f225,f1799])).
fof(f1799,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))) | $spl0),
  inference(superposition,[],[f200,f1352])).
fof(f200,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2] : s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ! [X11,X7] : s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7))) = s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X11),s(fun(X6,bool),X7)))))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.pred_set.FINITE_INSERT')).
fof(f225,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))) | $spl2),
  inference(resolution,[],[f220,f150])).
fof(f192,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ! [X2] : s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => ! [X2] : s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => ! [X11] : s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X11),s(fun(X6,bool),X7))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7))),s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))))))))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.pred_set.CARD_DEF|split|1')).
fof(f10006,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(resolution,[],[f197,f156])).
fof(f156,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.bool.EQ_CLAUSES|split|3')).
fof(f197,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X11] : ~p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.pred_set.NOT_IN_EMPTY')).
fof(f9827,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9190,f9793])).
fof(f9793,plain,(
  ( ! [X6,X5] : (s('type.num.num',X6) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X5)))) ) | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9792,f187])).
fof(f187,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X12] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X12)))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.num.NOT_SUC')).
fof(f9792,plain,(
  ( ! [X6,X5] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) | s('type.num.num',X6) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X5)))) ) | ($spl2 | $spl5 | $spl16)),
  inference(forward_demodulation,[],[f9681,f9586])).
fof(f9586,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1)))) ) | $spl16),
  inference(subsumption_resolution,[],[f9492,f9579])).
fof(f9579,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | $spl16),
  inference(subsumption_resolution,[],[f9162,f8622])).
fof(f8622,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl16),
  inference(subsumption_resolution,[],[f149,f8619])).
fof(f149,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f104])).
fof(f9162,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) ) | $spl16),
  inference(backward_demodulation,[],[f8619,f189])).
fof(f189,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,sK2(X1,X0)),s(fun(X0,bool),sK3(X1,X0))))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,sK2(X1,X0)),s(fun(X0,bool),sK3(X1,X0)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK2(X1,X0)),s(fun(X0,bool),sK3(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f67])).
fof(f67,plain,(
  ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | ? [X2,X3] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | ? [X2,X3] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : ! [X7] : (s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7) | ? [X11,X10] : (s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X11),s(fun(X6,bool),X10))) & ~p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.pred_set.SET_CASES')).
fof(f9492,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1)) ) | $spl16),
  inference(subsumption_resolution,[],[f9071,f8622])).
fof(f9071,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1)) ) | $spl16),
  inference(backward_demodulation,[],[f8619,f191])).
fof(f191,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) & (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)))),
  inference(ennf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1)))) => (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) <=> s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X7)))) => (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(X6,bool),X7))) <=> s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X7)))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.pred_set.CARD_EQ_|48|')).
fof(f9681,plain,(
  ( ! [X6,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num',X6) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X5)))) ) | ($spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f9579,f9445])).
fof(f9445,plain,(
  ( ! [X6,X5] : (s('type.num.num',X6) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X5))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))) ) | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f8964,f9256])).
fof(f9256,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1)) ) | $spl16),
  inference(subsumption_resolution,[],[f9255,f8622])).
fof(f9255,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X1))) ) | $spl16),
  inference(subsumption_resolution,[],[f204,f8622])).
fof(f204,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.bool.IMP_ANTISYM_AX')).
fof(f8964,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') != s(bool,g(s(sK0,sK5(X5,sK3(sK1,sK0),bool,sK0)))) | s('type.num.num',X6) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X5))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))) ) | ($spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f4328])).
fof(f4328,plain,(
  ( ! [X6,X5] : (s('type.num.num',X6) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X5))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(bool,f(s(sK0,sK5(X5,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X5,sK3(sK1,sK0),bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f706,f207])).
fof(f207,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.SING_IFF_CARD|49|',aHL_EXT)).
fof(f706,plain,(
  ( ! [X0] : (s('type.num.num',X0) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f582,f206])).
fof(f206,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X13,X12] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))) <=> s('type.num.num',X12) = s('type.num.num',X13))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.prim_rec.INV_SUC_EQ')).
fof(f582,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f578,f139])).
fof(f578,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))))) | ($spl2 | $spl5)),
  inference(backward_demodulation,[],[f576,f519])).
fof(f519,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f518,f379])).
fof(f379,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(trivial_inequality_removal,[],[f374])).
fof(f374,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.EMPTY_0') | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f372,f188])).
fof(f188,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,sK2(X1,X0)),s(fun(X0,bool),sK3(X1,X0)))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1)) )),
  inference(cnf_transformation,[],[f128])).
fof(f372,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f224,f369])).
fof(f224,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | $spl2),
  inference(resolution,[],[f220,f191])).
fof(f518,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f466,f199])).
fof(f466,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f192,f393])).
fof(f393,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(resolution,[],[f378,f154])).
fof(f154,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.bool.EQ_CLAUSES|split|2')).
fof(f378,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))))) | ($spl2 | $spl5)),
  inference(trivial_inequality_removal,[],[f375])).
fof(f375,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.EMPTY_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f372,f189])).
fof(f576,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f546,f225])).
fof(f546,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f200,f379])).
fof(f9190,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))) ) | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f4506])).
fof(f4506,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))) ) | ($spl0 | $spl2 | $spl5)),
  inference(forward_demodulation,[],[f4505,f4413])).
fof(f4413,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))))) | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f4409,f139])).
fof(f4409,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f4407,f4076])).
fof(f4076,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | ($spl0 | $spl2 | $spl5)),
  inference(forward_demodulation,[],[f4075,f3225])).
fof(f3225,plain,(
  s(fun(sK0,bool),sK3(sK1,sK0)) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f736,f3222])).
fof(f736,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK3(sK1,sK0)) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f699,f144])).
fof(f699,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(fun(sK0,bool),sK3(sK1,sK0)) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f582,f188])).
fof(f4075,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | ($spl0 | $spl2 | $spl5)),
  inference(forward_demodulation,[],[f4023,f199])).
fof(f4023,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f192,f3807])).
fof(f3807,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl0 | $spl2 | $spl5)),
  inference(resolution,[],[f3224,f154])).
fof(f3224,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f737,f3222])).
fof(f737,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f700,f144])).
fof(f700,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f582,f189])).
fof(f4407,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl0 | $spl2 | $spl5)),
  inference(forward_demodulation,[],[f4376,f576])).
fof(f4376,plain,(
  s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f200,f3225])).
fof(f4505,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))))))) ) | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f4451,f139])).
fof(f4451,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))))))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f192,f4407])).
fof(f9826,plain,(
  $false | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9189,f9793])).
fof(f9189,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK1,sK0)))))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f675])).
fof(f675,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK1,sK0)))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK1,sK0)))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) ) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f674,f582])).
fof(f674,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK1,sK0)))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK1,sK0)))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))))) ) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f618,f139])).
fof(f618,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK1,sK0)))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK1,sK0)))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f192,f576])).
fof(f9825,plain,(
  $false | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9824,f9793])).
fof(f9824,plain,(
  ( ! [X2] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X2),s(fun(sK0,bool),sK1)))))) ) | ($spl2 | $spl16)),
  inference(subsumption_resolution,[],[f9186,f8622])).
fof(f9186,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X2),s(fun(sK0,bool),sK1)))))) ) | ($spl2 | $spl16)),
  inference(backward_demodulation,[],[f8619,f5420])).
fof(f5420,plain,(
  ( ! [X2] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X2),s(fun(sK0,bool),sK1))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X2),s(fun(sK0,bool),sK1))))) ) | $spl2),
  inference(forward_demodulation,[],[f5361,f198])).
fof(f198,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X8,X9] : s(X6,X8) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X8),s(X6,X9)))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.bool.COND_CLAUSES|split|0')).
fof(f5361,plain,(
  ( ! [X2] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X2),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X2),s(fun(sK0,bool),sK1))))) ) | $spl2),
  inference(superposition,[],[f222,f150])).
fof(f222,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) ) | $spl2),
  inference(resolution,[],[f220,f192])).
fof(f9813,plain,(
  $false | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9812,f8622])).
fof(f9812,plain,(
  ( ! [X11] : (~p(s(bool,X11))) ) | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9811,f9793])).
fof(f9811,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X10),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) | ~p(s(bool,X11))) ) | ($spl2 | $spl16)),
  inference(subsumption_resolution,[],[f9179,f8622])).
fof(f9179,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X10),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) | ~p(s(bool,X11))) ) | ($spl2 | $spl16)),
  inference(backward_demodulation,[],[f8619,f5367])).
fof(f5367,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X10),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.bool.COND_3'(s(bool,X11),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,X10),s(fun(sK0,bool),sK1)))) | ~p(s(bool,X11))) ) | $spl2),
  inference(superposition,[],[f222,f204])).
fof(f9810,plain,(
  $false | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9177,f9793])).
fof(f9177,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) ) | ($spl2 | $spl16)),
  inference(backward_demodulation,[],[f8619,f222])).
fof(f9809,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9176,f9793])).
fof(f9176,plain,(
  ( ! [X0] : (s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl0 | $spl2 | $spl16)),
  inference(backward_demodulation,[],[f8619,f3489])).
fof(f9805,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9174,f8622])).
fof(f9174,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f5723])).
fof(f5723,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0))))) | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5684,f372])).
fof(f5684,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0))))) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f3227,f189])).
fof(f3227,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3221,f139])).
fof(f3221,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl5)),
  inference(backward_demodulation,[],[f1353,f603])).
fof(f603,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f557,f379])).
fof(f557,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl2 | $spl5)),
  inference(superposition,[],[f195,f379])).
fof(f9804,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9173,f8622])).
fof(f9173,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f4703])).
fof(f4703,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0))))) | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f4665,f372])).
fof(f4665,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0))))) | $spl0),
  inference(superposition,[],[f1843,f189])).
fof(f1843,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl0),
  inference(forward_demodulation,[],[f1842,f1352])).
fof(f1842,plain,(
  s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl0),
  inference(subsumption_resolution,[],[f1841,f139])).
fof(f1841,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl0),
  inference(forward_demodulation,[],[f1810,f1353])).
fof(f1810,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl0),
  inference(superposition,[],[f195,f1352])).
fof(f9803,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9172,f8622])).
fof(f9172,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f4405])).
fof(f4405,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))),sK0))))) | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f4372,f3223])).
fof(f3223,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK3(sK1,sK0)) | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f739,f3222])).
fof(f739,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK3(sK1,sK0)) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f738,f139])).
fof(f738,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK3(sK1,sK0)) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f703,f576])).
fof(f703,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK3(sK1,sK0)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f582,f191])).
fof(f4372,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK3(sK1,sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))),sK0))))) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f3225,f189])).
fof(f9802,plain,(
  $false | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9171,f8622])).
fof(f9171,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f575])).
fof(f575,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f543,f372])).
fof(f543,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f379,f189])).
fof(f9801,plain,(
  $false | ($spl0 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9800,f8622])).
fof(f9800,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X2)),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X2)),sK0)))))) ) | ($spl0 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9170,f8622])).
fof(f9170,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X2)),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X2)),sK0)))))) ) | ($spl0 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f6412])).
fof(f6412,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X2)),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X2)),sK0))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(X2,sK0)),s(fun(sK0,bool),sK3(X2,sK0)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f1834,f189])).
fof(f1834,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0))))) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f1795,f1833])).
fof(f1833,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f1832,f143])).
fof(f143,plain,(
  ( ! [X0] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.pred_set.FINITE_EMPTY')).
fof(f1832,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f1831,f200])).
fof(f1831,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f1830,f1352])).
fof(f1830,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f1805,f369])).
fof(f1805,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | $spl0),
  inference(superposition,[],[f191,f1352])).
fof(f1795,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0))))) | $spl0),
  inference(superposition,[],[f1352,f189])).
fof(f9796,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f9163,f8622])).
fof(f9163,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f8619,f197])).
fof(f9789,plain,(
  $false | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9788,f9739])).
fof(f9739,plain,(
  ( ! [X9] : (s('type.num.num',X9) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl2 | $spl5 | $spl16)),
  inference(forward_demodulation,[],[f9738,f9586])).
fof(f9738,plain,(
  ( ! [X9] : (s('type.num.num',X9) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9737,f9712])).
fof(f9712,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | ($spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9711,f205])).
fof(f205,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f134])).
fof(f9711,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) ) | ($spl5 | $spl16)),
  inference(forward_demodulation,[],[f9591,f9586])).
fof(f9591,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) ) | ($spl5 | $spl16)),
  inference(backward_demodulation,[],[f9579,f387])).
fof(f387,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) ) | $spl5),
  inference(superposition,[],[f369,f205])).
fof(f9737,plain,(
  ( ! [X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) | s('type.num.num',X9) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(forward_demodulation,[],[f9642,f9586])).
fof(f9642,plain,(
  ( ! [X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s('type.num.num',X9) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f9579,f9284])).
fof(f9284,plain,(
  ( ! [X9] : (s('type.num.num',X9) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f9273,f721])).
fof(f721,plain,(
  ( ! [X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))) | s('type.num.num',X9) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f206,f582])).
fof(f9273,plain,(
  ( ! [X1] : (s(fun(sK0,bool),sK1) = s(fun(sK0,bool),X1)) ) | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f8912,f9256])).
fof(f8912,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(sK0,sK5(X1,'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),bool,sK0)))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),X1)) ) | ($spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f545])).
fof(f545,plain,(
  ( ! [X1] : (s(fun(sK0,bool),sK1) = s(fun(sK0,bool),X1) | s(bool,f(s(sK0,sK5(X1,'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f379,f207])).
fof(f9788,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl2 | $spl5 | $spl16)),
  inference(forward_demodulation,[],[f9679,f9586])).
fof(f9679,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1)))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f9579,f9438])).
fof(f9438,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1)))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f8958,f9256])).
fof(f8958,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1)))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f702])).
fof(f702,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))) | s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f582,f207])).
fof(f9740,plain,(
  $false | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f141,f9739])).
fof(f9683,plain,(
  $false | ($spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9222,f9586])).
fof(f9222,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X2)))) ) | ($spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9221,f139])).
fof(f9221,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X2)))) ) | ($spl5 | $spl16)),
  inference(forward_demodulation,[],[f8663,f8619])).
fof(f8663,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X2))) | ~p(s(bool,g(s(sK0,sK5(sK1,X2,bool,sK0)))))) ) | ($spl5 | $spl16)),
  inference(subsumption_resolution,[],[f5630,f8619])).
fof(f5630,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,X2,bool,sK0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X2))) | ~p(s(bool,g(s(sK0,sK5(sK1,X2,bool,sK0)))))) ) | $spl5),
  inference(superposition,[],[f384,f150])).
fof(f384,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK5(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,X0,bool,sK0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X0)))) ) | $spl5),
  inference(superposition,[],[f369,f207])).
fof(f9689,plain,(
  $false | ($spl5 | $spl16)),
  inference(trivial_inequality_removal,[],[f9685])).
fof(f9685,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl5 | $spl16)),
  inference(backward_demodulation,[],[f9586,f9416])).
fof(f9416,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1)))))) ) | ($spl5 | $spl16)),
  inference(subsumption_resolution,[],[f8932,f9256])).
fof(f8932,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1)))))) ) | ($spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f4116])).
fof(f4116,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))) | s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0))))) ) | $spl5),
  inference(superposition,[],[f4109,f207])).
fof(f4109,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))) | $spl5),
  inference(forward_demodulation,[],[f4108,f144])).
fof(f4108,plain,(
  ( ! [X4] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X4,bool),'const.pred_set.EMPTY_0')))))) ) | $spl5),
  inference(trivial_inequality_removal,[],[f4097])).
fof(f4097,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X4,bool),'const.pred_set.EMPTY_0')))))) ) | $spl5),
  inference(superposition,[],[f387,f144])).
fof(f9580,plain,(
  $false | ($spl0 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9224,f9579])).
fof(f9224,plain,(
  ( ! [X2] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X2)) ) | ($spl0 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9223,f139])).
fof(f9223,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X2)) ) | ($spl0 | $spl5 | $spl16)),
  inference(forward_demodulation,[],[f8659,f8619])).
fof(f8659,plain,(
  ( ! [X2] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X2) | ~p(s(bool,g(s(sK0,sK5(sK1,X2,bool,sK0)))))) ) | ($spl0 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f5568,f8619])).
fof(f5568,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,X2,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X2) | ~p(s(bool,g(s(sK0,sK5(sK1,X2,bool,sK0)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f1849,f150])).
fof(f1849,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK5(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,X0,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X0)) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f1833,f207])).
fof(f9489,plain,(
  $false | ($spl0 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9488,f8622])).
fof(f9488,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),X0))))) ) | ($spl0 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9054,f9256])).
fof(f9054,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,g(s(sK0,sK5(sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0),X0,bool,sK0)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),X0))))) ) | ($spl0 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f6421])).
fof(f6421,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),X0)))) | s(bool,f(s(sK0,sK5(sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0),X0,bool,sK0))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f1834,f207])).
fof(f9487,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(subsumption_resolution,[],[f9486,f8622])).
fof(f9486,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK3(sK1,sK0),sK0),sK0)),s(fun(sK0,bool),X0))))) ) | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(subsumption_resolution,[],[f9053,f9256])).
fof(f9053,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,g(s(sK0,sK5(sK3(sK3(sK3(sK1,sK0),sK0),sK0),X0,bool,sK0)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK3(sK1,sK0),sK0),sK0)),s(fun(sK0,bool),X0))))) ) | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(backward_demodulation,[],[f8619,f5026])).
fof(f5026,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK3(sK1,sK0),sK0),sK0)),s(fun(sK0,bool),X0)))) | s(bool,f(s(sK0,sK5(sK3(sK3(sK3(sK1,sK0),sK0),sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK3(sK3(sK1,sK0),sK0),sK0),X0,bool,sK0))))) ) | ($spl0 | $spl2 | $spl5 | $spl11)),
  inference(superposition,[],[f4991,f207])).
fof(f4991,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK3(sK1,sK0),sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK3(sK1,sK0),sK0),sK0))))) | ($spl0 | $spl2 | $spl5 | $spl11)),
  inference(trivial_inequality_removal,[],[f4988])).
fof(f4988,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.EMPTY_0') | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK3(sK1,sK0),sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK3(sK1,sK0),sK0),sK0))))) | ($spl0 | $spl2 | $spl5 | $spl11)),
  inference(superposition,[],[f4985,f189])).
fof(f4985,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)) | ($spl0 | $spl2 | $spl5 | $spl11)),
  inference(subsumption_resolution,[],[f4508,f4982])).
fof(f4508,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)) | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f4453,f139])).
fof(f4453,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f191,f4407])).
fof(f9481,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9480,f8622])).
fof(f9480,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),X0))))) ) | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9043,f9256])).
fof(f9043,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,g(s(sK0,sK5(sK3(sK3(sK1,sK0),sK0),X0,bool,sK0)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),X0))))) ) | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f3817])).
fof(f3817,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),X0)))) | s(bool,f(s(sK0,sK5(sK3(sK3(sK1,sK0),sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK3(sK1,sK0),sK0),X0,bool,sK0))))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f3224,f207])).
fof(f9477,plain,(
  $false | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9476,f8622])).
fof(f9476,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X0))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9030,f9256])).
fof(f9030,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X0))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f403])).
fof(f403,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X0)))) | s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f378,f207])).
fof(f9460,plain,(
  $false | ($spl0 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9459,f8622])).
fof(f9459,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X3)),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X3)),sK0)))))) ) | ($spl0 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f8987,f9256])).
fof(f8987,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,g(s(sK0,sK5('const.pred_set.EMPTY_0',X3,bool,sK0)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X3)),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X3)),sK0)))))) ) | ($spl0 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f6416])).
fof(f6416,plain,(
  ( ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X3)),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X3)),sK0))))) | s(bool,f(s(sK0,sK5('const.pred_set.EMPTY_0',X3,bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.EMPTY_0',X3,bool,sK0))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f1834,f207])).
fof(f9456,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(subsumption_resolution,[],[f9455,f8622])).
fof(f9455,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK3(sK1,sK0),sK0),sK0)),s(fun(sK0,bool),X1))))) ) | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(subsumption_resolution,[],[f8979,f9256])).
fof(f8979,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(sK0,sK5(X1,sK3(sK3(sK3(sK1,sK0),sK0),sK0),bool,sK0)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK3(sK1,sK0),sK0),sK0)),s(fun(sK0,bool),X1))))) ) | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(backward_demodulation,[],[f8619,f5027])).
fof(f5027,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK3(sK1,sK0),sK0),sK0)),s(fun(sK0,bool),X1)))) | s(bool,f(s(sK0,sK5(X1,sK3(sK3(sK3(sK1,sK0),sK0),sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK3(sK3(sK1,sK0),sK0),sK0),bool,sK0))))) ) | ($spl0 | $spl2 | $spl5 | $spl11)),
  inference(superposition,[],[f4991,f207])).
fof(f9449,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9448,f8622])).
fof(f9448,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),X1))))) ) | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f8969,f9256])).
fof(f8969,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(sK0,sK5(X1,sK3(sK3(sK1,sK0),sK0),bool,sK0)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),X1))))) ) | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f3818])).
fof(f3818,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),X1)))) | s(bool,f(s(sK0,sK5(X1,sK3(sK3(sK1,sK0),sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK3(sK1,sK0),sK0),bool,sK0))))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f3224,f207])).
fof(f9437,plain,(
  $false | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9436,f8622])).
fof(f9436,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X1))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f8956,f9256])).
fof(f8956,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X1))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f404])).
fof(f404,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X1)))) | s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f378,f207])).
fof(f9435,plain,(
  $false | ($spl0 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9434,f8622])).
fof(f9434,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),X1))))) ) | ($spl0 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f8955,f9256])).
fof(f8955,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(sK0,sK5(X1,sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0),bool,sK0)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),X1))))) ) | ($spl0 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f6422])).
fof(f6422,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),X1)))) | s(bool,f(s(sK0,sK5(X1,sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0),bool,sK0))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f1834,f207])).
fof(f9272,plain,(
  $false | ($spl0 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9271,f8622])).
fof(f9271,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X4)),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X4)),sK0)))))) ) | ($spl0 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f8911,f9256])).
fof(f8911,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,g(s(sK0,sK5(X4,'const.pred_set.EMPTY_0',bool,sK0)))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X4)),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X4)),sK0)))))) ) | ($spl0 | $spl5 | $spl16)),
  inference(backward_demodulation,[],[f8619,f6417])).
fof(f6417,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X4)),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X4)),sK0))))) | s(bool,f(s(sK0,sK5(X4,'const.pred_set.EMPTY_0',bool,sK0)))) != s(bool,g(s(sK0,sK5(X4,'const.pred_set.EMPTY_0',bool,sK0))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f1834,f207])).
fof(f9263,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f8721,f8622])).
fof(f8721,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f8619,f213])).
fof(f213,plain,(
  ~p(s(bool,'F_0'))),
  inference(condensation,[],[f185])).
fof(f185,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.bool.IMP_CLAUSES|split|4')).
fof(f9262,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f8720,f8622])).
fof(f8720,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f8619,f208])).
fof(f208,plain,(
  ~p(s(bool,'F_0'))),
  inference(condensation,[],[f148])).
fof(f148,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'F_0'))) => ~p(s(bool,X10)))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.bool.IMP_F')).
fof(f9261,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f9260,f8622])).
fof(f9260,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl16),
  inference(subsumption_resolution,[],[f8713,f9256])).
fof(f8713,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl16),
  inference(backward_demodulation,[],[f8619,f155])).
fof(f155,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f107])).
fof(f9259,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f9258,f8622])).
fof(f9258,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl16),
  inference(subsumption_resolution,[],[f8711,f9256])).
fof(f8711,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl16),
  inference(backward_demodulation,[],[f8619,f153])).
fof(f153,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f106])).
fof(f9257,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f8709,f8622])).
fof(f8709,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f8619,f142])).
fof(f142,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.SING_IFF_CARD|49|',aHL_FALSITY)).
fof(f9209,plain,(
  $false | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9208,f8619])).
fof(f9208,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(forward_demodulation,[],[f8696,f8619])).
fof(f8696,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f431,f8619])).
fof(f431,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f391,f207])).
fof(f391,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(resolution,[],[f378,f149])).
fof(f9207,plain,(
  $false | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9206,f8619])).
fof(f9206,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(forward_demodulation,[],[f8697,f8619])).
fof(f8697,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f432,f8619])).
fof(f432,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f391,f207])).
fof(f9205,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9204,f8619])).
fof(f9204,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(X1,sK3(sK3(sK1,sK0),sK0),bool,sK0))))) ) | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(forward_demodulation,[],[f8700,f8619])).
fof(f8700,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK5(X1,sK3(sK3(sK1,sK0),sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK3(sK1,sK0),sK0),bool,sK0))))) ) | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f3989,f8619])).
fof(f3989,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),X1))) | s(bool,f(s(sK0,sK5(X1,sK3(sK3(sK1,sK0),sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK3(sK1,sK0),sK0),bool,sK0))))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f3805,f207])).
fof(f3805,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl0 | $spl2 | $spl5)),
  inference(resolution,[],[f3224,f149])).
fof(f9203,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f9202,f8619])).
fof(f9202,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK3(sK3(sK1,sK0),sK0),X0,bool,sK0))))) ) | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(forward_demodulation,[],[f8701,f8619])).
fof(f8701,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK5(sK3(sK3(sK1,sK0),sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK3(sK1,sK0),sK0),X0,bool,sK0))))) ) | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f3988,f8619])).
fof(f3988,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),X0))) | s(bool,f(s(sK0,sK5(sK3(sK3(sK1,sK0),sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK3(sK1,sK0),sK0),X0,bool,sK0))))) ) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f3805,f207])).
fof(f9201,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(subsumption_resolution,[],[f9200,f214])).
fof(f214,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f193])).
fof(f193,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X11] : (s(X6,X11) = s(X6,X11) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.bool.REFL_CLAUSE')).
fof(f9200,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(forward_demodulation,[],[f8703,f8619])).
fof(f8703,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK3(sK3(sK1,sK0),sK0),sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK3(sK3(sK1,sK0),sK0),sK0),sK0))))) | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(subsumption_resolution,[],[f5054,f8619])).
fof(f5054,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK3(sK1,sK0),sK0),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK3(sK3(sK1,sK0),sK0),sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK3(sK3(sK1,sK0),sK0),sK0),sK0))))) | ($spl0 | $spl2 | $spl5 | $spl11)),
  inference(superposition,[],[f5014,f189])).
fof(f5014,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK3(sK1,sK0),sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK3(sK1,sK0),sK0),sK0)))) | ($spl0 | $spl2 | $spl5 | $spl11)),
  inference(resolution,[],[f4991,f149])).
fof(f9197,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(subsumption_resolution,[],[f9196,f8619])).
fof(f9196,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(X1,sK3(sK3(sK3(sK1,sK0),sK0),sK0),bool,sK0))))) ) | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(forward_demodulation,[],[f8706,f8619])).
fof(f8706,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK5(X1,sK3(sK3(sK3(sK1,sK0),sK0),sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK3(sK3(sK1,sK0),sK0),sK0),bool,sK0))))) ) | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(subsumption_resolution,[],[f5056,f8619])).
fof(f5056,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK3(sK1,sK0),sK0),sK0)),s(fun(sK0,bool),X1))) | s(bool,f(s(sK0,sK5(X1,sK3(sK3(sK3(sK1,sK0),sK0),sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK3(sK3(sK1,sK0),sK0),sK0),bool,sK0))))) ) | ($spl0 | $spl2 | $spl5 | $spl11)),
  inference(superposition,[],[f5014,f207])).
fof(f9195,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(subsumption_resolution,[],[f9194,f8619])).
fof(f9194,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK3(sK3(sK3(sK1,sK0),sK0),sK0),X0,bool,sK0))))) ) | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(forward_demodulation,[],[f8707,f8619])).
fof(f8707,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK5(sK3(sK3(sK3(sK1,sK0),sK0),sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK3(sK3(sK1,sK0),sK0),sK0),X0,bool,sK0))))) ) | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(subsumption_resolution,[],[f5055,f8619])).
fof(f5055,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK3(sK1,sK0),sK0),sK0)),s(fun(sK0,bool),X0))) | s(bool,f(s(sK0,sK5(sK3(sK3(sK3(sK1,sK0),sK0),sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK3(sK3(sK1,sK0),sK0),sK0),X0,bool,sK0))))) ) | ($spl0 | $spl2 | $spl5 | $spl11)),
  inference(superposition,[],[f5014,f207])).
fof(f8702,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl11 | $spl16)),
  inference(subsumption_resolution,[],[f5014,f8619])).
fof(f8699,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f3805,f8619])).
fof(f8698,plain,(
  $false | ($spl2 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f391,f8619])).
fof(f8695,plain,(
  $false | ($spl0 | $spl5 | $spl16)),
  inference(subsumption_resolution,[],[f6401,f8619])).
fof(f6401,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)))) | ($spl0 | $spl5)),
  inference(resolution,[],[f1834,f149])).
fof(f8623,plain,(
  $false | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f1383,f8619])).
fof(f1383,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl0),
  inference(forward_demodulation,[],[f1355,f1353])).
fof(f1355,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))) | $spl0),
  inference(resolution,[],[f218,f153])).
fof(f8618,plain,(
  ( ! [X21] : (~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X21)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X21))))) ) | $spl14),
  inference(cnf_transformation,[],[f8618_D])).
fof(f8618_D,plain,(
  ( ! [X21] : (~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X21)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X21))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f8619,plain,(
  ( ! [X22] : (s(bool,'T_0') = s(bool,X22)) ) | $spl16),
  inference(cnf_transformation,[],[f8619_D])).
fof(f8619_D,plain,(
  ( ! [X22] : (s(bool,'T_0') = s(bool,X22)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f8614,plain,(
  ( ! [X21,X22] : (s(bool,'T_0') = s(bool,X22) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X21)))) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X21))))) ) | ($spl0 | $spl2)),
  inference(duplicate_literal_removal,[],[f8613])).
fof(f8613,plain,(
  ( ! [X21,X22] : (s(bool,'T_0') = s(bool,X22) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X21)))) | s(bool,'T_0') = s(bool,X22) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X21))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f8612,f3219])).
fof(f8612,plain,(
  ( ! [X21,X22] : (s(bool,X22) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X21)))) | s(bool,'T_0') = s(bool,X22) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X21))))) ) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f8569,f200])).
fof(f8569,plain,(
  ( ! [X21,X22] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X21)))) | s(bool,'T_0') = s(bool,X22) | s(bool,X22) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(X21,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X21))))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f8306,f195])).
fof(f8306,plain,(
  ( ! [X17,X18] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X17)))) | s(bool,'T_0') = s(bool,X18) | s(bool,X18) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X17)))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f8207,f1383])).
fof(f8207,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X18) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X17))) | s(bool,'T_0') = s(bool,X18) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X17))))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f7038,f154])).
fof(f7038,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X0))) | s(bool,X1) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X0))) | s(bool,'T_0') = s(bool,X1)) ) | ($spl0 | $spl2)),
  inference(resolution,[],[f6527,f150])).
fof(f6527,plain,(
  ( ! [X8,X7] : (p(s(bool,X8)) | s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X7))) | s(bool,X8) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X7)))) ) | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f6460,f6517])).
fof(f6517,plain,(
  ( ! [X3] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(X3,sK0))))) | s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X3)))) ) | ($spl0 | $spl2)),
  inference(duplicate_literal_removal,[],[f6456])).
fof(f6456,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X3))) | s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X3))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(X3,sK0)))))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f6369,f150])).
fof(f6369,plain,(
  ( ! [X3] : (s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X3))) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(X3,sK0)))) | s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X3)))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f200,f3428])).
fof(f3428,plain,(
  ( ! [X1] : (s(fun(sK0,bool),X1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(X1,sK0)),s(fun(sK0,bool),sK3(X1,sK0)))) | s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1)))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f3219,f188])).
fof(f6460,plain,(
  ( ! [X8,X7] : (s(bool,X8) = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X7))) | s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X7))) | p(s(bool,X8)) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(X7,sK0)))))) ) | ($spl0 | $spl2)),
  inference(superposition,[],[f6369,f201])).
fof(f201,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f99])).
fof(f4982,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | $spl11),
  inference(cnf_transformation,[],[f4982_D])).
fof(f4982_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f4984,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)) | $spl12),
  inference(cnf_transformation,[],[f4984_D])).
fof(f4984_D,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f4507,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl0 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f4452,f139])).
fof(f4452,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl0 | $spl2 | $spl5)),
  inference(superposition,[],[f190,f4407])).
fof(f190,plain,(
  ( ! [X0,X1] : (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(X0,bool),X1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f129])).
fof(f3209,plain,(
  $false | ($spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3208,f144])).
fof(f3208,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl5 | $spl6)),
  inference(forward_demodulation,[],[f369,f371])).
fof(f3207,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f372,f371])).
fof(f3206,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3205,f143])).
fof(f3205,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3204,f371])).
fof(f3204,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3203,f144])).
fof(f3203,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3008,f371])).
fof(f3008,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl5)),
  inference(trivial_inequality_removal,[],[f373])).
fof(f373,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.EMPTY_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f372,f190])).
fof(f3202,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3201,f599])).
fof(f599,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f598,f369])).
fof(f598,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f597,f379])).
fof(f597,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f596,f139])).
fof(f596,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f554,f225])).
fof(f554,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f191,f379])).
fof(f3201,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3009,f371])).
fof(f3009,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(trivial_inequality_removal,[],[f374])).
fof(f3199,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3198,f599])).
fof(f3198,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f379,f371])).
fof(f3197,plain,(
  $false | ($spl0 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f369,f387,f218,f1352,f1353,f1383,f1366,f1367,f1449,f1450,f1786,f1787,f1791,f1792,f1796,f1797,f1798,f1833,f1835,f1843,f1785,f1799,f2794,f2795,f2796,f2798,f2800,f2802,f2804,f2805,f2806,f2808,f2810,f2812,f1364,f1365,f1447,f1448,f3040,f1975,f1804,f3042,f1795,f1825,f3044,f3046,f3048,f1446,f1363,f371,f3049,f3062,f3065,f3068,f3070,f3071,f3073,f3074,f3075,f3076,f3077,f3078,f3079,f3080,f3098,f1850,f1849,f1834,f3184,f385,f384,f3196])).
fof(f3196,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)))) | ($spl0 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f1835,f371])).
fof(f385,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1)))) ) | $spl5),
  inference(superposition,[],[f369,f207])).
fof(f3184,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | ($spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3183,f205])).
fof(f3183,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) ) | ($spl5 | $spl6)),
  inference(forward_demodulation,[],[f3182,f144])).
fof(f3182,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) ) | ($spl5 | $spl6)),
  inference(forward_demodulation,[],[f387,f371])).
fof(f1850,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X1)) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f1833,f207])).
fof(f3098,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl0 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f1852,f371])).
fof(f1852,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl0 | $spl5)),
  inference(trivial_inequality_removal,[],[f1847])).
fof(f1847,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.EMPTY_0') | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl0 | $spl5)),
  inference(superposition,[],[f1833,f188])).
fof(f3080,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X0))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X0))))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f1785,f371])).
fof(f3079,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),X1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(X1,sK0)),s(fun(sK0,bool),sK3(X1,sK0))))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f1786,f371])).
fof(f3078,plain,(
  ( ! [X2] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X2))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(X2,sK0)),s(fun(sK0,bool),sK3(X2,sK0)))))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f1787,f371])).
fof(f3077,plain,(
  ( ! [X3] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X3))) | s(bool,f(s(sK0,sK5('const.pred_set.EMPTY_0',X3,bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.EMPTY_0',X3,bool,sK0))))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f1791,f371])).
fof(f3076,plain,(
  ( ! [X4] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X4))) | s(bool,f(s(sK0,sK5(X4,'const.pred_set.EMPTY_0',bool,sK0)))) != s(bool,g(s(sK0,sK5(X4,'const.pred_set.EMPTY_0',bool,sK0))))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f1792,f371])).
fof(f3075,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),X0) | s(bool,f(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),X0,bool,sK0))))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f1796,f371])).
fof(f3074,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),X1) | s(bool,f(s(sK0,sK5(X1,'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0))))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f1797,f371])).
fof(f3073,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X0) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X0))))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f1798,f371])).
fof(f3071,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f1843,f371])).
fof(f3070,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f3056,f144])).
fof(f3056,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))))) ) | ($spl0 | $spl6)),
  inference(backward_demodulation,[],[f371,f3040])).
fof(f3068,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3067,f3049])).
fof(f3067,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3055,f143])).
fof(f3055,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl0 | $spl6)),
  inference(backward_demodulation,[],[f371,f1804])).
fof(f3065,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3064,f143])).
fof(f3064,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s(bool,'T_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f3063,f371])).
fof(f3063,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3053,f144])).
fof(f3053,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl0 | $spl6)),
  inference(backward_demodulation,[],[f371,f1446])).
fof(f3062,plain,(
  p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3061,f143])).
fof(f3061,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f3060,f371])).
fof(f3060,plain,(
  p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3051,f144])).
fof(f3051,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl0 | $spl6)),
  inference(backward_demodulation,[],[f371,f1363])).
fof(f3049,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f3042,f371])).
fof(f1363,plain,(
  p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | $spl0),
  inference(superposition,[],[f218,f190])).
fof(f1446,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | $spl0),
  inference(superposition,[],[f1353,f190])).
fof(f3048,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl0),
  inference(subsumption_resolution,[],[f3047,f143])).
fof(f3047,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl0),
  inference(forward_demodulation,[],[f1802,f200])).
fof(f1802,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | $spl0),
  inference(superposition,[],[f190,f1352])).
fof(f3046,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | $spl0),
  inference(subsumption_resolution,[],[f3045,f143])).
fof(f3045,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | $spl0),
  inference(forward_demodulation,[],[f1828,f200])).
fof(f1828,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | $spl0),
  inference(forward_demodulation,[],[f1803,f1352])).
fof(f1803,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | $spl0),
  inference(superposition,[],[f190,f1352])).
fof(f3044,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl0),
  inference(subsumption_resolution,[],[f3043,f143])).
fof(f3043,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl0),
  inference(forward_demodulation,[],[f1805,f200])).
fof(f1825,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | $spl0),
  inference(forward_demodulation,[],[f1794,f1352])).
fof(f1794,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)))) | $spl0),
  inference(superposition,[],[f1352,f188])).
fof(f3042,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl0),
  inference(subsumption_resolution,[],[f3041,f143])).
fof(f3041,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl0),
  inference(forward_demodulation,[],[f1806,f200])).
fof(f1806,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | $spl0),
  inference(superposition,[],[f191,f1352])).
fof(f1804,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl0),
  inference(superposition,[],[f190,f1352])).
fof(f1975,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl0),
  inference(forward_demodulation,[],[f1807,f1352])).
fof(f1807,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl0),
  inference(superposition,[],[f191,f1352])).
fof(f3040,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f3039,f143])).
fof(f3039,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1)))))))) ) | $spl0),
  inference(forward_demodulation,[],[f1808,f200])).
fof(f1808,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) ) | $spl0),
  inference(superposition,[],[f192,f1352])).
fof(f1448,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))))) | $spl0),
  inference(superposition,[],[f1353,f189])).
fof(f1447,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | $spl0),
  inference(superposition,[],[f1353,f188])).
fof(f1365,plain,(
  p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))))) | $spl0),
  inference(superposition,[],[f218,f189])).
fof(f1364,plain,(
  p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | $spl0),
  inference(superposition,[],[f218,f188])).
fof(f2812,plain,(
  ( ! [X10,X11] : (s(bool,X11) != s(bool,f(s(sK0,sK5(X10,sK1,bool,sK0)))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X10)))) | ~p(s(bool,g(s(sK0,sK5(X10,sK1,bool,sK0))))) | ~p(s(bool,X11))) ) | $spl0),
  inference(superposition,[],[f1367,f204])).
fof(f2810,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,f(s(sK0,sK5(X6,sK1,bool,sK0)))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X6)))) | p(s(bool,X7)) | p(s(bool,g(s(sK0,sK5(X6,sK1,bool,sK0)))))) ) | $spl0),
  inference(superposition,[],[f1367,f201])).
fof(f2808,plain,(
  ( ! [X4] : (p(s(bool,g(s(sK0,sK5(X4,sK1,bool,sK0))))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X4)))) | s(bool,'F_0') != s(bool,f(s(sK0,sK5(X4,sK1,bool,sK0))))) ) | $spl0),
  inference(superposition,[],[f1367,f154])).
fof(f2806,plain,(
  ( ! [X2] : (~p(s(bool,g(s(sK0,sK5(X2,sK1,bool,sK0))))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X2)))) | s(bool,'T_0') != s(bool,f(s(sK0,sK5(X2,sK1,bool,sK0))))) ) | $spl0),
  inference(superposition,[],[f1367,f150])).
fof(f2805,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X1)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0))))) ) | $spl0),
  inference(superposition,[],[f1367,f145])).
fof(f145,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.SING_IFF_CARD|49|',aHL_BOOL_CASES)).
fof(f2804,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(X0,sK1,bool,sK0)))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK5(X0,sK1,bool,sK0))))) ) | $spl0),
  inference(superposition,[],[f1367,f145])).
fof(f2802,plain,(
  ( ! [X10,X11] : (s(bool,X11) != s(bool,f(s(sK0,sK5(sK1,X10,bool,sK0)))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X10)))) | ~p(s(bool,g(s(sK0,sK5(sK1,X10,bool,sK0))))) | ~p(s(bool,X11))) ) | $spl0),
  inference(superposition,[],[f1366,f204])).
fof(f2800,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,f(s(sK0,sK5(sK1,X6,bool,sK0)))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X6)))) | p(s(bool,X7)) | p(s(bool,g(s(sK0,sK5(sK1,X6,bool,sK0)))))) ) | $spl0),
  inference(superposition,[],[f1366,f201])).
fof(f2798,plain,(
  ( ! [X4] : (p(s(bool,g(s(sK0,sK5(sK1,X4,bool,sK0))))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X4)))) | s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK1,X4,bool,sK0))))) ) | $spl0),
  inference(superposition,[],[f1366,f154])).
fof(f2796,plain,(
  ( ! [X2] : (~p(s(bool,g(s(sK0,sK5(sK1,X2,bool,sK0))))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X2)))) | s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,X2,bool,sK0))))) ) | $spl0),
  inference(superposition,[],[f1366,f150])).
fof(f2795,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK1,X1,bool,sK0)))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X1)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK1,X1,bool,sK0))))) ) | $spl0),
  inference(superposition,[],[f1366,f145])).
fof(f2794,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,X0,bool,sK0)))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK5(sK1,X0,bool,sK0))))) ) | $spl0),
  inference(superposition,[],[f1366,f145])).
fof(f1785,plain,(
  ( ! [X0] : (s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X0))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X0))))) ) | $spl0),
  inference(superposition,[],[f1352,f190])).
fof(f1835,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)))) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f1825,f1833])).
fof(f1798,plain,(
  ( ! [X0] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X0) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X0))))) ) | $spl0),
  inference(superposition,[],[f196,f1352])).
fof(f196,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SING_1'(s(fun(X0,bool),X1)))) | s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(cnf_transformation,[],[f133])).
fof(f1797,plain,(
  ( ! [X1] : (s(fun(sK0,bool),sK1) = s(fun(sK0,bool),X1) | s(bool,f(s(sK0,sK5(X1,'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0))))) ) | $spl0),
  inference(superposition,[],[f1352,f207])).
fof(f1796,plain,(
  ( ! [X0] : (s(fun(sK0,bool),sK1) = s(fun(sK0,bool),X0) | s(bool,f(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),X0,bool,sK0))))) ) | $spl0),
  inference(superposition,[],[f1352,f207])).
fof(f1792,plain,(
  ( ! [X4] : (s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X4))) | s(bool,f(s(sK0,sK5(X4,'const.pred_set.EMPTY_0',bool,sK0)))) != s(bool,g(s(sK0,sK5(X4,'const.pred_set.EMPTY_0',bool,sK0))))) ) | $spl0),
  inference(superposition,[],[f1352,f207])).
fof(f1791,plain,(
  ( ! [X3] : (s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X3))) | s(bool,f(s(sK0,sK5('const.pred_set.EMPTY_0',X3,bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.EMPTY_0',X3,bool,sK0))))) ) | $spl0),
  inference(superposition,[],[f1352,f207])).
fof(f1787,plain,(
  ( ! [X2] : (s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X2))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(X2,sK0)),s(fun(sK0,bool),sK3(X2,sK0)))))) ) | $spl0),
  inference(superposition,[],[f1352,f189])).
fof(f1786,plain,(
  ( ! [X1] : (s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),X1))) | s(fun(sK0,bool),X1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(X1,sK0)),s(fun(sK0,bool),sK3(X1,sK0))))) ) | $spl0),
  inference(superposition,[],[f1352,f188])).
fof(f1450,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X1)))) ) | $spl0),
  inference(superposition,[],[f1353,f207])).
fof(f1449,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK5(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,X0,bool,sK0)))) | s(bool,'T_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X0)))) ) | $spl0),
  inference(superposition,[],[f1353,f207])).
fof(f1367,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0)))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X1))))) ) | $spl0),
  inference(superposition,[],[f218,f207])).
fof(f1366,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK5(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,X0,bool,sK0)))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X0))))) ) | $spl0),
  inference(superposition,[],[f218,f207])).
fof(f3195,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f387,f393,f428,f452,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f603,f663,f607,f608,f675,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110,f3114,f3118,f3120,f3122,f3126,f3128,f3130,f3132,f3134,f3038,f677,f676,f3139,f456,f455,f454,f453,f762,f757,f756,f3022,f432,f431,f430,f429,f3021,f610,f609,f3148,f750,f749,f748,f599,f747,f404,f403,f3166,f3170,f3174,f584,f391,f583,f3184,f385,f384,f3194])).
fof(f3194,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3193,f762])).
fof(f3193,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f581,f371])).
fof(f583,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f579,f139])).
fof(f579,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(backward_demodulation,[],[f576,f452])).
fof(f584,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f580,f139])).
fof(f580,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(backward_demodulation,[],[f576,f428])).
fof(f3174,plain,(
  s(fun(sK0,bool),sK3(sK1,sK0)) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3173,f187])).
fof(f3173,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s(fun(sK0,bool),sK3(sK1,sK0)) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3172,f144])).
fof(f3172,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),sK3(sK1,sK0)) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f736,f371])).
fof(f3170,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3169,f187])).
fof(f3169,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3168,f144])).
fof(f3168,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f737,f371])).
fof(f3166,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK3(sK1,sK0)) | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3165,f187])).
fof(f3165,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK3(sK1,sK0)) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3164,f144])).
fof(f3164,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK3(sK1,sK0)) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f739,f371])).
fof(f747,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),sK3(sK1,sK0)) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f599,f188])).
fof(f748,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f599,f189])).
fof(f749,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f599,f207])).
fof(f750,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f599,f207])).
fof(f3148,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3147,f139])).
fof(f3147,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f3146,f576])).
fof(f3146,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f3020,f200])).
fof(f3020,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))))))) | ($spl2 | $spl5)),
  inference(trivial_inequality_removal,[],[f753])).
fof(f753,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.EMPTY_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f599,f190])).
fof(f609,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X0))) | s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f576,f207])).
fof(f610,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))) | s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f576,f207])).
fof(f3021,plain,(
  s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)))) | ($spl2 | $spl5)),
  inference(trivial_inequality_removal,[],[f754])).
fof(f754,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.EMPTY_0') | s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f599,f188])).
fof(f429,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),sK3(sK1,sK0)) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f391,f188])).
fof(f430,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f391,f189])).
fof(f3022,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(trivial_inequality_removal,[],[f433])).
fof(f433,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f391,f145])).
fof(f756,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X0) | s(bool,f(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),X0,bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f599,f207])).
fof(f757,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X1) | s(bool,f(s(sK0,sK5(X1,'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f599,f207])).
fof(f762,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f761,f139])).
fof(f761,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f746,f576])).
fof(f746,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f599,f190])).
fof(f453,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),sK3(sK1,sK0)) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f393,f188])).
fof(f454,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f393,f189])).
fof(f455,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f393,f207])).
fof(f456,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f393,f207])).
fof(f3139,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK1,sK0)))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK1,sK0)))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))),s('type.num.num','const.num.|48|_0')))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3138,f144])).
fof(f3138,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK1,sK0)))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK3(sK1,sK0)))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f675,f371])).
fof(f676,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK3(sK1,sK0)) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f619,f139])).
fof(f619,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK3(sK1,sK0)) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f190,f576])).
fof(f677,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK3(sK1,sK0)) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f620,f139])).
fof(f620,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK3(sK1,sK0)) | ($spl2 | $spl5)),
  inference(superposition,[],[f191,f576])).
fof(f3038,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))))) | ($spl2 | $spl5)),
  inference(condensation,[],[f533])).
fof(f533,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))))) ) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f495,f153])).
fof(f495,plain,(
  ( ! [X5] : (s(bool,'F_0') = s(bool,X5) | ~p(s(bool,X5)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f204,f393])).
fof(f3134,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),sK3(sK1,sK0)) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f537,f371])).
fof(f3132,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f538,f371])).
fof(f3130,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f539,f371])).
fof(f3128,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f540,f371])).
fof(f3126,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3125,f187])).
fof(f3125,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X0))))) | s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3124,f144])).
fof(f3124,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X0))))) | s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f701,f371])).
fof(f3122,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),X0) | s(bool,f(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),X0,bool,sK0))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f544,f371])).
fof(f3120,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),X1) | s(bool,f(s(sK0,sK5(X1,'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),bool,sK0))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f545,f371])).
fof(f3118,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3117,f187])).
fof(f3117,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))) | s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3116,f144])).
fof(f3116,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))))) | s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f702,f371])).
fof(f3114,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3113,f187])).
fof(f3113,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3112,f144])).
fof(f3112,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f704,f371])).
fof(f3110,plain,(
  ( ! [X0] : (s('type.num.num',X0) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3109,f187])).
fof(f3109,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | s('type.num.num',X0) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3108,f144])).
fof(f3108,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num',X0) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f706,f371])).
fof(f3104,plain,(
  ( ! [X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) | s('type.num.num',X5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3103,f144])).
fof(f3103,plain,(
  ( ! [X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s('type.num.num',X5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f717,f371])).
fof(f3101,plain,(
  ( ! [X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) | s('type.num.num',X9) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3100,f144])).
fof(f3100,plain,(
  ( ! [X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s('type.num.num',X9) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))) ) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f721,f371])).
fof(f3095,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f574,f371])).
fof(f3088,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f225,f371])).
fof(f2792,plain,(
  ( ! [X10,X11] : (s(bool,X11) != s(bool,f(s(sK0,sK5(X10,sK1,bool,sK0)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X10)))) | ~p(s(bool,g(s(sK0,sK5(X10,sK1,bool,sK0))))) | ~p(s(bool,X11))) ) | $spl2),
  inference(superposition,[],[f239,f204])).
fof(f2790,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,f(s(sK0,sK5(X6,sK1,bool,sK0)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X6)))) | p(s(bool,X7)) | p(s(bool,g(s(sK0,sK5(X6,sK1,bool,sK0)))))) ) | $spl2),
  inference(superposition,[],[f239,f201])).
fof(f2788,plain,(
  ( ! [X4] : (p(s(bool,g(s(sK0,sK5(X4,sK1,bool,sK0))))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X4)))) | s(bool,'F_0') != s(bool,f(s(sK0,sK5(X4,sK1,bool,sK0))))) ) | $spl2),
  inference(superposition,[],[f239,f154])).
fof(f2786,plain,(
  ( ! [X2] : (~p(s(bool,g(s(sK0,sK5(X2,sK1,bool,sK0))))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X2)))) | s(bool,'T_0') != s(bool,f(s(sK0,sK5(X2,sK1,bool,sK0))))) ) | $spl2),
  inference(superposition,[],[f239,f150])).
fof(f2785,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0))))) ) | $spl2),
  inference(superposition,[],[f239,f145])).
fof(f2784,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(X0,sK1,bool,sK0)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK5(X0,sK1,bool,sK0))))) ) | $spl2),
  inference(superposition,[],[f239,f145])).
fof(f238,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK5(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,X0,bool,sK0)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X0))))) ) | $spl2),
  inference(superposition,[],[f220,f207])).
fof(f239,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1))))) ) | $spl2),
  inference(superposition,[],[f220,f207])).
fof(f283,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK5(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,X0,bool,sK0)))) | s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X0)))) ) | $spl2),
  inference(superposition,[],[f225,f207])).
fof(f284,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1)))) ) | $spl2),
  inference(superposition,[],[f225,f207])).
fof(f1137,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,X0,bool,sK0)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X0)))) | s(bool,'F_0') = s(bool,g(s(sK0,sK5(sK1,X0,bool,sK0))))) ) | $spl2),
  inference(superposition,[],[f238,f145])).
fof(f1138,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK1,X1,bool,sK0)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X1)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK5(sK1,X1,bool,sK0))))) ) | $spl2),
  inference(superposition,[],[f238,f145])).
fof(f1139,plain,(
  ( ! [X2] : (~p(s(bool,g(s(sK0,sK5(sK1,X2,bool,sK0))))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X2)))) | s(bool,'T_0') != s(bool,f(s(sK0,sK5(sK1,X2,bool,sK0))))) ) | $spl2),
  inference(superposition,[],[f238,f150])).
fof(f1141,plain,(
  ( ! [X4] : (p(s(bool,g(s(sK0,sK5(sK1,X4,bool,sK0))))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X4)))) | s(bool,'F_0') != s(bool,f(s(sK0,sK5(sK1,X4,bool,sK0))))) ) | $spl2),
  inference(superposition,[],[f238,f154])).
fof(f1143,plain,(
  ( ! [X6,X7] : (s(bool,X7) != s(bool,f(s(sK0,sK5(sK1,X6,bool,sK0)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X6)))) | p(s(bool,X7)) | p(s(bool,g(s(sK0,sK5(sK1,X6,bool,sK0)))))) ) | $spl2),
  inference(superposition,[],[f238,f201])).
fof(f1145,plain,(
  ( ! [X10,X11] : (s(bool,X11) != s(bool,f(s(sK0,sK5(sK1,X10,bool,sK0)))) | p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X10)))) | ~p(s(bool,g(s(sK0,sK5(sK1,X10,bool,sK0))))) | ~p(s(bool,X11))) ) | $spl2),
  inference(superposition,[],[f238,f204])).
fof(f717,plain,(
  ( ! [X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))))) | s('type.num.num',X5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f205,f582])).
fof(f711,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s('type.num.num',X1) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f205,f582])).
fof(f704,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f582,f205])).
fof(f701,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X0))))) | s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f582,f207])).
fof(f735,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f734,f139])).
fof(f734,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f733,f576])).
fof(f733,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f698,f144])).
fof(f698,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f582,f190])).
fof(f608,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f576,f189])).
fof(f607,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),sK3(sK1,sK0)) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f576,f188])).
fof(f663,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f662,f139])).
fof(f662,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f606,f576])).
fof(f606,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f576,f190])).
fof(f581,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f577,f139])).
fof(f577,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5)),
  inference(backward_demodulation,[],[f576,f536])).
fof(f544,plain,(
  ( ! [X0] : (s(fun(sK0,bool),sK1) = s(fun(sK0,bool),X0) | s(bool,f(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),X0,bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f379,f207])).
fof(f574,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f573,f379])).
fof(f573,plain,(
  s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f542,f372])).
fof(f542,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f379,f188])).
fof(f540,plain,(
  ( ! [X1] : (s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f379,f207])).
fof(f539,plain,(
  ( ! [X0] : (s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f379,f207])).
fof(f538,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f379,f189])).
fof(f537,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),sK3(sK1,sK0)) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f379,f188])).
fof(f536,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f379,f190])).
fof(f452,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f393,f190])).
fof(f428,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | ($spl2 | $spl5)),
  inference(superposition,[],[f391,f190])).
fof(f377,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X1) | s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f372,f207])).
fof(f376,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X0) | s(bool,f(s(sK0,sK5(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,X0,bool,sK0))))) ) | ($spl2 | $spl5)),
  inference(superposition,[],[f372,f207])).
fof(f282,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))))) | $spl2),
  inference(superposition,[],[f225,f189])).
fof(f281,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | $spl2),
  inference(superposition,[],[f225,f188])).
fof(f337,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | $spl2),
  inference(subsumption_resolution,[],[f336,f139])).
fof(f336,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | $spl2),
  inference(forward_demodulation,[],[f280,f225])).
fof(f280,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | $spl2),
  inference(superposition,[],[f225,f190])).
fof(f255,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl2),
  inference(forward_demodulation,[],[f227,f225])).
fof(f227,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))) | $spl2),
  inference(resolution,[],[f220,f153])).
fof(f3192,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3191,f599])).
fof(f3191,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl5 | $spl6)),
  inference(forward_demodulation,[],[f3190,f371])).
fof(f3190,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | $spl5),
  inference(subsumption_resolution,[],[f382,f144])).
fof(f382,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | $spl5),
  inference(superposition,[],[f369,f188])).
fof(f3187,plain,(
  $false | ($spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3186,f143])).
fof(f3186,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl5 | $spl6)),
  inference(forward_demodulation,[],[f3185,f371])).
fof(f3185,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3011,f371])).
fof(f3011,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | $spl5),
  inference(trivial_inequality_removal,[],[f386])).
fof(f386,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | $spl5),
  inference(superposition,[],[f369,f191])).
fof(f3181,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3180,f187])).
fof(f3180,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3179,f144])).
fof(f3179,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f582,f371])).
fof(f3178,plain,(
  $false | ($spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3177,f143])).
fof(f3177,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl5 | $spl6)),
  inference(forward_demodulation,[],[f3176,f371])).
fof(f3176,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl5 | $spl6)),
  inference(subsumption_resolution,[],[f389,f371])).
fof(f389,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | $spl5),
  inference(trivial_inequality_removal,[],[f386])).
fof(f3175,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f428,f452,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f603,f663,f607,f608,f675,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110,f3114,f3118,f3120,f3122,f3126,f3128,f3130,f3132,f3134,f3038,f677,f676,f3139,f456,f455,f454,f453,f762,f757,f756,f3022,f432,f431,f430,f429,f3021,f610,f609,f3148,f750,f749,f748,f599,f747,f404,f403,f3166,f3170,f3174])).
fof(f3171,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f428,f452,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f603,f663,f607,f608,f675,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110,f3114,f3118,f3120,f3122,f3126,f3128,f3130,f3132,f3134,f3038,f677,f676,f3139,f456,f455,f454,f453,f762,f757,f756,f3022,f432,f431,f430,f429,f3021,f610,f609,f3148,f750,f749,f748,f599,f747,f404,f403,f3166,f3170])).
fof(f3167,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f428,f452,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f603,f663,f607,f608,f675,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110,f3114,f3118,f3120,f3122,f3126,f3128,f3130,f3132,f3134,f3038,f677,f676,f3139,f456,f455,f454,f453,f762,f757,f756,f3022,f432,f431,f430,f429,f3021,f610,f609,f3148,f750,f749,f748,f599,f747,f404,f403,f3166])).
fof(f3163,plain,(
  $false | ($spl0 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3162,f143])).
fof(f3162,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3161,f371])).
fof(f3161,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl0 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3160,f144])).
fof(f3160,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl0 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3012,f371])).
fof(f3012,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl0 | $spl5)),
  inference(trivial_inequality_removal,[],[f1846])).
fof(f1846,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.EMPTY_0') | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl0 | $spl5)),
  inference(superposition,[],[f1833,f190])).
fof(f3159,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3158,f599])).
fof(f3158,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl0 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3013,f371])).
fof(f3013,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl0 | $spl5)),
  inference(trivial_inequality_removal,[],[f1847])).
fof(f3153,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f3152,f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f403,f404,f428,f452,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f599,f603,f663,f607,f608,f675,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f747,f748,f749,f750,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110,f3114,f3118,f3120,f3122,f3126,f3128,f3130,f3132,f3134,f3038,f677,f676,f3139,f456,f455,f454,f453,f762,f757,f756,f3022,f432,f431,f430,f429,f3021,f610,f609,f3148])).
fof(f3152,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f3151,f139])).
fof(f3151,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3150,f3065])).
fof(f3150,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f3149,f371])).
fof(f3149,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f603,f371])).
fof(f3140,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f403,f404,f428,f429,f430,f431,f432,f452,f453,f454,f455,f456,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f599,f603,f663,f607,f608,f609,f610,f675,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f762,f747,f748,f749,f750,f756,f757,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110,f3114,f3118,f3120,f3122,f3126,f3128,f3130,f3132,f3134,f3038,f677,f676,f3139])).
fof(f3135,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f403,f404,f428,f429,f430,f431,f432,f452,f453,f454,f455,f456,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f599,f603,f663,f607,f608,f609,f610,f675,f676,f677,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f762,f747,f748,f749,f750,f756,f757,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110,f3114,f3118,f3120,f3122,f3126,f3128,f3130,f3132,f3134])).
fof(f3133,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f403,f404,f428,f429,f430,f431,f432,f452,f453,f454,f455,f456,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f599,f603,f663,f607,f608,f609,f610,f675,f676,f677,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f762,f747,f748,f749,f750,f756,f757,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110,f3114,f3118,f3120,f3122,f3126,f3128,f3130,f3132])).
fof(f3131,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f403,f404,f428,f429,f430,f431,f432,f452,f453,f454,f455,f456,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f599,f603,f663,f607,f608,f609,f610,f675,f676,f677,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f762,f747,f748,f749,f750,f756,f757,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110,f3114,f3118,f3120,f3122,f3126,f3128,f3130])).
fof(f3129,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f403,f404,f428,f429,f430,f431,f432,f452,f453,f454,f455,f456,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f599,f603,f663,f607,f608,f609,f610,f675,f676,f677,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f762,f747,f748,f749,f750,f756,f757,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110,f3114,f3118,f3120,f3122,f3126,f3128])).
fof(f3127,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f403,f404,f428,f429,f430,f431,f432,f452,f453,f454,f455,f456,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f599,f603,f663,f607,f608,f609,f610,f675,f676,f677,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f762,f747,f748,f749,f750,f756,f757,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110,f3114,f3118,f3120,f3122,f3126])).
fof(f3123,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f403,f404,f428,f429,f430,f431,f432,f452,f453,f454,f455,f456,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f599,f603,f663,f607,f608,f609,f610,f675,f676,f677,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f762,f747,f748,f749,f750,f756,f757,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110,f3114,f3118,f3120,f3122])).
fof(f3121,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f403,f404,f428,f429,f430,f431,f432,f452,f453,f454,f455,f456,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f599,f603,f663,f607,f608,f609,f610,f675,f676,f677,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f762,f747,f748,f749,f750,f756,f757,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110,f3114,f3118,f3120])).
fof(f3119,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f403,f404,f428,f429,f430,f431,f432,f452,f453,f454,f455,f456,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f599,f603,f663,f607,f608,f609,f610,f675,f676,f677,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f762,f747,f748,f749,f750,f756,f757,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110,f3114,f3118])).
fof(f3115,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f403,f404,f428,f429,f430,f431,f432,f452,f453,f454,f455,f456,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f599,f603,f663,f607,f608,f609,f610,f675,f676,f677,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f762,f747,f748,f749,f750,f756,f757,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110,f3114])).
fof(f3111,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f403,f404,f428,f429,f430,f431,f432,f452,f453,f454,f455,f456,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f599,f603,f663,f607,f608,f609,f610,f675,f676,f677,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f762,f747,f748,f749,f750,f756,f757,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104,f3110])).
fof(f3105,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f403,f404,f428,f429,f430,f431,f432,f452,f453,f454,f455,f456,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f599,f603,f663,f607,f608,f609,f610,f675,f676,f677,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f762,f747,f748,f749,f750,f756,f757,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101,f3104])).
fof(f3102,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f403,f404,f428,f429,f430,f431,f432,f452,f453,f454,f455,f456,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f599,f603,f663,f607,f608,f609,f610,f675,f676,f677,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f762,f747,f748,f749,f750,f756,f757,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095,f3101])).
fof(f3099,plain,(
  $false | ($spl0 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f369,f384,f385,f387,f218,f1352,f1353,f1383,f1366,f1367,f1449,f1450,f1786,f1787,f1791,f1792,f1796,f1797,f1798,f1833,f1835,f1834,f1843,f1849,f1850,f1785,f1799,f2794,f2795,f2796,f2798,f2800,f2802,f2804,f2805,f2806,f2808,f2810,f2812,f1364,f1365,f1447,f1448,f3040,f1975,f1804,f3042,f1795,f1825,f3044,f3046,f3048,f1446,f1363,f371,f3049,f3062,f3065,f3068,f3070,f3071,f3073,f3074,f3075,f3076,f3077,f3078,f3079,f3080,f3098])).
fof(f3097,plain,(
  $false | ($spl0 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f1833,f371])).
fof(f3096,plain,(
  $false | ($spl2 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f369,f372,f379,f378,f376,f377,f384,f385,f387,f391,f393,f403,f404,f428,f429,f430,f431,f432,f452,f453,f454,f455,f456,f519,f536,f537,f538,f539,f540,f574,f544,f545,f581,f582,f583,f584,f599,f603,f663,f607,f608,f609,f610,f675,f676,f677,f735,f736,f737,f701,f702,f739,f704,f706,f711,f717,f721,f762,f747,f748,f749,f750,f756,f757,f1145,f1143,f1141,f1139,f1138,f1137,f284,f283,f239,f238,f2784,f2785,f2786,f2788,f2790,f2792,f371,f3088,f576,f575,f3095])).
fof(f2999,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f2216,f2995])).
fof(f2995,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f2994,f143])).
fof(f2994,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f2993,f200])).
fof(f2993,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f2958,f187])).
fof(f2958,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) ) | ($spl2 | $spl4)),
  inference(superposition,[],[f2910,f191])).
fof(f2910,plain,(
  ( ! [X4] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X4),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f2909,f199])).
fof(f2909,plain,(
  ( ! [X4] : (s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X4),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f2876,f197])).
fof(f2876,plain,(
  ( ! [X4] : (s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X4),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | p(s(bool,'const.bool.IN_2'(s(sK0,X4),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl2 | $spl4)),
  inference(superposition,[],[f2095,f154])).
fof(f2095,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl2 | $spl4)),
  inference(backward_demodulation,[],[f2083,f2081])).
fof(f2081,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl2 | $spl4)),
  inference(forward_demodulation,[],[f222,f368])).
fof(f368,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | $spl4),
  inference(cnf_transformation,[],[f368_D])).
fof(f368_D,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f2083,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f223,f368])).
fof(f2216,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f2163,f139])).
fof(f2163,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl4)),
  inference(superposition,[],[f195,f2085])).
fof(f2085,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl4)),
  inference(backward_demodulation,[],[f2083,f1353])).
fof(f2998,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f2455,f2995])).
fof(f2455,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f2454,f2216])).
fof(f2454,plain,(
  s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f2453,f139])).
fof(f2453,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f2430,f2085])).
fof(f2430,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl4)),
  inference(superposition,[],[f195,f2216])).
fof(f2997,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f2097,f2995])).
fof(f2097,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f1843,f2083])).
fof(f2996,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f2096,f2995])).
fof(f2096,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f1352,f2083])).
fof(f2985,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f2984,f187])).
fof(f2984,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f2951,f144])).
fof(f2951,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl4)),
  inference(superposition,[],[f2910,f2216])).
fof(f2983,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f2982,f187])).
fof(f2982,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f2950,f144])).
fof(f2950,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl4)),
  inference(superposition,[],[f2910,f2096])).
fof(f2080,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl8)),
  inference(subsumption_resolution,[],[f2079,f368])).
fof(f2079,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f223,f1987])).
fof(f1987,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f1986,f143])).
fof(f1986,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f1348,f1799])).
fof(f1348,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | $spl8),
  inference(subsumption_resolution,[],[f959,f187])).
fof(f959,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | $spl8),
  inference(superposition,[],[f191,f766])).
fof(f2078,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl8)),
  inference(global_subsumption,[],[f2077,f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f337,f281,f282,f766,f767,f768,f781,f783,f793,f794,f826,f827,f951,f952,f1000,f1348,f957,f978,f1350,f1351,f1639,f1641,f1651,f1652,f1677,f1678,f1679,f1680,f1701,f1702,f1703,f1704,f1768,f1700,f1676,f368,f779,f1145,f1143,f1141,f2055,f1139,f1138,f1018,f1017,f1137,f1015,f1010,f1009,f1005,f1004,f284,f283,f999,f998,f994,f993,f992,f239,f238,f852,f851])).
fof(f851,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X0))) | s(bool,f(s(sK0,sK5(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,X0,bool,sK0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f783,f207])).
fof(f852,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X1))) | s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f783,f207])).
fof(f992,plain,(
  ( ! [X0,X1] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),X0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X0))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),X0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f779,f190])).
fof(f993,plain,(
  ( ! [X2,X3] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X3),s(fun(sK0,bool),X2))) | s(fun(sK0,bool),X2) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(X2,sK0)),s(fun(sK0,bool),sK3(X2,sK0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f779,f188])).
fof(f994,plain,(
  ( ! [X4,X5] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X5),s(fun(sK0,bool),X4))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(X4,sK0)),s(fun(sK0,bool),sK3(X4,sK0)))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f779,f189])).
fof(f998,plain,(
  ( ! [X10,X9] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X10),s(fun(sK0,bool),X9))) | s(bool,f(s(sK0,sK5('const.pred_set.EMPTY_0',X9,bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.EMPTY_0',X9,bool,sK0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f779,f207])).
fof(f999,plain,(
  ( ! [X12,X11] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X12),s(fun(sK0,bool),X11))) | s(bool,f(s(sK0,sK5(X11,'const.pred_set.EMPTY_0',bool,sK0)))) != s(bool,g(s(sK0,sK5(X11,'const.pred_set.EMPTY_0',bool,sK0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f779,f207])).
fof(f1004,plain,(
  ( ! [X4,X5] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X5) | s(bool,f(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,X4),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),X5,bool,sK0)))) != s(bool,g(s(sK0,sK5('const.pred_set.INSERT_2'(s(sK0,X4),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),X5,bool,sK0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f779,f207])).
fof(f1005,plain,(
  ( ! [X6,X7] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X7) | s(bool,f(s(sK0,sK5(X7,'const.pred_set.INSERT_2'(s(sK0,X6),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0)))) != s(bool,g(s(sK0,sK5(X7,'const.pred_set.INSERT_2'(s(sK0,X6),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),bool,sK0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f779,f207])).
fof(f1009,plain,(
  ( ! [X4,X3] : (s(fun(sK0,bool),X3) != s(fun(sK0,bool),X4) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X4)))) | s(bool,f(s(sK0,sK5(sK1,X3,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,X3,bool,sK0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f1000,f207])).
fof(f1010,plain,(
  ( ! [X6,X5] : (s(fun(sK0,bool),X5) != s(fun(sK0,bool),X6) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X6)))) | s(bool,f(s(sK0,sK5(X5,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(X5,sK1,bool,sK0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f1000,f207])).
fof(f1015,plain,(
  ( ! [X6] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X6) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(X6,sK0)),s(fun(sK0,bool),sK3(X6,sK0))))))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),X6)) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f1000,f188])).
fof(f1017,plain,(
  ( ! [X8,X9] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X9) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X8)))) | s(bool,f(s(sK0,sK5(X8,X9,bool,sK0)))) != s(bool,g(s(sK0,sK5(X8,X9,bool,sK0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f1000,f207])).
fof(f1018,plain,(
  ( ! [X10,X11] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X11) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X10)))) | s(bool,f(s(sK0,sK5(X11,X10,bool,sK0)))) != s(bool,g(s(sK0,sK5(X11,X10,bool,sK0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f1000,f207])).
fof(f2055,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f337,f368])).
fof(f779,plain,(
  ( ! [X0] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl2 | $spl8)),
  inference(resolution,[],[f767,f196])).
fof(f1676,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | $spl8),
  inference(superposition,[],[f1639,f190])).
fof(f1700,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | $spl8),
  inference(superposition,[],[f1641,f190])).
fof(f1768,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | $spl8),
  inference(forward_demodulation,[],[f1767,f766])).
fof(f1767,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | $spl8),
  inference(forward_demodulation,[],[f1766,f1350])).
fof(f1766,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | $spl8),
  inference(forward_demodulation,[],[f1714,f199])).
fof(f1714,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK3(sK1,sK0))))) | $spl8),
  inference(superposition,[],[f192,f1641])).
fof(f1704,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | $spl8),
  inference(superposition,[],[f1641,f207])).
fof(f1703,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | $spl8),
  inference(superposition,[],[f1641,f207])).
fof(f1702,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | $spl8),
  inference(superposition,[],[f1641,f189])).
fof(f1701,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),sK3(sK1,sK0)) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | $spl8),
  inference(superposition,[],[f1641,f188])).
fof(f1680,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X1))) | s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | $spl8),
  inference(superposition,[],[f1639,f207])).
fof(f1679,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X0))) | s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | $spl8),
  inference(superposition,[],[f1639,f207])).
fof(f1678,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0))))) | $spl8),
  inference(superposition,[],[f1639,f189])).
fof(f1677,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),sK3(sK1,sK0)) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK3(sK1,sK0),sK0)),s(fun(sK0,bool),sK3(sK3(sK1,sK0),sK0)))) | $spl8),
  inference(superposition,[],[f1639,f188])).
fof(f1652,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X1)))) | s(bool,f(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK3(sK1,sK0),bool,sK0))))) ) | $spl8),
  inference(superposition,[],[f978,f207])).
fof(f1651,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),X0)))) | s(bool,f(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK3(sK1,sK0),X0,bool,sK0))))) ) | $spl8),
  inference(superposition,[],[f978,f207])).
fof(f1641,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | $spl8),
  inference(resolution,[],[f978,f154])).
fof(f1639,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | $spl8),
  inference(resolution,[],[f978,f149])).
fof(f1351,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | $spl8),
  inference(subsumption_resolution,[],[f215,f766])).
fof(f1350,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | $spl8),
  inference(subsumption_resolution,[],[f1349,f187])).
fof(f1349,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | $spl8),
  inference(forward_demodulation,[],[f949,f144])).
fof(f949,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0)))) | $spl8),
  inference(superposition,[],[f766,f188])).
fof(f978,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))))) | $spl8),
  inference(subsumption_resolution,[],[f977,f187])).
fof(f977,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))))) | $spl8),
  inference(forward_demodulation,[],[f950,f144])).
fof(f950,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,sK0)),s(fun(sK0,bool),sK3(sK1,sK0))))) | $spl8),
  inference(superposition,[],[f766,f189])).
fof(f957,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))))) ) | $spl8),
  inference(superposition,[],[f192,f766])).
fof(f1000,plain,(
  ( ! [X0] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X0) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f779,f195])).
fof(f952,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X1))) | s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0))))) ) | $spl8),
  inference(superposition,[],[f766,f207])).
fof(f951,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),X0))) | s(bool,f(s(sK0,sK5(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,X0,bool,sK0))))) ) | $spl8),
  inference(superposition,[],[f766,f207])).
fof(f827,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X1))) | s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f781,f207])).
fof(f826,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X0))) | s(bool,f(s(sK0,sK5(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,X0,bool,sK0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f781,f207])).
fof(f794,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X1)))) | s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f767,f207])).
fof(f793,plain,(
  ( ! [X0] : (~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X0)))) | s(bool,f(s(sK0,sK5(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,X0,bool,sK0))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f767,f207])).
fof(f783,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))) | ($spl2 | $spl8)),
  inference(resolution,[],[f767,f154])).
fof(f781,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))) | ($spl2 | $spl8)),
  inference(resolution,[],[f767,f149])).
fof(f768,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f766,f222])).
fof(f767,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f221,f766])).
fof(f2077,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl4 | $spl8)),
  inference(subsumption_resolution,[],[f2076,f139])).
fof(f2076,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl4 | $spl8)),
  inference(forward_demodulation,[],[f2075,f2056])).
fof(f2056,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))) | ($spl0 | $spl2 | $spl4)),
  inference(backward_demodulation,[],[f2055,f1799])).
fof(f2075,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl4 | $spl8)),
  inference(subsumption_resolution,[],[f848,f368])).
fof(f848,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(superposition,[],[f783,f190])).
fof(f2069,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f238,f239,f337,f281,f282,f766,f767,f768,f781,f783,f793,f794,f826,f827,f851,f852,f951,f952,f992,f993,f994,f998,f999,f1000,f1348,f957,f978,f1350,f1351,f218,f1352,f1353,f1383,f1366,f1367,f1449,f1450,f1639,f1641,f1651,f1652,f1677,f1678,f1679,f1680,f1701,f1702,f1703,f1704,f1768,f1786,f1787,f1791,f1792,f1825,f1795,f1796,f1797,f1798,f1839,f1843,f1785,f1700,f1676,f1978,f1979,f1799,f1987,f1988,f1989,f2014,f2016,f2037,f2039,f779,f1145,f1143,f1141,f1139,f1138,f1018,f1017,f1137,f1015,f1010,f1009,f1005,f1004,f284,f283,f2068])).
fof(f2068,plain,(
  ( ! [X0] : (s(fun(sK0,bool),sK1) != s(fun(sK0,bool),X0)) ) | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f1000,f1798])).
fof(f2039,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK5(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,X0,bool,sK0))))) ) | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f793,f1366])).
fof(f2037,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0))))) ) | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f794,f1367])).
fof(f2016,plain,(
  ( ! [X0] : (s(bool,f(s(sK0,sK5(sK1,X0,bool,sK0)))) != s(bool,g(s(sK0,sK5(sK1,X0,bool,sK0))))) ) | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f826,f1449])).
fof(f2014,plain,(
  ( ! [X1] : (s(bool,f(s(sK0,sK5(X1,sK1,bool,sK0)))) != s(bool,g(s(sK0,sK5(X1,sK1,bool,sK0))))) ) | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f827,f1450])).
fof(f1989,plain,(
  s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)))) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f1825,f1987])).
fof(f1988,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0)),s(fun(sK0,bool),sK3('const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')),sK0))))) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f1795,f1987])).
fof(f1979,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl0 | $spl8)),
  inference(global_subsumption,[],[f215,f137,f216,f766,f951,f952,f1348,f957,f978,f1350,f1351,f218,f1352,f1353,f1383,f1366,f1367,f1449,f1450,f1639,f1641,f1651,f1652,f1677,f1678,f1679,f1680,f1701,f1702,f1703,f1704,f1768,f1786,f1787,f1791,f1792,f1825,f1795,f1796,f1797,f1798,f1839,f1843,f1785,f1700,f1676,f1978])).
fof(f1978,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),sK1) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f1977,f1352])).
fof(f1977,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f1976,f187])).
fof(f1976,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f1975,f766])).
fof(f1839,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f1838,f143])).
fof(f1838,plain,(
  ( ! [X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) ) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f1837,f200])).
fof(f1837,plain,(
  ( ! [X1] : (s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X1),s(fun(sK0,bool),sK1))))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.bool.IN_2'(s(sK0,X1),s(fun(sK0,bool),sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK4(sK1,sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))) ) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f1808,f766])).
fof(f2067,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2066,f214])).
fof(f2066,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f1953,f1353])).
fof(f1953,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(trivial_inequality_removal,[],[f873])).
fof(f873,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(superposition,[],[f153,f783])).
fof(f2065,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl8)),
  inference(global_subsumption,[],[f2064,f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f238,f239,f337,f281,f282,f283,f284,f766,f767,f768,f781,f783,f793,f794,f826,f827,f851,f852,f951,f952,f992,f993,f994,f998,f999,f1000,f1004,f1005,f1348,f957,f978,f1350,f1351,f1639,f1641,f1651,f1652,f1677,f1678,f1679,f1680,f1701,f1702,f1703,f1704,f1768,f1700,f1676,f368,f779,f1145,f1143,f1141,f2055,f1139,f1138,f1018,f1017,f1137,f1015,f1010,f1009])).
fof(f2064,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X0) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X0))))) ) | ($spl0 | $spl2 | $spl4 | $spl8)),
  inference(subsumption_resolution,[],[f2063,f139])).
fof(f2063,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X0) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X0))))) ) | ($spl0 | $spl2 | $spl4 | $spl8)),
  inference(forward_demodulation,[],[f2062,f2056])).
fof(f2062,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X0) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X0)))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))))) ) | ($spl2 | $spl4 | $spl8)),
  inference(subsumption_resolution,[],[f1006,f368])).
fof(f1006,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X0) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f1000,f190])).
fof(f2061,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2060,f214])).
fof(f2060,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f1954,f1353])).
fof(f1954,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(trivial_inequality_removal,[],[f876])).
fof(f876,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(superposition,[],[f155,f783])).
fof(f2050,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f1843,f779])).
fof(f2049,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f1352,f779])).
fof(f2048,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f781,f1353])).
fof(f2047,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2046,f1383])).
fof(f2046,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f783,f1353])).
fof(f2045,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl8)),
  inference(global_subsumption,[],[f2044,f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f238,f239,f337,f281,f282,f283,f284,f766,f767,f768,f779,f781,f783,f793,f794,f826,f827,f851,f852,f951,f952,f992,f993,f994,f998,f999,f1000,f1004,f1005,f1009,f1010,f1015,f1017,f1018,f1137,f1138,f1139,f1141,f1143,f1145,f1348,f957,f978,f1350,f1351,f1639,f1641,f1651,f1652,f1677,f1678,f1679,f1680,f1701,f1702,f1703,f1704,f1768,f1700,f1676,f368])).
fof(f2044,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl2 | $spl4 | $spl8)),
  inference(subsumption_resolution,[],[f2043,f143])).
fof(f2043,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl2 | $spl4 | $spl8)),
  inference(forward_demodulation,[],[f2042,f1799])).
fof(f2042,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl4 | $spl8)),
  inference(subsumption_resolution,[],[f790,f368])).
fof(f790,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(superposition,[],[f767,f190])).
fof(f2040,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f238,f239,f337,f281,f282,f283,f284,f766,f767,f768,f779,f781,f783,f793,f794,f826,f827,f851,f852,f951,f952,f992,f993,f994,f998,f999,f1000,f1004,f1005,f1009,f1010,f1015,f1017,f1018,f1137,f1138,f1139,f1141,f1143,f1145,f1348,f957,f978,f1350,f1351,f218,f1352,f1353,f1383,f1366,f1367,f1449,f1450,f1639,f1641,f1651,f1652,f1677,f1678,f1679,f1680,f1701,f1702,f1703,f1704,f1768,f1786,f1787,f1791,f1792,f1825,f1795,f1796,f1797,f1798,f1839,f1843,f1785,f1700,f1676,f1978,f1979,f1799,f1987,f1988,f1989,f2014,f2016,f2037,f2039])).
fof(f2038,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f238,f239,f337,f281,f282,f283,f284,f766,f767,f768,f779,f781,f783,f793,f794,f826,f827,f851,f852,f951,f952,f992,f993,f994,f998,f999,f1000,f1004,f1005,f1009,f1010,f1015,f1017,f1018,f1137,f1138,f1139,f1141,f1143,f1145,f1348,f957,f978,f1350,f1351,f218,f1352,f1353,f1383,f1366,f1367,f1449,f1450,f1639,f1641,f1651,f1652,f1677,f1678,f1679,f1680,f1701,f1702,f1703,f1704,f1768,f1786,f1787,f1791,f1792,f1825,f1795,f1796,f1797,f1798,f1839,f1843,f1785,f1700,f1676,f1978,f1979,f1799,f1987,f1988,f1989,f2014,f2016,f2037])).
fof(f2036,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2035,f214])).
fof(f2035,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f1961,f1353])).
fof(f1961,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(condensation,[],[f909])).
fof(f909,plain,(
  ( ! [X2] : (~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ~p(s(bool,X2))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f858,f153])).
fof(f858,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,X2) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ~p(s(bool,X2))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f783,f204])).
fof(f2034,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2033,f214])).
fof(f2033,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f1962,f1353])).
fof(f1962,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(condensation,[],[f912])).
fof(f912,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f859,f153])).
fof(f859,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f783,f204])).
fof(f2032,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2031,f214])).
fof(f2031,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f1964,f1353])).
fof(f1964,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(condensation,[],[f1963])).
fof(f1963,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))))) ) | ($spl2 | $spl8)),
  inference(duplicate_literal_removal,[],[f803])).
fof(f803,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ~p(s(bool,X2))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f767,f204])).
fof(f2030,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2029,f214])).
fof(f2029,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f1966,f1353])).
fof(f1966,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(condensation,[],[f1965])).
fof(f1965,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))))) ) | ($spl2 | $spl8)),
  inference(duplicate_literal_removal,[],[f804])).
fof(f804,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f767,f204])).
fof(f2028,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2027,f214])).
fof(f2027,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f1967,f1353])).
fof(f1967,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(condensation,[],[f805])).
fof(f805,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))))) ) | ($spl2 | $spl8)),
  inference(duplicate_literal_removal,[],[f804])).
fof(f2026,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2025,f214])).
fof(f2025,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f1968,f1353])).
fof(f1968,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(condensation,[],[f807])).
fof(f807,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))))) ) | ($spl2 | $spl8)),
  inference(duplicate_literal_removal,[],[f803])).
fof(f2024,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2023,f214])).
fof(f2023,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f1969,f1353])).
fof(f1969,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(condensation,[],[f925])).
fof(f925,plain,(
  ( ! [X5] : (~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ~p(s(bool,X5))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f890,f153])).
fof(f890,plain,(
  ( ! [X5] : (s(bool,'F_0') = s(bool,X5) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ~p(s(bool,X5))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f204,f783])).
fof(f2022,plain,(
  $false | ($spl0 | $spl2 | $spl4 | $spl8)),
  inference(global_subsumption,[],[f2021,f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f238,f239,f337,f281,f282,f283,f284,f766,f767,f768,f779,f781,f783,f793,f794,f826,f827,f851,f852,f951,f952,f992,f993,f994,f998,f999,f1000,f1004,f1005,f1009,f1010,f1015,f1017,f1018,f1137,f1138,f1139,f1141,f1143,f1145,f1348,f957,f978,f1350,f1351,f1639,f1641,f1651,f1652,f1677,f1678,f1679,f1680,f1701,f1702,f1703,f1704,f1768,f1700,f1676,f368])).
fof(f2021,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl4 | $spl8)),
  inference(subsumption_resolution,[],[f2020,f143])).
fof(f2020,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | s(bool,'T_0') != s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl2 | $spl4 | $spl8)),
  inference(forward_demodulation,[],[f2019,f1799])).
fof(f2019,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl4 | $spl8)),
  inference(subsumption_resolution,[],[f823,f368])).
fof(f823,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(superposition,[],[f781,f190])).
fof(f2017,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f238,f239,f337,f281,f282,f283,f284,f766,f767,f768,f779,f781,f783,f793,f794,f826,f827,f851,f852,f951,f952,f992,f993,f994,f998,f999,f1000,f1004,f1005,f1009,f1010,f1015,f1017,f1018,f1137,f1138,f1139,f1141,f1143,f1145,f1348,f957,f978,f1350,f1351,f218,f1352,f1353,f1383,f1366,f1367,f1449,f1450,f1639,f1641,f1651,f1652,f1677,f1678,f1679,f1680,f1701,f1702,f1703,f1704,f1768,f1786,f1787,f1791,f1792,f1825,f1795,f1796,f1797,f1798,f1839,f1843,f1785,f1700,f1676,f1978,f1979,f1799,f1987,f1988,f1989,f2014,f2016])).
fof(f2015,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(global_subsumption,[],[f215,f137,f216,f220,f221,f222,f223,f224,f225,f255,f238,f239,f337,f281,f282,f283,f284,f766,f767,f768,f779,f781,f783,f793,f794,f826,f827,f851,f852,f951,f952,f992,f993,f994,f998,f999,f1000,f1004,f1005,f1009,f1010,f1015,f1017,f1018,f1137,f1138,f1139,f1141,f1143,f1145,f1348,f957,f978,f1350,f1351,f218,f1352,f1353,f1383,f1366,f1367,f1449,f1450,f1639,f1641,f1651,f1652,f1677,f1678,f1679,f1680,f1701,f1702,f1703,f1704,f1768,f1786,f1787,f1791,f1792,f1825,f1795,f1796,f1797,f1798,f1839,f1843,f1785,f1700,f1676,f1978,f1979,f1799,f1987,f1988,f1989,f2014])).
fof(f2013,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2012,f1383])).
fof(f2012,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f1970,f1353])).
fof(f1970,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))) | ($spl2 | $spl8)),
  inference(trivial_inequality_removal,[],[f828])).
fof(f828,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))) | ($spl2 | $spl8)),
  inference(superposition,[],[f781,f145])).
fof(f2011,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2010,f214])).
fof(f2010,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f1971,f1353])).
fof(f1971,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(condensation,[],[f928])).
fof(f928,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f891,f153])).
fof(f891,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,X6) | ~p(s(bool,X6)) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))))) ) | ($spl2 | $spl8)),
  inference(superposition,[],[f204,f783])).
fof(f2009,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2008,f214])).
fof(f2008,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f1972,f1353])).
fof(f1972,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(trivial_inequality_removal,[],[f830])).
fof(f830,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(superposition,[],[f781,f150])).
fof(f2007,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2006,f214])).
fof(f2006,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f1973,f1353])).
fof(f1973,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(trivial_inequality_removal,[],[f831])).
fof(f831,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl2 | $spl8)),
  inference(superposition,[],[f781,f152])).
fof(f152,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.SING_IFF_CARD|49|','thm.bool.EQ_CLAUSES|split|0')).
fof(f2005,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f2004,f187])).
fof(f2004,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl4 | $spl8)),
  inference(backward_demodulation,[],[f368,f766])).
fof(f2003,plain,(
  $false | ($spl0 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f371,f1987])).
fof(f1985,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f1984,f214])).
fof(f1984,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f1983,f1353])).
fof(f1983,plain,(
  ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f1982,f143])).
fof(f1982,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f1351,f1799])).
fof(f1981,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f1980,f143])).
fof(f1980,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl8)),
  inference(backward_demodulation,[],[f1799,f1979])).
fof(f1944,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1943,f143])).
fof(f1943,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(condensation,[],[f1942])).
fof(f1942,plain,(
  ( ! [X6] : (~p(s(bool,X6)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1902,f153])).
fof(f1902,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,X6) | ~p(s(bool,X6)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl3)),
  inference(superposition,[],[f204,f1826])).
fof(f1826,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f1799,f1411])).
fof(f1411,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1))) | $spl3),
  inference(resolution,[],[f219,f154])).
fof(f219,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | $spl3),
  inference(cnf_transformation,[],[f219_D])).
fof(f219_D,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f1941,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1940,f143])).
fof(f1940,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(condensation,[],[f1939])).
fof(f1939,plain,(
  ( ! [X5] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X5))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1901,f153])).
fof(f1901,plain,(
  ( ! [X5] : (s(bool,'F_0') = s(bool,X5) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X5))) ) | ($spl0 | $spl3)),
  inference(superposition,[],[f204,f1826])).
fof(f1934,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1913,f143])).
fof(f1913,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(trivial_inequality_removal,[],[f1887])).
fof(f1887,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(superposition,[],[f155,f1826])).
fof(f1933,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1914,f143])).
fof(f1914,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(trivial_inequality_removal,[],[f1884])).
fof(f1884,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(superposition,[],[f153,f1826])).
fof(f1932,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1931,f143])).
fof(f1931,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1882,f1383])).
fof(f1882,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(superposition,[],[f152,f1826])).
fof(f1930,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1929,f143])).
fof(f1929,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1878,f1383])).
fof(f1878,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(superposition,[],[f150,f1826])).
fof(f1927,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1870,f213])).
fof(f1870,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3)),
  inference(superposition,[],[f143,f1826])).
fof(f1926,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1925,f143])).
fof(f1925,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(condensation,[],[f1924])).
fof(f1924,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1869,f153])).
fof(f1869,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) ) | ($spl0 | $spl3)),
  inference(superposition,[],[f1826,f204])).
fof(f1923,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1922,f143])).
fof(f1922,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(condensation,[],[f1921])).
fof(f1921,plain,(
  ( ! [X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X2))) ) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1868,f153])).
fof(f1868,plain,(
  ( ! [X2] : (s(bool,'F_0') = s(bool,X2) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,X2))) ) | ($spl0 | $spl3)),
  inference(superposition,[],[f1826,f204])).
fof(f1920,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1919,f143])).
fof(f1919,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1865,f1383])).
fof(f1865,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(superposition,[],[f1826,f152])).
fof(f1918,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1917,f143])).
fof(f1917,plain,(
  ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f1864,f1383])).
fof(f1864,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | ($spl0 | $spl3)),
  inference(superposition,[],[f1826,f150])).
fof(f1347,plain,(
  $false | ($spl2 | $spl5 | $spl8)),
  inference(subsumption_resolution,[],[f1337,f770])).
fof(f770,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))))) | ($spl2 | $spl5 | $spl8)),
  inference(backward_demodulation,[],[f766,f582])).
fof(f1337,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))))) | ($spl2 | $spl5 | $spl8)),
  inference(equality_resolution,[],[f945])).
fof(f945,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))) ) | ($spl2 | $spl5 | $spl8)),
  inference(superposition,[],[f809,f205])).
fof(f809,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0)))) | ($spl2 | $spl5 | $spl8)),
  inference(subsumption_resolution,[],[f581,f779])).
fof(f1342,plain,(
  $false | ($spl2 | $spl5 | $spl8)),
  inference(trivial_inequality_removal,[],[f1341])).
fof(f1341,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ($spl2 | $spl5 | $spl8)),
  inference(forward_demodulation,[],[f1340,f144])).
fof(f1340,plain,(
  ( ! [X4] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X4,bool),'const.pred_set.EMPTY_0')))))) ) | ($spl2 | $spl5 | $spl8)),
  inference(forward_demodulation,[],[f1339,f770])).
fof(f1339,plain,(
  ( ! [X4] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X4,bool),'const.pred_set.EMPTY_0'))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))) ) | ($spl2 | $spl5 | $spl8)),
  inference(trivial_inequality_removal,[],[f1327])).
fof(f1327,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0') | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(X4,bool),'const.pred_set.EMPTY_0'))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK3(sK1,sK0))))))) ) | ($spl2 | $spl5 | $spl8)),
  inference(superposition,[],[f945,f144])).
fof(f218,plain,(
  p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) | $spl0),
  inference(cnf_transformation,[],[f218_D])).
fof(f218_D,plain,(
  p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f766,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | $spl8),
  inference(cnf_transformation,[],[f766_D])).
fof(f766_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f216,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))))),
  inference(forward_demodulation,[],[f136,f141])).
fof(f136,plain,(
  s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))))),
  inference(cnf_transformation,[],[f103])).
fof(f369,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | $spl5),
  inference(cnf_transformation,[],[f369_D])).
fof(f369_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f371,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | $spl6),
  inference(cnf_transformation,[],[f371_D])).
fof(f371_D,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f223,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),sK1) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.pred_set.CARD_1'(s(fun(sK0,bool),sK1))) | $spl2),
  inference(resolution,[],[f220,f190])).
fof(f220,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | $spl2),
  inference(cnf_transformation,[],[f220_D])).
fof(f220_D,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f137,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK1)))) | p(s(bool,'const.pred_set.SING_1'(s(fun(sK0,bool),sK1))))),
  inference(cnf_transformation,[],[f103])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 335
% Passive clauses: 2857
% Generated clauses: 15777
% Final active clauses: 85
% Final passive clauses: 1147
% Input formulas: 37
% Initial clauses: 72
% 
% Duplicate literals: 315
% Trivial inequalities: 425
% Fw subsumption resolutions: 1202
% Bw subsumption resolutions: 125
% Fw demodulations: 2318
% Bw demodulations: 772
% Condensations: 472
% Global subsumptions: 33
% 
% Simple tautologies: 459
% Equational tautologies: 11
% Forward subsumptions: 7579
% Backward subsumptions: 414
% Fw demodulations to eq. taut.: 7
% 
% Binary resolution: 817
% Factoring: 4
% Forward superposition: 5132
% Backward superposition: 2694
% Self superposition: 48
% Equality factoring: 12
% Equality resolution: 56
% 
% Split clauses: 5
% Split components: 10
% SAT solver clauses: 3025
% SAT solver unit clauses: 20
% SAT solver binary clauses: 182
% 
% Sat splits: 5
% Sat splitting refutations: 128
% 
% Memory used [KB]: 8955
% Time elapsed: 0.615 s
% ------------------------------
% ------------------------------
% Success in time 16.995 s
