Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.transc
% (111006)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 262
% Passive clauses: 1958
% Generated clauses: 27417
% Final active clauses: 205
% Final passive clauses: 1300
% Input formulas: 26
% Initial clauses: 45
% 
% Duplicate literals: 1464
% Trivial inequalities: 257
% Fw subsumption resolutions: 3795
% Fw demodulations: 5914
% Bw demodulations: 40
% 
% Simple tautologies: 485
% Equational tautologies: 25
% Forward subsumptions: 12221
% Fw demodulations to eq. taut.: 1
% 
% Binary resolution: 642
% Forward superposition: 10053
% Backward superposition: 5090
% Self superposition: 156
% Equality factoring: 38
% Equality resolution: 56
% 
% Split clauses: 724
% Split components: 1458
% SAT solver clauses: 1575
% SAT solver unit clauses: 21
% SAT solver binary clauses: 95
% 
% Sat splits: 724
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1573
% TWLsolver calls for satisfiability: 272
% 
% Memory used [KB]: 12537
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.transc
% (111023)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 245
% Passive clauses: 1785
% Generated clauses: 13772
% Final active clauses: 178
% Final passive clauses: 1396
% Input formulas: 26
% Initial clauses: 45
% 
% Duplicate literals: 341
% Trivial inequalities: 310
% Fw subsumption resolutions: 1001
% Fw demodulations: 2959
% Bw demodulations: 194
% Condensations: 429
% Global subsumptions: 4
% 
% Simple tautologies: 281
% Equational tautologies: 21
% Forward subsumptions: 4696
% Fw demodulations to eq. taut.: 5
% Bw demodulations to eq. taut.: 7
% 
% Binary resolution: 348
% Forward superposition: 3240
% Backward superposition: 2952
% Self superposition: 36
% Equality factoring: 30
% Equality resolution: 61
% 
% Split clauses: 49
% Split components: 98
% SAT solver clauses: 1927
% SAT solver unit clauses: 60
% SAT solver binary clauses: 42
% SAT solver learnt clauses: 9
% SAT solver learnt literals: 1
% 
% Sat splits: 49
% Sat splitting refutations: 13
% 
% TWLsolver clauses: 1910
% TWLsolver calls for satisfiability: 327
% 
% Memory used [KB]: 14839
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.transc
% (111040)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 591
% Passive clauses: 5053
% Generated clauses: 138489
% Final active clauses: 137
% Final passive clauses: 580
% Input formulas: 26
% Initial clauses: 45
% 
% Duplicate literals: 6575
% Trivial inequalities: 1267
% Fw subsumption resolutions: 14052
% Fw demodulations: 18698
% Bw demodulations: 3337
% Condensations: 13319
% Global subsumptions: 2149
% 
% Simple tautologies: 1249
% Equational tautologies: 2460
% Forward subsumptions: 71079
% Fw demodulations to eq. taut.: 32
% Bw demodulations to eq. taut.: 51
% 
% Unit resulting resolution: 8292
% Forward superposition: 32356
% Backward superposition: 23061
% Self superposition: 626
% Equality factoring: 377
% Equality resolution: 91
% 
% Split clauses: 214
% Split components: 442
% SAT solver clauses: 9926
% SAT solver unit clauses: 136
% SAT solver binary clauses: 566
% 
% Sat splits: 214
% Sat splitting refutations: 2294
% 
% Memory used [KB]: 65244
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.transc
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.transc
% SZS output start Proof for thm.transc
fof(f37473,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f138,f161,f159,f167,f168,f178,f177,f171,f172,f592,f851,f852,f979,f1002,f1000,f1426,f1362,f1430,f1390,f1431,f1391,f1392,f1393,f181,f1405,f1432,f1433,f1516,f1539,f1575,f1597,f1667,f1427,f1688,f1726,f1746,f1775,f1773,f2082,f1953,f2098,f1993,f184,f183,f2320,f2331,f2330,f2516,f2529,f2561,f2559,f2598,f2629,f2627,f2612,f2631,f173,f3074,f3087,f3136,f3135,f3259,f3260,f3324,f3357,f3414,f3416,f3433,f3508,f3415,f3542,f3359,f3559,f175,f3813,f3849,f3847,f4193,f4228,f4226,f4229,f4231,f4260,f4262,f4391,f4393,f4482,f4348,f4307,f4248,f4216,f4494,f4233,f4171,f3905,f3949,f4497,f3873,f3888,f4498,f3837,f4499,f3857,f3855,f4500,f3725,f3717,f3667,f4510,f3703,f4512,f3011,f2987,f2941,f4514,f2951,f2905,f4515,f4516,f1550,f2645,f4517,f1637,f2617,f179,f2578,f4518,f1588,f2549,f4495,f2565,f4519,f2222,f4513,f2114,f1950,f4524,f1928,f4529,f2076,f2074,f2068,f1910,f4530,f1887,f4531,f2050,f2083,f1763,f4496,f4532,f3261,f1685,f1670,f4535,f1536,f4538,f1523,f1429,f4539,f1519,f4540,f1428,f4561,f4628,f4485,f5941,f5986,f5984,f6708,f6749,f6747,f8016,f8066,f8064,f8912,f8897,f8914,f8916,f11018,f11064,f11062,f12248,f12303,f12301,f16155,f16158,f16156,f16315,f16300,f16317,f16319,f16623,f16608,f16625,f16627,f19424,f19433,f19435,f19417,f19426,f19428,f26126,f26130,f26128,f26134,f26138,f26136,f26183,f26193,f26204,f26184,f26205,f26206,f26244,f24438,f1434,f19822,f26827,f19532,f4493,f26836,f14352,f27219,f4484,f28494,f28496,f28495,f28835,f28849,f28850,f28852,f28836,f28854,f28855,f28858,f28859,f28906,f28907,f28909,f28910,f28949,f28951,f28962,f28964,f28989,f28995,f28993,f28997,f28991,f29005,f29007,f29013,f29015,f29043,f29044,f29047,f29048,f29050,f29051,f29055,f29056,f29061,f29062,f29071,f29142,f29150,f29180,f29186,f29188,f29192,f29148,f29196,f28953,f29194,f29214,f29220,f29216,f29224,f29226,f29182,f29230,f29232,f29236,f29240,f29190,f29145,f29242,f29244,f29246,f29228,f29144,f29262,f29269,f26203,f26202,f29273,f26245,f26186,f29419,f29448,f29446,f29777,f29778,f29810,f26236,f26234,f29811,f26240,f26238,f30782,f30784,f30805,f26379,f30812,f30783,f31664,f29152,f29222,f31665,f29146,f29218,f31714,f29073,f31716,f31748,f28999,f35843,f35866,f36840,f36872,f36893])).
fof(f36893,plain,(
  $false | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f36892,f127])).
fof(f127,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.transc.ROOT_INV',aHL_TRUTH)).
fof(f36892,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f36891,f921])).
fof(f921,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1))) | $spl0),
  inference(resolution,[],[f159,f101])).
fof(f101,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.transc.ROOT_INV','thm.bool.EQ_CLAUSES|split|1')).
fof(f36891,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1)))) | $spl10),
  inference(subsumption_resolution,[],[f36533,f106])).
fof(f106,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))),
  file('thm.transc.ROOT_INV','thm.real.POW_|48|')).
fof(f36533,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1)))) | $spl10),
  inference(superposition,[],[f1741,f35847])).
fof(f35847,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))) | $spl10),
  inference(trivial_inequality_removal,[],[f35846])).
fof(f35846,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))) | $spl10),
  inference(forward_demodulation,[],[f35824,f135])).
fof(f135,plain,(
  ( ! [X1] : (s('type.realax.real',sK1) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(resolution,[],[f88,f91])).
fof(f91,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real',X1) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) => s('type.realax.real',X1) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) => s('type.realax.real',X7) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.realax.real',X7))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))),
  file('thm.transc.ROOT_INV','thm.transc.ROOT_POW_POS')).
fof(f88,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1))))),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1)))) & s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f56])).
fof(f56,plain,(
  ? [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) & s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))),
  inference(ennf_transformation,[],[f28])).
fof(f28,plain,(
  ~! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) => s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) => s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.realax.real',X7))))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) => s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.realax.real',X7))))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7))))))),
  file('thm.transc.ROOT_INV','thm.transc.ROOT_INV')).
fof(f35824,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))) | $spl10),
  inference(superposition,[],[f26147,f104])).
fof(f104,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X1))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ! [X1] : s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X1))))),
  inference(ennf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => ! [X1] : s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X1))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : (~s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => ! [X1] : s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X1))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X11] : (~s('type.realax.real',X11) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => ! [X6] : s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X11),s('type.num.num',X6))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X11))),s('type.num.num',X6))))),
  file('thm.transc.ROOT_INV','thm.real.POW_INV')).
fof(f26147,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | $spl10),
  inference(subsumption_resolution,[],[f26146,f132])).
fof(f132,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f94])).
fof(f94,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X9] : ! [X7] : (s(X9,X7) = s(X9,X7) <=> p(s(bool,'T_0')))),
  file('thm.transc.ROOT_INV','thm.bool.REFL_CLAUSE')).
fof(f26146,plain,(
  ~p(s(bool,'T_0')) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | $spl10),
  inference(forward_demodulation,[],[f26106,f2924])).
fof(f2924,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.realax.real',sK1)))))))) )),
  inference(resolution,[],[f2205,f101])).
fof(f2205,plain,(
  ( ! [X10] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.realax.real',sK1))))))))) )),
  inference(resolution,[],[f133,f108])).
fof(f108,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0)))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X7)))))))),
  file('thm.transc.ROOT_INV','thm.real.REAL_LE_INV')).
fof(f133,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',sK1))))))) )),
  inference(resolution,[],[f88,f109])).
fof(f109,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1)))))))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1)))))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6,X7] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.realax.real',X7)))))))),
  file('thm.transc.ROOT_INV','thm.transc.ROOT_POS')).
fof(f26106,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1)))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | $spl10),
  inference(equality_resolution,[],[f178])).
fof(f1741,plain,(
  ( ! [X59,X58] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X58))),s('type.realax.real',X59))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X58))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X59))))) )),
  inference(subsumption_resolution,[],[f1721,f129])).
fof(f129,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & s('type.realax.real',X0) != s('type.realax.real',X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) = s('type.realax.real',X1))) & ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & s('type.realax.real',X0) != s('type.realax.real',X1)) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) <=> (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) = s('type.realax.real',X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X7,X10] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X7),s('type.realax.real',X10)))) <=> (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X7),s('type.realax.real',X10)))) | s('type.realax.real',X7) = s('type.realax.real',X10)))),
  file('thm.transc.ROOT_INV','thm.real.REAL_LE_LT')).
fof(f1721,plain,(
  ( ! [X59,X58] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X59)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X58))),s('type.realax.real',X59))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X58))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X59))))) )),
  inference(superposition,[],[f131,f91])).
fof(f131,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) => s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) => ~s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X7] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) => ~s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.transc.ROOT_INV','thm.real.REAL_POS_NZ')).
fof(f36872,plain,(
  $false | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f36871,f127])).
fof(f36871,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl10)),
  inference(forward_demodulation,[],[f36870,f141])).
fof(f141,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1)))),
  inference(resolution,[],[f88,f101])).
fof(f36870,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1)))) | ($spl3 | $spl10)),
  inference(subsumption_resolution,[],[f36869,f160])).
fof(f160,plain,(
  s('type.realax.real',sK1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(cnf_transformation,[],[f160_D])).
fof(f160_D,plain,(
  s('type.realax.real',sK1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f36869,plain,(
  s('type.realax.real',sK1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1)))) | $spl10),
  inference(forward_demodulation,[],[f36522,f106])).
fof(f36522,plain,(
  s('type.realax.real',sK1) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1)))) | $spl10),
  inference(superposition,[],[f91,f35847])).
fof(f36840,plain,(
  $false | ($spl10 | $spl15)),
  inference(subsumption_resolution,[],[f36505,f107])).
fof(f107,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f19])).
fof(f19,axiom,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.transc.ROOT_INV','thm.real.REAL_INV_|48|')).
fof(f36505,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl10 | $spl15)),
  inference(backward_demodulation,[],[f35847,f181])).
fof(f35866,plain,(
  $false | $spl9),
  inference(subsumption_resolution,[],[f35865,f127])).
fof(f35865,plain,(
  ~p(s(bool,'T_0')) | $spl9),
  inference(forward_demodulation,[],[f177,f982])).
fof(f982,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))),
  inference(resolution,[],[f134,f101])).
fof(f134,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),
  inference(resolution,[],[f88,f108])).
fof(f35843,plain,(
  $false | ($spl10 | $spl92)),
  inference(trivial_inequality_removal,[],[f35842])).
fof(f35842,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | ($spl10 | $spl92)),
  inference(forward_demodulation,[],[f35822,f135])).
fof(f35822,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ($spl10 | $spl92)),
  inference(superposition,[],[f26147,f4484])).
fof(f28999,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | $spl232),
  inference(cnf_transformation,[],[f28999_D])).
fof(f28999_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f31748,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | $spl284),
  inference(cnf_transformation,[],[f31748_D])).
fof(f31748_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f31716,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | $spl180),
  inference(resolution,[],[f26185,f128])).
fof(f128,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) = s('type.realax.real',X1) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f87])).
fof(f26185,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | $spl180),
  inference(cnf_transformation,[],[f26185_D])).
fof(f26185_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f29073,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | $spl248),
  inference(cnf_transformation,[],[f29073_D])).
fof(f29073_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f31714,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | $spl180),
  inference(resolution,[],[f26185,f4585])).
fof(f4585,plain,(
  ( ! [X6] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6))))) )),
  inference(superposition,[],[f4174,f128])).
fof(f4174,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(forward_demodulation,[],[f4173,f107])).
fof(f4173,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))),
  inference(forward_demodulation,[],[f4168,f107])).
fof(f4168,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))),
  inference(trivial_inequality_removal,[],[f4152])).
fof(f4152,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))),
  inference(superposition,[],[f1176,f107])).
fof(f1176,plain,(
  ( ! [X1] : (s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))))))) )),
  inference(resolution,[],[f108,f130])).
fof(f130,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | s('type.realax.real',X0) != s('type.realax.real',X1)) )),
  inference(cnf_transformation,[],[f87])).
fof(f29218,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1))) | $spl268),
  inference(cnf_transformation,[],[f29218_D])).
fof(f29218_D,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1))) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f29146,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | $spl252),
  inference(cnf_transformation,[],[f29146_D])).
fof(f29146_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f31665,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1))) | $spl181),
  inference(resolution,[],[f26186,f7187])).
fof(f7187,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',sK1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK1)))) )),
  inference(resolution,[],[f150,f97])).
fof(f97,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.transc.ROOT_INV','thm.bool.EQ_CLAUSES|split|3')).
fof(f150,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',sK1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) )),
  inference(superposition,[],[f88,f128])).
fof(f29222,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | $spl270),
  inference(cnf_transformation,[],[f29222_D])).
fof(f29222_D,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f29152,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | $spl256),
  inference(cnf_transformation,[],[f29152_D])).
fof(f29152_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f31664,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | $spl181),
  inference(resolution,[],[f26186,f7446])).
fof(f7446,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X2)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X2)))) )),
  inference(resolution,[],[f151,f97])).
fof(f151,plain,(
  ( ! [X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) )),
  inference(superposition,[],[f88,f128])).
fof(f30783,plain,(
  ( ! [X6] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6)))))) ) | $spl280),
  inference(cnf_transformation,[],[f30783_D])).
fof(f30783_D,plain,(
  ( ! [X6] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6)))))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f30812,plain,(
  ( ! [X23,X22] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X23))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X23))))) | s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X22))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X22)))))) ) | ($spl24 | $spl32)),
  inference(forward_demodulation,[],[f30811,f107])).
fof(f30811,plain,(
  ( ! [X23,X22] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X23))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X23))) | s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X22))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X22)))))) ) | ($spl24 | $spl32)),
  inference(forward_demodulation,[],[f30810,f107])).
fof(f30810,plain,(
  ( ! [X23,X22] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X22))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X22))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X23))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X23)))) ) | ($spl24 | $spl32)),
  inference(forward_demodulation,[],[f30707,f107])).
fof(f30707,plain,(
  ( ! [X23,X22] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X22))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X22))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X23))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',X23)))) ) | ($spl24 | $spl32)),
  inference(resolution,[],[f1432,f1427])).
fof(f26379,plain,(
  ( ! [X14] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X14)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl196),
  inference(cnf_transformation,[],[f26379_D])).
fof(f26379_D,plain,(
  ( ! [X14] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X14)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f30805,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X19),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X18))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X18))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19))))))) ) | $spl32),
  inference(forward_demodulation,[],[f30804,f107])).
fof(f30804,plain,(
  ( ! [X19,X18] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X18))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X18))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X19),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl32),
  inference(forward_demodulation,[],[f30803,f107])).
fof(f30803,plain,(
  ( ! [X19,X18] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X18))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X18))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X19),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl32),
  inference(forward_demodulation,[],[f30705,f107])).
fof(f30705,plain,(
  ( ! [X19,X18] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X18))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X18))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X19)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X19),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X19),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl32),
  inference(resolution,[],[f1432,f1189])).
fof(f1189,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) )),
  inference(superposition,[],[f108,f128])).
fof(f30784,plain,(
  ( ! [X8,X7] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl282),
  inference(cnf_transformation,[],[f30784_D])).
fof(f30784_D,plain,(
  ( ! [X8,X7] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f30782,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))) | s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))))) ) | ($spl32 | $spl48)),
  inference(forward_demodulation,[],[f30781,f107])).
fof(f30781,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real',X7) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))))) ) | ($spl32 | $spl48)),
  inference(forward_demodulation,[],[f30780,f4625])).
fof(f4625,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl48),
  inference(equality_resolution,[],[f2083])).
fof(f30780,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real',X7) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))))) ) | $spl32),
  inference(forward_demodulation,[],[f30779,f107])).
fof(f30779,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6))))) | s('type.realax.real',X7) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))))) ) | $spl32),
  inference(forward_demodulation,[],[f30699,f107])).
fof(f30699,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X6))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X6))) | s('type.realax.real',X7) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X7),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))))) ) | $spl32),
  inference(resolution,[],[f1432,f90])).
fof(f90,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X2) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | s('type.realax.real',X2) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) | s('type.realax.real',X2) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) & s('type.realax.real',X1) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) => s('type.realax.real',X2) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6,X7,X10] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X10)))) & s('type.realax.real',X7) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X10),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) => s('type.realax.real',X10) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.realax.real',X7))))),
  file('thm.transc.ROOT_INV','thm.transc.ROOT_POS_UNIQ')).
fof(f26238,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1))) | $spl190),
  inference(cnf_transformation,[],[f26238_D])).
fof(f26238_D,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f26240,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | $spl192),
  inference(cnf_transformation,[],[f26240_D])).
fof(f26240_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f29811,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1))) | ($spl0 | $spl183)),
  inference(resolution,[],[f26207,f9345])).
fof(f9345,plain,(
  ( ! [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',sK1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK1)))) ) | $spl0),
  inference(resolution,[],[f928,f97])).
fof(f928,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',sK1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',sK1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) ) | $spl0),
  inference(superposition,[],[f159,f128])).
fof(f26207,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | $spl183),
  inference(resolution,[],[f26202,f129])).
fof(f26234,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | $spl186),
  inference(cnf_transformation,[],[f26234_D])).
fof(f26234_D,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f26236,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | $spl188),
  inference(cnf_transformation,[],[f26236_D])).
fof(f26236_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f29810,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ($spl0 | $spl183)),
  inference(resolution,[],[f26207,f9458])).
fof(f9458,plain,(
  ( ! [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X2)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X2)))) ) | $spl0),
  inference(resolution,[],[f929,f97])).
fof(f929,plain,(
  ( ! [X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | $spl0),
  inference(superposition,[],[f159,f128])).
fof(f29778,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1))) | $spl183),
  inference(resolution,[],[f26202,f7187])).
fof(f29777,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | $spl183),
  inference(resolution,[],[f26202,f7446])).
fof(f29446,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | $spl276),
  inference(cnf_transformation,[],[f29446_D])).
fof(f29446_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f29448,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | $spl278),
  inference(cnf_transformation,[],[f29448_D])).
fof(f29448_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f29419,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | $spl182),
  inference(resolution,[],[f26201,f4585])).
fof(f26201,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | $spl182),
  inference(cnf_transformation,[],[f26201_D])).
fof(f26201_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f26186,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | $spl181),
  inference(cnf_transformation,[],[f26186_D])).
fof(f26186_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) <=> ~$spl181),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl181])])).
fof(f26245,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1)))) ) | $spl194),
  inference(cnf_transformation,[],[f26245_D])).
fof(f26245_D,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1)))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f29273,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | ($spl92 | $spl177)),
  inference(superposition,[],[f26171,f90])).
fof(f26171,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ($spl92 | $spl177)),
  inference(superposition,[],[f26138,f4484])).
fof(f26202,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | $spl183),
  inference(cnf_transformation,[],[f26202_D])).
fof(f26202_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) <=> ~$spl183),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl183])])).
fof(f26203,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) | $spl184),
  inference(cnf_transformation,[],[f26203_D])).
fof(f26203_D,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f29269,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | ($spl92 | $spl177)),
  inference(superposition,[],[f26171,f90])).
fof(f29262,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | ($spl3 | $spl177)),
  inference(forward_demodulation,[],[f29261,f2381])).
fof(f2381,plain,(
  ( ! [X6] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X6))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X6)))) ) | $spl3),
  inference(forward_demodulation,[],[f2380,f135])).
fof(f2380,plain,(
  ( ! [X6,X5] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))),s('type.num.num',X6))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))))),s('type.num.num',X6)))) ) | $spl3),
  inference(subsumption_resolution,[],[f2341,f160])).
fof(f2341,plain,(
  ( ! [X6,X5] : (s('type.realax.real',sK1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))),s('type.num.num',X6))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))))),s('type.num.num',X6)))) )),
  inference(superposition,[],[f135,f104])).
fof(f29261,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | ($spl3 | $spl177)),
  inference(forward_demodulation,[],[f29254,f2381])).
fof(f29254,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | $spl177),
  inference(superposition,[],[f26138,f90])).
fof(f29144,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | $spl251),
  inference(cnf_transformation,[],[f29144_D])).
fof(f29144_D,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) <=> ~$spl251),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl251])])).
fof(f29228,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | $spl273),
  inference(cnf_transformation,[],[f29228_D])).
fof(f29228_D,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) <=> ~$spl273),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl273])])).
fof(f29246,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f29245,f2381])).
fof(f29245,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f28821,f2381])).
fof(f28821,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl165 | $spl176)),
  inference(superposition,[],[f24778,f26137])).
fof(f26137,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | $spl176),
  inference(cnf_transformation,[],[f26137_D])).
fof(f26137_D,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f24778,plain,(
  ( ! [X1] : (s('type.realax.real',X1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))) | s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X1))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X1))))) ) | $spl165),
  inference(resolution,[],[f17118,f130])).
fof(f17118,plain,(
  ( ! [X4] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X4))),s('type.realax.real',X4)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X4)))) | s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X4)))) ) | $spl165),
  inference(resolution,[],[f16642,f100])).
fof(f100,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f74])).
fof(f16642,plain,(
  ( ! [X3] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X3)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))))) ) | $spl165),
  inference(superposition,[],[f16627,f128])).
fof(f29244,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | ($spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f29243,f2381])).
fof(f29243,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f28820,f2381])).
fof(f28820,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl161 | $spl176)),
  inference(superposition,[],[f24208,f26137])).
fof(f24208,plain,(
  ( ! [X1] : (s('type.realax.real',X1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X1))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X1))))) ) | $spl161),
  inference(resolution,[],[f16945,f130])).
fof(f16945,plain,(
  ( ! [X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X2)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X2)))) ) | $spl161),
  inference(resolution,[],[f16331,f97])).
fof(f16331,plain,(
  ( ! [X3] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X3)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))))))) ) | $spl161),
  inference(superposition,[],[f16319,f128])).
fof(f29242,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | ($spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f29241,f2381])).
fof(f29241,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | ($spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f28819,f2381])).
fof(f28819,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | ($spl161 | $spl176)),
  inference(superposition,[],[f23901,f26137])).
fof(f23901,plain,(
  ( ! [X1] : (s('type.realax.real',X1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',sK1))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',sK1))))) ) | $spl161),
  inference(resolution,[],[f16836,f130])).
fof(f16836,plain,(
  ( ! [X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',sK1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK1)))) ) | $spl161),
  inference(resolution,[],[f16330,f97])).
fof(f16330,plain,(
  ( ! [X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',sK1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))))) ) | $spl161),
  inference(superposition,[],[f16319,f128])).
fof(f29145,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | $spl253),
  inference(cnf_transformation,[],[f29145_D])).
fof(f29145_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) <=> ~$spl253),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl253])])).
fof(f29190,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | $spl262),
  inference(cnf_transformation,[],[f29190_D])).
fof(f29190_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f29240,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f29239,f2381])).
fof(f29239,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f28830,f2381])).
fof(f28830,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl165 | $spl176)),
  inference(trivial_inequality_removal,[],[f28814])).
fof(f28814,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl165 | $spl176)),
  inference(superposition,[],[f23189,f26137])).
fof(f23189,plain,(
  ( ! [X1] : (s('type.realax.real',X1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',sK1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',X1))))) ) | $spl165),
  inference(resolution,[],[f16663,f130])).
fof(f16663,plain,(
  ( ! [X4] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',X4)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',X4)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X4),s('type.realax.real',sK1))))) ) | $spl165),
  inference(superposition,[],[f16628,f128])).
fof(f16628,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',sK1)))) | $spl165),
  inference(resolution,[],[f16627,f129])).
fof(f29236,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ($spl3 | $spl20 | $spl176)),
  inference(forward_demodulation,[],[f29235,f2381])).
fof(f29235,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl3 | $spl20 | $spl176)),
  inference(forward_demodulation,[],[f28831,f2381])).
fof(f28831,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl20 | $spl176)),
  inference(trivial_inequality_removal,[],[f28795])).
fof(f28795,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl20 | $spl176)),
  inference(superposition,[],[f19549,f26137])).
fof(f19549,plain,(
  ( ! [X46] : (s('type.realax.real',X46) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X46)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X46))))) ) | $spl20),
  inference(forward_demodulation,[],[f19548,f1737])).
fof(f1737,plain,(
  ( ! [X1] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(subsumption_resolution,[],[f1694,f132])).
fof(f1694,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f91,f982])).
fof(f19548,plain,(
  ( ! [X45,X46] : (s('type.realax.real',X46) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X46)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X46))))) ) | $spl20),
  inference(subsumption_resolution,[],[f19547,f132])).
fof(f19547,plain,(
  ( ! [X45,X46] : (~p(s(bool,'T_0')) | s('type.realax.real',X46) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X46)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X46))))) ) | $spl20),
  inference(forward_demodulation,[],[f19514,f1460])).
fof(f1460,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | $spl20),
  inference(resolution,[],[f1000,f101])).
fof(f19514,plain,(
  ( ! [X45,X46] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | s('type.realax.real',X46) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X45))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X46)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X46))))) )),
  inference(superposition,[],[f1141,f1737])).
fof(f1141,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3)))) | s('type.realax.real',X2) != s('type.realax.real',X3) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) )),
  inference(superposition,[],[f131,f128])).
fof(f29232,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | $spl275),
  inference(cnf_transformation,[],[f29232_D])).
fof(f29232_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) <=> ~$spl275),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl275])])).
fof(f29230,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | ($spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f29229,f2381])).
fof(f29229,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1))) | ($spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f28783,f2381])).
fof(f28783,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1))) | ($spl161 | $spl176)),
  inference(superposition,[],[f17214,f26137])).
fof(f17214,plain,(
  ( ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',sK1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK1)))) ) | $spl161),
  inference(resolution,[],[f16408,f97])).
fof(f16408,plain,(
  ( ! [X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',sK1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))))))) ) | $spl161),
  inference(superposition,[],[f16320,f128])).
fof(f16320,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | $spl161),
  inference(resolution,[],[f16319,f129])).
fof(f29182,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | $spl259),
  inference(cnf_transformation,[],[f29182_D])).
fof(f29182_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) <=> ~$spl259),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl259])])).
fof(f29226,plain,(
  s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f29225,f2381])).
fof(f29225,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f28781,f2381])).
fof(f28781,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ($spl165 | $spl176)),
  inference(superposition,[],[f17118,f26137])).
fof(f29224,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f29223,f2381])).
fof(f29223,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f28780,f2381])).
fof(f28780,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ($spl165 | $spl176)),
  inference(superposition,[],[f17116,f26137])).
fof(f17116,plain,(
  ( ! [X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X2)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X2)))) ) | $spl165),
  inference(resolution,[],[f16642,f97])).
fof(f29216,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | $spl267),
  inference(cnf_transformation,[],[f29216_D])).
fof(f29216_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) <=> ~$spl267),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl267])])).
fof(f29220,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | ($spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f29219,f2381])).
fof(f29219,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ($spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f28777,f2381])).
fof(f28777,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ($spl161 | $spl176)),
  inference(superposition,[],[f16945,f26137])).
fof(f29214,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | ($spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f29213,f2381])).
fof(f29213,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1))) | ($spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f28774,f2381])).
fof(f28774,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1))) | ($spl161 | $spl176)),
  inference(superposition,[],[f16836,f26137])).
fof(f29194,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | $spl265),
  inference(cnf_transformation,[],[f29194_D])).
fof(f29194_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) <=> ~$spl265),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl265])])).
fof(f28953,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | $spl223),
  inference(cnf_transformation,[],[f28953_D])).
fof(f28953_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) <=> ~$spl223),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl223])])).
fof(f29196,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f29195,f2381])).
fof(f29195,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f28764,f2381])).
fof(f28764,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl165 | $spl176)),
  inference(superposition,[],[f16665,f26137])).
fof(f16665,plain,(
  ( ! [X3] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X3))))) ) | $spl165),
  inference(superposition,[],[f16628,f128])).
fof(f29148,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | $spl255),
  inference(cnf_transformation,[],[f29148_D])).
fof(f29148_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) <=> ~$spl255),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl255])])).
fof(f29192,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f29191,f2381])).
fof(f29191,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f28763,f2381])).
fof(f28763,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | ($spl165 | $spl176)),
  inference(superposition,[],[f16664,f26137])).
fof(f16664,plain,(
  ( ! [X2] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',sK1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK1))))) ) | $spl165),
  inference(superposition,[],[f16628,f128])).
fof(f29188,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | $spl261),
  inference(cnf_transformation,[],[f29188_D])).
fof(f29188_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) <=> ~$spl261),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl261])])).
fof(f29186,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f29185,f2381])).
fof(f29185,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f28762,f2381])).
fof(f28762,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | ($spl165 | $spl176)),
  inference(superposition,[],[f16663,f26137])).
fof(f29180,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f29179,f2381])).
fof(f29179,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | ($spl3 | $spl165 | $spl176)),
  inference(forward_demodulation,[],[f28759,f2381])).
fof(f28759,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | ($spl165 | $spl176)),
  inference(superposition,[],[f16641,f26137])).
fof(f16641,plain,(
  ( ! [X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',sK1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK1))))) ) | $spl165),
  inference(superposition,[],[f16627,f128])).
fof(f29150,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | ($spl3 | $spl157 | $spl176)),
  inference(forward_demodulation,[],[f29149,f2381])).
fof(f29149,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl3 | $spl157 | $spl176)),
  inference(forward_demodulation,[],[f28738,f2381])).
fof(f28738,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl157 | $spl176)),
  inference(superposition,[],[f16160,f26137])).
fof(f16160,plain,(
  ( ! [X3] : (s('type.realax.real',X3) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X3))))) ) | $spl157),
  inference(superposition,[],[f16158,f128])).
fof(f29142,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | ($spl3 | $spl157 | $spl176)),
  inference(forward_demodulation,[],[f29141,f2381])).
fof(f29141,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',sK1)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | ($spl3 | $spl157 | $spl176)),
  inference(forward_demodulation,[],[f28737,f2381])).
fof(f28737,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | ($spl157 | $spl176)),
  inference(superposition,[],[f16159,f26137])).
fof(f16159,plain,(
  ( ! [X2] : (s('type.realax.real',X2) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',sK1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',sK1))))) ) | $spl157),
  inference(superposition,[],[f16158,f128])).
fof(f29071,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f29070,f2381])).
fof(f29070,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f29069,f2381])).
fof(f29069,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | $spl176),
  inference(subsumption_resolution,[],[f29068,f132])).
fof(f29068,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | $spl176),
  inference(forward_demodulation,[],[f28665,f982])).
fof(f28665,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | $spl176),
  inference(superposition,[],[f4585,f26137])).
fof(f29062,plain,(
  ( ! [X79] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X79)))))))) ) | $spl246),
  inference(cnf_transformation,[],[f29062_D])).
fof(f29062_D,plain,(
  ( ! [X79] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X79)))))))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f29061,plain,(
  ( ! [X79] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X79)))))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f29060,f2381])).
fof(f29060,plain,(
  ( ! [X79] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f29059,f2381])).
fof(f29059,plain,(
  ( ! [X79] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))))))) | s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28657,f2381])).
fof(f28657,plain,(
  ( ! [X79] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))))) | s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X79))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | $spl176),
  inference(superposition,[],[f1874,f26137])).
fof(f1874,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) | s('type.realax.real',X0) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real',X0))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0))))) )),
  inference(factoring,[],[f90])).
fof(f29056,plain,(
  ( ! [X74,X75] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X75))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X75))))),s('type.realax.real',X74)))) | s('type.realax.real',X74) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))) ) | $spl244),
  inference(cnf_transformation,[],[f29056_D])).
fof(f29056_D,plain,(
  ( ! [X74,X75] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X75))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X75))))),s('type.realax.real',X74)))) | s('type.realax.real',X74) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f29055,plain,(
  ( ! [X74,X75] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X75))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X75))))),s('type.realax.real',X74)))) | s('type.realax.real',X74) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f29054,f2381])).
fof(f29054,plain,(
  ( ! [X74,X75] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X75))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X75))))),s('type.realax.real',X74)))) | s('type.realax.real',X74) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28653,f2381])).
fof(f28653,plain,(
  ( ! [X74,X75] : (s('type.realax.real',X74) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X75))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X75))))),s('type.realax.real',X74)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | $spl176),
  inference(superposition,[],[f1720,f26137])).
fof(f1720,plain,(
  ( ! [X57,X56,X55] : (s('type.realax.real',X56) != s('type.realax.real',X57) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X55))),s('type.realax.real',X56))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X55))))),s('type.realax.real',X57)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X56))))) )),
  inference(superposition,[],[f130,f91])).
fof(f29051,plain,(
  ( ! [X68,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X68),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X69))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X69)))))))) | s('type.realax.real',X68) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))) ) | $spl242),
  inference(cnf_transformation,[],[f29051_D])).
fof(f29051_D,plain,(
  ( ! [X68,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X68),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X69))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X69)))))))) | s('type.realax.real',X68) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f29050,plain,(
  ( ! [X68,X69] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X68),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X69))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X69)))))))) | s('type.realax.real',X68) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f29049,f2381])).
fof(f29049,plain,(
  ( ! [X68,X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X68),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X69))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X69)))))))) | s('type.realax.real',X68) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28650,f2381])).
fof(f28650,plain,(
  ( ! [X68,X69] : (s('type.realax.real',X68) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X68),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X69))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X69)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | $spl176),
  inference(superposition,[],[f1719,f26137])).
fof(f1719,plain,(
  ( ! [X54,X52,X53] : (s('type.realax.real',X53) != s('type.realax.real',X54) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X54),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X52))),s('type.realax.real',X53))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X52)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X53))))) )),
  inference(superposition,[],[f130,f91])).
fof(f29048,plain,(
  ( ! [X64] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',X64)))) | s('type.realax.real',X64) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))) ) | $spl240),
  inference(cnf_transformation,[],[f29048_D])).
fof(f29048_D,plain,(
  ( ! [X64] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',X64)))) | s('type.realax.real',X64) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f29047,plain,(
  ( ! [X64] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',X64)))) | s('type.realax.real',X64) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f29046,f2381])).
fof(f29046,plain,(
  ( ! [X64] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real',X64)))) | s('type.realax.real',X64) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f29045,f2381])).
fof(f29045,plain,(
  ( ! [X64] : (s('type.realax.real',X64) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',X64)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | $spl176),
  inference(forward_demodulation,[],[f28649,f1737])).
fof(f28649,plain,(
  ( ! [X64,X65] : (s('type.realax.real',X64) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X65))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X65))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',X64)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | $spl176),
  inference(superposition,[],[f1718,f26137])).
fof(f1718,plain,(
  ( ! [X50,X51,X49] : (s('type.realax.real',X51) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))),s('type.realax.real',X50))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X49))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X50),s('type.realax.real',X51)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X50))))) )),
  inference(superposition,[],[f130,f91])).
fof(f29044,plain,(
  ( ! [X60] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X60),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real',X60) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))) ) | $spl238),
  inference(cnf_transformation,[],[f29044_D])).
fof(f29044_D,plain,(
  ( ! [X60] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X60),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real',X60) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f29043,plain,(
  ( ! [X60] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X60),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real',X60) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f29042,f2381])).
fof(f29042,plain,(
  ( ! [X60] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X60),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real',X60) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f29041,f2381])).
fof(f29041,plain,(
  ( ! [X60] : (s('type.realax.real',X60) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X60),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | $spl176),
  inference(forward_demodulation,[],[f28648,f1737])).
fof(f28648,plain,(
  ( ! [X61,X60] : (s('type.realax.real',X60) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X61))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X61))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X60),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | $spl176),
  inference(superposition,[],[f1717,f26137])).
fof(f1717,plain,(
  ( ! [X47,X48,X46] : (s('type.realax.real',X48) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X46))),s('type.realax.real',X47))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X46))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X48),s('type.realax.real',X47)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X47))))) )),
  inference(superposition,[],[f130,f91])).
fof(f29015,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | $spl236),
  inference(cnf_transformation,[],[f29015_D])).
fof(f29015_D,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f29013,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | ($spl0 | $spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f29012,f2381])).
fof(f29012,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl0 | $spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f29011,f2381])).
fof(f29011,plain,(
  s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl0 | $spl161 | $spl176)),
  inference(subsumption_resolution,[],[f29010,f111])).
fof(f111,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X8)))),
  file('thm.transc.ROOT_INV','thm.bool.IMP_CLAUSES|split|4')).
fof(f29010,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl0 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f28627,f16400])).
fof(f16400,plain,(
  s(bool,'F_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | $spl161),
  inference(resolution,[],[f16320,f97])).
fof(f28627,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl0 | $spl176)),
  inference(superposition,[],[f1272,f26137])).
fof(f1272,plain,(
  ( ! [X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X1)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X1))))) ) | $spl0),
  inference(superposition,[],[f921,f128])).
fof(f29007,plain,(
  s(bool,'F_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | $spl235),
  inference(cnf_transformation,[],[f29007_D])).
fof(f29007_D,plain,(
  s(bool,'F_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f29005,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s(bool,'F_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | ($spl0 | $spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f29004,f2381])).
fof(f29004,plain,(
  s(bool,'F_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl0 | $spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f29003,f2381])).
fof(f29003,plain,(
  s(bool,'F_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl0 | $spl161 | $spl176)),
  inference(subsumption_resolution,[],[f29002,f111])).
fof(f29002,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl0 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f28625,f16400])).
fof(f28625,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | s(bool,'F_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl0 | $spl176)),
  inference(superposition,[],[f1257,f26137])).
fof(f1257,plain,(
  ( ! [X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X1)))) | s(bool,'F_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X1))))) ) | $spl0),
  inference(superposition,[],[f919,f128])).
fof(f919,plain,(
  s(bool,'F_0') != s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1))) | $spl0),
  inference(resolution,[],[f159,f96])).
fof(f96,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f72])).
fof(f28991,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | $spl227),
  inference(cnf_transformation,[],[f28991_D])).
fof(f28991_D,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) <=> ~$spl227),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl227])])).
fof(f28997,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28996,f2381])).
fof(f28996,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28623,f2381])).
fof(f28623,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | $spl176),
  inference(superposition,[],[f1238,f26137])).
fof(f1238,plain,(
  ( ! [X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) )),
  inference(superposition,[],[f1154,f128])).
fof(f1154,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(equality_resolution,[],[f131])).
fof(f28993,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl228),
  inference(cnf_transformation,[],[f28993_D])).
fof(f28993_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f28995,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl231),
  inference(cnf_transformation,[],[f28995_D])).
fof(f28995_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl231),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl231])])).
fof(f28989,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28988,f2381])).
fof(f28988,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28622,f2381])).
fof(f28622,plain,(
  ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl176),
  inference(superposition,[],[f1237,f26137])).
fof(f1237,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(superposition,[],[f1154,f128])).
fof(f28964,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | $spl224),
  inference(cnf_transformation,[],[f28964_D])).
fof(f28964_D,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f28962,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | ($spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f28961,f2381])).
fof(f28961,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f28960,f2381])).
fof(f28960,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl161 | $spl176)),
  inference(subsumption_resolution,[],[f28959,f111])).
fof(f28959,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl161 | $spl176)),
  inference(forward_demodulation,[],[f28603,f16400])).
fof(f28603,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | $spl176),
  inference(superposition,[],[f1078,f26137])).
fof(f1078,plain,(
  ( ! [X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X1)))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X1))))) )),
  inference(superposition,[],[f141,f128])).
fof(f28951,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | $spl221),
  inference(cnf_transformation,[],[f28951_D])).
fof(f28951_D,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) <=> ~$spl221),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl221])])).
fof(f28949,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | ($spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f28948,f2381])).
fof(f28948,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl3 | $spl161 | $spl176)),
  inference(forward_demodulation,[],[f28947,f2381])).
fof(f28947,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl161 | $spl176)),
  inference(subsumption_resolution,[],[f28946,f111])).
fof(f28946,plain,(
  p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ($spl161 | $spl176)),
  inference(forward_demodulation,[],[f28599,f16400])).
fof(f28599,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | $spl176),
  inference(superposition,[],[f1007,f26137])).
fof(f1007,plain,(
  ( ! [X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real',X1)))) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X1))))) )),
  inference(superposition,[],[f139,f128])).
fof(f139,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1)))),
  inference(resolution,[],[f88,f96])).
fof(f28910,plain,(
  ( ! [X34] : (s('type.realax.real',sK1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X34)))))) ) | $spl218),
  inference(cnf_transformation,[],[f28910_D])).
fof(f28910_D,plain,(
  ( ! [X34] : (s('type.realax.real',sK1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X34)))))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f28909,plain,(
  ( ! [X34] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real',sK1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X34)))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28908,f2381])).
fof(f28908,plain,(
  ( ! [X34] : (s('type.realax.real',sK1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28580,f2381])).
fof(f28580,plain,(
  ( ! [X34] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))))) | s('type.realax.real',sK1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X34))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | $spl176),
  inference(superposition,[],[f137,f26137])).
fof(f137,plain,(
  ( ! [X4,X5] : (s('type.realax.real',X5) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))) | s('type.realax.real',sK1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.realax.real',X5))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5))))) )),
  inference(resolution,[],[f88,f90])).
fof(f28907,plain,(
  ( ! [X33] : (s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.realax.real',sK1))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33)))))))) ) | $spl216),
  inference(cnf_transformation,[],[f28907_D])).
fof(f28907_D,plain,(
  ( ! [X33] : (s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.realax.real',sK1))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33)))))))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f28906,plain,(
  ( ! [X33] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.realax.real',sK1))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33)))))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28905,f2381])).
fof(f28905,plain,(
  ( ! [X33] : (s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.realax.real',sK1))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28904,f2381])).
fof(f28904,plain,(
  ( ! [X33] : (s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))))))) | s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.realax.real',sK1))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28579,f2381])).
fof(f28579,plain,(
  ( ! [X33] : (s('type.realax.real',sK1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))))) | s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.realax.real',sK1))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | $spl176),
  inference(superposition,[],[f136,f26137])).
fof(f136,plain,(
  ( ! [X2,X3] : (s('type.realax.real',sK1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) | s('type.realax.real',X2) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.realax.real',sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) )),
  inference(resolution,[],[f88,f90])).
fof(f28859,plain,(
  ( ! [X8,X7] : (s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.realax.real',X7))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) | s('type.realax.real',X7) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))) ) | $spl214),
  inference(cnf_transformation,[],[f28859_D])).
fof(f28859_D,plain,(
  ( ! [X8,X7] : (s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.realax.real',X7))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) | s('type.realax.real',X7) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f28858,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.realax.real',X7))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | s('type.realax.real',X7) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28857,f2381])).
fof(f28857,plain,(
  ( ! [X8,X7] : (s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.realax.real',X7))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | s('type.realax.real',X7) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28856,f2381])).
fof(f28856,plain,(
  ( ! [X8,X7] : (s('type.realax.real',X7) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))) | s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.realax.real',X7))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28552,f2381])).
fof(f28552,plain,(
  ( ! [X8,X7] : (s('type.realax.real',X7) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))) | s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.realax.real',X7))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7))))) ) | $spl176),
  inference(superposition,[],[f90,f26137])).
fof(f28855,plain,(
  ( ! [X6,X5] : (s('type.realax.real',X5) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) ) | $spl212),
  inference(cnf_transformation,[],[f28855_D])).
fof(f28855_D,plain,(
  ( ! [X6,X5] : (s('type.realax.real',X5) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f28854,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real',X5) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28853,f2381])).
fof(f28853,plain,(
  ( ! [X6,X5] : (s('type.realax.real',X5) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28551,f2381])).
fof(f28551,plain,(
  ( ! [X6,X5] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))) | s('type.realax.real',X5) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | $spl176),
  inference(superposition,[],[f90,f26137])).
fof(f28836,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1)))) ) | $spl208),
  inference(cnf_transformation,[],[f28836_D])).
fof(f28836_D,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1)))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f28852,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real',X4) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.realax.real',X4))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28851,f2381])).
fof(f28851,plain,(
  ( ! [X4,X3] : (s('type.realax.real',X4) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.realax.real',X4))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28550,f2381])).
fof(f28550,plain,(
  ( ! [X4,X3] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.realax.real',X4))) | s('type.realax.real',X4) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X4))))) ) | $spl176),
  inference(superposition,[],[f90,f26137])).
fof(f28850,plain,(
  ( ! [X2,X1] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real',X1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) ) | $spl210),
  inference(cnf_transformation,[],[f28850_D])).
fof(f28850_D,plain,(
  ( ! [X2,X1] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real',X1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f28849,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | s('type.realax.real',X1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28848,f2381])).
fof(f28848,plain,(
  ( ! [X2,X1] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | s('type.realax.real',X1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28549,f2381])).
fof(f28549,plain,(
  ( ! [X2,X1] : (s('type.realax.real',X1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | s('type.realax.real','const.real.pow_2'(s('type.realax.real',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))))) ) | $spl176),
  inference(superposition,[],[f90,f26137])).
fof(f28835,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28834,f2381])).
fof(f28834,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | ($spl3 | $spl176)),
  inference(forward_demodulation,[],[f28540,f2381])).
fof(f28540,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | $spl176),
  inference(superposition,[],[f26137,f90])).
fof(f28495,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X2))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2)))))))) ) | $spl204),
  inference(cnf_transformation,[],[f28495_D])).
fof(f28495_D,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X2))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2)))))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f28496,plain,(
  ( ! [X3] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X3))))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X3)))) ) | $spl206),
  inference(cnf_transformation,[],[f28496_D])).
fof(f28496_D,plain,(
  ( ! [X3] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X3))))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X3)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f28494,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X3))))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X3))) | s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X2))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2)))))))) ) | ($spl3 | $spl4)),
  inference(forward_demodulation,[],[f5065,f4702])).
fof(f4702,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) ) | ($spl3 | $spl4)),
  inference(forward_demodulation,[],[f173,f2381])).
fof(f5065,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X2))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X3))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X3)))) ) | ($spl3 | $spl4)),
  inference(superposition,[],[f4702,f104])).
fof(f4484,plain,(
  ( ! [X15,X13] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X13),s('type.num.num',X15))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X13))),s('type.num.num',X15)))) ) | $spl92),
  inference(cnf_transformation,[],[f4484_D])).
fof(f4484_D,plain,(
  ( ! [X15,X13] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X13),s('type.num.num',X15))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X13))),s('type.num.num',X15)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f27219,plain,(
  ( ! [X14] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl202),
  inference(cnf_transformation,[],[f27219_D])).
fof(f27219_D,plain,(
  ( ! [X14] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f14352,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X13),s('type.num.num',X15))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X13))),s('type.num.num',X15)))) )),
  inference(trivial_inequality_removal,[],[f14330])).
fof(f14330,plain,(
  ( ! [X14,X15,X13] : (s('type.realax.real',X13) != s('type.realax.real',X13) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X13),s('type.num.num',X15))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X13))),s('type.num.num',X15)))) )),
  inference(superposition,[],[f1137,f104])).
fof(f1137,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X3) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) )),
  inference(superposition,[],[f131,f128])).
fof(f26836,plain,(
  ( ! [X14] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14))))) ) | $spl200),
  inference(cnf_transformation,[],[f26836_D])).
fof(f26836_D,plain,(
  ( ! [X14] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f4493,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X9),s('type.realax.real',X9)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X9),s('type.num.num',X10))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.num.num',X10)))) ) | $spl96),
  inference(cnf_transformation,[],[f4493_D])).
fof(f4493_D,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X9),s('type.realax.real',X9)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X9),s('type.num.num',X10))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.num.num',X10)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f19532,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real',X13)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X13),s('type.num.num',X15))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X13))),s('type.num.num',X15)))) )),
  inference(subsumption_resolution,[],[f19501,f128])).
fof(f19501,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real',X13)))) | s('type.realax.real',X14) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X13),s('type.num.num',X15))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X13))),s('type.num.num',X15)))) )),
  inference(superposition,[],[f1141,f104])).
fof(f26827,plain,(
  ( ! [X14] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14))))) ) | $spl198),
  inference(cnf_transformation,[],[f26827_D])).
fof(f26827_D,plain,(
  ( ! [X14] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f19822,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X13),s('type.num.num',X15))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X13))),s('type.num.num',X15)))) )),
  inference(trivial_inequality_removal,[],[f19800])).
fof(f19800,plain,(
  ( ! [X14,X15,X13] : (s('type.realax.real',X13) != s('type.realax.real',X13) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X14)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X13),s('type.num.num',X15))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X13))),s('type.num.num',X15)))) )),
  inference(superposition,[],[f1143,f104])).
fof(f1143,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X3) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X2),s('type.realax.real',X3)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) )),
  inference(superposition,[],[f131,f128])).
fof(f1434,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X29),s('type.realax.real',X29)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X29),s('type.num.num',X30))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X29))),s('type.num.num',X30)))) ) | $spl34),
  inference(cnf_transformation,[],[f1434_D])).
fof(f1434_D,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X29),s('type.realax.real',X29)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X29),s('type.num.num',X30))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X29))),s('type.num.num',X30)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f24438,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X13)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X14)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X14),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X13),s('type.num.num',X15))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X13))),s('type.num.num',X15)))) )),
  inference(superposition,[],[f1189,f104])).
fof(f26244,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | ($spl3 | $spl177)),
  inference(forward_demodulation,[],[f26181,f2381])).
fof(f26181,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | ($spl3 | $spl177)),
  inference(forward_demodulation,[],[f26174,f2381])).
fof(f26174,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | $spl177),
  inference(superposition,[],[f26138,f90])).
fof(f26206,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1)))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.realax.real',sK1))) | $spl183),
  inference(resolution,[],[f26202,f7187])).
fof(f26205,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))) | $spl183),
  inference(resolution,[],[f26202,f7446])).
fof(f26184,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1)))) ) | $spl178),
  inference(cnf_transformation,[],[f26184_D])).
fof(f26184_D,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1)))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f26204,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | ($spl92 | $spl177)),
  inference(forward_demodulation,[],[f26197,f4484])).
fof(f26197,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | ($spl92 | $spl177)),
  inference(superposition,[],[f26171,f90])).
fof(f26193,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | ($spl92 | $spl177)),
  inference(superposition,[],[f26171,f90])).
fof(f26183,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))))) | s('type.realax.real',X1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | ($spl3 | $spl92 | $spl177)),
  inference(forward_demodulation,[],[f26182,f2381])).
fof(f26182,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | ($spl3 | $spl92 | $spl177)),
  inference(forward_demodulation,[],[f26181,f4484])).
fof(f26136,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | $spl175),
  inference(cnf_transformation,[],[f26136_D])).
fof(f26136_D,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) <=> ~$spl175),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl175])])).
fof(f26138,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | $spl177),
  inference(cnf_transformation,[],[f26138_D])).
fof(f26138_D,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) <=> ~$spl177),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl177])])).
fof(f26134,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | $spl10),
  inference(forward_demodulation,[],[f26133,f1737])).
fof(f26133,plain,(
  ( ! [X27] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f26132,f132])).
fof(f26132,plain,(
  ( ! [X27] : (~p(s(bool,'T_0')) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))) ) | $spl10),
  inference(forward_demodulation,[],[f26131,f982])).
fof(f26131,plain,(
  ( ! [X27] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))) ) | $spl10),
  inference(forward_demodulation,[],[f26094,f1737])).
fof(f26094,plain,(
  ( ! [X27] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))) ) | $spl10),
  inference(superposition,[],[f178,f1737])).
fof(f26128,plain,(
  s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | $spl171),
  inference(cnf_transformation,[],[f26128_D])).
fof(f26128_D,plain,(
  s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) <=> ~$spl171),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl171])])).
fof(f26130,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | $spl173),
  inference(cnf_transformation,[],[f26130_D])).
fof(f26130_D,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) <=> ~$spl173),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl173])])).
fof(f26126,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | $spl10),
  inference(forward_demodulation,[],[f26125,f135])).
fof(f26125,plain,(
  ( ! [X26] : (s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))) ) | $spl10),
  inference(subsumption_resolution,[],[f26124,f132])).
fof(f26124,plain,(
  ( ! [X26] : (~p(s(bool,'T_0')) | s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))) ) | $spl10),
  inference(forward_demodulation,[],[f26123,f141])).
fof(f26123,plain,(
  ( ! [X26] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1)))) | s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))) ) | $spl10),
  inference(forward_demodulation,[],[f26093,f135])).
fof(f26093,plain,(
  ( ! [X26] : (s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26)))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))) ) | $spl10),
  inference(superposition,[],[f178,f135])).
fof(f19428,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl169),
  inference(cnf_transformation,[],[f19428_D])).
fof(f19428_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl169),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl169])])).
fof(f19426,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl166),
  inference(cnf_transformation,[],[f19426_D])).
fof(f19426_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f19417,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl20),
  inference(equality_resolution,[],[f1486])).
fof(f1486,plain,(
  ( ! [X2] : (s('type.realax.real',X2) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',X2))))) ) | $spl20),
  inference(superposition,[],[f1456,f128])).
fof(f1456,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | $spl20),
  inference(resolution,[],[f1000,f131])).
fof(f19435,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl20),
  inference(forward_demodulation,[],[f19434,f106])).
fof(f19434,plain,(
  ( ! [X16] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))))))))) ) | $spl20),
  inference(forward_demodulation,[],[f19418,f106])).
fof(f19418,plain,(
  ( ! [X16] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))))))))) ) | $spl20),
  inference(trivial_inequality_removal,[],[f19404])).
fof(f19404,plain,(
  ( ! [X16] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))))))))) ) | $spl20),
  inference(superposition,[],[f1486,f106])).
fof(f19433,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl20 | $spl48)),
  inference(forward_demodulation,[],[f19432,f4625])).
fof(f19432,plain,(
  ( ! [X15] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl20 | $spl48)),
  inference(forward_demodulation,[],[f19419,f4625])).
fof(f19419,plain,(
  ( ! [X15] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl20 | $spl48)),
  inference(trivial_inequality_removal,[],[f19403])).
fof(f19403,plain,(
  ( ! [X15] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl20 | $spl48)),
  inference(superposition,[],[f1486,f4625])).
fof(f19424,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl20),
  inference(forward_demodulation,[],[f19423,f107])).
fof(f19423,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl20),
  inference(forward_demodulation,[],[f19420,f107])).
fof(f19420,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl20),
  inference(trivial_inequality_removal,[],[f19400])).
fof(f19400,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl20),
  inference(superposition,[],[f1486,f107])).
fof(f16627,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',sK1)))) | $spl165),
  inference(cnf_transformation,[],[f16627_D])).
fof(f16627_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',sK1)))) <=> ~$spl165),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl165])])).
fof(f16625,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',sK1)))) | $spl162),
  inference(cnf_transformation,[],[f16625_D])).
fof(f16625_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',sK1)))) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f16608,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',sK1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',sK1)))) | $spl157),
  inference(equality_resolution,[],[f16163])).
fof(f16163,plain,(
  ( ! [X4] : (s('type.realax.real',sK1) != s('type.realax.real',X4) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',X4)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',X4))))) ) | $spl157),
  inference(superposition,[],[f16158,f128])).
fof(f16623,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',sK1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',sK1)))) | $spl157),
  inference(forward_demodulation,[],[f16622,f135])).
fof(f16622,plain,(
  ( ! [X27] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real',sK1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))))))) ) | $spl157),
  inference(forward_demodulation,[],[f16609,f135])).
fof(f16609,plain,(
  ( ! [X27] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))))))) ) | $spl157),
  inference(trivial_inequality_removal,[],[f16606])).
fof(f16606,plain,(
  ( ! [X27] : (s('type.realax.real',sK1) != s('type.realax.real',sK1) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))))))) ) | $spl157),
  inference(superposition,[],[f16163,f135])).
fof(f16319,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | $spl161),
  inference(cnf_transformation,[],[f16319_D])).
fof(f16319_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f16317,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | $spl158),
  inference(cnf_transformation,[],[f16317_D])).
fof(f16317_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f16300,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | $spl157),
  inference(equality_resolution,[],[f16162])).
fof(f16162,plain,(
  ( ! [X3] : (s('type.realax.real',sK1) != s('type.realax.real',X3) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) ) | $spl157),
  inference(superposition,[],[f16158,f128])).
fof(f16315,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | $spl157),
  inference(forward_demodulation,[],[f16314,f135])).
fof(f16314,plain,(
  ( ! [X27] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) ) | $spl157),
  inference(forward_demodulation,[],[f16301,f135])).
fof(f16301,plain,(
  ( ! [X27] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) ) | $spl157),
  inference(trivial_inequality_removal,[],[f16298])).
fof(f16298,plain,(
  ( ! [X27] : (s('type.realax.real',sK1) != s('type.realax.real',sK1) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) ) | $spl157),
  inference(superposition,[],[f16162,f135])).
fof(f16156,plain,(
  ( ! [X12] : (s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) ) | $spl154),
  inference(cnf_transformation,[],[f16156_D])).
fof(f16156_D,plain,(
  ( ! [X12] : (s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f16158,plain,(
  s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | $spl157),
  inference(cnf_transformation,[],[f16158_D])).
fof(f16158_D,plain,(
  s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f16155,plain,(
  ( ! [X12] : (s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) )),
  inference(subsumption_resolution,[],[f16154,f132])).
fof(f16154,plain,(
  ( ! [X12] : (~p(s(bool,'T_0')) | s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) )),
  inference(forward_demodulation,[],[f16114,f2889])).
fof(f2889,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) )),
  inference(resolution,[],[f1501,f101])).
fof(f1501,plain,(
  ( ! [X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))) )),
  inference(resolution,[],[f109,f134])).
fof(f16114,plain,(
  ( ! [X12] : (s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))) )),
  inference(superposition,[],[f136,f1737])).
fof(f12301,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))))))))) | $spl150),
  inference(cnf_transformation,[],[f12301_D])).
fof(f12301_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))))))))) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f12303,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))))))) | $spl152),
  inference(cnf_transformation,[],[f12303_D])).
fof(f12303_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))))))) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f12248,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))))))))))),
  inference(resolution,[],[f8015,f128])).
fof(f8015,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))))))))))),
  inference(resolution,[],[f5940,f108])).
fof(f5940,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))))))))),
  inference(resolution,[],[f3812,f108])).
fof(f3812,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))))))),
  inference(resolution,[],[f2528,f108])).
fof(f2528,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))))),
  inference(resolution,[],[f1745,f108])).
fof(f1745,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))),
  inference(resolution,[],[f1177,f108])).
fof(f1177,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))),
  inference(resolution,[],[f108,f134])).
fof(f11062,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))))) | $spl146),
  inference(cnf_transformation,[],[f11062_D])).
fof(f11062_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))))) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f11064,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))) | $spl148),
  inference(cnf_transformation,[],[f11064_D])).
fof(f11064_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f11018,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))))) | $spl16),
  inference(resolution,[],[f6707,f128])).
fof(f6707,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))))) | $spl16),
  inference(resolution,[],[f4192,f108])).
fof(f4192,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))) | $spl16),
  inference(resolution,[],[f2597,f108])).
fof(f2597,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))) | $spl16),
  inference(resolution,[],[f182,f108])).
fof(f182,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | $spl16),
  inference(cnf_transformation,[],[f182_D])).
fof(f182_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f8916,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl145),
  inference(cnf_transformation,[],[f8916_D])).
fof(f8916_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl145),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl145])])).
fof(f8914,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl142),
  inference(cnf_transformation,[],[f8914_D])).
fof(f8914_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f8897,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl3),
  inference(equality_resolution,[],[f896])).
fof(f896,plain,(
  ( ! [X0] : (s('type.realax.real',sK1) != s('type.realax.real',X0) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl3),
  inference(superposition,[],[f160,f128])).
fof(f8912,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl3),
  inference(forward_demodulation,[],[f8911,f135])).
fof(f8911,plain,(
  ( ! [X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',sK1),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl3),
  inference(forward_demodulation,[],[f8898,f135])).
fof(f8898,plain,(
  ( ! [X25] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f8895])).
fof(f8895,plain,(
  ( ! [X25] : (s('type.realax.real',sK1) != s('type.realax.real',sK1) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))),s('type.realax.real',sK1))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X25))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl3),
  inference(superposition,[],[f896,f135])).
fof(f8064,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))))))) | $spl138),
  inference(cnf_transformation,[],[f8064_D])).
fof(f8064_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))))))) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f8066,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))))) | $spl140),
  inference(cnf_transformation,[],[f8066_D])).
fof(f8066_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f8016,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))))))))),
  inference(resolution,[],[f5940,f128])).
fof(f6747,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))) | $spl134),
  inference(cnf_transformation,[],[f6747_D])).
fof(f6747_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f6749,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))) | $spl136),
  inference(cnf_transformation,[],[f6749_D])).
fof(f6749_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f6708,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))) | $spl16),
  inference(resolution,[],[f4192,f128])).
fof(f5984,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))))) | $spl130),
  inference(cnf_transformation,[],[f5984_D])).
fof(f5984_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f5986,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))) | $spl132),
  inference(cnf_transformation,[],[f5986_D])).
fof(f5986_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f5941,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))))))),
  inference(resolution,[],[f3812,f128])).
fof(f4485,plain,(
  ( ! [X14] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl94),
  inference(cnf_transformation,[],[f4485_D])).
fof(f4485_D,plain,(
  ( ! [X14] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f4628,plain,(
  ( ! [X14,X15,X13] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X13),s('type.num.num',X15))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X13))),s('type.num.num',X15)))) ) | $spl48),
  inference(trivial_inequality_removal,[],[f4613])).
fof(f4613,plain,(
  ( ! [X14,X15,X13] : (s('type.realax.real',X13) != s('type.realax.real',X13) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X13),s('type.num.num',X15))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X13))),s('type.num.num',X15)))) ) | $spl48),
  inference(superposition,[],[f2083,f104])).
fof(f4561,plain,(
  $false | $spl27),
  inference(trivial_inequality_removal,[],[f4542])).
fof(f4542,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl27),
  inference(resolution,[],[f1428,f130])).
fof(f1428,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl27),
  inference(cnf_transformation,[],[f1428_D])).
fof(f1428_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f4540,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.realax.real',X9)))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X9),s('type.num.num',X11))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.num.num',X11)))) ) | $spl128),
  inference(cnf_transformation,[],[f4540_D])).
fof(f4540_D,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.realax.real',X9)))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X9),s('type.num.num',X11))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.num.num',X11)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f1519,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.realax.real',X9)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X9),s('type.num.num',X11))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.num.num',X11)))) )),
  inference(superposition,[],[f109,f104])).
fof(f4539,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))))),s('type.num.num',X2)))) ) | $spl126),
  inference(cnf_transformation,[],[f4539_D])).
fof(f4539_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))))),s('type.num.num',X2)))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f1429,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl26),
  inference(cnf_transformation,[],[f1429_D])).
fof(f1429_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1523,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f109,f104])).
fof(f4538,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real',X0)))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X2)))) ) | $spl124),
  inference(cnf_transformation,[],[f4538_D])).
fof(f4538_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real',X0)))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X2)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f1536,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real',X0)))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f109,f104])).
fof(f4535,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X0) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X2)))) ) | $spl122),
  inference(cnf_transformation,[],[f4535_D])).
fof(f4535_D,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X0) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X2)))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f1670,plain,(
  ( ! [X10,X11,X9] : (s('type.realax.real',X9) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.realax.real',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X9),s('type.num.num',X11))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.num.num',X11)))) )),
  inference(superposition,[],[f91,f104])).
fof(f1685,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f91,f104])).
fof(f3261,plain,(
  ( ! [X3] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X3))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X3)))) ) | $spl72),
  inference(cnf_transformation,[],[f3261_D])).
fof(f3261_D,plain,(
  ( ! [X3] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X3))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X3)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4532,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X0))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X0))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl3),
  inference(forward_demodulation,[],[f1776,f2381])).
fof(f1776,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X0))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f1760,f107])).
fof(f1760,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X0))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X0)))) )),
  inference(superposition,[],[f1177,f104])).
fof(f4496,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2)))) ) | $spl102),
  inference(cnf_transformation,[],[f4496_D])).
fof(f4496_D,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1763,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X0))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X0)))) )),
  inference(superposition,[],[f1177,f104])).
fof(f2083,plain,(
  ( ! [X19,X18] : (s('type.realax.real',X19) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))),s('type.realax.real',X19)))) ) | $spl48),
  inference(cnf_transformation,[],[f2083_D])).
fof(f2083_D,plain,(
  ( ! [X19,X18] : (s('type.realax.real',X19) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))),s('type.realax.real',X19)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2050,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X29) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.realax.real',X29)))) )),
  inference(forward_demodulation,[],[f2049,f106])).
fof(f2049,plain,(
  ( ! [X28,X29,X27] : (s('type.realax.real',X29) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.realax.real',X29))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))))))) )),
  inference(subsumption_resolution,[],[f2048,f130])).
fof(f2048,plain,(
  ( ! [X28,X29,X27] : (s('type.realax.real',X29) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.realax.real',X29))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X29))))) )),
  inference(forward_demodulation,[],[f2047,f106])).
fof(f2047,plain,(
  ( ! [X28,X29,X27] : (s('type.realax.real',X29) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.realax.real',X29))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X29))))) )),
  inference(forward_demodulation,[],[f1880,f106])).
fof(f1880,plain,(
  ( ! [X28,X29,X27] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.realax.real',X29))) | s('type.realax.real',X29) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X29))))) )),
  inference(superposition,[],[f90,f106])).
fof(f4531,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.realax.real',X18) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.realax.real',X17))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X17),s('type.num.num',X20))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))),s('type.num.num',X20))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X18)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) ) | $spl120),
  inference(cnf_transformation,[],[f4531_D])).
fof(f4531_D,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.realax.real',X18) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.realax.real',X17))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X17),s('type.num.num',X20))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))),s('type.num.num',X20))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X18)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f1887,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.realax.real',X18) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.realax.real',X17))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X18)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X17),s('type.num.num',X20))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))),s('type.num.num',X20)))) )),
  inference(superposition,[],[f90,f104])).
fof(f4530,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.realax.real',X18) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X17),s('type.num.num',X20))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))),s('type.num.num',X20))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X18)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.realax.real',X18)))) ) | $spl118),
  inference(cnf_transformation,[],[f4530_D])).
fof(f4530_D,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.realax.real',X18) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X17),s('type.num.num',X20))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))),s('type.num.num',X20))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X18)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.realax.real',X18)))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f1910,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.realax.real',X18) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.realax.real',X18))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X18)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X17),s('type.num.num',X20))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))),s('type.num.num',X20)))) )),
  inference(superposition,[],[f90,f104])).
fof(f2068,plain,(
  ( ! [X21,X22] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.realax.real',X21))) | s('type.realax.real',X21) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2067,f107])).
fof(f2067,plain,(
  ( ! [X21,X22] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.realax.real',X21))) | s('type.realax.real',X21) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(forward_demodulation,[],[f2066,f107])).
fof(f2066,plain,(
  ( ! [X21,X22] : (s('type.realax.real',X21) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.realax.real',X21))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(subsumption_resolution,[],[f2065,f130])).
fof(f2065,plain,(
  ( ! [X21,X22] : (s('type.realax.real',X21) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.realax.real',X21))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X21))))) )),
  inference(forward_demodulation,[],[f1911,f106])).
fof(f1911,plain,(
  ( ! [X21,X22] : (s('type.realax.real',X21) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.realax.real',X21))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X21))))) )),
  inference(superposition,[],[f90,f107])).
fof(f2074,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))),s('type.realax.real',X29))) | s('type.realax.real',X29) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2073,f106])).
fof(f2073,plain,(
  ( ! [X30,X28,X29] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))),s('type.realax.real',X29))) | s('type.realax.real',X29) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))))))))) )),
  inference(forward_demodulation,[],[f2072,f106])).
fof(f2072,plain,(
  ( ! [X30,X28,X29] : (s('type.realax.real',X29) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))),s('type.realax.real',X29))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))))))))) )),
  inference(subsumption_resolution,[],[f2071,f130])).
fof(f2071,plain,(
  ( ! [X30,X28,X29] : (s('type.realax.real',X29) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))),s('type.realax.real',X29))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X28)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X29))))) )),
  inference(forward_demodulation,[],[f1913,f106])).
fof(f1913,plain,(
  ( ! [X30,X28,X29] : (s('type.realax.real',X29) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))))) | s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))),s('type.realax.real',X29))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X28)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X29))))) )),
  inference(superposition,[],[f90,f106])).
fof(f2076,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(subsumption_resolution,[],[f1916,f130])).
fof(f1916,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) )),
  inference(superposition,[],[f90,f106])).
fof(f4529,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.realax.real',X5) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.realax.real',X7))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X5),s('type.num.num',X8))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.num.num',X8))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) | s('type.realax.real',X7) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) ) | $spl116),
  inference(cnf_transformation,[],[f4529_D])).
fof(f4529_D,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.realax.real',X5) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.realax.real',X7))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X5),s('type.num.num',X8))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.num.num',X8))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) | s('type.realax.real',X7) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f1928,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X5) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.realax.real',X7))) | s('type.realax.real',X7) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X5),s('type.num.num',X8))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.num.num',X8)))) )),
  inference(superposition,[],[f90,f104])).
fof(f4524,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.realax.real',X6) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.realax.real',X5))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X5),s('type.num.num',X8))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.num.num',X8))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) | s('type.realax.real',X5) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))) ) | $spl114),
  inference(cnf_transformation,[],[f4524_D])).
fof(f4524_D,plain,(
  ( ! [X6,X8,X7,X5] : (s('type.realax.real',X6) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.realax.real',X5))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X5),s('type.num.num',X8))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.num.num',X8))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) | s('type.realax.real',X5) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f1950,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X6) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.realax.real',X5))) | s('type.realax.real',X5) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X6)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X5),s('type.num.num',X8))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.num.num',X8)))) )),
  inference(superposition,[],[f90,f104])).
fof(f2114,plain,(
  ( ! [X169,X168] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X169) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X168))),s('type.realax.real',X169)))) )),
  inference(forward_demodulation,[],[f2113,f106])).
fof(f2113,plain,(
  ( ! [X167,X169,X168] : (s('type.realax.real',X169) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X168))),s('type.realax.real',X169))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X167))))))))) )),
  inference(subsumption_resolution,[],[f2112,f130])).
fof(f2112,plain,(
  ( ! [X167,X169,X168] : (s('type.realax.real',X169) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X168))),s('type.realax.real',X169))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X167)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X169))))) )),
  inference(forward_demodulation,[],[f2111,f106])).
fof(f2111,plain,(
  ( ! [X167,X169,X168] : (s('type.realax.real',X169) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X168))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X168))),s('type.realax.real',X169))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X167)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X169))))) )),
  inference(forward_demodulation,[],[f1996,f106])).
fof(f1996,plain,(
  ( ! [X167,X169,X168] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X168))),s('type.realax.real',X169))) | s('type.realax.real',X169) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X167))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X168))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X167)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X169))))) )),
  inference(superposition,[],[f106,f90])).
fof(f4513,plain,(
  ( ! [X6,X5] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))),s('type.num.num',X6))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))))),s('type.num.num',X6)))) ) | $spl108),
  inference(cnf_transformation,[],[f4513_D])).
fof(f4513_D,plain,(
  ( ! [X6,X5] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))),s('type.num.num',X6))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))))),s('type.num.num',X6)))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f2222,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))),s('type.num.num',X6))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))))),s('type.num.num',X6)))) )),
  inference(superposition,[],[f133,f104])).
fof(f4519,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl3),
  inference(forward_demodulation,[],[f2563,f2381])).
fof(f2563,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f2562,f107])).
fof(f2562,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f2543,f107])).
fof(f2543,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f1745,f104])).
fof(f2565,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f2546,f107])).
fof(f2546,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f1745,f104])).
fof(f4495,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))),s('type.num.num',X2)))) ) | $spl100),
  inference(cnf_transformation,[],[f4495_D])).
fof(f4495_D,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))),s('type.num.num',X2)))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f2549,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f1745,f104])).
fof(f1588,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl39),
  inference(cnf_transformation,[],[f1588_D])).
fof(f1588_D,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f4518,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl3),
  inference(forward_demodulation,[],[f2590,f2381])).
fof(f2590,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f2575,f107])).
fof(f2575,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f1747,f104])).
fof(f1747,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))),
  inference(resolution,[],[f1177,f96])).
fof(f2578,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f1747,f104])).
fof(f179,plain,(
  ( ! [X3] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X3))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X3)))) ) | $spl12),
  inference(cnf_transformation,[],[f179_D])).
fof(f179_D,plain,(
  ( ! [X3] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X3))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X3)))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2617,plain,(
  ( ! [X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X3))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X3)))) ) | $spl16),
  inference(superposition,[],[f182,f104])).
fof(f1637,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl40),
  inference(cnf_transformation,[],[f1637_D])).
fof(f1637_D,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4517,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl3),
  inference(forward_demodulation,[],[f2687,f2381])).
fof(f2687,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f2642,f107])).
fof(f2642,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f1749,f104])).
fof(f1749,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))),
  inference(resolution,[],[f1177,f101])).
fof(f2645,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f1749,f104])).
fof(f1550,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl36),
  inference(cnf_transformation,[],[f1550_D])).
fof(f1550_D,plain,(
  ( ! [X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4516,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X3))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X3))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl3),
  inference(forward_demodulation,[],[f2900,f2381])).
fof(f2900,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X3))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X3)))) )),
  inference(superposition,[],[f1501,f104])).
fof(f4515,plain,(
  ( ! [X6,X5] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X6))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X6)))) ) | $spl112),
  inference(cnf_transformation,[],[f4515_D])).
fof(f4515_D,plain,(
  ( ! [X6,X5] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X6))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X6)))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f2905,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X6))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X6)))) )),
  inference(superposition,[],[f1501,f104])).
fof(f2951,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))),s('type.num.num',X6))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))))),s('type.num.num',X6)))) )),
  inference(forward_demodulation,[],[f2937,f107])).
fof(f2937,plain,(
  ( ! [X6,X5] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))),s('type.num.num',X6))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))))),s('type.num.num',X6)))) )),
  inference(superposition,[],[f2205,f104])).
fof(f4514,plain,(
  ( ! [X4,X3] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.realax.real',sK1))))),s('type.num.num',X4))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.realax.real',sK1))))))),s('type.num.num',X4)))) ) | $spl110),
  inference(cnf_transformation,[],[f4514_D])).
fof(f4514_D,plain,(
  ( ! [X4,X3] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.realax.real',sK1))))),s('type.num.num',X4))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.realax.real',sK1))))))),s('type.num.num',X4)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f2941,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.realax.real',sK1))))),s('type.num.num',X4))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.realax.real',sK1))))))),s('type.num.num',X4)))) )),
  inference(superposition,[],[f2205,f104])).
fof(f2987,plain,(
  ( ! [X6,X5] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))),s('type.num.num',X6))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))))),s('type.num.num',X6)))) )),
  inference(superposition,[],[f2207,f104])).
fof(f2207,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))),s('type.realax.real',sK1)))))) )),
  inference(resolution,[],[f133,f96])).
fof(f3011,plain,(
  ( ! [X6,X5] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))),s('type.num.num',X6))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))),s('type.realax.real',sK1))))),s('type.num.num',X6)))) )),
  inference(superposition,[],[f2209,f104])).
fof(f2209,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.realax.real',sK1)))))) )),
  inference(resolution,[],[f133,f101])).
fof(f4512,plain,(
  ( ! [X2,X1] : (s('type.realax.real',X2) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real',X2))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl70),
  inference(subsumption_resolution,[],[f4511,f130])).
fof(f4511,plain,(
  ( ! [X2,X1] : (s('type.realax.real',X2) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real',X2))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) ) | $spl70),
  inference(forward_demodulation,[],[f3573,f3567])).
fof(f3567,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))))) ) | $spl70),
  inference(superposition,[],[f3260,f106])).
fof(f3573,plain,(
  ( ! [X2,X1] : (s('type.realax.real',X2) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real',X2))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) ) | $spl70),
  inference(superposition,[],[f90,f3260])).
fof(f3703,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X3) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.realax.real',X3)))) ) | $spl70),
  inference(forward_demodulation,[],[f3702,f3567])).
fof(f3702,plain,(
  ( ! [X2,X3,X1] : (s('type.realax.real',X3) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.realax.real',X3))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))))))) ) | $spl70),
  inference(subsumption_resolution,[],[f3701,f130])).
fof(f3701,plain,(
  ( ! [X2,X3,X1] : (s('type.realax.real',X3) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.realax.real',X3))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3))))) ) | $spl70),
  inference(forward_demodulation,[],[f3700,f106])).
fof(f3700,plain,(
  ( ! [X2,X3,X1] : (s('type.realax.real',X3) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.realax.real',X3))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3))))) ) | $spl70),
  inference(forward_demodulation,[],[f3663,f3567])).
fof(f3663,plain,(
  ( ! [X2,X3,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.realax.real',X3))) | s('type.realax.real',X3) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X3))))) ) | $spl70),
  inference(superposition,[],[f3567,f90])).
fof(f4510,plain,(
  ( ! [X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))))) ) | $spl106),
  inference(cnf_transformation,[],[f4510_D])).
fof(f4510_D,plain,(
  ( ! [X2] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f3667,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))))))) ) | $spl70),
  inference(superposition,[],[f108,f3567])).
fof(f3717,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real',X8) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.realax.real',X8)))) ) | $spl70),
  inference(forward_demodulation,[],[f3716,f3567])).
fof(f3716,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real',X8) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.realax.real',X8))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))))))) ) | $spl70),
  inference(subsumption_resolution,[],[f3715,f130])).
fof(f3715,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real',X8) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.realax.real',X8))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))))) ) | $spl70),
  inference(forward_demodulation,[],[f3714,f106])).
fof(f3714,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real',X8) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.realax.real',X8))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))))) ) | $spl70),
  inference(forward_demodulation,[],[f3669,f3567])).
fof(f3669,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.realax.real',X8))) | s('type.realax.real',X8) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X8))))) ) | $spl70),
  inference(superposition,[],[f90,f3567])).
fof(f3725,plain,(
  ( ! [X14,X13] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.realax.real',X13))) | s('type.realax.real',X13) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl70),
  inference(forward_demodulation,[],[f3724,f3567])).
fof(f3724,plain,(
  ( ! [X14,X12,X13] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.realax.real',X13))) | s('type.realax.real',X13) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))))))))))) ) | $spl70),
  inference(forward_demodulation,[],[f3723,f3567])).
fof(f3723,plain,(
  ( ! [X14,X12,X13] : (s('type.realax.real',X13) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.realax.real',X13))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))))))))))) ) | $spl70),
  inference(subsumption_resolution,[],[f3722,f130])).
fof(f3722,plain,(
  ( ! [X14,X12,X13] : (s('type.realax.real',X13) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.realax.real',X13))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X12)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X13))))) ) | $spl70),
  inference(forward_demodulation,[],[f3671,f106])).
fof(f3671,plain,(
  ( ! [X14,X12,X13] : (s('type.realax.real',X13) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))))) | s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.realax.real',X13))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X12)))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X13))))) ) | $spl70),
  inference(superposition,[],[f90,f3567])).
fof(f4500,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl3),
  inference(forward_demodulation,[],[f3852,f2381])).
fof(f3852,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f3851,f107])).
fof(f3851,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f3850,f107])).
fof(f3850,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f3827,f107])).
fof(f3827,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f2528,f104])).
fof(f3855,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f3854,f107])).
fof(f3854,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f3830,f107])).
fof(f3830,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f2528,f104])).
fof(f3857,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f3833,f107])).
fof(f3833,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f2528,f104])).
fof(f4499,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))),s('type.num.num',X2)))) ) | $spl104),
  inference(cnf_transformation,[],[f4499_D])).
fof(f4499_D,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))),s('type.num.num',X2)))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f3837,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f2528,f104])).
fof(f4498,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl3),
  inference(forward_demodulation,[],[f3886,f2381])).
fof(f3886,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f3885,f107])).
fof(f3885,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f3867,f107])).
fof(f3867,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f2530,f104])).
fof(f2530,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))),
  inference(resolution,[],[f1745,f96])).
fof(f3888,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f3870,f107])).
fof(f3870,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f2530,f104])).
fof(f3873,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f2530,f104])).
fof(f4497,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl3),
  inference(forward_demodulation,[],[f3948,f2381])).
fof(f3948,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f3947,f107])).
fof(f3947,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f3899,f107])).
fof(f3899,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f2532,f104])).
fof(f2532,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))),
  inference(resolution,[],[f1745,f101])).
fof(f3949,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2)))) )),
  inference(forward_demodulation,[],[f3902,f107])).
fof(f3902,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f2532,f104])).
fof(f3905,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f2532,f104])).
fof(f4171,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X9),s('type.num.num',X10))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.num.num',X10)))) )),
  inference(forward_demodulation,[],[f4169,f107])).
fof(f4169,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X9),s('type.num.num',X10))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.num.num',X10)))) )),
  inference(trivial_inequality_removal,[],[f4151])).
fof(f4151,plain,(
  ( ! [X10,X9] : (s('type.realax.real',X9) != s('type.realax.real',X9) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X9),s('type.num.num',X10))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.num.num',X10)))) )),
  inference(superposition,[],[f1176,f104])).
fof(f4233,plain,(
  ( ! [X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X3))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X3)))) ) | $spl16),
  inference(forward_demodulation,[],[f4212,f107])).
fof(f4212,plain,(
  ( ! [X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X3))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X3)))) ) | $spl16),
  inference(superposition,[],[f2597,f104])).
fof(f4494,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))),s('type.num.num',X2)))) ) | $spl98),
  inference(cnf_transformation,[],[f4494_D])).
fof(f4494_D,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))),s('type.num.num',X2)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f4216,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))),s('type.num.num',X2)))) ) | $spl16),
  inference(superposition,[],[f2597,f104])).
fof(f4248,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X3))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X3)))) ) | $spl16),
  inference(superposition,[],[f2599,f104])).
fof(f2599,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) | $spl16),
  inference(resolution,[],[f182,f96])).
fof(f4307,plain,(
  ( ! [X10,X9] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X9),s('type.realax.real',X9)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X9),s('type.num.num',X10))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.num.num',X10)))) )),
  inference(forward_demodulation,[],[f4284,f107])).
fof(f4284,plain,(
  ( ! [X10,X9] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X9),s('type.realax.real',X9)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X9),s('type.num.num',X10))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.num.num',X10)))) )),
  inference(superposition,[],[f1175,f104])).
fof(f1175,plain,(
  ( ! [X0] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))))))) )),
  inference(resolution,[],[f108,f129])).
fof(f4348,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X3))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X3)))) ) | $spl16),
  inference(superposition,[],[f2601,f104])).
fof(f2601,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) | $spl16),
  inference(resolution,[],[f182,f101])).
fof(f4482,plain,(
  ( ! [X14,X15,X13] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X13),s('type.num.num',X15))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X13))),s('type.num.num',X15)))) ) | $spl26),
  inference(trivial_inequality_removal,[],[f4464])).
fof(f4464,plain,(
  ( ! [X14,X15,X13] : (s('type.realax.real',X13) != s('type.realax.real',X13) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X14))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X13),s('type.num.num',X15))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X13))),s('type.num.num',X15)))) ) | $spl26),
  inference(superposition,[],[f2051,f104])).
fof(f2051,plain,(
  ( ! [X28,X29] : (s('type.realax.real',X29) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X28))),s('type.realax.real',X29)))) ) | $spl26),
  inference(subsumption_resolution,[],[f2050,f1429])).
fof(f4393,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl90),
  inference(cnf_transformation,[],[f4393_D])).
fof(f4393_D,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f4391,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | ($spl3 | $spl16)),
  inference(forward_demodulation,[],[f4343,f2381])).
fof(f4343,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) ) | $spl16),
  inference(superposition,[],[f2601,f104])).
fof(f4262,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl89),
  inference(cnf_transformation,[],[f4262_D])).
fof(f4262_D,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f4260,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) ) | ($spl3 | $spl16)),
  inference(forward_demodulation,[],[f4243,f2381])).
fof(f4243,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) ) | $spl16),
  inference(superposition,[],[f2599,f104])).
fof(f4231,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) | $spl86),
  inference(cnf_transformation,[],[f4231_D])).
fof(f4231_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f4229,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))) ) | ($spl3 | $spl16)),
  inference(forward_demodulation,[],[f4207,f2381])).
fof(f4207,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) ) | $spl16),
  inference(superposition,[],[f2597,f104])).
fof(f4226,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))) | $spl82),
  inference(cnf_transformation,[],[f4226_D])).
fof(f4226_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f4228,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) | $spl84),
  inference(cnf_transformation,[],[f4228_D])).
fof(f4228_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f4193,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))) | $spl16),
  inference(resolution,[],[f2597,f128])).
fof(f3847,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))) | $spl78),
  inference(cnf_transformation,[],[f3847_D])).
fof(f3847_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f3849,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))) | $spl80),
  inference(cnf_transformation,[],[f3849_D])).
fof(f3849_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f3813,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))))),
  inference(resolution,[],[f2528,f128])).
fof(f175,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl7),
  inference(cnf_transformation,[],[f175_D])).
fof(f175_D,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f3559,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl3),
  inference(forward_demodulation,[],[f3551,f2381])).
fof(f3551,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f89,f104])).
fof(f89,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))),
  inference(cnf_transformation,[],[f69])).
fof(f3359,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl74),
  inference(cnf_transformation,[],[f3359_D])).
fof(f3359_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3542,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl3 | $spl60)),
  inference(forward_demodulation,[],[f3528,f2381])).
fof(f3528,plain,(
  ( ! [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) ) | $spl60),
  inference(superposition,[],[f2627,f104])).
fof(f3415,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl77),
  inference(cnf_transformation,[],[f3415_D])).
fof(f3415_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f3508,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl3 | $spl63)),
  inference(forward_demodulation,[],[f3499,f2381])).
fof(f3499,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) ) | $spl63),
  inference(superposition,[],[f2628,f104])).
fof(f2628,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | $spl63),
  inference(cnf_transformation,[],[f2628_D])).
fof(f2628_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f3433,plain,(
  $false | ($spl20 | $spl62)),
  inference(subsumption_resolution,[],[f3432,f127])).
fof(f3432,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl62)),
  inference(forward_demodulation,[],[f3431,f982])).
fof(f3431,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | ($spl20 | $spl62)),
  inference(subsumption_resolution,[],[f3430,f1456])).
fof(f3430,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | $spl62),
  inference(forward_demodulation,[],[f3379,f106])).
fof(f3379,plain,(
  s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | $spl62),
  inference(superposition,[],[f91,f2629])).
fof(f3416,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl76),
  inference(cnf_transformation,[],[f3416_D])).
fof(f3416_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f3414,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl3 | $spl62)),
  inference(forward_demodulation,[],[f3370,f2381])).
fof(f3370,plain,(
  ( ! [X2] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) ) | $spl62),
  inference(superposition,[],[f2629,f104])).
fof(f3357,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl3 | $spl60)),
  inference(forward_demodulation,[],[f3343,f2381])).
fof(f3343,plain,(
  ( ! [X2] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) ) | $spl60),
  inference(superposition,[],[f2627,f104])).
fof(f3324,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl3),
  inference(forward_demodulation,[],[f3316,f2381])).
fof(f3316,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f89,f104])).
fof(f3260,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X2)))) ) | $spl70),
  inference(cnf_transformation,[],[f3260_D])).
fof(f3260_D,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X2)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f3259,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X3))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X3))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X2)))) ) | $spl3),
  inference(forward_demodulation,[],[f3220,f2381])).
fof(f3220,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X2))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X3))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X3)))) ) | $spl3),
  inference(superposition,[],[f2381,f104])).
fof(f3135,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))) | $spl66),
  inference(cnf_transformation,[],[f3135_D])).
fof(f3135_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f3136,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))),s('type.num.num',X0)))))) ) | $spl68),
  inference(cnf_transformation,[],[f3136_D])).
fof(f3136_D,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))),s('type.num.num',X0)))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f3087,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))),s('type.num.num',X0))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1)))) ) | $spl14),
  inference(superposition,[],[f104,f180])).
fof(f180,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | $spl14),
  inference(cnf_transformation,[],[f180_D])).
fof(f180_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f3074,plain,(
  $false | $spl53),
  inference(subsumption_resolution,[],[f2330,f2205])).
fof(f173,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) ) | $spl4),
  inference(cnf_transformation,[],[f173_D])).
fof(f173_D,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f2631,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl64),
  inference(cnf_transformation,[],[f2631_D])).
fof(f2631_D,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2612,plain,(
  ( ! [X2] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) ) | $spl16),
  inference(superposition,[],[f182,f104])).
fof(f2627,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | $spl60),
  inference(cnf_transformation,[],[f2627_D])).
fof(f2627_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f2629,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | $spl62),
  inference(cnf_transformation,[],[f2629_D])).
fof(f2629_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f2598,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | $spl16),
  inference(resolution,[],[f182,f128])).
fof(f2559,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))) | $spl56),
  inference(cnf_transformation,[],[f2559_D])).
fof(f2559_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2561,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) | $spl58),
  inference(cnf_transformation,[],[f2561_D])).
fof(f2561_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2529,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))))),
  inference(resolution,[],[f1745,f128])).
fof(f2516,plain,(
  $false | $spl17),
  inference(subsumption_resolution,[],[f2515,f127])).
fof(f2515,plain,(
  ~p(s(bool,'T_0')) | $spl17),
  inference(forward_demodulation,[],[f2483,f982])).
fof(f2483,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | $spl17),
  inference(resolution,[],[f183,f109])).
fof(f2330,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1)))))))) | $spl53),
  inference(cnf_transformation,[],[f2330_D])).
fof(f2330_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1)))))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f2331,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) | $spl54),
  inference(cnf_transformation,[],[f2331_D])).
fof(f2331_D,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1)))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2320,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X1))))) ) | $spl15),
  inference(superposition,[],[f181,f90])).
fof(f183,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | $spl17),
  inference(cnf_transformation,[],[f183_D])).
fof(f183_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f184,plain,(
  ( ! [X4,X5] : (s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.realax.real',X5))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) | s('type.realax.real',X5) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))) ) | $spl18),
  inference(cnf_transformation,[],[f184_D])).
fof(f184_D,plain,(
  ( ! [X4,X5] : (s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.realax.real',X5))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5)))) | s('type.realax.real',X5) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1993,plain,(
  ( ! [X158,X157] : (s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X157))),s('type.realax.real',X158))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | s('type.realax.real',X158) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X157))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X158))))) )),
  inference(superposition,[],[f89,f90])).
fof(f2098,plain,(
  ( ! [X19,X18] : (s('type.realax.real',X18) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X18)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) ) | $spl50),
  inference(cnf_transformation,[],[f2098_D])).
fof(f2098_D,plain,(
  ( ! [X19,X18] : (s('type.realax.real',X18) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X18)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19)))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1953,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',X17)))) | s('type.realax.real',X18) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X18),s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X18)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X17),s('type.num.num',X20))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))),s('type.num.num',X20)))) )),
  inference(superposition,[],[f90,f104])).
fof(f2082,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.realax.real',X19) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',X17)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))),s('type.realax.real',X19))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X17),s('type.num.num',X20))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))),s('type.num.num',X20)))) )),
  inference(subsumption_resolution,[],[f2081,f130])).
fof(f2081,plain,(
  ( ! [X19,X17,X20,X18] : (s('type.realax.real',X19) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',X17)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))),s('type.realax.real',X19))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X19)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X17),s('type.num.num',X20))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))),s('type.num.num',X20)))) )),
  inference(forward_demodulation,[],[f1931,f106])).
fof(f1931,plain,(
  ( ! [X19,X17,X20,X18] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',X17)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))),s('type.realax.real',X19))) | s('type.realax.real',X19) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X19)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X17),s('type.num.num',X20))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X17))),s('type.num.num',X20)))) )),
  inference(superposition,[],[f90,f104])).
fof(f1773,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | $spl44),
  inference(cnf_transformation,[],[f1773_D])).
fof(f1773_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1775,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | $spl46),
  inference(cnf_transformation,[],[f1775_D])).
fof(f1775_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1746,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))))),
  inference(resolution,[],[f1177,f128])).
fof(f1726,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) ) | $spl42),
  inference(cnf_transformation,[],[f1726_D])).
fof(f1726_D,plain,(
  ( ! [X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1688,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X9)))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X9),s('type.num.num',X11))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.num.num',X11)))) )),
  inference(superposition,[],[f91,f104])).
fof(f1427,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X1)))) ) | $spl24),
  inference(cnf_transformation,[],[f1427_D])).
fof(f1427_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X1)))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1667,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f91,f104])).
fof(f1597,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X0))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X0)))) )),
  inference(superposition,[],[f982,f104])).
fof(f1575,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X0))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X0)))) )),
  inference(superposition,[],[f980,f104])).
fof(f980,plain,(
  s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))),
  inference(resolution,[],[f134,f96])).
fof(f1539,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X9)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X9),s('type.num.num',X11))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X9))),s('type.num.num',X11)))) )),
  inference(superposition,[],[f109,f104])).
fof(f1516,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f109,f104])).
fof(f1433,plain,(
  ( ! [X30,X29] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X29),s('type.realax.real',X29)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X29),s('type.num.num',X30))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X29))),s('type.num.num',X30)))) )),
  inference(forward_demodulation,[],[f1406,f107])).
fof(f1406,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X29),s('type.realax.real',X29)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X29),s('type.num.num',X30))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X29))),s('type.num.num',X30)))) )),
  inference(superposition,[],[f108,f104])).
fof(f1432,plain,(
  ( ! [X28,X27] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X27),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X27)))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X27),s('type.num.num',X28))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X27))),s('type.num.num',X28)))) ) | $spl32),
  inference(cnf_transformation,[],[f1432_D])).
fof(f1432_D,plain,(
  ( ! [X28,X27] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X27),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X27)))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X27),s('type.num.num',X28))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X27))),s('type.num.num',X28)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1405,plain,(
  ( ! [X28,X27] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X27),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X27)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X27),s('type.num.num',X28))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X27))),s('type.num.num',X28)))) )),
  inference(superposition,[],[f108,f104])).
fof(f181,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | $spl15),
  inference(cnf_transformation,[],[f181_D])).
fof(f181_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f1393,plain,(
  ( ! [X70] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X70))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X70)))) )),
  inference(superposition,[],[f89,f104])).
fof(f1392,plain,(
  ( ! [X69] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X69))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X69)))) )),
  inference(superposition,[],[f134,f104])).
fof(f1391,plain,(
  ( ! [X68] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X68))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X68)))) )),
  inference(superposition,[],[f89,f104])).
fof(f1431,plain,(
  ( ! [X66,X67] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X66)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X66))),s('type.num.num',X67))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X66))))),s('type.num.num',X67)))) ) | $spl30),
  inference(cnf_transformation,[],[f1431_D])).
fof(f1431_D,plain,(
  ( ! [X66,X67] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X66)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X66))),s('type.num.num',X67))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X66))))),s('type.num.num',X67)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1390,plain,(
  ( ! [X66,X67] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X66)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X66))),s('type.num.num',X67))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X66))))),s('type.num.num',X67)))) )),
  inference(superposition,[],[f108,f104])).
fof(f1430,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2)))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X2),s('type.num.num',X3))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.num.num',X3)))) ) | $spl28),
  inference(cnf_transformation,[],[f1430_D])).
fof(f1430_D,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2)))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X2),s('type.num.num',X3))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.num.num',X3)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1362,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2)))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X2),s('type.num.num',X3))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X2))),s('type.num.num',X3)))) )),
  inference(superposition,[],[f108,f104])).
fof(f1426,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f1361,f107])).
fof(f1361,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X0))),s('type.num.num',X1)))) )),
  inference(superposition,[],[f108,f104])).
fof(f1000,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | $spl20),
  inference(cnf_transformation,[],[f1000_D])).
fof(f1000_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1002,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | $spl22),
  inference(cnf_transformation,[],[f1002_D])).
fof(f1002_D,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f979,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),
  inference(resolution,[],[f134,f128])).
fof(f852,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f820,f190])).
fof(f190,plain,(
  s('type.realax.real',sK1) = s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | $spl2),
  inference(backward_demodulation,[],[f161,f107])).
fof(f820,plain,(
  s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | $spl2),
  inference(backward_demodulation,[],[f812,f255])).
fof(f255,plain,(
  s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | $spl2),
  inference(backward_demodulation,[],[f190,f89])).
fof(f812,plain,(
  ( ! [X0] : (s('type.realax.real',sK1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',sK1)))) ) | $spl2),
  inference(equality_resolution,[],[f794])).
fof(f794,plain,(
  ( ! [X10,X11] : (s('type.realax.real',sK1) != s('type.realax.real',X10) | s('type.realax.real',sK1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.realax.real',X10)))) ) | $spl2),
  inference(subsumption_resolution,[],[f793,f132])).
fof(f793,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s('type.realax.real',sK1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.realax.real',X10))) | s('type.realax.real',sK1) != s('type.realax.real',X10)) ) | $spl2),
  inference(forward_demodulation,[],[f792,f241])).
fof(f241,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',sK1))) | $spl2),
  inference(resolution,[],[f185,f101])).
fof(f185,plain,(
  p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',sK1)))) | $spl2),
  inference(backward_demodulation,[],[f161,f88])).
fof(f792,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',sK1)))) | s('type.realax.real',sK1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.realax.real',X10))) | s('type.realax.real',sK1) != s('type.realax.real',X10)) ) | $spl2),
  inference(forward_demodulation,[],[f791,f190])).
fof(f791,plain,(
  ( ! [X10,X11] : (s('type.realax.real',sK1) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.realax.real',X10))) | s('type.realax.real',sK1) != s('type.realax.real',X10) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) ) | $spl2),
  inference(forward_demodulation,[],[f790,f190])).
fof(f790,plain,(
  ( ! [X10,X11] : (s('type.realax.real',sK1) != s('type.realax.real',X10) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.realax.real',X10))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) ) | $spl2),
  inference(subsumption_resolution,[],[f789,f130])).
fof(f789,plain,(
  ( ! [X10,X11] : (s('type.realax.real',sK1) != s('type.realax.real',X10) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X10)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.realax.real',X10))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) ) | $spl2),
  inference(forward_demodulation,[],[f771,f189])).
fof(f189,plain,(
  ( ! [X0] : (s('type.realax.real',sK1) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) | $spl2),
  inference(backward_demodulation,[],[f161,f106])).
fof(f771,plain,(
  ( ! [X10,X11] : (s('type.realax.real',X10) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',sK1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X10)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.realax.real',X10))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) ) | $spl2),
  inference(superposition,[],[f210,f190])).
fof(f210,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X1)))) | s('type.realax.real',X2) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X2))))) ) | $spl2),
  inference(forward_demodulation,[],[f186,f161])).
fof(f186,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',X1)))) | s('type.realax.real',X2) = s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.realax.real',X1))) | s('type.realax.real',X1) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))))) ) | $spl2),
  inference(backward_demodulation,[],[f161,f90])).
fof(f851,plain,(
  $false | ($spl2 | $spl15)),
  inference(subsumption_resolution,[],[f819,f190])).
fof(f819,plain,(
  s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) | ($spl2 | $spl15)),
  inference(backward_demodulation,[],[f812,f646])).
fof(f646,plain,(
  s('type.realax.real',sK1) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | ($spl2 | $spl15)),
  inference(forward_demodulation,[],[f181,f161])).
fof(f592,plain,(
  $false | ($spl2 | $spl9)),
  inference(subsumption_resolution,[],[f591,f127])).
fof(f591,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl9)),
  inference(forward_demodulation,[],[f590,f241])).
fof(f590,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',sK1),s('type.realax.real',sK1)))) | ($spl2 | $spl9)),
  inference(forward_demodulation,[],[f589,f161])).
fof(f589,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1)))) | ($spl2 | $spl9)),
  inference(forward_demodulation,[],[f177,f190])).
fof(f172,plain,(
  ( ! [X4,X5] : (s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.realax.real',X5))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | s('type.realax.real',X5) != s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X5))))) )),
  inference(superposition,[],[f89,f90])).
fof(f171,plain,(
  ( ! [X3] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X3))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))),s('type.num.num',X3)))) )),
  inference(superposition,[],[f89,f104])).
fof(f177,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) | $spl9),
  inference(cnf_transformation,[],[f177_D])).
fof(f177_D,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1)))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f178,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))) ) | $spl10),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f168,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))))) )),
  inference(superposition,[],[f89,f90])).
fof(f167,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real',sK1))))) != s('type.realax.real','const.transc.root_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))),s('type.num.num',X2))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',sK1))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f89,f104])).
fof(f159,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1)))) | $spl0),
  inference(cnf_transformation,[],[f159_D])).
fof(f159_D,plain,(
  p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f161,plain,(
  s('type.realax.real',sK1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(cnf_transformation,[],[f161_D])).
fof(f161_D,plain,(
  s('type.realax.real',sK1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f138,plain,(
  s('type.realax.real',sK1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',sK1))))),
  inference(resolution,[],[f88,f128])).
% SZS output end Proof for thm.transc
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 533
% Passive clauses: 8568
% Generated clauses: 39348
% Final active clauses: 385
% Final passive clauses: 6788
% Input formulas: 26
% Initial clauses: 44
% 
% Selected by SInE selection: 25
% SInE iterations: 1
% 
% Duplicate literals: 231
% Trivial inequalities: 698
% Fw subsumption resolutions: 3572
% Fw demodulations: 13320
% Bw demodulations: 102
% 
% Simple tautologies: 287
% Equational tautologies: 9
% Forward subsumptions: 12880
% Fw demodulations to eq. taut.: 78
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 2952
% Factoring: 1
% Forward superposition: 11434
% Backward superposition: 4676
% Self superposition: 29
% Equality resolution: 92
% 
% Split clauses: 174
% Split components: 364
% SAT solver clauses: 188
% SAT solver unit clauses: 5
% SAT solver binary clauses: 75
% 
% Sat splits: 174
% Sat splitting refutations: 12
% 
% Memory used [KB]: 13432
% Time elapsed: 0.898 s
% ------------------------------
% ------------------------------
% Success in time 12.616 s
