Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quantHeuristics
% (63206)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 759
% Passive clauses: 10728
% Generated clauses: 20717
% Final active clauses: 722
% Final passive clauses: 9279
% Input formulas: 56
% Initial clauses: 229
% 
% Duplicate literals: 335
% Trivial inequalities: 52
% Fw subsumption resolutions: 1773
% Fw demodulations: 617
% Bw demodulations: 9
% 
% Simple tautologies: 270
% Equational tautologies: 57
% Forward subsumptions: 6616
% 
% Binary resolution: 11598
% Factoring: 1021
% Forward superposition: 2642
% Backward superposition: 2207
% Self superposition: 27
% Equality factoring: 26
% Equality resolution: 18
% 
% Split clauses: 340
% Split components: 684
% SAT solver clauses: 663
% SAT solver unit clauses: 16
% SAT solver binary clauses: 172
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 339
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 662
% TWLsolver calls for satisfiability: 387
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quantHeuristics
% (63214)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 464
% Passive clauses: 2802
% Generated clauses: 13427
% Final active clauses: 453
% Final passive clauses: 2316
% Input formulas: 56
% Initial clauses: 229
% 
% Duplicate literals: 225
% Trivial inequalities: 36
% Fw subsumption resolutions: 983
% Fw demodulations: 817
% Bw demodulations: 33
% Condensations: 326
% Global subsumptions: 13
% 
% Simple tautologies: 296
% Equational tautologies: 22
% Forward subsumptions: 4417
% 
% Binary resolution: 3422
% Factoring: 179
% Forward superposition: 2403
% Backward superposition: 1844
% Self superposition: 41
% Equality factoring: 26
% Equality resolution: 9
% 
% Split clauses: 74
% Split components: 150
% SAT solver clauses: 3629
% SAT solver unit clauses: 211
% SAT solver binary clauses: 79
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 74
% 
% TWLsolver clauses: 3577
% TWLsolver calls for satisfiability: 2118
% 
% Memory used [KB]: 14456
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quantHeuristics
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quantHeuristics
% SZS output start Proof for thm.quantHeuristics
fof(f58419,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f326,f578,f325,f580,f576,f574,f324,f581,f323,f1672,f1707,f1705,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4672,f4578,f4674,f4582,f4677,f4586,f4679,f4589,f4681,f4592,f4683,f4595,f4685,f4598,f4687,f4601,f4689,f4604,f4691,f4607,f4693,f4610,f4695,f4613,f4697,f4616,f4699,f4619,f4701,f4622,f4703,f4625,f4934,f4935,f4936,f4937,f4938,f4939,f4940,f4941,f4942,f4943,f4944,f4945,f4946,f4947,f4948,f4949,f4965,f4971,f4977,f4981,f4984,f5000,f5002,f5004,f5006,f5008,f5010,f5012,f5014,f5016,f5018,f5020,f5022,f5024,f5026,f5028,f5030,f5032,f5034,f5036,f5041,f5042,f5043,f5044,f7656,f7667,f7671,f7674,f7677,f7681,f7684,f7690,f7693,f7697,f7700,f7705,f7709,f7713,f7745,f7760,f7807,f7843,f7869,f7877,f7879,f7883,f7891,f7893,f7901,f7903,f7907,f7913,f7921,f7923,f7927,f7933,f7941,f7943,f7948,f7953,f7961,f7963,f8417,f8418,f7982,f7991,f8001,f8009,f8011,f8017,f8019,f8023,f8025,f8031,f8033,f8037,f8039,f8043,f8046,f8048,f8051,f8053,f8092,f8095,f8100,f8105,f8108,f8120,f8123,f8129,f8133,f8136,f8146,f8148,f8152,f8162,f8167,f8172,f8179,f8183,f8190,f8194,f8201,f8206,f8210,f8216,f8221,f8225,f8227,f8234,f8236,f8239,f8241,f8263,f8273,f8283,f8291,f8293,f8298,f8305,f8308,f8315,f8321,f8333,f8335,f8339,f8342,f8347,f8351,f8362,f8368,f8374,f8382,f8386,f8390,f8393,f8403,f8421,f8422,f8423,f8424,f8425,f8426,f8431,f8432,f8433,f8434,f8435,f8436,f8437,f8438,f8439,f8440,f8441,f8442,f8443,f8444,f8445,f8446,f8447,f8448,f8449,f8450,f8451,f8452,f8453,f8454,f8455,f8456,f8457,f8458,f8459,f8460,f8461,f8462,f8463,f8464,f8465,f8466,f8467,f8468,f8469,f8470,f8471,f8472,f8473,f8474,f8475,f8476,f8477,f8478,f8479,f8480,f8481,f8482,f8483,f8484,f8485,f8486,f8487,f8488,f8489,f8490,f8491,f8492,f8493,f8494,f8495,f8496,f8497,f8498,f8499,f8500,f8501,f8502,f8503,f8504,f8505,f8506,f8507,f8508,f8509,f8510,f8511,f8512,f8513,f8514,f8515,f8516,f8517,f8518,f8519,f8520,f8521,f8522,f8523,f8524,f8525,f8526,f8527,f8528,f8529,f8530,f8531,f8532,f8533,f8534,f8535,f8536,f8537,f8538,f8539,f8540,f8541,f8542,f8543,f8544,f8545,f8546,f8547,f8548,f8549,f8550,f8553,f8554,f8555,f8556,f8557,f6932,f2417,f8559,f2389,f8561,f573,f9331,f2549,f575,f9334,f11141,f11243,f13710,f13713,f13711,f13753,f13756,f13754,f13924,f13927,f13925,f13960,f13963,f13961,f14003,f14006,f14004,f14309,f14312,f14310,f14362,f14365,f14363,f14433,f14436,f14434,f14471,f14474,f14472,f14530,f14533,f14531,f14574,f14577,f14575,f14640,f14643,f14641,f14676,f14679,f14677,f15695,f15698,f15696,f15727,f15729,f15728,f15731,f15733,f15735,f15741,f15743,f15745,f15747,f15751,f15753,f15755,f15757,f15759,f15761,f15777,f15779,f15795,f15797,f15814,f15831,f15841,f15843,f15854,f15865,f15875,f15877,f15887,f15889,f15899,f15901,f15935,f15992,f16083,f16085,f16641,f16645,f16654,f16658,f16662,f16666,f16714,f16718,f16724,f16729,f16734,f16738,f16743,f16762,f16774,f16788,f16800,f16802,f16816,f16828,f16842,f16854,f16856,f16858,f16860,f16862,f19056,f19057,f19058,f19059,f19060,f19061,f16914,f16916,f16918,f16920,f16922,f16924,f16926,f16929,f16931,f16933,f16935,f16937,f16939,f16941,f16998,f17000,f17002,f17004,f17036,f17068,f17108,f17110,f17117,f17119,f17121,f17123,f17136,f17138,f17140,f17142,f17155,f17157,f17159,f17161,f17163,f17165,f17167,f17169,f17210,f17212,f17219,f17221,f17223,f17225,f17238,f17240,f17242,f17244,f17257,f17259,f17261,f17263,f17265,f17267,f17269,f17271,f17286,f17298,f17312,f17324,f17326,f17368,f17370,f17396,f17398,f17440,f17442,f17468,f17470,f17520,f17522,f17548,f17550,f17572,f17604,f17626,f17658,f17692,f17726,f17792,f17794,f17796,f17798,f17800,f17806,f17880,f17883,f17885,f17887,f17889,f17891,f17893,f17911,f17913,f17915,f17937,f17939,f17961,f17963,f17985,f17987,f18009,f18011,f18046,f18076,f18116,f18118,f18120,f18122,f18124,f18130,f18148,f18150,f18152,f18154,f18156,f18158,f18160,f18168,f18170,f18172,f18258,f18260,f18262,f18264,f18266,f18268,f18270,f18272,f18274,f18276,f18325,f18327,f18329,f18331,f18380,f18382,f18384,f18386,f18410,f18412,f18436,f18438,f18440,f18464,f18488,f18512,f18514,f18538,f18540,f18564,f18594,f18626,f18656,f18667,f18673,f18676,f18700,f18702,f18704,f18706,f18708,f18710,f18712,f18714,f18716,f18718,f18740,f18742,f18744,f18746,f18748,f18750,f18752,f18754,f18756,f18758,f18760,f18762,f18764,f18766,f18768,f18784,f18787,f18799,f18803,f18814,f18817,f18829,f18833,f18844,f18847,f18859,f18863,f18868,f18871,f18878,f18881,f18912,f18915,f18926,f18928,f18931,f18933,f18948,f18950,f18952,f18954,f18956,f18971,f18986,f19000,f19006,f19031,f19062,f19063,f19064,f19065,f19066,f19067,f19068,f19069,f19070,f19071,f19072,f19073,f19074,f19075,f19076,f19077,f19078,f19079,f19080,f19081,f19082,f19083,f19084,f19085,f19086,f19087,f19088,f19089,f19090,f19091,f19092,f19093,f19094,f19095,f19096,f19097,f19098,f19099,f19100,f19101,f19102,f19103,f19104,f19105,f19106,f19107,f19108,f19109,f19110,f19111,f19112,f19113,f19114,f19115,f19116,f19117,f19118,f19119,f19120,f19121,f19122,f19123,f19124,f19125,f19126,f19127,f19128,f19129,f19130,f19131,f19132,f19133,f19134,f19135,f19136,f19137,f19138,f19139,f19140,f19141,f19142,f19143,f19144,f19145,f19146,f19147,f19148,f19149,f19150,f19151,f19152,f19153,f19154,f19155,f19156,f19157,f19158,f19159,f19160,f19161,f19162,f19163,f19164,f19165,f19166,f19167,f19168,f19169,f19170,f19171,f19172,f19173,f19174,f19175,f19176,f19177,f19178,f19179,f19180,f19181,f19182,f19183,f19184,f19185,f19186,f19187,f19188,f19189,f19190,f19191,f19192,f19193,f19194,f19195,f19196,f19197,f19198,f19199,f19200,f19201,f19202,f19203,f19204,f19205,f19206,f19207,f19208,f19209,f19210,f19211,f19212,f19213,f19214,f19215,f19216,f19217,f19218,f19219,f19220,f19221,f19222,f19223,f19224,f19225,f19226,f19227,f19228,f19229,f19230,f19231,f19232,f19233,f19234,f19235,f19236,f19237,f19238,f19239,f19240,f19241,f19242,f19243,f19244,f19245,f19246,f19247,f19248,f19249,f19250,f19251,f19252,f19253,f19254,f19255,f19256,f19257,f19258,f19259,f19260,f19261,f19262,f19263,f19264,f19265,f19266,f19267,f19268,f19269,f19270,f19271,f19272,f19273,f19274,f19275,f19276,f19277,f19278,f19279,f19280,f19281,f19282,f19283,f19284,f19285,f19286,f19287,f19288,f19289,f19290,f19291,f19292,f19293,f19294,f19295,f19296,f19297,f19298,f19299,f19300,f19301,f19302,f19303,f19304,f19305,f19306,f19307,f19308,f19309,f19310,f19311,f19312,f19313,f19314,f19315,f19316,f19317,f19318,f19319,f19320,f19321,f19322,f19323,f19324,f19325,f19326,f19327,f19328,f19329,f19330,f19331,f19332,f19333,f19334,f19335,f19336,f19337,f19338,f19339,f19340,f19341,f19342,f19343,f19344,f19345,f19346,f19347,f19348,f19349,f19350,f19351,f19352,f19353,f19354,f19355,f19356,f19357,f19358,f19359,f19360,f19361,f19362,f19363,f19364,f19365,f19366,f19367,f19368,f19369,f19370,f19371,f19372,f19373,f19374,f19375,f19376,f19377,f19378,f19379,f19380,f19381,f19382,f19383,f19384,f19385,f19386,f19387,f19388,f19389,f19390,f19391,f19392,f19393,f19394,f19395,f19396,f19397,f19398,f19399,f19400,f19401,f19402,f19403,f19404,f19405,f19406,f19407,f19408,f19409,f19410,f19411,f19412,f19413,f19414,f19415,f19416,f19417,f19418,f19419,f19420,f19421,f19422,f19423,f19424,f19425,f19426,f19427,f19428,f19429,f19430,f19431,f19432,f19433,f19434,f19435,f19436,f19437,f19438,f19439,f19440,f19441,f19442,f19443,f19444,f19445,f19446,f19447,f19448,f19449,f19450,f19451,f19452,f19453,f19454,f19455,f19456,f19457,f19458,f19459,f19460,f19461,f19462,f19463,f19464,f19465,f19466,f19467,f19468,f19469,f19470,f19471,f19472,f19473,f19474,f19475,f19476,f19477,f19478,f19479,f19480,f19481,f19482,f19483,f19484,f19485,f19486,f19487,f19488,f19489,f19490,f19491,f19492,f19493,f19494,f19495,f19496,f19497,f19498,f19499,f19500,f19501,f19502,f19503,f19504,f19505,f19506,f19507,f19508,f19509,f19510,f19511,f19512,f19513,f19514,f19515,f19516,f19517,f19518,f19519,f19520,f19521,f19522,f19523,f19524,f19525,f19526,f19527,f19528,f19529,f19530,f19531,f19532,f19533,f19534,f19535,f19536,f19537,f19538,f19539,f19540,f19541,f19542,f19543,f19544,f19545,f19546,f19547,f19548,f19549,f19550,f19551,f19552,f19553,f19554,f19555,f19556,f19557,f19558,f19559,f19560,f19561,f19562,f19563,f19564,f19565,f19566,f19567,f19568,f19569,f19570,f19571,f19572,f19573,f19574,f19575,f19576,f19577,f19578,f19579,f19580,f19581,f19582,f19583,f19584,f19585,f19586,f19587,f19588,f19589,f19590,f19591,f19592,f19593,f19594,f19595,f19596,f19597,f19598,f19599,f19600,f19601,f19602,f19603,f19604,f19605,f19606,f19607,f19608,f19609,f19610,f19611,f19612,f19613,f19614,f19615,f19616,f19617,f19618,f19619,f19620,f19621,f19622,f19623,f19624,f19625,f19626,f19627,f19628,f19629,f19630,f19631,f19632,f19633,f19634,f19635,f19636,f19637,f19638,f19639,f19640,f19641,f19642,f19643,f19644,f19645,f19646,f19647,f19648,f19649,f19650,f19651,f19652,f19653,f19654,f19655,f19656,f19657,f19658,f19659,f19660,f19661,f19662,f19663,f19664,f19665,f19666,f19667,f19668,f19669,f19670,f19671,f19672,f19673,f19674,f19675,f19676,f19677,f19678,f19679,f19680,f19681,f19682,f19683,f19684,f19685,f19686,f19687,f19688,f19689,f19690,f19691,f19692,f19693,f19694,f19695,f19696,f19697,f19698,f19699,f19700,f19701,f19702,f19703,f19704,f19705,f19706,f19707,f19708,f19709,f19710,f19711,f19712,f19713,f19714,f19715,f19716,f19717,f19718,f19719,f19720,f19721,f19722,f19723,f19724,f19725,f19726,f19727,f19728,f19729,f19730,f19731,f19732,f19733,f19734,f19735,f19736,f19737,f19738,f19739,f19740,f19741,f19742,f19743,f19744,f19745,f19746,f19747,f19748,f19749,f19750,f19751,f19752,f19753,f19832,f19835,f19833,f19897,f19900,f19898,f20184,f20185,f20192,f20189,f20186,f20187,f20195,f20188,f20197,f20190,f20194,f21347,f21349,f21348,f22355,f22356,f22368,f22366,f22591,f22593,f22592,f23313,f23315,f23314,f23341,f23343,f15697,f23342,f37230,f37232,f37231,f39651,f39653,f39652,f47720,f47723,f47721,f47807,f47810,f47808,f47898,f47901,f47899,f47983,f47986,f47991,f47992,f47995,f47996,f47984,f48063,f48066,f48064,f48124,f48127,f4579,f48125,f56901,f57138])).
fof(f57138,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f57137,f578])).
fof(f57137,plain,(
  s(sK14,sK15) = s(sK14,sK17) | ($spl0 | $spl2 | $spl6)),
  inference(subsumption_resolution,[],[f56890,f327])).
fof(f327,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|',aHL_TRUTH)).
fof(f56890,plain,(
  ~p(s(bool,'T_0')) | s(sK14,sK15) = s(sK14,sK17) | ($spl0 | $spl2 | $spl6)),
  inference(superposition,[],[f56884,f12312])).
fof(f12312,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f327,f11256,f394])).
fof(f394,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.bool.IMP_ANTISYM_AX')).
fof(f11256,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl2),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f334,f342,f341,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f575,f585,f584,f587,f586,f589,f588,f590,f592,f595,f594,f596,f598,f601,f600,f603,f602,f605,f604,f607,f606,f608,f610,f613,f612,f615,f614,f616,f618,f621,f620,f622,f625,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f643,f676,f677,f680,f683,f686,f689,f692,f703,f710,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f789,f790,f794,f795,f801,f802,f805,f809,f810,f814,f817,f821,f822,f843,f842,f847,f848,f867,f866,f864,f863,f871,f872,f890,f934,f954,f991,f989,f1002,f1001,f999,f998,f1017,f1013,f1030,f1029,f1027,f1026,f1078,f1088,f1089,f1106,f1116,f1117,f1121,f1129,f1142,f1148,f1147,f1146,f1166,f1176,f1177,f1181,f1207,f1208,f1241,f1255,f1266,f1267,f1271,f1282,f1294,f1315,f1316,f1319,f1337,f1352,f1360,f1373,f1383,f1384,f1388,f1403,f1431,f1432,f1436,f1449,f1448,f1446,f1445,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1506,f1534,f1533,f1525,f1524,f1522,f1521,f1519,f1518,f1517,f1543,f1545,f1540,f1584,f1583,f1570,f1569,f1568,f1637,f1636,f1628,f1667,f1666,f1665,f1664,f1671,f2065,f2089,f2309,f2342,f2582,f2603,f2604,f2586,f2587,f2589,f2591,f2618,f2622,f2623,f2625,f2627,f2668,f2669,f2648,f2655,f2656,f2658,f2682,f2686,f2687,f2689,f2714,f2713,f2711,f2710,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3092,f3091,f3087,f3083,f3079,f3078,f3077,f3076,f3075,f3074,f3073,f3072,f3071,f3070,f3069,f3068,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2910,f2909,f2903,f2901,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2888,f2887,f2886,f2885,f2883,f2882,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2841,f2839,f2837,f2835,f2833,f2831,f2829,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3510,f3509,f3505,f3501,f3497,f3496,f3495,f3494,f3493,f3492,f3491,f3490,f3489,f3488,f3487,f3486,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3328,f3327,f3321,f3319,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3306,f3305,f3304,f3303,f3301,f3300,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3259,f3257,f3255,f3253,f3251,f3249,f3247,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3710,f3717,f3739,f3792,f3794,f3810,f3822,f3824,f3869,f3870,f3871,f3872,f3938,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4199,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f2257,f2260,f2076,f2261,f441,f2231,f2234,f2055,f2235,f435,f2458,f2442,f3821,f3791,f2690,f2659,f2626,f2590,f2258,f2232,f1700,f1699,f1697,f1696,f1695,f1694,f1693,f1692,f1691,f1690,f1689,f1688,f1687,f1686,f1684,f1683,f1682,f1681,f2334,f4206,f4205,f4202,f4193,f2333,f4198,f4197,f4196,f4195,f3945,f3944,f3941,f3937,f3936,f3935,f3934,f3932,f3931,f4192,f2090,f443,f4661,f2077,f440,f4660,f2056,f434,f2057,f3930,f4191,f369,f4200,f4207,f3939,f3946,f2460,f2459,f446,f2461,f2444,f2443,f432,f4655,f4664,f4663,f2341,f445,f444,f431,f4666,f4665,f4662,f4659,f4658,f4657,f4656,f4204,f4201,f4190,f4189,f4187,f3943,f3940,f3929,f3928,f3926,f2563,f2561,f2516,f2515,f2497,f2496,f2456,f2440,f2336,f2167,f2139,f2121,f2104,f442,f439,f433,f389,f387,f386,f384,f3933,f4194,f4653,f4652,f4650,f4649,f2416,f2412,f2388,f2384,f1657,f376,f375,f374,f373,f372,f2301,f4203,f4188,f3942,f3927,f2302,f2304,f4654,f4648,f2547,f2545,f2067,f2066,f438,f437,f436,f3924,f4185,f2723,f2685,f2654,f2651,f2621,f2585,f2579,f2455,f2454,f2439,f2438,f2335,f2303,f2259,f2248,f2233,f2222,f2211,f2210,f2194,f2163,f2135,f2117,f2100,f2087,f2074,f2063,f2053,f1670,f1669,f1462,f582,f346,f345,f8557,f6932,f2417,f2389,f2549,f11255])).
fof(f11255,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0))))) )),
  inference(forward_demodulation,[],[f11254,f428])).
fof(f11254,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0)))) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,happ(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(sK13,sK29(X0,sK16,'_dst_x0_1'(s(sK14,sK15)),sK13,sK14)))))))) )),
  inference(subsumption_resolution,[],[f9681,f327])).
fof(f9681,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0)))) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,happ(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(sK13,sK29(X0,sK16,'_dst_x0_1'(s(sK14,sK15)),sK13,sK14)))))))) )),
  inference(superposition,[],[f436,f644])).
fof(f345,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.bool.EQ_CLAUSES|split|3')).
fof(f346,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f213])).
fof(f582,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f330,f327])).
fof(f330,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f21])).
fof(f21,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.bool.NOT_CLAUSES|split|2')).
fof(f1462,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2))) | s(bool,'F_0') = s(bool,X1)) )),
  inference(superposition,[],[f428,f334])).
fof(f1669,plain,(
  ( ! [X0] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X0))) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X0)))) )),
  inference(superposition,[],[f369,f334])).
fof(f1670,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,happ(s(fun(X2,bool),X3),s(X2,'const.min.@_1'(s(fun(X2,bool),X3))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),X3)))) )),
  inference(superposition,[],[f369,f334])).
fof(f2053,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X2,bool),X1),s(fun(bool,bool),X0)))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK28(X0,X1,X2,bool))) )),
  inference(superposition,[],[f434,f334])).
fof(f2063,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X2,bool),X1),s(fun(bool,bool),X0)))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK30(X0,X1,X2,bool))) )),
  inference(superposition,[],[f437,f334])).
fof(f2074,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X2,bool),X1),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK32(X0,X1,X2,bool))) )),
  inference(superposition,[],[f440,f334])).
fof(f2087,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,bool),X1),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK34(X0,X1,X2,bool))) )),
  inference(superposition,[],[f443,f334])).
fof(f2100,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,sK21(X0,X1,bool))) )),
  inference(superposition,[],[f384,f334])).
fof(f2117,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,sK22(X0,X1,bool))) )),
  inference(superposition,[],[f386,f334])).
fof(f2135,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,sK23(X0,X1,bool))) )),
  inference(superposition,[],[f387,f334])).
fof(f2163,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X2)))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,sK24(X0,X1,bool))) )),
  inference(superposition,[],[f389,f334])).
fof(f2194,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,sK19(X0,bool)) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0)))) | ~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK20(X0,bool))) )),
  inference(superposition,[],[f376,f334])).
fof(f2210,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK25(X0,X1,X2,bool))) )),
  inference(superposition,[],[f430,f334])).
fof(f2211,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK25(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK25(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f430,f334])).
fof(f2222,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,happ(s(fun(X2,bool),X1),s(X2,X3))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X2,bool),X1),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK28(X0,X1,X2,bool))) )),
  inference(superposition,[],[f435,f334])).
fof(f2233,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,sK28(X7,X5,X4,bool)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X4,bool),X5),s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6)))) )),
  inference(superposition,[],[f435,f334])).
fof(f2248,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,happ(s(fun(X2,bool),X1),s(X2,X3))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X2,bool),X1),s(fun(bool,bool),X0)))) | s(bool,'T_0') = s(bool,sK32(X0,X1,X2,bool))) )),
  inference(superposition,[],[f441,f334])).
fof(f2259,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'F_0') != s(bool,sK32(X7,X5,X4,bool)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X4,bool),X5),s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6)))) )),
  inference(superposition,[],[f441,f334])).
fof(f2303,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,'F_0')))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,bool),X5),s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6)))) )),
  inference(superposition,[],[f438,f334])).
fof(f2335,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,bool),X5),s(fun(bool,bool),X7)))) | p(s(bool,happ(s(fun(bool,bool),X7),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(X4,bool),X5),s(X4,X6)))) )),
  inference(superposition,[],[f444,f334])).
fof(f2438,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X2,bool),X0),s(X2,happ(s(fun(bool,X2),X1),s(bool,'F_0')))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(bool,X2),X1),s(fun(X2,bool),X0)))) | s(bool,'T_0') = s(bool,sK26(X0,X1,bool,X2))) )),
  inference(superposition,[],[f432,f334])).
fof(f2439,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0')))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,bool),X1),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK26(X2,X1,X0,bool))))) )),
  inference(superposition,[],[f432,f334])).
fof(f2454,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X2,bool),X0),s(X2,happ(s(fun(bool,X2),X1),s(bool,'F_0')))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(bool,X2),X1),s(fun(X2,bool),X0)))) | s(bool,'T_0') = s(bool,sK35(X0,X1,bool,X2))) )),
  inference(superposition,[],[f446,f334])).
fof(f2455,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,bool),X1),s(fun(bool,bool),X2)))) | p(s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK35(X2,X1,X0,bool))))) )),
  inference(superposition,[],[f446,f334])).
fof(f2579,plain,(
  ( ! [X0,X1] : (~sP3('F_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f596,f590,f1445,f475])).
fof(f2585,plain,(
  ( ! [X2,X0,X1] : (~sP6('F_0',happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X1,X2)),X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1445,f334])).
fof(f2621,plain,(
  ( ! [X2,X0,X1] : (sP6('T_0',happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X1,X2)),X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1446,f334])).
fof(f2651,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f598,f592,f1448,f474])).
fof(f2654,plain,(
  ( ! [X2,X0,X1] : (~sP6('F_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1448,f334])).
fof(f2685,plain,(
  ( ! [X2,X0,X1] : (sP6('T_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X1,X2))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1449,f334])).
fof(f2723,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X2,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(superposition,[],[f1671,f334])).
fof(f4185,plain,(
  ( ! [X308,X310,X309,X311] : (s(bool,'T_0') != s(bool,f(s(X308,sK25(X309,X310,bool,X308)))) | s(fun(X308,bool),X309) = s(fun(X308,bool),X310) | s(bool,X311) = s(bool,g(s(X308,sK25(X309,X310,bool,X308)))) | s(bool,'T_0') = s(bool,X311)) )),
  inference(superposition,[],[f430,f890])).
fof(f3924,plain,(
  ( ! [X315,X316,X314,X317] : (s(bool,X317) != s(bool,f(s(X314,sK25(X315,X316,bool,X314)))) | s(fun(X314,bool),X315) = s(fun(X314,bool),X316) | s(bool,'T_0') = s(bool,X317) | s(bool,'T_0') = s(bool,g(s(X314,sK25(X315,X316,bool,X314))))) )),
  inference(superposition,[],[f430,f890])).
fof(f436,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK29(X6,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK29(X6,X3,X2,X1,X0))))))))) & ((p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK30(X3,X2,X1,X0))))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK29,sK30])],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(ennf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X7,X6),X8),s(fun(X6,bool),X9)))) <=> ! [X10] : (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) => ? [X21] : p(s(bool,happ(s(fun(X6,bool),X9),s(X6,happ(s(fun(X7,X6),X8),s(X7,X21))))))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.quantHeuristics.GUESS_EXISTS_FORALL_REWRITES|split|0')).
fof(f437,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK30(X3,X2,X1,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f266])).
fof(f438,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f266])).
fof(f2066,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK30(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X5,X0),X4),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f2062,f428])).
fof(f2062,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK30(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X5,X0),X4),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f437,f428])).
fof(f2067,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X3,X0),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK30(X1,X2,X3,X0))))) )),
  inference(subsumption_resolution,[],[f2064,f582])).
fof(f2064,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X3,X0),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK30(X1,X2,X3,X0))))) )),
  inference(superposition,[],[f437,f334])).
fof(f2545,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK29(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6))))) )),
  inference(forward_demodulation,[],[f2541,f428])).
fof(f2541,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6)))) | p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(X0,sK29(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)))))))) )),
  inference(superposition,[],[f436,f428])).
fof(f2547,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK29(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0)))))))) )),
  inference(forward_demodulation,[],[f2546,f428])).
fof(f2546,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK29(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0)))))))) )),
  inference(forward_demodulation,[],[f2542,f428])).
fof(f2542,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X6)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK29(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0)))))))) )),
  inference(superposition,[],[f436,f428])).
fof(f4648,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (~p(s(bool,happ(s(fun(X134,bool),X136),s(X134,X137)))) | p(s(bool,happ(s(fun(X134,bool),X136),s(X134,happ(s(fun(X133,X134),X135),s(X133,sK29(X137,X136,X135,X133,X134))))))) | s(bool,X138) = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X133,X134),X135),s(fun(X134,bool),X136))) | s(bool,'T_0') = s(bool,X138)) )),
  inference(subsumption_resolution,[],[f4144,f327])).
fof(f4144,plain,(
  ( ! [X134,X138,X136,X135,X133,X137] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X134,bool),X136),s(X134,X137)))) | p(s(bool,happ(s(fun(X134,bool),X136),s(X134,happ(s(fun(X133,X134),X135),s(X133,sK29(X137,X136,X135,X133,X134))))))) | s(bool,X138) = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X133,X134),X135),s(fun(X134,bool),X136))) | s(bool,'T_0') = s(bool,X138)) )),
  inference(superposition,[],[f436,f890])).
fof(f4654,plain,(
  ( ! [X167,X165,X169,X166,X170,X168] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X167,X165),X168),s(fun(X165,bool),X166)))) | s(bool,X170) = s(bool,happ(s(fun(X165,bool),X166),s(X165,happ(s(fun(X167,X165),X168),s(X167,X169))))) | s(bool,'T_0') = s(bool,X170)) )),
  inference(subsumption_resolution,[],[f4153,f327])).
fof(f4153,plain,(
  ( ! [X167,X165,X169,X166,X170,X168] : (~p(s(bool,'T_0')) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X167,X165),X168),s(fun(X165,bool),X166)))) | s(bool,X170) = s(bool,happ(s(fun(X165,bool),X166),s(X165,happ(s(fun(X167,X165),X168),s(X167,X169))))) | s(bool,'T_0') = s(bool,X170)) )),
  inference(superposition,[],[f438,f890])).
fof(f2304,plain,(
  ( ! [X12,X10,X8,X9] : (~p(s(bool,happ(s(fun(X8,bool),X12),s(X8,X10)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X9,X8),'_dst_x0_1'(s(X8,X10))),s(fun(X8,bool),X12))))) )),
  inference(superposition,[],[f438,f428])).
fof(f2302,plain,(
  ( ! [X2,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,bool),X2),s(fun(bool,bool),X3))))) )),
  inference(superposition,[],[f438,f369])).
fof(f3927,plain,(
  ( ! [X327,X329,X331,X328,X330] : (~p(s(bool,happ(s(fun(bool,bool),X331),s(bool,X330)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X327,bool),X328),s(fun(bool,bool),X331)))) | s(bool,'T_0') = s(bool,X330) | s(bool,'T_0') = s(bool,happ(s(fun(X327,bool),X328),s(X327,X329)))) )),
  inference(superposition,[],[f438,f890])).
fof(f3942,plain,(
  ( ! [X391,X393,X390,X392] : (p(s(bool,happ(s(fun(bool,bool),X390),s(bool,X393)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X392,bool),X391),s(fun(bool,bool),X390)))) | s(bool,'T_0') = s(bool,X393) | s(bool,'T_0') = s(bool,sK30(X390,X391,X392,bool))) )),
  inference(superposition,[],[f437,f890])).
fof(f4188,plain,(
  ( ! [X325,X321,X323,X324,X322] : (~p(s(bool,happ(s(fun(bool,bool),X324),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X321,bool),X322),s(fun(bool,bool),X324)))) | s(bool,X325) = s(bool,happ(s(fun(X321,bool),X322),s(X321,X323))) | s(bool,'T_0') = s(bool,X325)) )),
  inference(superposition,[],[f438,f890])).
fof(f4203,plain,(
  ( ! [X383,X385,X384,X386] : (p(s(bool,happ(s(fun(bool,bool),X383),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X385,bool),X384),s(fun(bool,bool),X383)))) | s(bool,X386) = s(bool,sK30(X383,X384,X385,bool)) | s(bool,'T_0') = s(bool,X386)) )),
  inference(superposition,[],[f437,f890])).
fof(f2301,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X0))))) ) | $spl2),
  inference(superposition,[],[f438,f643])).
fof(f372,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X1,X0)))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X1,X0))))) & ! [X5,X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | s(X0,X5) = s(X0,X6)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X1,X0))))) & s(X0,sK19(X1,X0)) != s(X0,sK20(X1,X0))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19,sK20])],[f230])).
fof(f230,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) & ! [X5,X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | s(X0,X5) = s(X0,X6)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f229])).
fof(f229,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3)))) & ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ? [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & s(X0,X2) != s(X0,X3))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s(X0,X2) = s(X0,X3))))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => s(X0,X2) = s(X0,X3))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & ! [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => s(X0,X2) = s(X0,X3))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X7] : ! [X9] : (p(s(bool,'const.bool.?|33|_1'(s(fun(X7,bool),X9)))) <=> (? [X14] : p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X14)))) & ! [X14,X16] : ((p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X14)))) & p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X16))))) => s(X7,X14) = s(X7,X16))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.bool.EXISTS_UNIQUE_THM')).
fof(f373,plain,(
  ( ! [X6,X0,X5,X1] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | s(X0,X5) = s(X0,X6)) )),
  inference(cnf_transformation,[],[f231])).
fof(f374,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X1,X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f375,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X1,X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f376,plain,(
  ( ! [X4,X0,X1] : (s(X0,sK19(X1,X0)) != s(X0,sK20(X1,X0)) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f1657,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK18(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)))))) )),
  inference(forward_demodulation,[],[f1654,f428])).
fof(f1654,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,sK18(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0)))))) )),
  inference(superposition,[],[f372,f428])).
fof(f2384,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK19(X1,X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))) )),
  inference(superposition,[],[f374,f369])).
fof(f2388,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f2387,f428])).
fof(f2387,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(forward_demodulation,[],[f2375,f428])).
fof(f2375,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,sK19(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f374,f428])).
fof(f2412,plain,(
  ( ! [X0,X1] : (~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X1,X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X1))))) )),
  inference(superposition,[],[f375,f369])).
fof(f2416,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f2415,f428])).
fof(f2415,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(forward_demodulation,[],[f2403,f428])).
fof(f2403,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,sK20(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X0))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f375,f428])).
fof(f4649,plain,(
  ( ! [X144,X142,X145,X143,X141] : (~p(s(bool,happ(s(fun(X141,bool),X142),s(X141,X143)))) | ~p(s(bool,happ(s(fun(X141,bool),X142),s(X141,X144)))) | s(X141,X143) = s(X141,X144) | s(bool,X145) = s(bool,'const.bool.?|33|_1'(s(fun(X141,bool),X142))) | s(bool,'T_0') = s(bool,X145)) )),
  inference(subsumption_resolution,[],[f4147,f327])).
fof(f4147,plain,(
  ( ! [X144,X142,X145,X143,X141] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X141,bool),X142),s(X141,X143)))) | ~p(s(bool,happ(s(fun(X141,bool),X142),s(X141,X144)))) | s(X141,X143) = s(X141,X144) | s(bool,X145) = s(bool,'const.bool.?|33|_1'(s(fun(X141,bool),X142))) | s(bool,'T_0') = s(bool,X145)) )),
  inference(superposition,[],[f373,f890])).
fof(f4650,plain,(
  ( ! [X146,X147,X148] : (p(s(bool,happ(s(fun(X146,bool),X147),s(X146,sK18(X147,X146))))) | s(bool,X148) = s(bool,'const.bool.?|33|_1'(s(fun(X146,bool),X147))) | s(bool,'T_0') = s(bool,X148)) )),
  inference(subsumption_resolution,[],[f4148,f327])).
fof(f4148,plain,(
  ( ! [X146,X147,X148] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X146,bool),X147),s(X146,sK18(X147,X146))))) | s(bool,X148) = s(bool,'const.bool.?|33|_1'(s(fun(X146,bool),X147))) | s(bool,'T_0') = s(bool,X148)) )),
  inference(superposition,[],[f372,f890])).
fof(f4652,plain,(
  ( ! [X154,X152,X151,X153] : (p(s(bool,happ(s(fun(X151,bool),X152),s(X151,sK20(X152,X151))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X151,bool),X152)))) | s(bool,X154) = s(bool,happ(s(fun(X151,bool),X152),s(X151,X153))) | s(bool,'T_0') = s(bool,X154)) )),
  inference(subsumption_resolution,[],[f4150,f327])).
fof(f4150,plain,(
  ( ! [X154,X152,X151,X153] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X151,bool),X152),s(X151,sK20(X152,X151))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X151,bool),X152)))) | s(bool,X154) = s(bool,happ(s(fun(X151,bool),X152),s(X151,X153))) | s(bool,'T_0') = s(bool,X154)) )),
  inference(superposition,[],[f375,f890])).
fof(f4653,plain,(
  ( ! [X158,X156,X157,X155] : (p(s(bool,happ(s(fun(X155,bool),X156),s(X155,sK19(X156,X155))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X155,bool),X156)))) | s(bool,X158) = s(bool,happ(s(fun(X155,bool),X156),s(X155,X157))) | s(bool,'T_0') = s(bool,X158)) )),
  inference(subsumption_resolution,[],[f4151,f327])).
fof(f4151,plain,(
  ( ! [X158,X156,X157,X155] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X155,bool),X156),s(X155,sK19(X156,X155))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(X155,bool),X156)))) | s(bool,X158) = s(bool,happ(s(fun(X155,bool),X156),s(X155,X157))) | s(bool,'T_0') = s(bool,X158)) )),
  inference(superposition,[],[f374,f890])).
fof(f4194,plain,(
  ( ! [X349,X347,X348] : (s(bool,'T_0') != s(bool,sK19(X347,bool)) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X347)))) | ~p(s(bool,happ(s(fun(bool,bool),X347),s(bool,X348)))) | s(bool,X349) = s(bool,sK20(X347,bool)) | s(bool,'T_0') = s(bool,X349)) )),
  inference(superposition,[],[f376,f890])).
fof(f3933,plain,(
  ( ! [X356,X354,X355] : (s(bool,X355) != s(bool,sK19(X354,bool)) | p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X354)))) | ~p(s(bool,happ(s(fun(bool,bool),X354),s(bool,X356)))) | s(bool,'T_0') = s(bool,X355) | s(bool,'T_0') = s(bool,sK20(X354,bool))) )),
  inference(superposition,[],[f376,f890])).
fof(f384,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(X2,X1,X0))))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & ((~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0)))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21,sK22])],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ! [X5] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (? [X4] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) & ! [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ! [X4] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X7] : ! [X9,X17] : ((p(s(bool,X9)) & ! [X14] : p(s(bool,happ(s(fun(X7,bool),X17),s(X7,X14))))) <=> ! [X14] : (p(s(bool,X9)) & p(s(bool,happ(s(fun(X7,bool),X17),s(X7,X14))))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.bool.RIGHT_AND_FORALL_THM')).
fof(f386,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f240])).
fof(f387,plain,(
  ( ! [X2,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0))))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23,sK24])],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2] : ((? [X6] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) & ~p(s(bool,X1))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(rectify,[],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((? [X3] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,X1))) | (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1)))) & ((? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & ~p(s(bool,X1))) | ! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1)))))),
  inference(nnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) <=> (! [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,X1))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X7] : ! [X17,X9] : (! [X14] : (p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X14)))) | p(s(bool,X17))) <=> (! [X14] : p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X14)))) | p(s(bool,X17))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.bool.LEFT_FORALL_OR_THM')).
fof(f389,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f244])).
fof(f433,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X6,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X6,X3,X2,X1,X0)))))) & ((p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK28(X3,X2,X1,X0))))) & ! [X5] : s(X0,sK28(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27,sK28])],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f261])).
fof(f261,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))) & (? [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6,X7] : ! [X22,X23] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X7,X6),X22),s(fun(X6,bool),X23)))) <=> ! [X10] : (p(s(bool,happ(s(fun(X6,bool),X23),s(X6,X10)))) => ? [X21] : s(X6,X10) = s(X6,happ(s(fun(X7,X6),X22),s(X7,X21)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.quantHeuristics.GUESS_REWRITES|split|4')).
fof(f439,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK31(X6,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK31(X6,X3,X2,X1,X0)))))) & ((~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK32(X3,X2,X1,X0))))) & ! [X5] : s(X0,sK32(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK31,sK32])],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X7))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : s(X0,X4) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(ennf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : s(X0,X4) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6,X7] : ! [X22,X23] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X7,X6),X22),s(fun(X6,bool),X23)))) <=> ! [X10] : (~p(s(bool,happ(s(fun(X6,bool),X23),s(X6,X10)))) => ? [X21] : s(X6,X10) = s(X6,happ(s(fun(X7,X6),X22),s(X7,X21)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.quantHeuristics.GUESS_REWRITES|split|5')).
fof(f442,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK33(X6,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK33(X6,X3,X2,X1,X0))))))))) & ((~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK34(X3,X2,X1,X0))))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK33,sK34])],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X6] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X6)))) | ? [X7] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X7)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f270])).
fof(f270,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))))) & (? [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) & ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) | ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(ennf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))) => ? [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X7,X6),X8),s(fun(X6,bool),X9)))) <=> ! [X10] : (~p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))) => ? [X21] : ~p(s(bool,happ(s(fun(X6,bool),X9),s(X6,happ(s(fun(X7,X6),X8),s(X7,X21))))))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.quantHeuristics.GUESS_EXISTS_FORALL_REWRITES|split|1')).
fof(f2104,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,X4))) )),
  inference(forward_demodulation,[],[f2099,f428])).
fof(f2099,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK21(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X5)))) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f384,f428])).
fof(f2121,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | ~p(s(bool,X4))) )),
  inference(forward_demodulation,[],[f2116,f428])).
fof(f2116,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X5)))) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f386,f428])).
fof(f2139,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X4))) )),
  inference(forward_demodulation,[],[f2134,f428])).
fof(f2134,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X5)))) | p(s(bool,X4))) )),
  inference(superposition,[],[f387,f428])).
fof(f2167,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) | p(s(bool,X4))) )),
  inference(forward_demodulation,[],[f2162,f428])).
fof(f2162,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X5)))) | p(s(bool,X4))) )),
  inference(superposition,[],[f389,f428])).
fof(f2336,plain,(
  ( ! [X12,X10,X8,X9] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X9,X8),'_dst_x0_1'(s(X8,X10))),s(fun(X8,bool),X12)))) | p(s(bool,happ(s(fun(X8,bool),X12),s(X8,X10))))) )),
  inference(superposition,[],[f444,f428])).
fof(f2440,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,happ(s(fun(X3,bool),X6),s(X3,X5)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X4,X3),'_dst_x0_1'(s(X3,X5))),s(fun(X3,bool),X6))))) )),
  inference(superposition,[],[f432,f428])).
fof(f2456,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X4,X3),'_dst_x0_1'(s(X3,X5))),s(fun(X3,bool),X6)))) | p(s(bool,happ(s(fun(X3,bool),X6),s(X3,X5))))) )),
  inference(superposition,[],[f446,f428])).
fof(f2496,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,sK27(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6))))) )),
  inference(forward_demodulation,[],[f2493,f428])).
fof(f2493,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(X0,sK27(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))) )),
  inference(superposition,[],[f433,f428])).
fof(f2497,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK27(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) )),
  inference(forward_demodulation,[],[f2494,f428])).
fof(f2494,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK27(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) )),
  inference(superposition,[],[f433,f428])).
fof(f2515,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,sK31(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)))) | p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6))))) )),
  inference(forward_demodulation,[],[f2512,f428])).
fof(f2512,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6)))) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(X0,sK31(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))) )),
  inference(superposition,[],[f439,f428])).
fof(f2516,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK31(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) )),
  inference(forward_demodulation,[],[f2513,f428])).
fof(f2513,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X6)))) | s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK31(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) )),
  inference(superposition,[],[f439,f428])).
fof(f2561,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK33(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))))) | p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6))))) )),
  inference(forward_demodulation,[],[f2558,f428])).
fof(f2558,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | p(s(bool,happ(s(fun(X1,bool),X5),s(X1,X6)))) | ~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(X0,sK33(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)))))))) )),
  inference(superposition,[],[f442,f428])).
fof(f2563,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK33(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0)))))))) )),
  inference(forward_demodulation,[],[f2562,f428])).
fof(f2562,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK33(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0)))))))) )),
  inference(forward_demodulation,[],[f2559,f428])).
fof(f2559,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK33(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0)))))))) )),
  inference(superposition,[],[f442,f428])).
fof(f3926,plain,(
  ( ! [X325,X323,X324,X326,X322] : (p(s(bool,happ(s(fun(bool,bool),X326),s(bool,X325)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X322,bool),X323),s(fun(bool,bool),X326)))) | s(bool,'T_0') = s(bool,X325) | s(bool,'T_0') = s(bool,happ(s(fun(X322,bool),X323),s(X322,X324)))) )),
  inference(superposition,[],[f444,f890])).
fof(f3928,plain,(
  ( ! [X333,X335,X332,X334,X336] : (s(bool,X335) != s(bool,sK32(X336,X333,X332,bool)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X332,bool),X333),s(fun(bool,bool),X336)))) | s(bool,'T_0') = s(bool,X335) | s(bool,'T_0') = s(bool,happ(s(fun(X332,bool),X333),s(X332,X334)))) )),
  inference(superposition,[],[f441,f890])).
fof(f3929,plain,(
  ( ! [X341,X337,X339,X340,X338] : (s(bool,X340) != s(bool,sK28(X341,X338,X337,bool)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X337,bool),X338),s(fun(bool,bool),X341)))) | s(bool,'T_0') = s(bool,X340) | s(bool,'T_0') = s(bool,happ(s(fun(X337,bool),X338),s(X337,X339)))) )),
  inference(superposition,[],[f435,f890])).
fof(f3940,plain,(
  ( ! [X383,X385,X382,X384,X381] : (s(bool,X384) != s(bool,happ(s(fun(X383,bool),X382),s(X383,X385))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X383,bool),X382),s(fun(bool,bool),X381)))) | s(bool,'T_0') = s(bool,X384) | s(bool,'T_0') = s(bool,sK28(X381,X382,X383,bool))) )),
  inference(superposition,[],[f435,f890])).
fof(f3943,plain,(
  ( ! [X397,X395,X396,X398,X394] : (s(bool,X397) != s(bool,happ(s(fun(X396,bool),X395),s(X396,X398))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X396,bool),X395),s(fun(bool,bool),X394)))) | s(bool,'T_0') = s(bool,X397) | s(bool,'T_0') = s(bool,sK32(X394,X395,X396,bool))) )),
  inference(superposition,[],[f441,f890])).
fof(f4187,plain,(
  ( ! [X316,X318,X320,X317,X319] : (p(s(bool,happ(s(fun(bool,bool),X319),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X316,bool),X317),s(fun(bool,bool),X319)))) | s(bool,X320) = s(bool,happ(s(fun(X316,bool),X317),s(X316,X318))) | s(bool,'T_0') = s(bool,X320)) )),
  inference(superposition,[],[f444,f890])).
fof(f4189,plain,(
  ( ! [X327,X329,X326,X328,X330] : (s(bool,'T_0') != s(bool,sK32(X329,X327,X326,bool)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X326,bool),X327),s(fun(bool,bool),X329)))) | s(bool,X330) = s(bool,happ(s(fun(X326,bool),X327),s(X326,X328))) | s(bool,'T_0') = s(bool,X330)) )),
  inference(superposition,[],[f441,f890])).
fof(f4190,plain,(
  ( ! [X333,X335,X331,X332,X334] : (s(bool,'T_0') != s(bool,sK28(X334,X332,X331,bool)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X331,bool),X332),s(fun(bool,bool),X334)))) | s(bool,X335) = s(bool,happ(s(fun(X331,bool),X332),s(X331,X333))) | s(bool,'T_0') = s(bool,X335)) )),
  inference(superposition,[],[f435,f890])).
fof(f4201,plain,(
  ( ! [X377,X374,X376,X378,X375] : (s(bool,'T_0') != s(bool,happ(s(fun(X376,bool),X375),s(X376,X377))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X376,bool),X375),s(fun(bool,bool),X374)))) | s(bool,X378) = s(bool,sK28(X374,X375,X376,bool)) | s(bool,'T_0') = s(bool,X378)) )),
  inference(superposition,[],[f435,f890])).
fof(f4204,plain,(
  ( ! [X389,X391,X387,X388,X390] : (s(bool,'T_0') != s(bool,happ(s(fun(X389,bool),X388),s(X389,X390))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X389,bool),X388),s(fun(bool,bool),X387)))) | s(bool,X391) = s(bool,sK32(X387,X388,X389,bool)) | s(bool,'T_0') = s(bool,X391)) )),
  inference(superposition,[],[f441,f890])).
fof(f4656,plain,(
  ( ! [X187,X185,X188,X186,X184] : (p(s(bool,happ(s(fun(X184,bool),X185),s(X184,X187)))) | ~p(s(bool,X186)) | s(bool,X188) = s(bool,happ(s(fun(X184,bool),X185),s(X184,sK21(X185,X186,X184)))) | s(bool,'T_0') = s(bool,X188)) )),
  inference(subsumption_resolution,[],[f4157,f327])).
fof(f4157,plain,(
  ( ! [X187,X185,X188,X186,X184] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X184,bool),X185),s(X184,X187)))) | ~p(s(bool,X186)) | s(bool,X188) = s(bool,happ(s(fun(X184,bool),X185),s(X184,sK21(X185,X186,X184)))) | s(bool,'T_0') = s(bool,X188)) )),
  inference(superposition,[],[f384,f890])).
fof(f4657,plain,(
  ( ! [X191,X189,X192,X190,X193] : (p(s(bool,happ(s(fun(X189,bool),X190),s(X189,X192)))) | ~p(s(bool,X191)) | s(bool,X193) = s(bool,happ(s(fun(X189,bool),X190),s(X189,sK22(X190,X191,X189)))) | s(bool,'T_0') = s(bool,X193)) )),
  inference(subsumption_resolution,[],[f4158,f327])).
fof(f4158,plain,(
  ( ! [X191,X189,X192,X190,X193] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X189,bool),X190),s(X189,X192)))) | ~p(s(bool,X191)) | s(bool,X193) = s(bool,happ(s(fun(X189,bool),X190),s(X189,sK22(X190,X191,X189)))) | s(bool,'T_0') = s(bool,X193)) )),
  inference(superposition,[],[f386,f890])).
fof(f4658,plain,(
  ( ! [X198,X196,X194,X197,X195] : (p(s(bool,happ(s(fun(X194,bool),X195),s(X194,X197)))) | p(s(bool,X196)) | s(bool,X198) = s(bool,happ(s(fun(X194,bool),X195),s(X194,sK23(X195,X196,X194)))) | s(bool,'T_0') = s(bool,X198)) )),
  inference(subsumption_resolution,[],[f4159,f327])).
fof(f4159,plain,(
  ( ! [X198,X196,X194,X197,X195] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X194,bool),X195),s(X194,X197)))) | p(s(bool,X196)) | s(bool,X198) = s(bool,happ(s(fun(X194,bool),X195),s(X194,sK23(X195,X196,X194)))) | s(bool,'T_0') = s(bool,X198)) )),
  inference(superposition,[],[f387,f890])).
fof(f4659,plain,(
  ( ! [X202,X200,X199,X203,X201] : (p(s(bool,happ(s(fun(X199,bool),X200),s(X199,X202)))) | p(s(bool,X201)) | s(bool,X203) = s(bool,happ(s(fun(X199,bool),X200),s(X199,sK24(X200,X201,X199)))) | s(bool,'T_0') = s(bool,X203)) )),
  inference(subsumption_resolution,[],[f4160,f327])).
fof(f4160,plain,(
  ( ! [X202,X200,X199,X203,X201] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X199,bool),X200),s(X199,X202)))) | p(s(bool,X201)) | s(bool,X203) = s(bool,happ(s(fun(X199,bool),X200),s(X199,sK24(X200,X201,X199)))) | s(bool,'T_0') = s(bool,X203)) )),
  inference(superposition,[],[f389,f890])).
fof(f4662,plain,(
  ( ! [X231,X229,X233,X230,X234,X232] : (p(s(bool,happ(s(fun(X230,bool),X232),s(X230,X233)))) | ~p(s(bool,happ(s(fun(X230,bool),X232),s(X230,happ(s(fun(X229,X230),X231),s(X229,sK33(X233,X232,X231,X229,X230))))))) | s(bool,X234) = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X229,X230),X231),s(fun(X230,bool),X232))) | s(bool,'T_0') = s(bool,X234)) )),
  inference(subsumption_resolution,[],[f4168,f327])).
fof(f4168,plain,(
  ( ! [X231,X229,X233,X230,X234,X232] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X230,bool),X232),s(X230,X233)))) | ~p(s(bool,happ(s(fun(X230,bool),X232),s(X230,happ(s(fun(X229,X230),X231),s(X229,sK33(X233,X232,X231,X229,X230))))))) | s(bool,X234) = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X229,X230),X231),s(fun(X230,bool),X232))) | s(bool,'T_0') = s(bool,X234)) )),
  inference(superposition,[],[f442,f890])).
fof(f4665,plain,(
  ( ! [X251,X249,X252,X250,X248,X247] : (~p(s(bool,happ(s(fun(X248,bool),X250),s(X248,X251)))) | s(X248,X251) = s(X248,happ(s(fun(X247,X248),X249),s(X247,sK27(X251,X250,X249,X247,X248)))) | s(bool,X252) = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X247,X248),X249),s(fun(X248,bool),X250))) | s(bool,'T_0') = s(bool,X252)) )),
  inference(subsumption_resolution,[],[f4171,f327])).
fof(f4171,plain,(
  ( ! [X251,X249,X252,X250,X248,X247] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X248,bool),X250),s(X248,X251)))) | s(X248,X251) = s(X248,happ(s(fun(X247,X248),X249),s(X247,sK27(X251,X250,X249,X247,X248)))) | s(bool,X252) = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X247,X248),X249),s(fun(X248,bool),X250))) | s(bool,'T_0') = s(bool,X252)) )),
  inference(superposition,[],[f433,f890])).
fof(f4666,plain,(
  ( ! [X255,X253,X257,X254,X256,X258] : (p(s(bool,happ(s(fun(X254,bool),X256),s(X254,X257)))) | s(X254,X257) = s(X254,happ(s(fun(X253,X254),X255),s(X253,sK31(X257,X256,X255,X253,X254)))) | s(bool,X258) = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X253,X254),X255),s(fun(X254,bool),X256))) | s(bool,'T_0') = s(bool,X258)) )),
  inference(subsumption_resolution,[],[f4172,f327])).
fof(f4172,plain,(
  ( ! [X255,X253,X257,X254,X256,X258] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X254,bool),X256),s(X254,X257)))) | s(X254,X257) = s(X254,happ(s(fun(X253,X254),X255),s(X253,sK31(X257,X256,X255,X253,X254)))) | s(bool,X258) = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X253,X254),X255),s(fun(X254,bool),X256))) | s(bool,'T_0') = s(bool,X258)) )),
  inference(superposition,[],[f439,f890])).
fof(f431,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK26(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK26])],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f258])).
fof(f258,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f123])).
fof(f123,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6,X7] : ! [X22,X23] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X7,X6),X22),s(fun(X6,bool),X23)))) <=> ! [X21] : p(s(bool,happ(s(fun(X6,bool),X23),s(X6,happ(s(fun(X7,X6),X22),s(X7,X21)))))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.quantHeuristics.GUESS_REWRITES|split|2')).
fof(f444,plain,(
  ( ! [X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f272])).
fof(f445,plain,(
  ( ! [X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK35(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK35])],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))))))) & (? [X4] : p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))))),
  inference(nnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) <=> ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,X4)))))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6,X7] : ! [X22,X23] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X7,X6),X22),s(fun(X6,bool),X23)))) <=> ! [X21] : ~p(s(bool,happ(s(fun(X6,bool),X23),s(X6,happ(s(fun(X7,X6),X22),s(X7,X21)))))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.quantHeuristics.GUESS_REWRITES|split|3')).
fof(f2341,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X0),X3),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4)))))) )),
  inference(subsumption_resolution,[],[f2337,f582])).
fof(f2337,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X0),X3),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4)))))) )),
  inference(superposition,[],[f444,f334])).
fof(f4663,plain,(
  ( ! [X239,X237,X235,X240,X238,X236] : (p(s(bool,happ(s(fun(X236,bool),X238),s(X236,happ(s(fun(X235,X236),X237),s(X235,X239)))))) | s(bool,X240) = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X235,X236),X237),s(fun(X236,bool),X238))) | s(bool,'T_0') = s(bool,X240)) )),
  inference(subsumption_resolution,[],[f4169,f327])).
fof(f4169,plain,(
  ( ! [X239,X237,X235,X240,X238,X236] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X236,bool),X238),s(X236,happ(s(fun(X235,X236),X237),s(X235,X239)))))) | s(bool,X240) = s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X235,X236),X237),s(fun(X236,bool),X238))) | s(bool,'T_0') = s(bool,X240)) )),
  inference(superposition,[],[f431,f890])).
fof(f4664,plain,(
  ( ! [X243,X241,X246,X244,X242,X245] : (~p(s(bool,happ(s(fun(X242,bool),X244),s(X242,happ(s(fun(X241,X242),X243),s(X241,X245)))))) | s(bool,X246) = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X241,X242),X243),s(fun(X242,bool),X244))) | s(bool,'T_0') = s(bool,X246)) )),
  inference(subsumption_resolution,[],[f4170,f327])).
fof(f4170,plain,(
  ( ! [X243,X241,X246,X244,X242,X245] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X242,bool),X244),s(X242,happ(s(fun(X241,X242),X243),s(X241,X245)))))) | s(bool,X246) = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X241,X242),X243),s(fun(X242,bool),X244))) | s(bool,'T_0') = s(bool,X246)) )),
  inference(superposition,[],[f445,f890])).
fof(f4655,plain,(
  ( ! [X175,X173,X171,X174,X172] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X173,X171),X174),s(fun(X171,bool),X172)))) | s(bool,X175) = s(bool,happ(s(fun(X171,bool),X172),s(X171,happ(s(fun(X173,X171),X174),s(X173,sK26(X172,X174,X173,X171)))))) | s(bool,'T_0') = s(bool,X175)) )),
  inference(subsumption_resolution,[],[f4154,f327])).
fof(f4154,plain,(
  ( ! [X175,X173,X171,X174,X172] : (~p(s(bool,'T_0')) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X173,X171),X174),s(fun(X171,bool),X172)))) | s(bool,X175) = s(bool,happ(s(fun(X171,bool),X172),s(X171,happ(s(fun(X173,X171),X174),s(X173,sK26(X172,X174,X173,X171)))))) | s(bool,'T_0') = s(bool,X175)) )),
  inference(superposition,[],[f432,f890])).
fof(f432,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK26(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f260])).
fof(f2443,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK26(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f2436,f428])).
fof(f2436,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK26(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X4,X0),X5),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f432,f428])).
fof(f2444,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK26(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5))))) )),
  inference(forward_demodulation,[],[f2437,f428])).
fof(f2437,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK26(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(fun(X1,bool),X5))))) )),
  inference(superposition,[],[f432,f428])).
fof(f2461,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X2,X0),X3),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,sK35(X1,X3,X2,X0))))))) )),
  inference(subsumption_resolution,[],[f2457,f582])).
fof(f2457,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X2,X0),X3),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,sK35(X1,X3,X2,X0))))))) )),
  inference(superposition,[],[f446,f334])).
fof(f446,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,happ(s(fun(X1,X0),X2),s(X1,sK35(X3,X2,X1,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f275])).
fof(f2459,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK35(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f2452,f428])).
fof(f2452,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,happ(s(fun(X4,X0),X5),s(X4,sK35(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X4,X0),X5),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f446,f428])).
fof(f2460,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK35(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5))))) )),
  inference(forward_demodulation,[],[f2453,f428])).
fof(f2453,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),X5),s(X1,happ(s(fun(X0,X1),X3),s(X0,sK35(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(fun(X1,bool),X5))))) )),
  inference(superposition,[],[f446,f428])).
fof(f3946,plain,(
  ( ! [X408,X410,X407,X409] : (p(s(bool,happ(s(fun(X409,bool),X407),s(X409,happ(s(fun(bool,X409),X408),s(bool,X410)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(bool,X409),X408),s(fun(X409,bool),X407)))) | s(bool,'T_0') = s(bool,X410) | s(bool,'T_0') = s(bool,sK35(X407,X408,bool,X409))) )),
  inference(superposition,[],[f446,f890])).
fof(f3939,plain,(
  ( ! [X377,X379,X380,X378] : (~p(s(bool,happ(s(fun(X379,bool),X377),s(X379,happ(s(fun(bool,X379),X378),s(bool,X380)))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(bool,X379),X378),s(fun(X379,bool),X377)))) | s(bool,'T_0') = s(bool,X380) | s(bool,'T_0') = s(bool,sK26(X377,X378,bool,X379))) )),
  inference(superposition,[],[f432,f890])).
fof(f4207,plain,(
  ( ! [X401,X403,X400,X402] : (p(s(bool,happ(s(fun(X402,bool),X400),s(X402,happ(s(fun(bool,X402),X401),s(bool,'T_0')))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(bool,X402),X401),s(fun(X402,bool),X400)))) | s(bool,X403) = s(bool,sK35(X400,X401,bool,X402)) | s(bool,'T_0') = s(bool,X403)) )),
  inference(superposition,[],[f446,f890])).
fof(f4200,plain,(
  ( ! [X372,X370,X373,X371] : (~p(s(bool,happ(s(fun(X372,bool),X370),s(X372,happ(s(fun(bool,X372),X371),s(bool,'T_0')))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(bool,X372),X371),s(fun(X372,bool),X370)))) | s(bool,X373) = s(bool,sK26(X370,X371,bool,X372)) | s(bool,'T_0') = s(bool,X373)) )),
  inference(superposition,[],[f432,f890])).
fof(f369,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : ! [X1] : s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X5] : s(bool,'|dollar|exists_1'(s(fun(X7,bool),X5))) = s(bool,happ(s(fun(X7,bool),X5),s(X7,'const.min.@_1'(s(fun(X7,bool),X5)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.bool.EXISTS_DEF')).
fof(f4191,plain,(
  ( ! [X337,X339,X336,X338] : (~p(s(bool,happ(s(fun(bool,bool),X338),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X336,bool),X337),s(fun(bool,bool),X338)))) | s(bool,X339) = s(bool,happ(s(fun(X336,bool),X337),s(X336,sK26(X338,X337,X336,bool)))) | s(bool,'T_0') = s(bool,X339)) )),
  inference(superposition,[],[f432,f890])).
fof(f3930,plain,(
  ( ! [X344,X343,X345,X342] : (~p(s(bool,happ(s(fun(bool,bool),X344),s(bool,X345)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X342,bool),X343),s(fun(bool,bool),X344)))) | s(bool,'T_0') = s(bool,X345) | s(bool,'T_0') = s(bool,happ(s(fun(X342,bool),X343),s(X342,sK26(X344,X343,X342,bool))))) )),
  inference(superposition,[],[f432,f890])).
fof(f2057,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X3,X0),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X1,X2,X3,X0))))) )),
  inference(subsumption_resolution,[],[f2054,f582])).
fof(f2054,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X3,X0),X2),s(fun(X0,bool),X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X1,X2,X3,X0))))) )),
  inference(superposition,[],[f434,f334])).
fof(f434,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK28(X3,X2,X1,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f263])).
fof(f2056,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK28(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X5,X0),X4),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f2052,f428])).
fof(f2052,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK28(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X5,X0),X4),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f434,f428])).
fof(f4660,plain,(
  ( ! [X218,X216,X215,X217,X214] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X217,X214),X216),s(fun(X214,bool),X215)))) | s(bool,X218) = s(bool,happ(s(fun(X214,bool),X215),s(X214,sK32(X215,X216,X217,X214)))) | s(bool,'T_0') = s(bool,X218)) )),
  inference(subsumption_resolution,[],[f4163,f327])).
fof(f4163,plain,(
  ( ! [X218,X216,X215,X217,X214] : (~p(s(bool,'T_0')) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X217,X214),X216),s(fun(X214,bool),X215)))) | s(bool,X218) = s(bool,happ(s(fun(X214,bool),X215),s(X214,sK32(X215,X216,X217,X214)))) | s(bool,'T_0') = s(bool,X218)) )),
  inference(superposition,[],[f440,f890])).
fof(f440,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK32(X3,X2,X1,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f269])).
fof(f2077,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK32(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X5,X0),X4),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f2073,f428])).
fof(f2073,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK32(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X5,X0),X4),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f440,f428])).
fof(f4661,plain,(
  ( ! [X222,X220,X219,X223,X221] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X222,X219),X221),s(fun(X219,bool),X220)))) | s(bool,X223) = s(bool,happ(s(fun(X219,bool),X220),s(X219,sK34(X220,X221,X222,X219)))) | s(bool,'T_0') = s(bool,X223)) )),
  inference(subsumption_resolution,[],[f4164,f327])).
fof(f4164,plain,(
  ( ! [X222,X220,X219,X223,X221] : (~p(s(bool,'T_0')) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X222,X219),X221),s(fun(X219,bool),X220)))) | s(bool,X223) = s(bool,happ(s(fun(X219,bool),X220),s(X219,sK34(X220,X221,X222,X219)))) | s(bool,'T_0') = s(bool,X223)) )),
  inference(superposition,[],[f443,f890])).
fof(f443,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X3),s(X0,sK34(X3,X2,X1,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f272])).
fof(f2090,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK34(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X5,X0),X4),s(fun(X0,bool),X2))))) )),
  inference(forward_demodulation,[],[f2086,f428])).
fof(f2086,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK34(happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X4,X5,X0))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X5,X0),X4),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3))))))) )),
  inference(superposition,[],[f443,f428])).
fof(f4192,plain,(
  ( ! [X341,X343,X340,X342] : (p(s(bool,happ(s(fun(bool,bool),X342),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X340,bool),X341),s(fun(bool,bool),X342)))) | s(bool,X343) = s(bool,happ(s(fun(X340,bool),X341),s(X340,sK35(X342,X341,X340,bool)))) | s(bool,'T_0') = s(bool,X343)) )),
  inference(superposition,[],[f446,f890])).
fof(f3931,plain,(
  ( ! [X346,X349,X347,X348] : (p(s(bool,happ(s(fun(bool,bool),X348),s(bool,X349)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X346,bool),X347),s(fun(bool,bool),X348)))) | s(bool,'T_0') = s(bool,X349) | s(bool,'T_0') = s(bool,happ(s(fun(X346,bool),X347),s(X346,sK35(X348,X347,X346,bool))))) )),
  inference(superposition,[],[f446,f890])).
fof(f3932,plain,(
  ( ! [X350,X351] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X350))) = s(bool,happ(s(fun(bool,bool),X350),s(bool,X351))) | s(bool,'T_0') = s(bool,X351) | s(bool,'T_0') = s(bool,'const.min.@_1'(s(fun(bool,bool),X350)))) )),
  inference(superposition,[],[f369,f890])).
fof(f3934,plain,(
  ( ! [X358,X360,X357,X359] : (~p(s(bool,happ(s(fun(bool,bool),X357),s(bool,X359)))) | p(s(bool,happ(s(fun(bool,bool),X357),s(bool,X360)))) | ~p(s(bool,X358)) | s(bool,'T_0') = s(bool,X359) | s(bool,'T_0') = s(bool,sK21(X357,X358,bool))) )),
  inference(superposition,[],[f384,f890])).
fof(f3935,plain,(
  ( ! [X364,X362,X361,X363] : (~p(s(bool,happ(s(fun(bool,bool),X361),s(bool,X363)))) | p(s(bool,happ(s(fun(bool,bool),X361),s(bool,X364)))) | ~p(s(bool,X362)) | s(bool,'T_0') = s(bool,X363) | s(bool,'T_0') = s(bool,sK22(X361,X362,bool))) )),
  inference(superposition,[],[f386,f890])).
fof(f3936,plain,(
  ( ! [X366,X368,X365,X367] : (~p(s(bool,happ(s(fun(bool,bool),X365),s(bool,X367)))) | p(s(bool,happ(s(fun(bool,bool),X365),s(bool,X368)))) | p(s(bool,X366)) | s(bool,'T_0') = s(bool,X367) | s(bool,'T_0') = s(bool,sK23(X365,X366,bool))) )),
  inference(superposition,[],[f387,f890])).
fof(f3937,plain,(
  ( ! [X372,X370,X369,X371] : (~p(s(bool,happ(s(fun(bool,bool),X369),s(bool,X371)))) | p(s(bool,happ(s(fun(bool,bool),X369),s(bool,X372)))) | p(s(bool,X370)) | s(bool,'T_0') = s(bool,X371) | s(bool,'T_0') = s(bool,sK24(X369,X370,bool))) )),
  inference(superposition,[],[f389,f890])).
fof(f3941,plain,(
  ( ! [X389,X387,X388,X386] : (p(s(bool,happ(s(fun(bool,bool),X386),s(bool,X389)))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X388,bool),X387),s(fun(bool,bool),X386)))) | s(bool,'T_0') = s(bool,X389) | s(bool,'T_0') = s(bool,sK28(X386,X387,X388,bool))) )),
  inference(superposition,[],[f434,f890])).
fof(f3944,plain,(
  ( ! [X399,X401,X400,X402] : (~p(s(bool,happ(s(fun(bool,bool),X399),s(bool,X402)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X401,bool),X400),s(fun(bool,bool),X399)))) | s(bool,'T_0') = s(bool,X402) | s(bool,'T_0') = s(bool,sK32(X399,X400,X401,bool))) )),
  inference(superposition,[],[f440,f890])).
fof(f3945,plain,(
  ( ! [X405,X403,X404,X406] : (~p(s(bool,happ(s(fun(bool,bool),X403),s(bool,X406)))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X405,bool),X404),s(fun(bool,bool),X403)))) | s(bool,'T_0') = s(bool,X406) | s(bool,'T_0') = s(bool,sK34(X403,X404,X405,bool))) )),
  inference(superposition,[],[f443,f890])).
fof(f4195,plain,(
  ( ! [X350,X352,X351,X353] : (~p(s(bool,happ(s(fun(bool,bool),X350),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X350),s(bool,X352)))) | ~p(s(bool,X351)) | s(bool,X353) = s(bool,sK21(X350,X351,bool)) | s(bool,'T_0') = s(bool,X353)) )),
  inference(superposition,[],[f384,f890])).
fof(f4196,plain,(
  ( ! [X356,X354,X357,X355] : (~p(s(bool,happ(s(fun(bool,bool),X354),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X354),s(bool,X356)))) | ~p(s(bool,X355)) | s(bool,X357) = s(bool,sK22(X354,X355,bool)) | s(bool,'T_0') = s(bool,X357)) )),
  inference(superposition,[],[f386,f890])).
fof(f4197,plain,(
  ( ! [X358,X360,X359,X361] : (~p(s(bool,happ(s(fun(bool,bool),X358),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X358),s(bool,X360)))) | p(s(bool,X359)) | s(bool,X361) = s(bool,sK23(X358,X359,bool)) | s(bool,'T_0') = s(bool,X361)) )),
  inference(superposition,[],[f387,f890])).
fof(f4198,plain,(
  ( ! [X364,X362,X365,X363] : (~p(s(bool,happ(s(fun(bool,bool),X362),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X362),s(bool,X364)))) | p(s(bool,X363)) | s(bool,X365) = s(bool,sK24(X362,X363,bool)) | s(bool,'T_0') = s(bool,X365)) )),
  inference(superposition,[],[f389,f890])).
fof(f2333,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X0))))) ) | $spl2),
  inference(superposition,[],[f444,f643])).
fof(f4193,plain,(
  ( ! [X344,X345] : (s(bool,'|dollar|exists_1'(s(fun(bool,bool),X344))) = s(bool,happ(s(fun(bool,bool),X344),s(bool,'T_0'))) | s(bool,X345) = s(bool,'const.min.@_1'(s(fun(bool,bool),X344))) | s(bool,'T_0') = s(bool,X345)) )),
  inference(superposition,[],[f369,f890])).
fof(f4202,plain,(
  ( ! [X379,X380,X382,X381] : (p(s(bool,happ(s(fun(bool,bool),X379),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X381,bool),X380),s(fun(bool,bool),X379)))) | s(bool,X382) = s(bool,sK28(X379,X380,X381,bool)) | s(bool,'T_0') = s(bool,X382)) )),
  inference(superposition,[],[f434,f890])).
fof(f4205,plain,(
  ( ! [X393,X395,X392,X394] : (~p(s(bool,happ(s(fun(bool,bool),X392),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X394,bool),X393),s(fun(bool,bool),X392)))) | s(bool,X395) = s(bool,sK32(X392,X393,X394,bool)) | s(bool,'T_0') = s(bool,X395)) )),
  inference(superposition,[],[f440,f890])).
fof(f4206,plain,(
  ( ! [X397,X399,X396,X398] : (~p(s(bool,happ(s(fun(bool,bool),X396),s(bool,'T_0')))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X398,bool),X397),s(fun(bool,bool),X396)))) | s(bool,X399) = s(bool,sK34(X396,X397,X398,bool)) | s(bool,'T_0') = s(bool,X399)) )),
  inference(superposition,[],[f443,f890])).
fof(f2334,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2)))))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,bool),X2),s(fun(bool,bool),X3))))) )),
  inference(superposition,[],[f444,f369])).
fof(f1681,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,happ(s(fun(X19,bool),X20),s(X19,'const.min.@_1'(s(fun(X19,bool),X20))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X19,bool),X20))))) )),
  inference(superposition,[],[f400,f369])).
fof(f1682,plain,(
  ( ! [X24,X23,X22] : (sP0(happ(s(fun(X22,bool),X23),s(X22,'const.min.@_1'(s(fun(X22,bool),X23)))),X24) | ~p(s(bool,'|dollar|exists_1'(s(fun(X22,bool),X23))))) )),
  inference(superposition,[],[f401,f369])).
fof(f1683,plain,(
  ( ! [X26,X27,X25] : (sP1(happ(s(fun(X25,bool),X26),s(X25,'const.min.@_1'(s(fun(X25,bool),X26)))),X27) | p(s(bool,X27)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X25,bool),X26))))) )),
  inference(superposition,[],[f416,f369])).
fof(f1684,plain,(
  ( ! [X30,X28,X29] : (sP1(X30,happ(s(fun(X28,bool),X29),s(X28,'const.min.@_1'(s(fun(X28,bool),X29))))) | p(s(bool,X30)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X28,bool),X29))))) )),
  inference(superposition,[],[f419,f369])).
fof(f1686,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'|dollar|exists_1'(s(fun(X35,bool),X36)))) | sP6(X37,X38,happ(s(fun(X35,bool),X36),s(X35,'const.min.@_1'(s(fun(X35,bool),X36))))) | ~sP5(X37,happ(s(fun(X35,bool),X36),s(X35,'const.min.@_1'(s(fun(X35,bool),X36)))),X38) | ~sP4(X38,happ(s(fun(X35,bool),X36),s(X35,'const.min.@_1'(s(fun(X35,bool),X36)))),X37) | ~sP3(happ(s(fun(X35,bool),X36),s(X35,'const.min.@_1'(s(fun(X35,bool),X36)))),X38,X37)) )),
  inference(superposition,[],[f473,f369])).
fof(f1687,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'|dollar|exists_1'(s(fun(X39,bool),X40)))) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP5(X41,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))),X42)) )),
  inference(superposition,[],[f476,f369])).
fof(f1688,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'|dollar|exists_1'(s(fun(X43,bool),X44)))) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,happ(s(fun(X43,bool),X44),s(X43,'const.min.@_1'(s(fun(X43,bool),X44)))),X46)) )),
  inference(superposition,[],[f480,f369])).
fof(f1689,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'|dollar|exists_1'(s(fun(X47,bool),X48)))) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X49,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48)))),X50)) )),
  inference(superposition,[],[f499,f369])).
fof(f1690,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'|dollar|exists_1'(s(fun(X51,bool),X52)))) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP8(X54,X53,happ(s(fun(X51,bool),X52),s(X51,'const.min.@_1'(s(fun(X51,bool),X52)))))) )),
  inference(superposition,[],[f508,f369])).
fof(f1691,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'|dollar|exists_1'(s(fun(X55,bool),X56)))) | p(s(bool,X57)) | p(s(bool,X58)) | ~sP9(happ(s(fun(X55,bool),X56),s(X55,'const.min.@_1'(s(fun(X55,bool),X56)))),X57,X58)) )),
  inference(superposition,[],[f516,f369])).
fof(f1692,plain,(
  ( ! [X61,X59,X62,X60] : (sP10(happ(s(fun(X59,bool),X60),s(X59,'const.min.@_1'(s(fun(X59,bool),X60)))),X61,X62) | ~p(s(bool,X61)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X59,bool),X60))))) )),
  inference(superposition,[],[f522,f369])).
fof(f1693,plain,(
  ( ! [X66,X64,X65,X63] : (~sP10(happ(s(fun(X63,bool),X64),s(X63,'const.min.@_1'(s(fun(X63,bool),X64)))),X66,X65) | p(s(bool,X65)) | p(s(bool,X66)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X63,bool),X64))))) )),
  inference(superposition,[],[f523,f369])).
fof(f1694,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'|dollar|exists_1'(s(fun(X67,bool),X68)))) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~sP11(X70,X69,happ(s(fun(X67,bool),X68),s(X67,'const.min.@_1'(s(fun(X67,bool),X68)))))) )),
  inference(superposition,[],[f533,f369])).
fof(f1695,plain,(
  ( ! [X74,X72,X71,X73] : (sP12(happ(s(fun(X71,bool),X72),s(X71,'const.min.@_1'(s(fun(X71,bool),X72)))),X74,X73) | p(s(bool,X73)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X71,bool),X72))))) )),
  inference(superposition,[],[f538,f369])).
fof(f1696,plain,(
  ( ! [X78,X76,X77,X75] : (sP12(happ(s(fun(X75,bool),X76),s(X75,'const.min.@_1'(s(fun(X75,bool),X76)))),X77,X78) | ~p(s(bool,X77)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X75,bool),X76))))) )),
  inference(superposition,[],[f539,f369])).
fof(f1697,plain,(
  ( ! [X80,X81,X79,X82] : (~sP12(happ(s(fun(X79,bool),X80),s(X79,'const.min.@_1'(s(fun(X79,bool),X80)))),X82,X81) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~p(s(bool,'|dollar|exists_1'(s(fun(X79,bool),X80))))) )),
  inference(superposition,[],[f540,f369])).
fof(f1699,plain,(
  ( ! [X88,X87,X89,X86] : (sP12(happ(s(fun(X86,bool),X87),s(X86,'const.min.@_1'(s(fun(X86,bool),X87)))),X88,X89) | ~sP11(happ(s(fun(X86,bool),X87),s(X86,'const.min.@_1'(s(fun(X86,bool),X87)))),X88,X89) | ~p(s(bool,'|dollar|exists_1'(s(fun(X86,bool),X87))))) )),
  inference(superposition,[],[f563,f369])).
fof(f1700,plain,(
  ( ! [X92,X90,X93,X91] : (sP10(happ(s(fun(X90,bool),X91),s(X90,'const.min.@_1'(s(fun(X90,bool),X91)))),X92,X93) | ~sP9(happ(s(fun(X90,bool),X91),s(X90,'const.min.@_1'(s(fun(X90,bool),X91)))),X92,X93) | ~p(s(bool,'|dollar|exists_1'(s(fun(X90,bool),X91))))) )),
  inference(superposition,[],[f564,f369])).
fof(f2232,plain,(
  ( ! [X2,X3,X1] : (s(bool,sK28(X3,X2,X1,bool)) != s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,bool),X2),s(fun(bool,bool),X3))))) )),
  inference(superposition,[],[f435,f369])).
fof(f2258,plain,(
  ( ! [X2,X3,X1] : (s(bool,sK32(X3,X2,X1,bool)) != s(bool,'|dollar|exists_1'(s(fun(X1,bool),X2))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,bool),X2),s(fun(bool,bool),X3))))) )),
  inference(superposition,[],[f441,f369])).
fof(f2590,plain,(
  ( ! [X14,X12,X15,X13] : (~sP6('F_0',happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X13))))),s(X14,X15)),happ(s(fun(X12,bool),X13),s(X12,'const.min.@_1'(s(fun(X12,bool),X13)))))) )),
  inference(superposition,[],[f1445,f369])).
fof(f2626,plain,(
  ( ! [X14,X12,X15,X13] : (sP6('T_0',happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X13))))),s(X14,X15)),happ(s(fun(X12,bool),X13),s(X12,'const.min.@_1'(s(fun(X12,bool),X13)))))) )),
  inference(superposition,[],[f1446,f369])).
fof(f2659,plain,(
  ( ! [X14,X12,X15,X13] : (~sP6('F_0',happ(s(fun(X12,bool),X13),s(X12,'const.min.@_1'(s(fun(X12,bool),X13)))),happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X13))))),s(X14,X15)))) )),
  inference(superposition,[],[f1448,f369])).
fof(f2690,plain,(
  ( ! [X14,X12,X15,X13] : (sP6('T_0',happ(s(fun(X12,bool),X13),s(X12,'const.min.@_1'(s(fun(X12,bool),X13)))),happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(X12,bool),X13))))),s(X14,X15)))) )),
  inference(superposition,[],[f1449,f369])).
fof(f3791,plain,(
  ( ! [X24,X23,X25] : (s(bool,X25) = s(bool,'|dollar|exists_1'(s(fun(X23,bool),X24))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,happ(s(fun(X23,bool),X24),s(X23,'const.min.@_1'(s(fun(X23,bool),X24)))))) )),
  inference(superposition,[],[f890,f369])).
fof(f3821,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X15,bool),X16))) | s(bool,X17) = s(bool,happ(s(fun(X15,bool),X16),s(X15,'const.min.@_1'(s(fun(X15,bool),X16))))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f890,f369])).
fof(f2442,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X5)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_POINT_2'(s(fun(X6,X4),X7),s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))))))) )),
  inference(superposition,[],[f432,f428])).
fof(f2458,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X6,X4),X7),s(fun(X4,bool),'_dst_x0_1'(s(bool,X5)))))) | p(s(bool,X5))) )),
  inference(superposition,[],[f446,f428])).
fof(f435,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,sK28(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f263])).
fof(f2235,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | s(X1,happ(s(fun(X0,X1),X3),s(X0,X6))) != s(X1,sK28(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))) )),
  inference(forward_demodulation,[],[f2223,f428])).
fof(f2223,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),X3),s(X0,X6))) != s(X1,sK28(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(fun(X1,bool),X5))))) )),
  inference(superposition,[],[f435,f428])).
fof(f2055,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X7,X4),X6),s(fun(X4,bool),'_dst_x0_1'(s(bool,X5)))))) | p(s(bool,X5))) )),
  inference(superposition,[],[f434,f428])).
fof(f2234,plain,(
  ( ! [X12,X10,X8,X9] : (s(X8,X10) != s(X8,sK28(X12,'_dst_x0_1'(s(X8,X10)),X9,X8)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X9,X8),'_dst_x0_1'(s(X8,X10))),s(fun(X8,bool),X12))))) )),
  inference(superposition,[],[f435,f428])).
fof(f2231,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK28(X0,sK16,sK14,bool)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X0))))) ) | $spl2),
  inference(superposition,[],[f435,f643])).
fof(f441,plain,(
  ( ! [X2,X0,X5,X3,X1] : (s(X0,sK32(X3,X2,X1,X0)) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f269])).
fof(f2261,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5)))) | s(X1,happ(s(fun(X0,X1),X3),s(X0,X6))) != s(X1,sK32(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))) )),
  inference(forward_demodulation,[],[f2249,f428])).
fof(f2249,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,happ(s(fun(X0,X1),X3),s(X0,X6))) != s(X1,sK32(X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X0,X1),happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4))),s(fun(X1,bool),X5))))) )),
  inference(superposition,[],[f441,f428])).
fof(f2076,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X5)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X7,X4),X6),s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))))))) )),
  inference(superposition,[],[f440,f428])).
fof(f2260,plain,(
  ( ! [X12,X10,X8,X9] : (s(X8,X10) != s(X8,sK32(X12,'_dst_x0_1'(s(X8,X10)),X9,X8)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X9,X8),'_dst_x0_1'(s(X8,X10))),s(fun(X8,bool),X12))))) )),
  inference(superposition,[],[f441,f428])).
fof(f2257,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,sK32(X0,sK16,sK14,bool)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X0))))) ) | $spl2),
  inference(superposition,[],[f441,f643])).
fof(f4199,plain,(
  ( ! [X366,X368,X367,X369] : (s(X368,f(s(bool,'T_0'))) != s(X368,g(s(bool,'T_0'))) | s(fun(bool,X368),X366) = s(fun(bool,X368),X367) | s(bool,X369) = s(bool,sK25(X366,X367,X368,bool)) | s(bool,'T_0') = s(bool,X369)) )),
  inference(superposition,[],[f430,f890])).
fof(f4647,plain,(
  ( ! [X121,X118,X122,X120,X119] : (sP0(happ(s(fun(X119,bool),'_dst_x0_1'(s(bool,X118))),s(X119,X120)),X121) | s(bool,X118) = s(bool,X122) | s(bool,'T_0') = s(bool,X122)) )),
  inference(subsumption_resolution,[],[f4137,f327])).
fof(f4137,plain,(
  ( ! [X121,X118,X122,X120,X119] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X119,bool),'_dst_x0_1'(s(bool,X118))),s(X119,X120)),X121) | s(bool,X118) = s(bool,X122) | s(bool,'T_0') = s(bool,X122)) )),
  inference(superposition,[],[f1472,f890])).
fof(f4646,plain,(
  ( ! [X116,X114,X117,X115,X113] : (sP0(X114,happ(s(fun(X115,bool),'_dst_x0_1'(s(bool,X113))),s(X115,X116))) | s(bool,X113) = s(bool,X117) | s(bool,'T_0') = s(bool,X117)) )),
  inference(subsumption_resolution,[],[f4136,f327])).
fof(f4136,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,'T_0')) | sP0(X114,happ(s(fun(X115,bool),'_dst_x0_1'(s(bool,X113))),s(X115,X116))) | s(bool,X113) = s(bool,X117) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f1471,f890])).
fof(f4135,plain,(
  ( ! [X111,X109,X112,X110] : (sP6('T_0',X109,happ(s(fun(X110,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X110,X111))) | s(bool,X109) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f1449,f890])).
fof(f4134,plain,(
  ( ! [X107,X105,X108,X106] : (~sP6('F_0',X105,happ(s(fun(X106,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X106,X107))) | s(bool,X105) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f1448,f890])).
fof(f4133,plain,(
  ( ! [X103,X101,X102,X104] : (sP6('T_0',happ(s(fun(X102,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X102,X103)),X101) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f1446,f890])).
fof(f4132,plain,(
  ( ! [X99,X97,X100,X98] : (~sP6('F_0',happ(s(fun(X98,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X98,X99)),X97) | s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f1445,f890])).
fof(f4645,plain,(
  ( ! [X92,X90,X93,X91] : (~sP9(X90,X91,X92) | sP10(X90,X91,X92) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f4130,f327])).
fof(f4130,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,'T_0')) | ~sP9(X90,X91,X92) | sP10(X90,X91,X92) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f564,f890])).
fof(f4644,plain,(
  ( ! [X88,X87,X89,X86] : (~sP11(X86,X87,X88) | sP12(X86,X87,X88) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f4129,f327])).
fof(f4129,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,'T_0')) | ~sP11(X86,X87,X88) | sP12(X86,X87,X88) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f563,f890])).
fof(f4643,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP12(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f4127,f327])).
fof(f4127,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP12(X79,X81,X80) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f540,f890])).
fof(f4642,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | sP12(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f4126,f327])).
fof(f4126,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | sP12(X75,X76,X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f539,f890])).
fof(f4641,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X72)) | sP12(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f4125,f327])).
fof(f4125,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X72)) | sP12(X71,X73,X72) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f538,f890])).
fof(f4640,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP11(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f4124,f327])).
fof(f4124,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | ~p(s(bool,X69)) | ~sP11(X69,X68,X67) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f533,f890])).
fof(f4639,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | p(s(bool,X65)) | ~sP10(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f4123,f327])).
fof(f4123,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP10(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f523,f890])).
fof(f4638,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP10(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f4122,f327])).
fof(f4122,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP10(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f522,f890])).
fof(f4637,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X56)) | p(s(bool,X57)) | ~sP9(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f4121,f327])).
fof(f4121,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X56)) | p(s(bool,X57)) | ~sP9(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f516,f890])).
fof(f4636,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X53,X52,X51) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f4120,f327])).
fof(f4120,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X53,X52,X51) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f508,f890])).
fof(f4635,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X48,X47,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f4119,f327])).
fof(f4119,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X48,X47,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f499,f890])).
fof(f4634,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP4(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f4118,f327])).
fof(f4118,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP4(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f480,f890])).
fof(f4633,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP5(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f4117,f327])).
fof(f4117,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP5(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f476,f890])).
fof(f4632,plain,(
  ( ! [X37,X35,X38,X36] : (sP6(X36,X37,X35) | ~sP5(X36,X35,X37) | ~sP4(X37,X35,X36) | ~sP3(X35,X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f4116,f327])).
fof(f4116,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | sP6(X36,X37,X35) | ~sP5(X36,X35,X37) | ~sP4(X37,X35,X36) | ~sP3(X35,X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f473,f890])).
fof(f4631,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f4114,f327])).
fof(f4114,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f419,f890])).
fof(f4630,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f4113,f327])).
fof(f4113,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f416,f890])).
fof(f4629,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f4112,f327])).
fof(f4112,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f401,f890])).
fof(f4628,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f4111,f327])).
fof(f4111,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f400,f890])).
fof(f4627,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f4110,f327])).
fof(f4110,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f394,f890])).
fof(f4626,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4106,f327])).
fof(f4106,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f345,f890])).
fof(f4087,plain,(
  ( ! [X118,X116,X117,X115] : (sP6('T_0','T_0',happ(s(fun(X116,bool),'_dst_x0_1'(s(bool,X115))),s(X116,X117))) | s(bool,X115) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) )),
  inference(superposition,[],[f1449,f890])).
fof(f4086,plain,(
  ( ! [X111,X114,X112,X113] : (~sP6('F_0','T_0',happ(s(fun(X112,bool),'_dst_x0_1'(s(bool,X111))),s(X112,X113))) | s(bool,X111) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f1448,f890])).
fof(f4085,plain,(
  ( ! [X109,X107,X110,X108] : (sP6('T_0',happ(s(fun(X108,bool),'_dst_x0_1'(s(bool,X107))),s(X108,X109)),'T_0') | s(bool,X107) = s(bool,X110) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f1446,f890])).
fof(f4084,plain,(
  ( ! [X103,X105,X106,X104] : (~sP6('F_0',happ(s(fun(X104,bool),'_dst_x0_1'(s(bool,X103))),s(X104,X105)),'T_0') | s(bool,X103) = s(bool,X106) | s(bool,'T_0') = s(bool,X106)) )),
  inference(superposition,[],[f1445,f890])).
fof(f4055,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f341,f890])).
fof(f3938,plain,(
  ( ! [X374,X376,X373,X375] : (s(X375,f(s(bool,X376))) != s(X375,g(s(bool,X376))) | s(fun(bool,X375),X373) = s(fun(bool,X375),X374) | s(bool,'T_0') = s(bool,X376) | s(bool,'T_0') = s(bool,sK25(X373,X374,X375,bool))) )),
  inference(superposition,[],[f430,f890])).
fof(f3872,plain,(
  ( ! [X111,X109,X112,X110] : (sP6('T_0',X109,happ(s(fun(X111,bool),'_dst_x0_1'(s(bool,X110))),s(X111,X112))) | s(bool,'T_0') = s(bool,X110) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f1449,f890])).
fof(f3871,plain,(
  ( ! [X107,X105,X108,X106] : (~sP6('F_0',X105,happ(s(fun(X107,bool),'_dst_x0_1'(s(bool,X106))),s(X107,X108))) | s(bool,'T_0') = s(bool,X106) | s(bool,'T_0') = s(bool,X105)) )),
  inference(superposition,[],[f1448,f890])).
fof(f3870,plain,(
  ( ! [X103,X101,X102,X104] : (sP6('T_0',happ(s(fun(X103,bool),'_dst_x0_1'(s(bool,X102))),s(X103,X104)),X101) | s(bool,'T_0') = s(bool,X102) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f1446,f890])).
fof(f3869,plain,(
  ( ! [X99,X97,X100,X98] : (~sP6('F_0',happ(s(fun(X99,bool),'_dst_x0_1'(s(bool,X98))),s(X99,X100)),X97) | s(bool,'T_0') = s(bool,X98) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f1445,f890])).
fof(f3824,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,X23) | s(bool,X25) = s(bool,happ(s(fun(X22,bool),'_dst_x0_1'(s(bool,X23))),s(X22,X24))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f890,f428])).
fof(f3822,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,'|dollar|exists_1'(s(fun(X18,bool),'_dst_x0_1'(s(bool,X19))))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f890,f1671])).
fof(f3810,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f890,f890])).
fof(f3794,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'T_0') = s(bool,happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,X31))),s(X30,X32)))) )),
  inference(superposition,[],[f890,f428])).
fof(f3792,plain,(
  ( ! [X28,X26,X27] : (s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(X26,bool),'_dst_x0_1'(s(bool,X27)))))) )),
  inference(superposition,[],[f890,f1671])).
fof(f3739,plain,(
  ( ! [X2,X0,X1] : (sP0(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f327,f1472])).
fof(f3717,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f327,f1471])).
fof(f3710,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f1506,f528])).
fof(f3205,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f327,f2342,f394])).
fof(f3206,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3))))) = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2342,f2342,f394])).
fof(f3207,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f2342,f400])).
fof(f3208,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2342,f401])).
fof(f3209,plain,(
  ( ! [X2,X0,X3,X1] : (~sP1('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f2342,f415])).
fof(f3211,plain,(
  ( ! [X2,X0,X3,X1] : (~sP1('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f327,f2342,f415])).
fof(f3212,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP1('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2342,f2342,f415])).
fof(f3213,plain,(
  ( ! [X2,X0,X3,X1] : (sP1('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f2342,f416])).
fof(f3214,plain,(
  ( ! [X2,X0,X3,X1] : (sP1('F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f582,f2342,f419])).
fof(f3229,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f2342,f458])).
fof(f3234,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2('T_0',X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f327,f2342,f458])).
fof(f3235,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP2('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X5,X6),X7),s(fun(X6,bool),'_dst_x0_1'(s(bool,X8)))))) )),
  inference(unit_resulting_resolution,[],[f2342,f2342,f458])).
fof(f3241,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2342,f472])).
fof(f3242,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2342,f582,f2342,f472])).
fof(f3245,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('T_0','T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f600,f594,f588,f2342,f473])).
fof(f3247,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f601,f595,f588,f2342,f473])).
fof(f3249,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f630,f629,f588,f2342,f473])).
fof(f3251,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f600,f594,f589,f2342,f473])).
fof(f3253,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f601,f595,f589,f2342,f473])).
fof(f3255,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f630,f629,f589,f2342,f473])).
fof(f3257,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f600,f594,f628,f2342,f473])).
fof(f3259,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f601,f595,f628,f2342,f473])).
fof(f3261,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f630,f629,f628,f2342,f473])).
fof(f3262,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2342,f476])).
fof(f3270,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2342,f476])).
fof(f3271,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP5('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2342,f582,f2342,f476])).
fof(f3278,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(X0,X1,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2342,f477])).
fof(f3279,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2342,f480])).
fof(f3287,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2342,f480])).
fof(f3288,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP4('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2342,f582,f2342,f480])).
fof(f3295,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP4(X0,X1,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2342,f481])).
fof(f3296,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP3(X0,X1,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2342,f485])).
fof(f3297,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP3(X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2342,f486])).
fof(f3298,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP3('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,X5)) )),
  inference(unit_resulting_resolution,[],[f2342,f487])).
fof(f3299,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'|dollar|exists_1'(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5)))),X5)) )),
  inference(unit_resulting_resolution,[],[f1671,f2342,f489])).
fof(f3300,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f643,f2342,f489])).
fof(f3301,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f644,f2342,f489])).
fof(f3303,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(X4,bool),X5),s(X4,'const.min.@_1'(s(fun(X4,bool),X5)))),'|dollar|exists_1'(s(fun(X4,bool),X5)))) )),
  inference(unit_resulting_resolution,[],[f369,f2342,f489])).
fof(f3304,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))),s(X4,X6)),X5)) )),
  inference(unit_resulting_resolution,[],[f428,f2342,f489])).
fof(f3305,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'|dollar|exists_1'(s(fun(X5,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f1671,f2342,f489])).
fof(f3306,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f643,f2342,f489])).
fof(f3307,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f644,f2342,f489])).
fof(f3309,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'|dollar|exists_1'(s(fun(X4,bool),X5)),happ(s(fun(X4,bool),X5),s(X4,'const.min.@_1'(s(fun(X4,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f369,f2342,f489])).
fof(f3310,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,happ(s(fun(X5,bool),'_dst_x0_1'(s(bool,X4))),s(X5,X6)))) )),
  inference(unit_resulting_resolution,[],[f428,f2342,f489])).
fof(f3312,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f639,f2342,f490])).
fof(f3314,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f639,f2342,f490])).
fof(f3315,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP8('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f2342,f493])).
fof(f3316,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP8('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0',X4)) )),
  inference(unit_resulting_resolution,[],[f582,f2342,f494])).
fof(f3319,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f327,f607,f2342,f495])).
fof(f3321,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f633,f2342,f495])).
fof(f3327,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f327,f607,f2342,f495])).
fof(f3328,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2342,f607,f2342,f495])).
fof(f3329,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f327,f633,f2342,f495])).
fof(f3330,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2342,f633,f2342,f495])).
fof(f3333,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2342,f499])).
fof(f3341,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2342,f499])).
fof(f3342,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP7('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2342,f582,f2342,f499])).
fof(f3349,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7(X0,X1,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2342,f500])).
fof(f3350,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f327,f2342,f505])).
fof(f3356,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('T_0','T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f327,f327,f2342,f505])).
fof(f3358,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f327,f2342,f2342,f505])).
fof(f3362,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f327,f2342,f505])).
fof(f3363,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f2342,f327,f2342,f505])).
fof(f3364,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f327,f2342,f2342,f505])).
fof(f3365,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (sP8('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X8,X9),X10),s(fun(X9,bool),'_dst_x0_1'(s(bool,X11)))))) )),
  inference(unit_resulting_resolution,[],[f2342,f2342,f2342,f505])).
fof(f3366,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8('F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2342,f508])).
fof(f3371,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8('F_0','T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2342,f508])).
fof(f3372,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP8('F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2342,f582,f2342,f508])).
fof(f3376,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP10('F_0',X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f582,f2342,f509])).
fof(f3382,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP10('F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f582,f2342,f510])).
fof(f3392,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f582,f2342,f516])).
fof(f3399,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9(X0,X1,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2342,f517])).
fof(f3400,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9(X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2342,f518])).
fof(f3401,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('T_0',X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f327,f2342,f521])).
fof(f3406,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X5,X6),X7),s(fun(X6,bool),'_dst_x0_1'(s(bool,X8)))))) )),
  inference(unit_resulting_resolution,[],[f2342,f2342,f521])).
fof(f3407,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f327,f2342,f522])).
fof(f3412,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),X8)) )),
  inference(unit_resulting_resolution,[],[f2342,f2342,f522])).
fof(f3413,plain,(
  ( ! [X2,X0,X3,X1] : (~sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f582,f2342,f523])).
fof(f3414,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP12('F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f582,f2342,f527])).
fof(f3444,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f582,f327,f2342,f533])).
fof(f3452,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('T_0','F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f582,f327,f2342,f533])).
fof(f3453,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP11('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f582,f2342,f2342,f533])).
fof(f3460,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP11(X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2342,f535])).
fof(f3461,plain,(
  ( ! [X2,X0,X3,X1] : (sP12('F_0','F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f582,f582,f2342,f537])).
fof(f3464,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('T_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f327,f2342,f539])).
fof(f3470,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP12('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),X8)) )),
  inference(unit_resulting_resolution,[],[f2342,f2342,f539])).
fof(f3471,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12('T_0','F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2342,f540])).
fof(f3476,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2342,f540])).
fof(f3477,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP12('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2342,f582,f2342,f540])).
fof(f3478,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f582,f2342,f562])).
fof(f3481,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'F_0')) )),
  inference(unit_resulting_resolution,[],[f618,f2342,f563])).
fof(f3484,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0',X4)) )),
  inference(unit_resulting_resolution,[],[f620,f2342,f563])).
fof(f3485,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4)) )),
  inference(unit_resulting_resolution,[],[f636,f2342,f563])).
fof(f3486,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X4)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f621,f2342,f563])).
fof(f3487,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'T_0')) )),
  inference(unit_resulting_resolution,[],[f612,f2342,f564])).
fof(f3488,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f613,f2342,f564])).
fof(f3489,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f634,f2342,f564])).
fof(f3490,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0',X4)) )),
  inference(unit_resulting_resolution,[],[f614,f2342,f564])).
fof(f3491,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4)) )),
  inference(unit_resulting_resolution,[],[f635,f2342,f564])).
fof(f3492,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X4)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f615,f2342,f564])).
fof(f3493,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP2(X0,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2342,f566])).
fof(f3494,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2342,f789])).
fof(f3495,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2342,f790])).
fof(f3496,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2342,f821])).
fof(f3497,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f2342,f1089])).
fof(f3501,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f2342,f1117])).
fof(f3505,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f2342,f1177])).
fof(f3509,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8('F_0','const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f2342,f1208])).
fof(f3510,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f2342,f1384])).
fof(f2787,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f327,f2309,f394])).
fof(f2788,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3))))) = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2309,f2309,f394])).
fof(f2789,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0(X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f2309,f400])).
fof(f2790,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP0('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2309,f401])).
fof(f2791,plain,(
  ( ! [X2,X0,X3,X1] : (~sP1('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f2309,f415])).
fof(f2793,plain,(
  ( ! [X2,X0,X3,X1] : (~sP1('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f327,f2309,f415])).
fof(f2794,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP1('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2309,f2309,f415])).
fof(f2795,plain,(
  ( ! [X2,X0,X3,X1] : (sP1('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f2309,f416])).
fof(f2796,plain,(
  ( ! [X2,X0,X3,X1] : (sP1('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f582,f2309,f419])).
fof(f2811,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f2309,f458])).
fof(f2816,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP2('T_0',X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f327,f2309,f458])).
fof(f2817,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP2('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X5,X6),X7),s(fun(X6,bool),'_dst_x0_1'(s(bool,X8)))))) )),
  inference(unit_resulting_resolution,[],[f2309,f2309,f458])).
fof(f2823,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2309,f472])).
fof(f2824,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2309,f582,f2309,f472])).
fof(f2827,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f600,f594,f588,f2309,f473])).
fof(f2829,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f601,f595,f588,f2309,f473])).
fof(f2831,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f630,f629,f588,f2309,f473])).
fof(f2833,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f600,f594,f589,f2309,f473])).
fof(f2835,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f601,f595,f589,f2309,f473])).
fof(f2837,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f630,f629,f589,f2309,f473])).
fof(f2839,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f600,f594,f628,f2309,f473])).
fof(f2841,plain,(
  ( ! [X2,X0,X3,X1] : (sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f601,f595,f628,f2309,f473])).
fof(f2843,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f630,f629,f628,f2309,f473])).
fof(f2844,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2309,f476])).
fof(f2852,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2309,f476])).
fof(f2853,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP5('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2309,f582,f2309,f476])).
fof(f2860,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5(X0,X1,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2309,f477])).
fof(f2861,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2309,f480])).
fof(f2869,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2309,f480])).
fof(f2870,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP4('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2309,f582,f2309,f480])).
fof(f2877,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP4(X0,X1,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2309,f481])).
fof(f2878,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP3(X0,X1,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2309,f485])).
fof(f2879,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP3(X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2309,f486])).
fof(f2880,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP3('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,X5)) )),
  inference(unit_resulting_resolution,[],[f2309,f487])).
fof(f2881,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'|dollar|exists_1'(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5)))),X5)) )),
  inference(unit_resulting_resolution,[],[f1671,f2309,f489])).
fof(f2882,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f643,f2309,f489])).
fof(f2883,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f644,f2309,f489])).
fof(f2885,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(X4,bool),X5),s(X4,'const.min.@_1'(s(fun(X4,bool),X5)))),'|dollar|exists_1'(s(fun(X4,bool),X5)))) )),
  inference(unit_resulting_resolution,[],[f369,f2309,f489])).
fof(f2886,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))),s(X4,X6)),X5)) )),
  inference(unit_resulting_resolution,[],[f428,f2309,f489])).
fof(f2887,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'|dollar|exists_1'(s(fun(X5,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f1671,f2309,f489])).
fof(f2888,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f643,f2309,f489])).
fof(f2889,plain,(
  ( ! [X2,X0,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f644,f2309,f489])).
fof(f2891,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'|dollar|exists_1'(s(fun(X4,bool),X5)),happ(s(fun(X4,bool),X5),s(X4,'const.min.@_1'(s(fun(X4,bool),X5)))))) )),
  inference(unit_resulting_resolution,[],[f369,f2309,f489])).
fof(f2892,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,happ(s(fun(X5,bool),'_dst_x0_1'(s(bool,X4))),s(X5,X6)))) )),
  inference(unit_resulting_resolution,[],[f428,f2309,f489])).
fof(f2894,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f639,f2309,f490])).
fof(f2896,plain,(
  ( ! [X2,X0,X3,X1] : (~sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f639,f2309,f490])).
fof(f2897,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f2309,f493])).
fof(f2898,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0',X4)) )),
  inference(unit_resulting_resolution,[],[f582,f2309,f494])).
fof(f2901,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f327,f607,f2309,f495])).
fof(f2903,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f633,f2309,f495])).
fof(f2909,plain,(
  ( ! [X2,X0,X3,X1] : (sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f327,f607,f2309,f495])).
fof(f2910,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2309,f607,f2309,f495])).
fof(f2911,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f327,f633,f2309,f495])).
fof(f2912,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2309,f633,f2309,f495])).
fof(f2915,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2309,f499])).
fof(f2923,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2309,f499])).
fof(f2924,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP7('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f2309,f582,f2309,f499])).
fof(f2931,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7(X0,X1,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2309,f500])).
fof(f2932,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f327,f2309,f505])).
fof(f2938,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f327,f327,f2309,f505])).
fof(f2940,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f327,f2309,f2309,f505])).
fof(f2944,plain,(
  ( ! [X2,X0,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f327,f2309,f505])).
fof(f2945,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f2309,f327,f2309,f505])).
fof(f2946,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f327,f2309,f2309,f505])).
fof(f2947,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X11,X9] : (sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X8,X9),X10),s(fun(X9,bool),'_dst_x0_1'(s(bool,X11)))))) )),
  inference(unit_resulting_resolution,[],[f2309,f2309,f2309,f505])).
fof(f2948,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2309,f508])).
fof(f2953,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2309,f508])).
fof(f2954,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP8('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2309,f582,f2309,f508])).
fof(f2958,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP10('F_0',X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f582,f2309,f509])).
fof(f2964,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP10('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f582,f2309,f510])).
fof(f2974,plain,(
  ( ! [X2,X0,X3,X1] : (~sP9('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f582,f2309,f516])).
fof(f2981,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9(X0,X1,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X2,X3),X4),s(fun(X3,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(unit_resulting_resolution,[],[f2309,f517])).
fof(f2982,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP9(X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2309,f518])).
fof(f2983,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('T_0',X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))))) )),
  inference(unit_resulting_resolution,[],[f327,f2309,f521])).
fof(f2988,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X5,X6),X7),s(fun(X6,bool),'_dst_x0_1'(s(bool,X8)))))) )),
  inference(unit_resulting_resolution,[],[f2309,f2309,f521])).
fof(f2989,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f327,f2309,f522])).
fof(f2994,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),X8)) )),
  inference(unit_resulting_resolution,[],[f2309,f2309,f522])).
fof(f2995,plain,(
  ( ! [X2,X0,X3,X1] : (~sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f582,f2309,f523])).
fof(f2996,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP12('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f582,f2309,f527])).
fof(f3026,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f582,f327,f2309,f533])).
fof(f3034,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f582,f327,f2309,f533])).
fof(f3035,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP11('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f582,f2309,f2309,f533])).
fof(f3042,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP11(X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2309,f535])).
fof(f3043,plain,(
  ( ! [X2,X0,X3,X1] : (sP12('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f582,f582,f2309,f537])).
fof(f3046,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f327,f2309,f539])).
fof(f3052,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))),X8)) )),
  inference(unit_resulting_resolution,[],[f2309,f2309,f539])).
fof(f3053,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12('T_0','F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2309,f540])).
fof(f3058,plain,(
  ( ! [X2,X0,X3,X1] : (~sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f2309,f540])).
fof(f3059,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X4,X5),X6),s(fun(X5,bool),'_dst_x0_1'(s(bool,X7)))))) )),
  inference(unit_resulting_resolution,[],[f2309,f582,f2309,f540])).
fof(f3060,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))))) )),
  inference(unit_resulting_resolution,[],[f582,f2309,f562])).
fof(f3063,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'F_0')) )),
  inference(unit_resulting_resolution,[],[f618,f2309,f563])).
fof(f3066,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0',X4)) )),
  inference(unit_resulting_resolution,[],[f620,f2309,f563])).
fof(f3067,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4)) )),
  inference(unit_resulting_resolution,[],[f636,f2309,f563])).
fof(f3068,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X4)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f621,f2309,f563])).
fof(f3069,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'T_0')) )),
  inference(unit_resulting_resolution,[],[f612,f2309,f564])).
fof(f3070,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f613,f2309,f564])).
fof(f3071,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f634,f2309,f564])).
fof(f3072,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'T_0',X4)) )),
  inference(unit_resulting_resolution,[],[f614,f2309,f564])).
fof(f3073,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4)) )),
  inference(unit_resulting_resolution,[],[f635,f2309,f564])).
fof(f3074,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X4)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f615,f2309,f564])).
fof(f3075,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP2(X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,X2),X3),s(fun(X2,bool),'_dst_x0_1'(s(bool,X4)))),X5)) )),
  inference(unit_resulting_resolution,[],[f2309,f566])).
fof(f3076,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2309,f789])).
fof(f3077,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) )),
  inference(unit_resulting_resolution,[],[f2309,f790])).
fof(f3078,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),X4)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f2309,f821])).
fof(f3079,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f2309,f1089])).
fof(f3083,plain,(
  ( ! [X2,X0,X3,X1] : (~sP4('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f2309,f1117])).
fof(f3087,plain,(
  ( ! [X2,X0,X3,X1] : (~sP7('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f2309,f1177])).
fof(f3091,plain,(
  ( ! [X2,X0,X3,X1] : (~sP8('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f2309,f1208])).
fof(f3092,plain,(
  ( ! [X2,X0,X3,X1] : (~sP11('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'_dst_x0_1'(s(bool,X3)))),'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f2309,f1384])).
fof(f2755,plain,(
  ( ! [X109,X107,X108,X106] : (sP6('T_0','|dollar|exists_1'(s(fun(X106,bool),'_dst_x0_1'(s(bool,X107)))),happ(s(fun(X108,bool),'_dst_x0_1'(s(bool,X107))),s(X108,X109)))) )),
  inference(superposition,[],[f1449,f1671])).
fof(f2754,plain,(
  ( ! [X103,X105,X102,X104] : (~sP6('F_0','|dollar|exists_1'(s(fun(X102,bool),'_dst_x0_1'(s(bool,X103)))),happ(s(fun(X104,bool),'_dst_x0_1'(s(bool,X103))),s(X104,X105)))) )),
  inference(superposition,[],[f1448,f1671])).
fof(f2753,plain,(
  ( ! [X101,X99,X100,X98] : (sP6('T_0',happ(s(fun(X100,bool),'_dst_x0_1'(s(bool,X99))),s(X100,X101)),'|dollar|exists_1'(s(fun(X98,bool),'_dst_x0_1'(s(bool,X99)))))) )),
  inference(superposition,[],[f1446,f1671])).
fof(f2752,plain,(
  ( ! [X94,X97,X95,X96] : (~sP6('F_0',happ(s(fun(X96,bool),'_dst_x0_1'(s(bool,X95))),s(X96,X97)),'|dollar|exists_1'(s(fun(X94,bool),'_dst_x0_1'(s(bool,X95)))))) )),
  inference(superposition,[],[f1445,f1671])).
fof(f2751,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,X91)) | ~sP9('|dollar|exists_1'(s(fun(X90,bool),'_dst_x0_1'(s(bool,X91)))),X92,X93) | sP10('|dollar|exists_1'(s(fun(X90,bool),'_dst_x0_1'(s(bool,X91)))),X92,X93)) )),
  inference(superposition,[],[f564,f1671])).
fof(f2750,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,X87)) | ~sP11('|dollar|exists_1'(s(fun(X86,bool),'_dst_x0_1'(s(bool,X87)))),X88,X89) | sP12('|dollar|exists_1'(s(fun(X86,bool),'_dst_x0_1'(s(bool,X87)))),X88,X89)) )),
  inference(superposition,[],[f563,f1671])).
fof(f2748,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~sP12('|dollar|exists_1'(s(fun(X79,bool),'_dst_x0_1'(s(bool,X80)))),X82,X81)) )),
  inference(superposition,[],[f540,f1671])).
fof(f2747,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | ~p(s(bool,X77)) | sP12('|dollar|exists_1'(s(fun(X75,bool),'_dst_x0_1'(s(bool,X76)))),X77,X78)) )),
  inference(superposition,[],[f539,f1671])).
fof(f2746,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | sP12('|dollar|exists_1'(s(fun(X71,bool),'_dst_x0_1'(s(bool,X72)))),X74,X73)) )),
  inference(superposition,[],[f538,f1671])).
fof(f2745,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,X68)) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~sP11(X70,X69,'|dollar|exists_1'(s(fun(X67,bool),'_dst_x0_1'(s(bool,X68)))))) )),
  inference(superposition,[],[f533,f1671])).
fof(f2744,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,X64)) | p(s(bool,X65)) | p(s(bool,X66)) | ~sP10('|dollar|exists_1'(s(fun(X63,bool),'_dst_x0_1'(s(bool,X64)))),X66,X65)) )),
  inference(superposition,[],[f523,f1671])).
fof(f2743,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | ~p(s(bool,X61)) | sP10('|dollar|exists_1'(s(fun(X59,bool),'_dst_x0_1'(s(bool,X60)))),X61,X62)) )),
  inference(superposition,[],[f522,f1671])).
fof(f2742,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | p(s(bool,X57)) | p(s(bool,X58)) | ~sP9('|dollar|exists_1'(s(fun(X55,bool),'_dst_x0_1'(s(bool,X56)))),X57,X58)) )),
  inference(superposition,[],[f516,f1671])).
fof(f2741,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP8(X54,X53,'|dollar|exists_1'(s(fun(X51,bool),'_dst_x0_1'(s(bool,X52)))))) )),
  inference(superposition,[],[f508,f1671])).
fof(f2740,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X49,'|dollar|exists_1'(s(fun(X47,bool),'_dst_x0_1'(s(bool,X48)))),X50)) )),
  inference(superposition,[],[f499,f1671])).
fof(f2739,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,'|dollar|exists_1'(s(fun(X43,bool),'_dst_x0_1'(s(bool,X44)))),X46)) )),
  inference(superposition,[],[f480,f1671])).
fof(f2738,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP5(X41,'|dollar|exists_1'(s(fun(X39,bool),'_dst_x0_1'(s(bool,X40)))),X42)) )),
  inference(superposition,[],[f476,f1671])).
fof(f2737,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | sP6(X37,X38,'|dollar|exists_1'(s(fun(X35,bool),'_dst_x0_1'(s(bool,X36))))) | ~sP5(X37,'|dollar|exists_1'(s(fun(X35,bool),'_dst_x0_1'(s(bool,X36)))),X38) | ~sP4(X38,'|dollar|exists_1'(s(fun(X35,bool),'_dst_x0_1'(s(bool,X36)))),X37) | ~sP3('|dollar|exists_1'(s(fun(X35,bool),'_dst_x0_1'(s(bool,X36)))),X38,X37)) )),
  inference(superposition,[],[f473,f1671])).
fof(f2735,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1(X30,'|dollar|exists_1'(s(fun(X28,bool),'_dst_x0_1'(s(bool,X29)))))) )),
  inference(superposition,[],[f419,f1671])).
fof(f2734,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,X26)) | p(s(bool,X27)) | sP1('|dollar|exists_1'(s(fun(X25,bool),'_dst_x0_1'(s(bool,X26)))),X27)) )),
  inference(superposition,[],[f416,f1671])).
fof(f2733,plain,(
  ( ! [X24,X23,X22] : (sP0('|dollar|exists_1'(s(fun(X22,bool),'_dst_x0_1'(s(bool,X23)))),X24) | ~p(s(bool,X23))) )),
  inference(superposition,[],[f401,f1671])).
fof(f2732,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'|dollar|exists_1'(s(fun(X19,bool),'_dst_x0_1'(s(bool,X20))))) | ~p(s(bool,X20))) )),
  inference(superposition,[],[f400,f1671])).
fof(f2710,plain,(
  ( ! [X0,X1] : (~sP6('F_0','|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))),X1)) )),
  inference(unit_resulting_resolution,[],[f582,f1671,f491])).
fof(f2711,plain,(
  ( ! [X0,X1] : (sP6('T_0','|dollar|exists_1'(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1)))),X1)) )),
  inference(unit_resulting_resolution,[],[f327,f1671,f489])).
fof(f2713,plain,(
  ( ! [X0,X1] : (~sP6('F_0',X0,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f582,f1671,f491])).
fof(f2714,plain,(
  ( ! [X0,X1] : (sP6('T_0',X0,'|dollar|exists_1'(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f327,f1671,f489])).
fof(f2689,plain,(
  ( ! [X10,X11] : (sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)))) ) | $spl2),
  inference(superposition,[],[f1449,f643])).
fof(f2687,plain,(
  ( ! [X6,X7] : (sP6('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X6,X7)))) )),
  inference(superposition,[],[f1449,f644])).
fof(f2686,plain,(
  ( ! [X4,X5,X3] : (sP6('T_0','F_0',happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X3))),s(X4,X5))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1449,f334])).
fof(f2682,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f1449,f471])).
fof(f2658,plain,(
  ( ! [X10,X11] : (~sP6('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)))) ) | $spl2),
  inference(superposition,[],[f1448,f643])).
fof(f2656,plain,(
  ( ! [X6,X7] : (~sP6('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X6,X7)))) )),
  inference(superposition,[],[f1448,f644])).
fof(f2655,plain,(
  ( ! [X4,X5,X3] : (~sP6('F_0','F_0',happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X3))),s(X4,X5))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1448,f334])).
fof(f2648,plain,(
  ( ! [X0,X1] : (~sP4('T_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f602,f588,f1448,f475])).
fof(f2669,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) ) | $spl2),
  inference(forward_demodulation,[],[f2649,f643])).
fof(f2649,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))))),s(X0,X1)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f603,f589,f1448,f475])).
fof(f2668,plain,(
  ( ! [X0,X1] : (~sP4('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) )),
  inference(forward_demodulation,[],[f2650,f644])).
fof(f2650,plain,(
  ( ! [X0,X1] : (~sP4('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))))),s(X0,X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f631,f628,f1448,f475])).
fof(f2627,plain,(
  ( ! [X19,X17,X20,X18,X16] : (sP6('T_0',happ(s(fun(X19,bool),'_dst_x0_1'(s(bool,X17))),s(X19,X20)),happ(s(fun(X16,bool),'_dst_x0_1'(s(bool,X17))),s(X16,X18)))) )),
  inference(superposition,[],[f1446,f428])).
fof(f2625,plain,(
  ( ! [X10,X11] : (sP6('T_0',happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(superposition,[],[f1446,f643])).
fof(f2623,plain,(
  ( ! [X6,X7] : (sP6('T_0',happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X6,X7)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(superposition,[],[f1446,f644])).
fof(f2622,plain,(
  ( ! [X4,X5,X3] : (sP6('T_0',happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X3))),s(X4,X5)),'F_0') | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1446,f334])).
fof(f2618,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f1446,f471])).
fof(f2591,plain,(
  ( ! [X19,X17,X20,X18,X16] : (~sP6('F_0',happ(s(fun(X19,bool),'_dst_x0_1'(s(bool,X17))),s(X19,X20)),happ(s(fun(X16,bool),'_dst_x0_1'(s(bool,X17))),s(X16,X18)))) )),
  inference(superposition,[],[f1445,f428])).
fof(f2589,plain,(
  ( ! [X10,X11] : (~sP6('F_0',happ(s(fun(X10,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X10,X11)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(superposition,[],[f1445,f643])).
fof(f2587,plain,(
  ( ! [X6,X7] : (~sP6('F_0',happ(s(fun(X6,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X6,X7)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(superposition,[],[f1445,f644])).
fof(f2586,plain,(
  ( ! [X4,X5,X3] : (~sP6('F_0',happ(s(fun(X4,bool),'_dst_x0_1'(s(bool,X3))),s(X4,X5)),'F_0') | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1445,f334])).
fof(f2604,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) )),
  inference(forward_demodulation,[],[f2574,f644])).
fof(f2574,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))))),s(X0,X1)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f632,f592,f1445,f475])).
fof(f2603,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')) ) | $spl2),
  inference(forward_demodulation,[],[f2575,f643])).
fof(f2575,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))))),s(X0,X1)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f605,f592,f1445,f475])).
fof(f2582,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f604,f592,f1445,f473])).
fof(f2342,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X7,X5),X8),s(fun(X5,bool),'_dst_x0_1'(s(bool,X6))))))) )),
  inference(subsumption_resolution,[],[f2338,f2089])).
fof(f2338,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,X6)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X7,X5),X8),s(fun(X5,bool),'_dst_x0_1'(s(bool,X6))))))) )),
  inference(superposition,[],[f444,f428])).
fof(f2309,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X7,X5),X8),s(fun(X5,bool),'_dst_x0_1'(s(bool,X6))))))) )),
  inference(subsumption_resolution,[],[f2306,f2065])).
fof(f2306,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,X6)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X7,X5),X8),s(fun(X5,bool),'_dst_x0_1'(s(bool,X6))))))) )),
  inference(superposition,[],[f438,f428])).
fof(f2089,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,X5)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X7,X4),X6),s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))))))) )),
  inference(superposition,[],[f443,f428])).
fof(f2065,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X5)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X7,X4),X6),s(fun(X4,bool),'_dst_x0_1'(s(bool,X5))))))) )),
  inference(superposition,[],[f437,f428])).
fof(f1671,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,'|dollar|exists_1'(s(fun(X4,bool),'_dst_x0_1'(s(bool,X5)))))) )),
  inference(superposition,[],[f369,f428])).
fof(f1664,plain,(
  ( ! [X0,X1] : (~sP6('F_0',happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f582,f369,f491])).
fof(f1665,plain,(
  ( ! [X0,X1] : (sP6('T_0',happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))),'|dollar|exists_1'(s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f327,f369,f489])).
fof(f1666,plain,(
  ( ! [X0,X1] : (~sP6('F_0','|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f582,f369,f491])).
fof(f1667,plain,(
  ( ! [X0,X1] : (sP6('T_0','|dollar|exists_1'(s(fun(X0,bool),X1)),happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) )),
  inference(unit_resulting_resolution,[],[f327,f369,f489])).
fof(f1628,plain,(
  ~sP4('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')),
  inference(unit_resulting_resolution,[],[f632,f1030,f588,f582,f475])).
fof(f1636,plain,(
  sP6('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')),
  inference(unit_resulting_resolution,[],[f631,f596,f590,f582,f475])).
fof(f1637,plain,(
  sP6('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f603,f596,f590,f582,f475])).
fof(f1568,plain,(
  sP6('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f600,f594,f590,f582,f474])).
fof(f1569,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f601,f595,f590,f582,f474])).
fof(f1570,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f630,f629,f590,f582,f474])).
fof(f1583,plain,(
  sP6('F_0','F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f632,f598,f592,f582,f474])).
fof(f1584,plain,(
  sP6('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f605,f598,f592,f582,f474])).
fof(f1540,plain,(
  ( ! [X14,X12,X15,X13,X16] : (sP6(X15,X16,happ(s(fun(X12,bool),'_dst_x0_1'(s(bool,X13))),s(X12,X14))) | ~p(s(bool,X13)) | ~sP5(X15,happ(s(fun(X12,bool),'_dst_x0_1'(s(bool,X13))),s(X12,X14)),X16) | ~sP4(X16,happ(s(fun(X12,bool),'_dst_x0_1'(s(bool,X13))),s(X12,X14)),X15) | ~sP3(happ(s(fun(X12,bool),'_dst_x0_1'(s(bool,X13))),s(X12,X14)),X16,X15)) )),
  inference(superposition,[],[f473,f428])).
fof(f1545,plain,(
  ( ! [X10,X11] : (sP6(X10,X11,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ~sP5(X10,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X11) | ~sP4(X11,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X10)) ) | $spl2),
  inference(subsumption_resolution,[],[f1544,f605])).
fof(f1544,plain,(
  ( ! [X10,X11] : (sP6(X10,X11,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ~sP5(X10,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X11) | ~sP4(X11,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X10) | ~sP3(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X11,X10)) ) | $spl2),
  inference(subsumption_resolution,[],[f1539,f327])).
fof(f1539,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | sP6(X10,X11,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ~sP5(X10,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X11) | ~sP4(X11,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X10) | ~sP3(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X11,X10)) ) | $spl2),
  inference(superposition,[],[f473,f643])).
fof(f1543,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | ~sP5(X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7) | ~sP4(X7,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X6)) )),
  inference(subsumption_resolution,[],[f1542,f632])).
fof(f1542,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | ~sP5(X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7) | ~sP4(X7,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X6) | ~sP3('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7,X6)) )),
  inference(subsumption_resolution,[],[f1537,f327])).
fof(f1537,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | sP6(X6,X7,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | ~sP5(X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7) | ~sP4(X7,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X6) | ~sP3('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7,X6)) )),
  inference(superposition,[],[f473,f644])).
fof(f1517,plain,(
  sP6('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f604,f594,f588,f327,f473])).
fof(f1518,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f604,f595,f588,f327,f473])).
fof(f1519,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f604,f629,f588,f327,f473])).
fof(f1521,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f604,f595,f589,f327,f473])).
fof(f1522,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f604,f629,f589,f327,f473])).
fof(f1524,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f604,f595,f628,f327,f473])).
fof(f1525,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0')),
  inference(unit_resulting_resolution,[],[f604,f629,f628,f327,f473])).
fof(f1533,plain,(
  ~sP4('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f604,f592,f1027,f327,f473])).
fof(f1534,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f604,f592,f1026,f327,f473])).
fof(f1506,plain,(
  ~sP11('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f327,f1384])).
fof(f1489,plain,(
  ( ! [X116,X114,X112,X115,X113] : (~p(s(bool,X113)) | ~sP9(happ(s(fun(X112,bool),'_dst_x0_1'(s(bool,X113))),s(X112,X114)),X115,X116) | sP10(happ(s(fun(X112,bool),'_dst_x0_1'(s(bool,X113))),s(X112,X114)),X115,X116)) )),
  inference(superposition,[],[f564,f428])).
fof(f1488,plain,(
  ( ! [X111,X109,X107,X110,X108] : (~p(s(bool,X108)) | ~sP11(happ(s(fun(X107,bool),'_dst_x0_1'(s(bool,X108))),s(X107,X109)),X110,X111) | sP12(happ(s(fun(X107,bool),'_dst_x0_1'(s(bool,X108))),s(X107,X109)),X110,X111)) )),
  inference(superposition,[],[f563,f428])).
fof(f1486,plain,(
  ( ! [X101,X99,X102,X100,X98] : (~p(s(bool,X101)) | ~p(s(bool,X99)) | p(s(bool,X102)) | ~sP12(happ(s(fun(X98,bool),'_dst_x0_1'(s(bool,X99))),s(X98,X100)),X102,X101)) )),
  inference(superposition,[],[f540,f428])).
fof(f1485,plain,(
  ( ! [X94,X97,X95,X93,X96] : (~p(s(bool,X96)) | ~p(s(bool,X94)) | sP12(happ(s(fun(X93,bool),'_dst_x0_1'(s(bool,X94))),s(X93,X95)),X96,X97)) )),
  inference(superposition,[],[f539,f428])).
fof(f1484,plain,(
  ( ! [X92,X90,X88,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X91)) | sP12(happ(s(fun(X88,bool),'_dst_x0_1'(s(bool,X89))),s(X88,X90)),X92,X91)) )),
  inference(superposition,[],[f538,f428])).
fof(f1483,plain,(
  ( ! [X87,X85,X83,X86,X84] : (~p(s(bool,X87)) | p(s(bool,X86)) | ~p(s(bool,X84)) | ~sP11(X87,X86,happ(s(fun(X83,bool),'_dst_x0_1'(s(bool,X84))),s(X83,X85)))) )),
  inference(superposition,[],[f533,f428])).
fof(f1482,plain,(
  ( ! [X80,X78,X81,X79,X82] : (~p(s(bool,X79)) | p(s(bool,X81)) | p(s(bool,X82)) | ~sP10(happ(s(fun(X78,bool),'_dst_x0_1'(s(bool,X79))),s(X78,X80)),X82,X81)) )),
  inference(superposition,[],[f523,f428])).
fof(f1481,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~p(s(bool,X76)) | ~p(s(bool,X74)) | sP10(happ(s(fun(X73,bool),'_dst_x0_1'(s(bool,X74))),s(X73,X75)),X76,X77)) )),
  inference(superposition,[],[f522,f428])).
fof(f1480,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~p(s(bool,X69)) | p(s(bool,X71)) | p(s(bool,X72)) | ~sP9(happ(s(fun(X68,bool),'_dst_x0_1'(s(bool,X69))),s(X68,X70)),X71,X72)) )),
  inference(superposition,[],[f516,f428])).
fof(f1479,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,X66)) | ~p(s(bool,X64)) | p(s(bool,X67)) | ~sP8(X67,X66,happ(s(fun(X63,bool),'_dst_x0_1'(s(bool,X64))),s(X63,X65)))) )),
  inference(superposition,[],[f508,f428])).
fof(f1478,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,X61)) | ~p(s(bool,X59)) | p(s(bool,X62)) | ~sP7(X61,happ(s(fun(X58,bool),'_dst_x0_1'(s(bool,X59))),s(X58,X60)),X62)) )),
  inference(superposition,[],[f499,f428])).
fof(f1477,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,X56)) | ~p(s(bool,X54)) | p(s(bool,X57)) | ~sP4(X56,happ(s(fun(X53,bool),'_dst_x0_1'(s(bool,X54))),s(X53,X55)),X57)) )),
  inference(superposition,[],[f480,f428])).
fof(f1476,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,X51)) | ~p(s(bool,X49)) | p(s(bool,X52)) | ~sP5(X51,happ(s(fun(X48,bool),'_dst_x0_1'(s(bool,X49))),s(X48,X50)),X52)) )),
  inference(superposition,[],[f476,f428])).
fof(f1474,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X42)) | sP1(X42,happ(s(fun(X39,bool),'_dst_x0_1'(s(bool,X40))),s(X39,X41)))) )),
  inference(superposition,[],[f419,f428])).
fof(f1473,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X38)) | sP1(happ(s(fun(X35,bool),'_dst_x0_1'(s(bool,X36))),s(X35,X37)),X38)) )),
  inference(superposition,[],[f416,f428])).
fof(f1472,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | sP0(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,X32))),s(X31,X33)),X34)) )),
  inference(superposition,[],[f401,f428])).
fof(f1471,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X28)) | sP0(X30,happ(s(fun(X27,bool),'_dst_x0_1'(s(bool,X28))),s(X27,X29)))) )),
  inference(superposition,[],[f400,f428])).
fof(f1445,plain,(
  ( ! [X2,X0,X1] : (~sP6('F_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f582,f428,f491])).
fof(f1446,plain,(
  ( ! [X2,X0,X1] : (sP6('T_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) )),
  inference(unit_resulting_resolution,[],[f327,f428,f489])).
fof(f1448,plain,(
  ( ! [X2,X0,X1] : (~sP6('F_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f582,f428,f491])).
fof(f1449,plain,(
  ( ! [X2,X0,X1] : (sP6('T_0',X0,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) )),
  inference(unit_resulting_resolution,[],[f327,f428,f489])).
fof(f1436,plain,(
  ~sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f1403,f511])).
fof(f1432,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | ~sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X11,X10)) ) | $spl2),
  inference(subsumption_resolution,[],[f1426,f327])).
fof(f1426,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | p(s(bool,X11)) | ~sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X11,X10)) ) | $spl2),
  inference(superposition,[],[f540,f643])).
fof(f1431,plain,(
  ( ! [X6,X7] : (~sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7,X6) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1424,f327])).
fof(f1424,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7,X6)) )),
  inference(superposition,[],[f540,f644])).
fof(f1403,plain,(
  ~sP9(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f582,f582,f1267])).
fof(f1388,plain,(
  ~sP12('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1373,f528])).
fof(f1384,plain,(
  ( ! [X10,X11] : (~sP11(X11,X10,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ~p(s(bool,X11)) | p(s(bool,X10))) ) | $spl2),
  inference(subsumption_resolution,[],[f1378,f327])).
fof(f1378,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X10)) | ~p(s(bool,X11)) | ~sP11(X11,X10,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(superposition,[],[f533,f643])).
fof(f1383,plain,(
  ( ! [X6,X7] : (~sP11(X7,X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | ~p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1376,f327])).
fof(f1376,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | ~p(s(bool,X7)) | ~sP11(X7,X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(superposition,[],[f533,f644])).
fof(f1373,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f327,f582,f327,f533])).
fof(f1360,plain,(
  sP12('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f622,f327,f582,f532])).
fof(f1352,plain,(
  ~sP8('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f1337,f492])).
fof(f1337,plain,(
  ~sP7('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f327,f582,f1177])).
fof(f1319,plain,(
  ~sP4('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f327,f582,f1117])).
fof(f1316,plain,(
  ( ! [X10,X11] : (~sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X11,X10) | p(s(bool,X11)) | p(s(bool,X10))) ) | $spl2),
  inference(subsumption_resolution,[],[f1309,f327])).
fof(f1309,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X10)) | p(s(bool,X11)) | ~sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X11,X10)) ) | $spl2),
  inference(superposition,[],[f523,f643])).
fof(f1315,plain,(
  ( ! [X6,X7] : (~sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7,X6) | p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1307,f327])).
fof(f1307,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | p(s(bool,X7)) | ~sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7,X6)) )),
  inference(superposition,[],[f523,f644])).
fof(f1294,plain,(
  ~sP6('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f1282,f471])).
fof(f1282,plain,(
  ~sP5('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f327,f582,f1089])).
fof(f1271,plain,(
  ~sP10('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1255,f511])).
fof(f1267,plain,(
  ( ! [X10,X11] : (~sP9(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X10,X11) | p(s(bool,X11)) | p(s(bool,X10))) ) | $spl2),
  inference(subsumption_resolution,[],[f1260,f327])).
fof(f1260,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X10)) | p(s(bool,X11)) | ~sP9(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X10,X11)) ) | $spl2),
  inference(superposition,[],[f516,f643])).
fof(f1266,plain,(
  ( ! [X6,X7] : (~sP9('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X6,X7) | p(s(bool,X7)) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1258,f327])).
fof(f1258,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | p(s(bool,X7)) | ~sP9('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X6,X7)) )),
  inference(superposition,[],[f516,f644])).
fof(f1255,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f582,f582,f327,f516])).
fof(f1241,plain,(
  sP10('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f616,f582,f582,f515])).
fof(f1208,plain,(
  ( ! [X10,X11] : (~sP8(X11,X10,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | p(s(bool,X11)) | ~p(s(bool,X10))) ) | $spl2),
  inference(subsumption_resolution,[],[f1204,f327])).
fof(f1204,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | p(s(bool,X11)) | ~sP8(X11,X10,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(superposition,[],[f508,f643])).
fof(f1207,plain,(
  ( ! [X6,X7] : (~sP8(X7,X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1202,f327])).
fof(f1202,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP8(X7,X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(superposition,[],[f508,f644])).
fof(f1181,plain,(
  ~sP8('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1166,f492])).
fof(f1177,plain,(
  ( ! [X10,X11] : (~sP7(X10,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X11) | p(s(bool,X11)) | ~p(s(bool,X10))) ) | $spl2),
  inference(subsumption_resolution,[],[f1171,f327])).
fof(f1171,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | p(s(bool,X11)) | ~sP7(X10,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X11)) ) | $spl2),
  inference(superposition,[],[f499,f643])).
fof(f1176,plain,(
  ( ! [X6,X7] : (~sP7(X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1169,f327])).
fof(f1169,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP7(X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7)) )),
  inference(superposition,[],[f499,f644])).
fof(f1166,plain,(
  ~sP7('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f582,f327,f327,f499])).
fof(f1146,plain,(
  sP8('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f606,f327,f327,f495])).
fof(f1147,plain,(
  sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f607,f327,f327,f495])).
fof(f1148,plain,(
  sP8('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f633,f327,f327,f495])).
fof(f1142,plain,(
  ~sP6('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1129,f469])).
fof(f1129,plain,(
  ~sP3('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f582,f582,f582,f484])).
fof(f1121,plain,(
  ~sP6('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1106,f470])).
fof(f1117,plain,(
  ( ! [X10,X11] : (~sP4(X10,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X11) | p(s(bool,X11)) | ~p(s(bool,X10))) ) | $spl2),
  inference(subsumption_resolution,[],[f1111,f327])).
fof(f1111,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | p(s(bool,X11)) | ~sP4(X10,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X11)) ) | $spl2),
  inference(superposition,[],[f480,f643])).
fof(f1116,plain,(
  ( ! [X6,X7] : (~sP4(X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1109,f327])).
fof(f1109,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP4(X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7)) )),
  inference(superposition,[],[f480,f644])).
fof(f1106,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f582,f327,f327,f480])).
fof(f1089,plain,(
  ( ! [X10,X11] : (~sP5(X10,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X11) | p(s(bool,X11)) | ~p(s(bool,X10))) ) | $spl2),
  inference(subsumption_resolution,[],[f1083,f327])).
fof(f1083,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | p(s(bool,X11)) | ~sP5(X10,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X11)) ) | $spl2),
  inference(superposition,[],[f476,f643])).
fof(f1088,plain,(
  ( ! [X6,X7] : (~sP5(X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7) | p(s(bool,X7)) | ~p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f1081,f327])).
fof(f1081,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | p(s(bool,X7)) | ~sP5(X6,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X7)) )),
  inference(superposition,[],[f476,f644])).
fof(f1078,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f582,f327,f327,f476])).
fof(f1026,plain,(
  ~sP6('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f582,f643,f491])).
fof(f1027,plain,(
  ~sP6('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0')),
  inference(unit_resulting_resolution,[],[f582,f644,f491])).
fof(f1029,plain,(
  ~sP6('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f643,f491])).
fof(f1030,plain,(
  ~sP6('F_0','T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f582,f644,f491])).
fof(f1013,plain,(
  ~sP6('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f327,f639,f490])).
fof(f1017,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f327,f639,f490])).
fof(f998,plain,(
  sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f327,f643,f489])).
fof(f999,plain,(
  sP6('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0')),
  inference(unit_resulting_resolution,[],[f327,f644,f489])).
fof(f1001,plain,(
  sP6('T_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f327,f643,f489])).
fof(f1002,plain,(
  sP6('T_0','T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f327,f644,f489])).
fof(f989,plain,(
  sP6('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f582,f639,f488])).
fof(f991,plain,(
  sP6('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f582,f639,f488])).
fof(f954,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f944,f582])).
fof(f944,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f391,f334])).
fof(f934,plain,(
  ( ! [X0] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f327,f821])).
fof(f890,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f334,f334])).
fof(f872,plain,(
  ( ! [X4,X5] : (~sP9('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4,X5) | sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4,X5)) )),
  inference(subsumption_resolution,[],[f870,f327])).
fof(f870,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | ~sP9('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4,X5) | sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4,X5)) )),
  inference(superposition,[],[f564,f644])).
fof(f871,plain,(
  ( ! [X0,X1] : (~sP9(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,X1) | sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,X1)) ) | $spl2),
  inference(subsumption_resolution,[],[f868,f327])).
fof(f868,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP9(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,X1) | sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,X1)) ) | $spl2),
  inference(superposition,[],[f564,f643])).
fof(f863,plain,(
  ( ! [X0] : (sP10('T_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f613,f327,f564])).
fof(f864,plain,(
  ( ! [X0] : (sP10('T_0',X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f634,f327,f564])).
fof(f866,plain,(
  ( ! [X0] : (sP10('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X0)) )),
  inference(unit_resulting_resolution,[],[f635,f327,f564])).
fof(f867,plain,(
  ( ! [X0] : (sP10('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f615,f327,f564])).
fof(f848,plain,(
  ( ! [X4,X5] : (~sP11('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4,X5) | sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4,X5)) )),
  inference(subsumption_resolution,[],[f846,f327])).
fof(f846,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | ~sP11('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4,X5) | sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4,X5)) )),
  inference(superposition,[],[f563,f644])).
fof(f847,plain,(
  ( ! [X0,X1] : (~sP11(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,X1) | sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,X1)) ) | $spl2),
  inference(subsumption_resolution,[],[f844,f327])).
fof(f844,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP11(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,X1) | sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,X1)) ) | $spl2),
  inference(superposition,[],[f563,f643])).
fof(f842,plain,(
  ( ! [X0] : (sP12('T_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X0)) )),
  inference(unit_resulting_resolution,[],[f636,f327,f563])).
fof(f843,plain,(
  ( ! [X0] : (sP12('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f621,f327,f563])).
fof(f822,plain,(
  ( ! [X4,X5] : (sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4,X5) | ~p(s(bool,X4))) )),
  inference(subsumption_resolution,[],[f820,f327])).
fof(f820,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X4)) | sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4,X5)) )),
  inference(superposition,[],[f539,f644])).
fof(f821,plain,(
  ( ! [X0,X1] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,X1) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f818,f327])).
fof(f818,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,X1)) ) | $spl2),
  inference(superposition,[],[f539,f643])).
fof(f817,plain,(
  ( ! [X0] : (sP12('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f327,f327,f539])).
fof(f814,plain,(
  ( ! [X0] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f809])).
fof(f810,plain,(
  ( ! [X4,X5] : (sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X5,X4) | p(s(bool,X4))) )),
  inference(subsumption_resolution,[],[f808,f327])).
fof(f808,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | p(s(bool,X4)) | sP12('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X5,X4)) )),
  inference(superposition,[],[f538,f644])).
fof(f809,plain,(
  ( ! [X0,X1] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X1,X0) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f806,f327])).
fof(f806,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X1,X0)) ) | $spl2),
  inference(superposition,[],[f538,f643])).
fof(f805,plain,(
  ( ! [X0] : (sP12('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f327,f538])).
fof(f802,plain,(
  ( ! [X0] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f327,f790])).
fof(f801,plain,(
  ( ! [X0] : (~sP12('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f327,f582,f527])).
fof(f795,plain,(
  ( ! [X0] : (~sP12('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f582,f526])).
fof(f794,plain,(
  ( ! [X0] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f327,f789])).
fof(f790,plain,(
  ( ! [X4,X5] : (sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4,X5) | ~p(s(bool,X4))) )),
  inference(subsumption_resolution,[],[f788,f327])).
fof(f788,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X4)) | sP10('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X4,X5)) )),
  inference(superposition,[],[f522,f644])).
fof(f789,plain,(
  ( ! [X0,X1] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,X1) | ~p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f786,f327])).
fof(f786,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,X1)) ) | $spl2),
  inference(superposition,[],[f522,f643])).
fof(f785,plain,(
  ( ! [X0] : (sP10('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f327,f327,f522])).
fof(f778,plain,(
  ( ! [X0] : (sP10('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f327,f521])).
fof(f776,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f327,f582,f510])).
fof(f775,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f509])).
fof(f768,plain,(
  ( ! [X0] : (sP8('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f582,f582,f504])).
fof(f765,plain,(
  ( ! [X0] : (sP8('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f582,f503])).
fof(f762,plain,(
  ( ! [X0] : (~sP8('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f327,f582,f494])).
fof(f759,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f327,f582,f493])).
fof(f758,plain,(
  sP1('F_0','const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f582,f747])).
fof(f755,plain,(
  ( ! [X0] : (sP2('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f327,f458])).
fof(f750,plain,(
  sP1('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')),
  inference(unit_resulting_resolution,[],[f582,f746])).
fof(f747,plain,(
  ( ! [X3] : (sP1(X3,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | p(s(bool,X3))) )),
  inference(subsumption_resolution,[],[f739,f327])).
fof(f739,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,X3)) | sP1(X3,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(superposition,[],[f419,f644])).
fof(f746,plain,(
  ( ! [X2] : (sP1('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2) | p(s(bool,X2))) )),
  inference(subsumption_resolution,[],[f738,f327])).
fof(f738,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,X2)) | sP1('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X2)) )),
  inference(superposition,[],[f416,f644])).
fof(f720,plain,(
  ( ! [X0] : (~sP2('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f582,f582,f454])).
fof(f713,plain,(
  ( ! [X0] : (~sP2(X0,'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f582,f453])).
fof(f710,plain,(
  sP1('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f692])).
fof(f703,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f582,f582,f583])).
fof(f692,plain,(
  ( ! [X0] : (sP1(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f690,f327])).
fof(f690,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(superposition,[],[f419,f643])).
fof(f689,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f582,f327,f419])).
fof(f686,plain,(
  sP1(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f582,f683])).
fof(f683,plain,(
  ( ! [X0] : (sP1(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0) | p(s(bool,X0))) ) | $spl2),
  inference(subsumption_resolution,[],[f681,f327])).
fof(f681,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | $spl2),
  inference(superposition,[],[f416,f643])).
fof(f680,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f582,f327,f416])).
fof(f677,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f327,f327,f415])).
fof(f676,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f582,f582,f414])).
fof(f643,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f575,f342])).
fof(f644,plain,(
  s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))),
  inference(unit_resulting_resolution,[],[f322,f342])).
fof(f639,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f582,f341])).
fof(f626,plain,(
  ( ! [X0] : (sP0(X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f322,f400])).
fof(f627,plain,(
  ( ! [X0] : (sP0('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X0)) )),
  inference(unit_resulting_resolution,[],[f322,f401])).
fof(f628,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f322,f477])).
fof(f629,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f322,f481])).
fof(f630,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f322,f485])).
fof(f631,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X1)) )),
  inference(unit_resulting_resolution,[],[f322,f486])).
fof(f632,plain,(
  ( ! [X0,X1] : (sP3('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f322,f487])).
fof(f633,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f322,f500])).
fof(f634,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) )),
  inference(unit_resulting_resolution,[],[f322,f517])).
fof(f635,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X1)) )),
  inference(unit_resulting_resolution,[],[f322,f518])).
fof(f636,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X1)) )),
  inference(unit_resulting_resolution,[],[f322,f535])).
fof(f637,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),X1)) )),
  inference(unit_resulting_resolution,[],[f322,f566])).
fof(f624,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f327,f566])).
fof(f625,plain,(
  ( ! [X0,X1] : (sP2(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f575,f566])).
fof(f622,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f582,f536])).
fof(f620,plain,(
  ( ! [X0,X1] : (sP11(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f327,f535])).
fof(f621,plain,(
  ( ! [X0,X1] : (sP11(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f575,f535])).
fof(f618,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f582,f534])).
fof(f616,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f582,f519])).
fof(f614,plain,(
  ( ! [X0,X1] : (sP9(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f327,f518])).
fof(f615,plain,(
  ( ! [X0,X1] : (sP9(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f575,f518])).
fof(f612,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f517])).
fof(f613,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f575,f517])).
fof(f610,plain,(
  ( ! [X0,X1] : (sP7('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f582,f502])).
fof(f608,plain,(
  ( ! [X0,X1] : (sP7(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f582,f501])).
fof(f606,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f500])).
fof(f607,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f575,f500])).
fof(f604,plain,(
  ( ! [X0,X1] : (sP3('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f327,f487])).
fof(f605,plain,(
  ( ! [X0,X1] : (sP3(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f575,f487])).
fof(f602,plain,(
  ( ! [X0,X1] : (sP3(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f327,f486])).
fof(f603,plain,(
  ( ! [X0,X1] : (sP3(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f575,f486])).
fof(f600,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f485])).
fof(f601,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f575,f485])).
fof(f598,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f582,f483])).
fof(f596,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f582,f482])).
fof(f594,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f481])).
fof(f595,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f575,f481])).
fof(f592,plain,(
  ( ! [X0,X1] : (sP5('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f582,f479])).
fof(f590,plain,(
  ( ! [X0,X1] : (sP5(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f582,f478])).
fof(f588,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f477])).
fof(f589,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f575,f477])).
fof(f586,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f327,f401])).
fof(f587,plain,(
  ( ! [X0] : (sP0(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f575,f401])).
fof(f584,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f327,f400])).
fof(f585,plain,(
  ( ! [X0] : (sP0(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f575,f400])).
fof(f562,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f561])).
fof(f561,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f550])).
fof(f550,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X14,X19,X16,X20] : ((s(bool,X14) = s(bool,X19) & (p(s(bool,X19)) => s(bool,X16) = s(bool,X20))) => ((p(s(bool,X14)) => p(s(bool,X16))) <=> (p(s(bool,X19)) => p(s(bool,X20)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.bool.IMP_CONG')).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (sP12(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP12(X2,X0,X1)) & (~sP12(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP12(X2,X0,X1))),
  inference(definition_folding,[],[f182,f203,f202])).
fof(f202,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f203,plain,(
  ! [X2,X0,X1] : (sP12(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f182,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> (p(s(bool,X25)) => p(s(bool,X24)))) <=> ((p(s(bool,X26)) | p(s(bool,X25))) & (p(s(bool,X26)) | ~p(s(bool,X24))) & (~p(s(bool,X25)) | p(s(bool,X24)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.sat.sat_27_prove_ltb')).
fof(f538,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f319])).
fof(f539,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f319])).
fof(f540,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP12(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f319])).
fof(f533,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f317])).
fof(f317,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f316])).
fof(f316,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f202])).
fof(f534,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f317])).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f317])).
fof(f536,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f317])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : ((~sP12(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP11(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)))),
  inference(rectify,[],[f313])).
fof(f313,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1) | sP12(X2,X0,X1)))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X2,X0,X1] : ((~sP12(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP11(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP11(X2,X0,X1)) | sP12(X2,X0,X1)))),
  inference(nnf_transformation,[],[f203])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f314])).
fof(f528,plain,(
  ( ! [X2,X0,X1] : (~sP12(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f314])).
fof(f563,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP11(X0,X1,X2) | sP12(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f529])).
fof(f529,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f314])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (sP12(X0,X1,X2) | ~sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f314])).
fof(f520,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f310])).
fof(f310,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP10(X2,X0,X1)) & (~sP10(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP10(X2,X0,X1))),
  inference(definition_folding,[],[f143,f200,f199])).
fof(f199,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f200,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f143,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> (p(s(bool,X25)) | p(s(bool,X24)))) <=> ((p(s(bool,X26)) | ~p(s(bool,X25))) & (p(s(bool,X26)) | ~p(s(bool,X24))) & (p(s(bool,X25)) | p(s(bool,X24)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.sat.sat_25_prove_ltb')).
fof(f521,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f311])).
fof(f522,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f311])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP10(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f311])).
fof(f516,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f308])).
fof(f308,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f307])).
fof(f307,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f199])).
fof(f517,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f309])).
fof(f518,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f309])).
fof(f519,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f309])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP9(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f305])).
fof(f305,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f304])).
fof(f304,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP9(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP9(X2,X0,X1)) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f200])).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f306])).
fof(f511,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f306])).
fof(f564,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP9(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f512])).
fof(f512,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f306])).
fof(f513,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f306])).
fof(f514,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f306])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~sP9(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f306])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f302])).
fof(f302,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP8(X2,X0,X1)) & (~sP8(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP8(X2,X0,X1))),
  inference(definition_folding,[],[f141,f197,f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : (sP7(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f197,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f141,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> (p(s(bool,X25)) & p(s(bool,X24)))) <=> ((p(s(bool,X26)) | ~p(s(bool,X25)) | ~p(s(bool,X24))) & (p(s(bool,X25)) | ~p(s(bool,X26))) & (p(s(bool,X24)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.sat.sat_23_prove_ltb')).
fof(f504,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f303])).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (sP8(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f303])).
fof(f508,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP8(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f303])).
fof(f499,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(nnf_transformation,[],[f196])).
fof(f500,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f301])).
fof(f501,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f301])).
fof(f502,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f301])).
fof(f492,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | (sP7(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP7(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f298])).
fof(f298,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (sP7(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP7(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f197])).
fof(f493,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f299])).
fof(f494,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f299])).
fof(f495,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f299])).
fof(f565,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f498])).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP7(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f299])).
fof(f488,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP6(X2,X1,X0)) & (~sP6(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP6(X2,X1,X0))),
  inference(definition_folding,[],[f139,f194,f193,f192,f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : (sP3(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f192,plain,(
  ! [X1,X0,X2] : (sP4(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f193,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f194,plain,(
  ! [X2,X1,X0] : (sP6(X2,X1,X0) <=> (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f139,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X24,X25,X26] : ((p(s(bool,X26)) <=> s(bool,X24) = s(bool,X25)) <=> ((p(s(bool,X26)) | p(s(bool,X25)) | p(s(bool,X24))) & (p(s(bool,X26)) | ~p(s(bool,X24)) | ~p(s(bool,X25))) & (p(s(bool,X25)) | ~p(s(bool,X24)) | ~p(s(bool,X26))) & (p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.sat.sat_2|49|_prove_ltb')).
fof(f489,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f296])).
fof(f490,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f296])).
fof(f491,plain,(
  ( ! [X2,X0,X1] : (~sP6(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f296])).
fof(f484,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(flattening,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(nnf_transformation,[],[f191])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f295])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f295])).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f295])).
fof(f480,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f292])).
fof(f292,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(flattening,[],[f291])).
fof(f291,plain,(
  ! [X1,X0,X2] : ((~sP4(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP4(X1,X0,X2)))),
  inference(nnf_transformation,[],[f192])).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f293])).
fof(f482,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f293])).
fof(f483,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f293])).
fof(f476,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f289])).
fof(f289,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f288])).
fof(f288,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f193])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f290])).
fof(f478,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f290])).
fof(f479,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f290])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (sP3(X2,X1,X0) & sP4(X1,X2,X0) & sP5(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP3(X2,X1,X0) | ~sP4(X1,X2,X0) | ~sP5(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f286])).
fof(f286,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP6(X2,X1,X0)))),
  inference(flattening,[],[f285])).
fof(f285,plain,(
  ! [X2,X1,X0] : ((~sP6(X2,X1,X0) | (sP3(X0,X1,X2) & sP4(X1,X0,X2) & sP5(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP3(X0,X1,X2) | ~sP4(X1,X0,X2) | ~sP5(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP6(X2,X1,X0)))),
  inference(nnf_transformation,[],[f194])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f287])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | sP5(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f287])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f287])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP6(X0,X1,X2) | ~sP5(X0,X2,X1) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f287])).
fof(f474,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X2,X1) | p(s(bool,X1)) | sP6(X0,X1,X2) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f287])).
fof(f475,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X2,X1) | p(s(bool,X0)) | sP6(X0,X1,X2) | ~sP4(X1,X2,X0) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f287])).
fof(f453,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | ((p(s(bool,X1)) | p(s(bool,X2))) & (p(s(bool,X1)) | p(s(bool,X0))))) & ((~p(s(bool,X1)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f279])).
fof(f279,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2))))) & (((~p(s(bool,X0)) & ~p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X2)))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> ((p(s(bool,X0)) | p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f454,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f280])).
fof(f566,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(duplicate_literal_removal,[],[f455])).
fof(f455,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f280])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f280])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f280])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f280])).
fof(f430,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK25(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK25(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK25(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK25(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK25])],[f175])).
fof(f175,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f121])).
fof(f121,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|',aHL_EXT)).
fof(f428,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,happ(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(X0,X3)))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X27,X28] : ! [X29,X30] : s(X28,X29) = s(X28,happ(s(fun(X27,X28),'_dst_x0_1'(s(X28,X29))),s(X27,X30)))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','_dst_x0_ax')).
fof(f414,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f252])).
fof(f252,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f415,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f253])).
fof(f416,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f253])).
fof(f419,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f253])).
fof(f583,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f399,f582])).
fof(f399,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f245])).
fof(f245,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f400,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f246])).
fof(f401,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f246])).
fof(f391,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f161])).
fof(f341,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.bool.EQ_CLAUSES|split|0')).
fof(f342,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f211])).
fof(f334,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|',aHL_BOOL_CASES)).
fof(f322,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))))),
  inference(cnf_transformation,[],[f208])).
fof(f208,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) & (p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | (p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) & ! [X5] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X5)))) | s(sK14,sK15) = s(sK14,X5)))) & (~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | (p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) & s(sK14,sK15) != s(sK14,sK17)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16,sK17])],[f207])).
fof(f207,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(fun(X1,bool),X3)))) & (p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) | (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X2)))) & ! [X5] : (~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X5)))) | s(X1,X2) = s(X1,X5)))) & (~p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X2)))) | ? [X4] : (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) & s(X1,X2) != s(X1,X4))))),
  inference(rectify,[],[f206])).
fof(f206,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(fun(X1,bool),X3)))) & (p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) | (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X2)))) & ! [X4] : (~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) | s(X1,X2) = s(X1,X4)))) & (~p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) | ~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X2)))) | ? [X4] : (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) & s(X1,X2) != s(X1,X4))))),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(fun(X1,bool),X3)))) & ((p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) | (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X2)))) & ! [X4] : (~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) | s(X1,X2) = s(X1,X4)))) & (~p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) | (~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X2)))) | ? [X4] : (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) & s(X1,X2) != s(X1,X4))))))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(fun(X1,bool),X3)))) & (p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) <~> (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X2)))) & ! [X4] : (~p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) | s(X1,X2) = s(X1,X4)))))),
  inference(ennf_transformation,[],[f59])).
fof(f59,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(fun(X1,bool),X3)))) => (p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) <=> (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X2)))) & ! [X4] : (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) => s(X1,X2) = s(X1,X4)))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ~! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,X1),'_dst_x0_1'(s(X1,X2))),s(fun(X1,bool),X3)))) => (p(s(bool,'const.bool.?|33|_1'(s(fun(X1,bool),X3)))) <=> (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X2)))) & ! [X4] : (p(s(bool,happ(s(fun(X1,bool),X3),s(X1,X4)))) => s(X1,X2) = s(X1,X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X6,X7),'_dst_x0_1'(s(X7,X8))),s(fun(X7,bool),X9)))) => (p(s(bool,'const.bool.?|33|_1'(s(fun(X7,bool),X9)))) <=> (p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X8)))) & ! [X10] : (p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X10)))) => s(X7,X8) = s(X7,X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X6,X7),'_dst_x0_1'(s(X7,X8))),s(fun(X7,bool),X9)))) => (p(s(bool,'const.bool.?|33|_1'(s(fun(X7,bool),X9)))) <=> (p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X8)))) & ! [X10] : (p(s(bool,happ(s(fun(X7,bool),X9),s(X7,X10)))) => s(X7,X8) = s(X7,X10)))))),
  file('thm.quantHeuristics.GUESSES_UEXISTS_THM|49|','thm.quantHeuristics.GUESSES_UEXISTS_THM|49|')).
fof(f56884,plain,(
  ( ! [X2] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X2)))) | s(sK14,sK17) = s(sK14,X2)) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f56872,f327])).
fof(f56872,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X2)))) | s(sK14,sK17) = s(sK14,X2)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f11253,f11261])).
fof(f11261,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl6),
  inference(unit_resulting_resolution,[],[f327,f580,f394])).
fof(f11253,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X1)))) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0)))) | s(sK14,X0) = s(sK14,X1)) ) | $spl0),
  inference(subsumption_resolution,[],[f10005,f327])).
fof(f10005,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0)))) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X1)))) | s(sK14,X0) = s(sK14,X1)) ) | $spl0),
  inference(superposition,[],[f373,f8567])).
fof(f8567,plain,(
  s(bool,'T_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f327,f573,f394])).
fof(f56901,plain,(
  $false | ($spl0 | $spl2 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f56900,f327])).
fof(f56900,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl5 | $spl6)),
  inference(forward_demodulation,[],[f56887,f12312])).
fof(f56887,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl0 | $spl5 | $spl6)),
  inference(unit_resulting_resolution,[],[f578,f56884])).
fof(f48125,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl132),
  inference(cnf_transformation,[],[f48125_D])).
fof(f48125_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f4579,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f4579_D])).
fof(f4579_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f48127,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(condensation,[],[f48126])).
fof(f48126,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | ~sP12(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X30,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl0),
  inference(subsumption_resolution,[],[f48104,f954])).
fof(f48104,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | ~sP12(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X30,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl0),
  inference(superposition,[],[f14681,f890])).
fof(f14681,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP12(happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(subsumption_resolution,[],[f14667,f327])).
fof(f14667,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP12(happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(superposition,[],[f1486,f8567])).
fof(f48124,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(condensation,[],[f48123])).
fof(f48123,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | ~sP12(happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27)),X25,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl0),
  inference(subsumption_resolution,[],[f48103,f954])).
fof(f48103,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | ~sP12(happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27)),X25,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl0),
  inference(superposition,[],[f14681,f890])).
fof(f48064,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) ) | $spl130),
  inference(cnf_transformation,[],[f48064_D])).
fof(f48064_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f48066,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) ) | $spl0),
  inference(condensation,[],[f48065])).
fof(f48065,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X30,happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl0),
  inference(subsumption_resolution,[],[f48033,f954])).
fof(f48033,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X30,happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl0),
  inference(superposition,[],[f14645,f890])).
fof(f14645,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X50)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X50,happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X51))),s(X52,X53)))) ) | $spl0),
  inference(subsumption_resolution,[],[f14625,f327])).
fof(f14625,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X50)) | ~p(s(bool,X51)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X50,happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X51))),s(X52,X53)))) ) | $spl0),
  inference(superposition,[],[f1483,f8567])).
fof(f48063,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) ) | $spl0),
  inference(condensation,[],[f48062])).
fof(f48062,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X25,happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl0),
  inference(subsumption_resolution,[],[f48032,f954])).
fof(f48032,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | ~sP11('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X25,happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl0),
  inference(superposition,[],[f14645,f890])).
fof(f47984,plain,(
  ( ! [X2,X3,X1] : (~sP8(X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3))) | p(s(bool,X1))) ) | $spl124),
  inference(cnf_transformation,[],[f47984_D])).
fof(f47984_D,plain,(
  ( ! [X2,X3,X1] : (~sP8(X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3))) | p(s(bool,X1))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f47996,plain,(
  ( ! [X55] : (~p(s(bool,X55))) ) | $spl128),
  inference(cnf_transformation,[],[f47996_D])).
fof(f47996_D,plain,(
  ( ! [X55] : (~p(s(bool,X55))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f47995,plain,(
  ( ! [X59,X57,X58,X55] : (~p(s(bool,X55)) | ~sP8(X57,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X58,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X58,X59))) | p(s(bool,X57))) ) | $spl0),
  inference(subsumption_resolution,[],[f47994,f345])).
fof(f47994,plain,(
  ( ! [X59,X57,X58,X55] : (~p(s(bool,X55)) | ~sP8(X57,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X58,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X58,X59))) | p(s(bool,X57)) | s(bool,'F_0') = s(bool,X55)) ) | $spl0),
  inference(forward_demodulation,[],[f47948,f428])).
fof(f47948,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (~sP8(X57,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X58,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X58,X59))) | p(s(bool,X57)) | ~p(s(bool,happ(s(fun(X54,bool),'_dst_x0_1'(s(bool,X55))),s(X54,X56)))) | s(bool,'F_0') = s(bool,X55)) ) | $spl0),
  inference(superposition,[],[f14476,f1462])).
fof(f14476,plain,(
  ( ! [X52,X50,X53,X51] : (~sP8(X51,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53))) | p(s(bool,X51)) | ~p(s(bool,X50))) ) | $spl0),
  inference(subsumption_resolution,[],[f14461,f327])).
fof(f14461,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP8(X51,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)))) ) | $spl0),
  inference(superposition,[],[f1479,f8567])).
fof(f47992,plain,(
  ( ! [X45,X44] : (~p(s(bool,'|dollar|exists_1'(s(fun(X44,bool),X45))))) ) | $spl126),
  inference(cnf_transformation,[],[f47992_D])).
fof(f47992_D,plain,(
  ( ! [X45,X44] : (~p(s(bool,'|dollar|exists_1'(s(fun(X44,bool),X45))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f47991,plain,(
  ( ! [X47,X45,X48,X46,X44] : (~p(s(bool,'|dollar|exists_1'(s(fun(X44,bool),X45)))) | ~sP8(X46,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X47,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X47,X48))) | p(s(bool,X46))) ) | $spl0),
  inference(subsumption_resolution,[],[f47990,f345])).
fof(f47990,plain,(
  ( ! [X47,X45,X48,X46,X44] : (~p(s(bool,'|dollar|exists_1'(s(fun(X44,bool),X45)))) | ~sP8(X46,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X47,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X47,X48))) | p(s(bool,X46)) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X44,bool),X45)))) ) | $spl0),
  inference(forward_demodulation,[],[f47946,f369])).
fof(f47946,plain,(
  ( ! [X47,X45,X48,X46,X44] : (~sP8(X46,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X47,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X47,X48))) | p(s(bool,X46)) | ~p(s(bool,happ(s(fun(X44,bool),X45),s(X44,'const.min.@_1'(s(fun(X44,bool),X45)))))) | s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(X44,bool),X45)))) ) | $spl0),
  inference(superposition,[],[f14476,f1670])).
fof(f47986,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8(X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3))) | p(s(bool,X1))) ) | $spl0),
  inference(condensation,[],[f47985])).
fof(f47985,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~sP8(X6,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X7,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X7,X8))) | p(s(bool,X6)) | s(bool,X5) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl0),
  inference(subsumption_resolution,[],[f47937,f954])).
fof(f47937,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~sP8(X6,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X7,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X7,X8))) | p(s(bool,X6)) | ~p(s(bool,X5)) | s(bool,X5) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl0),
  inference(superposition,[],[f14476,f890])).
fof(f47983,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8(X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3))) | p(s(bool,X1))) ) | $spl0),
  inference(condensation,[],[f47982])).
fof(f47982,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP8(X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3))) | p(s(bool,X1)) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl0),
  inference(subsumption_resolution,[],[f47936,f954])).
fof(f47936,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP8(X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3))) | p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl0),
  inference(superposition,[],[f14476,f890])).
fof(f47899,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl122),
  inference(cnf_transformation,[],[f47899_D])).
fof(f47899_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f47901,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl0),
  inference(condensation,[],[f47900])).
fof(f47900,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl0),
  inference(subsumption_resolution,[],[f47868,f954])).
fof(f47868,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl0),
  inference(superposition,[],[f14438,f890])).
fof(f14438,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51)) ) | $spl0),
  inference(subsumption_resolution,[],[f14418,f327])).
fof(f14418,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51)) ) | $spl0),
  inference(superposition,[],[f1478,f8567])).
fof(f47898,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl0),
  inference(condensation,[],[f47897])).
fof(f47897,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27)),X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl0),
  inference(subsumption_resolution,[],[f47867,f954])).
fof(f47867,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | ~sP7('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27)),X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl0),
  inference(superposition,[],[f14438,f890])).
fof(f47808,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl120),
  inference(cnf_transformation,[],[f47808_D])).
fof(f47808_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f47810,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl0),
  inference(condensation,[],[f47809])).
fof(f47809,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl0),
  inference(subsumption_resolution,[],[f47777,f954])).
fof(f47777,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl0),
  inference(superposition,[],[f14367,f890])).
fof(f14367,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51)) ) | $spl0),
  inference(subsumption_resolution,[],[f14347,f327])).
fof(f14347,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51)) ) | $spl0),
  inference(superposition,[],[f1477,f8567])).
fof(f47807,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl0),
  inference(condensation,[],[f47806])).
fof(f47806,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27)),X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl0),
  inference(subsumption_resolution,[],[f47776,f954])).
fof(f47776,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27)),X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl0),
  inference(superposition,[],[f14367,f890])).
fof(f47721,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl118),
  inference(cnf_transformation,[],[f47721_D])).
fof(f47721_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f47723,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl0),
  inference(condensation,[],[f47722])).
fof(f47722,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl0),
  inference(subsumption_resolution,[],[f47685,f954])).
fof(f47685,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) ) | $spl0),
  inference(superposition,[],[f14314,f890])).
fof(f14314,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51)) ) | $spl0),
  inference(subsumption_resolution,[],[f14294,f327])).
fof(f14294,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X52,bool),'_dst_x0_1'(s(bool,X50))),s(X52,X53)),X51)) ) | $spl0),
  inference(superposition,[],[f1476,f8567])).
fof(f47720,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl0),
  inference(condensation,[],[f47719])).
fof(f47719,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27)),X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl0),
  inference(subsumption_resolution,[],[f47684,f954])).
fof(f47684,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27)),X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl0),
  inference(superposition,[],[f14314,f890])).
fof(f39652,plain,(
  ( ! [X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X1))))) ) | $spl116),
  inference(cnf_transformation,[],[f39652_D])).
fof(f39652_D,plain,(
  ( ! [X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X1))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f39653,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl6 | $spl13)),
  inference(subsumption_resolution,[],[f39649,f31002])).
fof(f31002,plain,(
  ( ! [X24,X23] : (s(bool,X23) != s(bool,sK28(X24,sK16,sK14,bool)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X24)))) | s(bool,'T_0') = s(bool,X23)) ) | $spl13),
  inference(superposition,[],[f2232,f3705])).
fof(f3705,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | s(bool,'T_0') = s(bool,X13)) ) | $spl13),
  inference(subsumption_resolution,[],[f3699,f582])).
fof(f3699,plain,(
  ( ! [X13] : (p(s(bool,'F_0')) | s(bool,X13) = s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | s(bool,'T_0') = s(bool,X13)) ) | $spl13),
  inference(superposition,[],[f954,f1720])).
fof(f1720,plain,(
  s(bool,'F_0') = s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl13),
  inference(unit_resulting_resolution,[],[f582,f1707,f391])).
fof(f39649,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X3)))) | s(bool,X4) = s(bool,sK28(X3,sK16,sK14,bool)) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(trivial_inequality_removal,[],[f39645])).
fof(f39645,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X3)))) | s(bool,X4) = s(bool,sK28(X3,sK16,sK14,bool)) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f12934,f890])).
fof(f12934,plain,(
  ( ! [X59] : (s(bool,'T_0') != s(bool,sK28(X59,sK16,sK14,bool)) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X59))))) ) | $spl6),
  inference(superposition,[],[f435,f11261])).
fof(f39651,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl6 | $spl13)),
  inference(subsumption_resolution,[],[f39650,f31002])).
fof(f39650,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK28(X1,sK16,sK14,bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(trivial_inequality_removal,[],[f39644])).
fof(f39644,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK28(X1,sK16,sK14,bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(superposition,[],[f12934,f890])).
fof(f37231,plain,(
  ( ! [X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X1))))) ) | $spl114),
  inference(cnf_transformation,[],[f37231_D])).
fof(f37231_D,plain,(
  ( ! [X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X1))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f37232,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl6 | $spl13)),
  inference(subsumption_resolution,[],[f37228,f31003])).
fof(f31003,plain,(
  ( ! [X26,X25] : (s(bool,X25) != s(bool,sK32(X26,sK16,sK14,bool)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X26)))) | s(bool,'T_0') = s(bool,X25)) ) | $spl13),
  inference(superposition,[],[f2258,f3705])).
fof(f37228,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X3)))) | s(bool,X4) = s(bool,sK32(X3,sK16,sK14,bool)) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(trivial_inequality_removal,[],[f37224])).
fof(f37224,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X3)))) | s(bool,X4) = s(bool,sK32(X3,sK16,sK14,bool)) | s(bool,'T_0') = s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f12933,f890])).
fof(f12933,plain,(
  ( ! [X58] : (s(bool,'T_0') != s(bool,sK32(X58,sK16,sK14,bool)) | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X58))))) ) | $spl6),
  inference(superposition,[],[f441,f11261])).
fof(f37230,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl6 | $spl13)),
  inference(subsumption_resolution,[],[f37229,f31003])).
fof(f37229,plain,(
  ( ! [X2,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK32(X1,sK16,sK14,bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(trivial_inequality_removal,[],[f37223])).
fof(f37223,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(sK14,bool),sK16),s(fun(bool,bool),X1)))) | s(bool,X2) = s(bool,sK32(X1,sK16,sK14,bool)) | s(bool,'T_0') = s(bool,X2)) ) | $spl6),
  inference(superposition,[],[f12933,f890])).
fof(f23342,plain,(
  ( ! [X23,X21,X19,X17,X22,X18,X16] : (s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK32(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19)))) ) | $spl112),
  inference(cnf_transformation,[],[f23342_D])).
fof(f23342_D,plain,(
  ( ! [X23,X21,X19,X17,X22,X18,X16] : (s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK32(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19)))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f15697,plain,(
  ( ! [X14] : (p(s(bool,X14))) ) | $spl82),
  inference(cnf_transformation,[],[f15697_D])).
fof(f15697_D,plain,(
  ( ! [X14] : (p(s(bool,X14))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f23343,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,X28)) | s(X25,happ(s(fun(X24,X25),X26),s(X24,X29))) != s(X25,sK32(X27,happ(s(fun(X30,fun(X24,X25)),'_dst_x0_1'(s(fun(X24,X25),X26))),s(X30,X31)),X24,X25)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X24,X25),X26),s(fun(X25,bool),X27)))) )),
  inference(subsumption_resolution,[],[f23338,f341])).
fof(f23338,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,X28)) | s(X25,happ(s(fun(X24,X25),X26),s(X24,X29))) != s(X25,sK32(X27,happ(s(fun(X30,fun(X24,X25)),'_dst_x0_1'(s(fun(X24,X25),X26))),s(X30,X31)),X24,X25)) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X24,X25),X26),s(fun(X25,bool),X27)))) )),
  inference(superposition,[],[f2261,f890])).
fof(f23341,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,X20)) | s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK32(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19)))) )),
  inference(subsumption_resolution,[],[f23337,f341])).
fof(f23337,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,X20)) | s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK32(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f2261,f890])).
fof(f23314,plain,(
  ( ! [X23,X21,X19,X17,X22,X18,X16] : (s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK28(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19)))) ) | $spl110),
  inference(cnf_transformation,[],[f23314_D])).
fof(f23314_D,plain,(
  ( ! [X23,X21,X19,X17,X22,X18,X16] : (s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK28(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19)))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f23315,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,X28)) | s(X25,happ(s(fun(X24,X25),X26),s(X24,X29))) != s(X25,sK28(X27,happ(s(fun(X30,fun(X24,X25)),'_dst_x0_1'(s(fun(X24,X25),X26))),s(X30,X31)),X24,X25)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X24,X25),X26),s(fun(X25,bool),X27)))) )),
  inference(subsumption_resolution,[],[f23310,f341])).
fof(f23310,plain,(
  ( ! [X30,X28,X26,X24,X31,X29,X27,X25] : (p(s(bool,X28)) | s(X25,happ(s(fun(X24,X25),X26),s(X24,X29))) != s(X25,sK28(X27,happ(s(fun(X30,fun(X24,X25)),'_dst_x0_1'(s(fun(X24,X25),X26))),s(X30,X31)),X24,X25)) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X24,X25),X26),s(fun(X25,bool),X27)))) )),
  inference(superposition,[],[f2235,f890])).
fof(f23313,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,X20)) | s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK28(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19)))) )),
  inference(subsumption_resolution,[],[f23309,f341])).
fof(f23309,plain,(
  ( ! [X23,X21,X19,X17,X22,X20,X18,X16] : (p(s(bool,X20)) | s(X17,happ(s(fun(X16,X17),X18),s(X16,X21))) != s(X17,sK28(X19,happ(s(fun(X22,fun(X16,X17)),'_dst_x0_1'(s(fun(X16,X17),X18))),s(X22,X23)),X16,X17)) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X16,X17),X18),s(fun(X17,bool),X19))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f2235,f890])).
fof(f22592,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,happ(s(fun(X10,X11),X12),s(X10,sK35(X13,X12,X10,X11)))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) ) | $spl108),
  inference(cnf_transformation,[],[f22592_D])).
fof(f22592_D,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,happ(s(fun(X10,X11),X12),s(X10,sK35(X13,X12,X10,X11)))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f22593,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X18),s(X16,happ(s(fun(X15,X16),X17),s(X15,sK35(X18,X17,X15,X16)))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X15,X16),X17),s(fun(X16,bool),X18)))) )),
  inference(subsumption_resolution,[],[f22587,f341])).
fof(f22587,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X18),s(X16,happ(s(fun(X15,X16),X17),s(X15,sK35(X18,X17,X15,X16)))))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X15,X16),X17),s(fun(X16,bool),X18)))) )),
  inference(superposition,[],[f2461,f890])).
fof(f22591,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,happ(s(fun(X10,X11),X12),s(X10,sK35(X13,X12,X10,X11)))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) )),
  inference(subsumption_resolution,[],[f22586,f341])).
fof(f22586,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,happ(s(fun(X10,X11),X12),s(X10,sK35(X13,X12,X10,X11)))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f2461,f890])).
fof(f22366,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK25(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK25(X0,X1,bool,bool))))) ) | $spl104),
  inference(cnf_transformation,[],[f22366_D])).
fof(f22366_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK25(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK25(X0,X1,bool,bool))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f22368,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl107),
  inference(cnf_transformation,[],[f22368_D])).
fof(f22368_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f22356,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK25(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK25(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2211,f890])).
fof(f22355,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK25(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK25(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2211,f890])).
fof(f21348,plain,(
  ( ! [X14,X12,X17,X15,X13] : (s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),X15),s(X13,happ(s(fun(X12,X13),X14),s(X12,X17))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X12,X13),X14),s(fun(X13,bool),X15)))) ) | $spl102),
  inference(cnf_transformation,[],[f21348_D])).
fof(f21348_D,plain,(
  ( ! [X14,X12,X17,X15,X13] : (s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),X15),s(X13,happ(s(fun(X12,X13),X14),s(X12,X17))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X12,X13),X14),s(fun(X13,bool),X15)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f21349,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,X22)) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X21),s(X19,happ(s(fun(X18,X19),X20),s(X18,X23))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X18,X19),X20),s(fun(X19,bool),X21)))) )),
  inference(subsumption_resolution,[],[f21343,f341])).
fof(f21343,plain,(
  ( ! [X23,X21,X19,X22,X20,X18] : (p(s(bool,X22)) | s(bool,'T_0') = s(bool,happ(s(fun(X19,bool),X21),s(X19,happ(s(fun(X18,X19),X20),s(X18,X23))))) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X18,X19),X20),s(fun(X19,bool),X21)))) )),
  inference(superposition,[],[f2341,f890])).
fof(f21347,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,X16)) | s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),X15),s(X13,happ(s(fun(X12,X13),X14),s(X12,X17))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X12,X13),X14),s(fun(X13,bool),X15)))) )),
  inference(subsumption_resolution,[],[f21342,f341])).
fof(f21342,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (p(s(bool,X16)) | s(bool,'T_0') = s(bool,happ(s(fun(X13,bool),X15),s(X13,happ(s(fun(X12,X13),X14),s(X12,X17))))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X12,X13),X14),s(fun(X13,bool),X15))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f2341,f890])).
fof(f20194,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl96),
  inference(cnf_transformation,[],[f20194_D])).
fof(f20194_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f20190,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK25(X0,X1,bool,bool))) ) | $spl92),
  inference(cnf_transformation,[],[f20190_D])).
fof(f20190_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK25(X0,X1,bool,bool))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f20197,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl101),
  inference(cnf_transformation,[],[f20197_D])).
fof(f20197_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f20188,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK25(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2210,f334])).
fof(f20195,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl98),
  inference(cnf_transformation,[],[f20195_D])).
fof(f20195_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f20187,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK25(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2210,f890])).
fof(f20186,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK25(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f2210,f890])).
fof(f20189,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl90),
  inference(cnf_transformation,[],[f20189_D])).
fof(f20189_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f20192,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl95),
  inference(cnf_transformation,[],[f20192_D])).
fof(f20192_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f20185,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK25(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2210,f890])).
fof(f20184,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK25(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2210,f890])).
fof(f19898,plain,(
  ( ! [X4,X2,X3,X1] : (~sP9(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP10(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) ) | $spl88),
  inference(cnf_transformation,[],[f19898_D])).
fof(f19898_D,plain,(
  ( ! [X4,X2,X3,X1] : (~sP9(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP10(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f19900,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP10(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) )),
  inference(condensation,[],[f19899])).
fof(f19899,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~sP9(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | sP10(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f19882,f954])).
fof(f19882,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~sP9(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | sP10(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1489,f890])).
fof(f19897,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP10(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) )),
  inference(condensation,[],[f19896])).
fof(f19896,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~sP9(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | sP10(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f19881,f954])).
fof(f19881,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~sP9(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | sP10(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1489,f890])).
fof(f19833,plain,(
  ( ! [X4,X2,X3,X1] : (~sP11(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP12(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) ) | $spl86),
  inference(cnf_transformation,[],[f19833_D])).
fof(f19833_D,plain,(
  ( ! [X4,X2,X3,X1] : (~sP11(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP12(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f19835,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP12(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) )),
  inference(condensation,[],[f19834])).
fof(f19834,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~sP11(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | sP12(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f19817,f954])).
fof(f19817,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~sP11(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | sP12(happ(s(fun(X36,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X36,X37)),X38,X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1488,f890])).
fof(f19832,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4) | sP12(happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X1,X2)),X3,X4)) )),
  inference(condensation,[],[f19831])).
fof(f19831,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~sP11(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | sP12(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f19816,f954])).
fof(f19816,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~sP11(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | sP12(happ(s(fun(X30,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X30,X31)),X32,X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1488,f890])).
fof(f19753,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f15734,f18999])).
fof(f18999,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,X2)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f15697,f15697,f373])).
fof(f15734,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f345])).
fof(f19752,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15740,f327])).
fof(f15740,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f394])).
fof(f1948,plain,(
  s(bool,'T_0') != s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl13),
  inference(unit_resulting_resolution,[],[f327,f1707,f562])).
fof(f19751,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f15742,f327])).
fof(f15742,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f639,f15697,f394])).
fof(f19750,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15744,f15697])).
fof(f15744,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f394])).
fof(f19749,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15750,f327])).
fof(f15750,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f394])).
fof(f19748,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f15752,f327])).
fof(f15752,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f639,f15697,f394])).
fof(f19747,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15754,f15697])).
fof(f15754,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f394])).
fof(f19746,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15763,f15697])).
fof(f15763,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1726,f15697,f415])).
fof(f1726,plain,(
  sP1('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl13),
  inference(unit_resulting_resolution,[],[f327,f1707,f416])).
fof(f19745,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f15764,f327])).
fof(f15764,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f689,f15697,f415])).
fof(f19744,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f15765,f15587])).
fof(f15587,plain,(
  ( ! [X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))))) ) | $spl2),
  inference(subsumption_resolution,[],[f15483,f327])).
fof(f15483,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X1,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))))) ) | $spl2),
  inference(superposition,[],[f2304,f12312])).
fof(f15765,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl82),
  inference(unit_resulting_resolution,[],[f758,f15697,f415])).
fof(f19743,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f15766,f15697])).
fof(f15766,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8576,f15697,f415])).
fof(f8576,plain,(
  sP1('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f582,f573,f419])).
fof(f19742,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f15767,f15697])).
fof(f15767,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11270,f15697,f415])).
fof(f11270,plain,(
  sP1('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl6),
  inference(unit_resulting_resolution,[],[f582,f580,f419])).
fof(f19741,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f15768,f15697])).
fof(f15768,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12321,f15697,f415])).
fof(f12321,plain,(
  sP1('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f11256,f419])).
fof(f19740,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f15769,f15697])).
fof(f15769,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11893,f15697,f415])).
fof(f11893,plain,(
  sP1('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f582,f8562,f419])).
fof(f8562,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f372])).
fof(f19739,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15775,f327])).
fof(f15775,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1727,f15697,f415])).
fof(f1727,plain,(
  sP1('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f327,f1707,f419])).
fof(f19738,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f15780,f327])).
fof(f15780,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f680,f15697,f415])).
fof(f19737,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15781,f327])).
fof(f15781,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1726,f15697,f415])).
fof(f19736,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f15788,f15587])).
fof(f15788,plain,(
  ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | $spl82),
  inference(unit_resulting_resolution,[],[f750,f15697,f415])).
fof(f19735,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f15789,f15697])).
fof(f15789,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8575,f15697,f415])).
fof(f8575,plain,(
  sP1('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f582,f573,f416])).
fof(f19734,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f15790,f15697])).
fof(f15790,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11269,f15697,f415])).
fof(f11269,plain,(
  sP1(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f582,f580,f416])).
fof(f19733,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f15791,f15697])).
fof(f15791,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12320,f15697,f415])).
fof(f12320,plain,(
  sP1(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f582,f11256,f416])).
fof(f19732,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f15792,f15697])).
fof(f15792,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11892,f15697,f415])).
fof(f11892,plain,(
  sP1(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f582,f8562,f416])).
fof(f19731,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15793,f15697])).
fof(f15793,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1727,f15697,f415])).
fof(f19730,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15833,f15697])).
fof(f15833,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1731,f15697,f456])).
fof(f1731,plain,(
  ( ! [X0] : (~sP2(X0,'F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl13),
  inference(unit_resulting_resolution,[],[f582,f1707,f453])).
fof(f19729,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f15834,f15697])).
fof(f15834,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f720,f15697,f456])).
fof(f19728,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15835,f15697])).
fof(f15835,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1736,f15697,f456])).
fof(f1736,plain,(
  ( ! [X0] : (~sP2('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0',X0)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f582,f1707,f454])).
fof(f19727,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15837,f15697])).
fof(f15837,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1732,f15697,f456])).
fof(f1732,plain,(
  ( ! [X0] : (~sP2(X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl13),
  inference(unit_resulting_resolution,[],[f1707,f1707,f453])).
fof(f19726,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15838,f15697])).
fof(f15838,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1741,f15697,f456])).
fof(f1741,plain,(
  ( ! [X0] : (~sP2('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f582,f1707,f454])).
fof(f19725,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15839,f15697])).
fof(f15839,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1742,f15697,f456])).
fof(f1742,plain,(
  ( ! [X0] : (~sP2('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f1707,f1707,f454])).
fof(f19724,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15848,f15697])).
fof(f15848,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1728,f15697,f456])).
fof(f1728,plain,(
  ( ! [X0] : (~sP2(X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | $spl13),
  inference(unit_resulting_resolution,[],[f582,f1707,f453])).
fof(f19723,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15849,f15697])).
fof(f15849,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1732,f15697,f456])).
fof(f19722,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15850,f15697])).
fof(f15850,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1741,f15697,f456])).
fof(f19721,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15851,f15697])).
fof(f15851,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1742,f15697,f456])).
fof(f19720,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15859,f15697])).
fof(f15859,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1728,f15697,f457])).
fof(f19719,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15860,f15697])).
fof(f15860,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1732,f15697,f457])).
fof(f19718,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15861,f15697])).
fof(f15861,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1741,f15697,f457])).
fof(f19717,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15862,f15697])).
fof(f15862,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1742,f15697,f457])).
fof(f19716,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f15866,f15697])).
fof(f15866,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f713,f15697,f457])).
fof(f19715,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15867,f15697])).
fof(f15867,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1731,f15697,f457])).
fof(f19714,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15869,f15697])).
fof(f15869,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1736,f15697,f457])).
fof(f19713,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15870,f15697])).
fof(f15870,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1728,f15697,f457])).
fof(f19712,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15871,f15697])).
fof(f15871,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1732,f15697,f457])).
fof(f19711,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15873,f15697])).
fof(f15873,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1742,f15697,f457])).
fof(f19710,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15879,f15697])).
fof(f15879,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1731,f15697,f458])).
fof(f19709,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f15880,f15697])).
fof(f15880,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f720,f15697,f458])).
fof(f19708,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15881,f15697])).
fof(f15881,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1736,f15697,f458])).
fof(f19707,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15883,f15697])).
fof(f15883,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1732,f15697,f458])).
fof(f19706,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15884,f15697])).
fof(f15884,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1741,f15697,f458])).
fof(f19705,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15885,f15697])).
fof(f15885,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1742,f15697,f458])).
fof(f19704,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f15890,f15697])).
fof(f15890,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f713,f15697,f458])).
fof(f19703,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15891,f15697])).
fof(f15891,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1731,f15697,f458])).
fof(f19702,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15893,f15697])).
fof(f15893,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1736,f15697,f458])).
fof(f19701,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15894,f15697])).
fof(f15894,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1728,f15697,f458])).
fof(f19700,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15895,f15697])).
fof(f15895,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1732,f15697,f458])).
fof(f19699,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15897,f15697])).
fof(f15897,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1742,f15697,f458])).
fof(f19698,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f15903,f15697])).
fof(f15903,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f1446,f582,f15697,f472])).
fof(f19697,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15904,f15697])).
fof(f15904,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1748,f582,f15697,f472])).
fof(f1748,plain,(
  sP6('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl13),
  inference(unit_resulting_resolution,[],[f600,f594,f590,f1707,f474])).
fof(f19696,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f15905,f327])).
fof(f15905,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f991,f582,f15697,f472])).
fof(f19695,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f15906,f15697])).
fof(f15906,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f9745,f582,f15697,f472])).
fof(f9745,plain,(
  sP6('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f582,f592,f596,f8650,f475])).
fof(f8650,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f486])).
fof(f19694,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f15907,f15697])).
fof(f15907,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11697,f582,f15697,f472])).
fof(f11697,plain,(
  sP6('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f582,f592,f596,f11344,f475])).
fof(f11344,plain,(
  ( ! [X0,X1] : (sP3(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f580,f486])).
fof(f19693,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f15908,f15697])).
fof(f15908,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f13041,f582,f15697,f472])).
fof(f13041,plain,(
  sP6('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f582,f592,f596,f12406,f475])).
fof(f12406,plain,(
  ( ! [X0,X1] : (sP3(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f11256,f486])).
fof(f19692,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f15909,f15697])).
fof(f15909,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f13861,f582,f15697,f472])).
fof(f13861,plain,(
  sP6('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f582,f592,f596,f11978,f475])).
fof(f11978,plain,(
  ( ! [X0,X1] : (sP3(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f8562,f486])).
fof(f19691,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f15911,f15697])).
fof(f15911,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1))))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8657,f582,f15697,f472])).
fof(f8657,plain,(
  ( ! [X2,X0,X1] : (sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f428,f573,f489])).
fof(f19690,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15912,f15697])).
fof(f15912,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f8655,f582,f15697,f472])).
fof(f8655,plain,(
  sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0') | ($spl0 | $spl13)),
  inference(unit_resulting_resolution,[],[f1720,f573,f489])).
fof(f19689,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f15914,f15697])).
fof(f15914,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1))))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11351,f582,f15697,f472])).
fof(f11351,plain,(
  ( ! [X2,X0,X1] : (sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,X1))),s(X0,X2)),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f428,f580,f489])).
fof(f19688,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f15917,f327])).
fof(f15917,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1816,f582,f15697,f472])).
fof(f1816,plain,(
  sP6('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','F_0') | $spl13),
  inference(unit_resulting_resolution,[],[f639,f1707,f488])).
fof(f19687,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16066,f327])).
fof(f16066,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f1568,f582,f15697,f472])).
fof(f19686,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16067,f327])).
fof(f16067,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f1446,f582,f15697,f472])).
fof(f19685,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16068,f327])).
fof(f16068,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1748,f582,f15697,f472])).
fof(f19684,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16074,f15697])).
fof(f16074,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f9736,f582,f15697,f472])).
fof(f9736,plain,(
  sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f582,f590,f598,f8649,f474])).
fof(f8649,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f485])).
fof(f19683,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16075,f15697])).
fof(f16075,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8657,f582,f15697,f472])).
fof(f19682,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16076,f15697])).
fof(f16076,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f8655,f582,f15697,f472])).
fof(f19681,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16077,f15697])).
fof(f16077,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11687,f582,f15697,f472])).
fof(f11687,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f582,f590,f598,f11343,f474])).
fof(f11343,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f580,f485])).
fof(f19680,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16078,f15697])).
fof(f16078,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11351,f582,f15697,f472])).
fof(f19679,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f16079,f15697])).
fof(f16079,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f13030,f582,f15697,f472])).
fof(f13030,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f582,f590,f598,f12405,f474])).
fof(f12405,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f11256,f485])).
fof(f19678,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16080,f15697])).
fof(f16080,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f13849,f582,f15697,f472])).
fof(f13849,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f582,f590,f598,f11977,f474])).
fof(f11977,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f8562,f485])).
fof(f19677,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16081,f15697])).
fof(f16081,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1816,f582,f15697,f472])).
fof(f19676,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16134,f16999])).
fof(f16999,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f639,f15697,f490])).
fof(f16134,plain,(
  ( ! [X0] : (sP6(X0,'T_0','F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f602,f596,f588,f15697,f473])).
fof(f19675,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16141,f16997])).
fof(f16997,plain,(
  ( ! [X0] : (~sP6(X0,'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f490])).
fof(f16141,plain,(
  ( ! [X0] : (sP6(X0,'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f1781,f588,f15697,f473])).
fof(f1781,plain,(
  ( ! [X0,X1] : (sP4(X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f1707,f482])).
fof(f19674,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16464,f16999])).
fof(f16464,plain,(
  ( ! [X0] : (sP6(X0,'T_0','F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f602,f596,f590,f15697,f473])).
fof(f19673,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16471,f1017])).
fof(f16471,plain,(
  sP6('T_0','T_0','F_0') | $spl82),
  inference(unit_resulting_resolution,[],[f602,f594,f590,f15697,f473])).
fof(f19672,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16473,f8607])).
fof(f8607,plain,(
  ~sP6('T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f573,f472])).
fof(f16473,plain,(
  sP6('T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8650,f594,f590,f15697,f473])).
fof(f19671,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16474,f11301])).
fof(f11301,plain,(
  ~sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f327,f582,f580,f472])).
fof(f16474,plain,(
  sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11344,f594,f590,f15697,f473])).
fof(f19670,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f16475,f12352])).
fof(f12352,plain,(
  ~sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f327,f582,f11256,f472])).
fof(f16475,plain,(
  sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12406,f594,f590,f15697,f473])).
fof(f19669,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16476,f11924])).
fof(f11924,plain,(
  ~sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f8562,f472])).
fof(f16476,plain,(
  sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11978,f594,f590,f15697,f473])).
fof(f19668,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16478,f16999])).
fof(f16478,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','F_0') | $spl82),
  inference(unit_resulting_resolution,[],[f602,f629,f590,f15697,f473])).
fof(f19667,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16485,f8665])).
fof(f8665,plain,(
  ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f639,f573,f490])).
fof(f16485,plain,(
  sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0','F_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f8648,f590,f15697,f473])).
fof(f8648,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f481])).
fof(f19666,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16487,f8608])).
fof(f8608,plain,(
  ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f573,f582,f573,f472])).
fof(f16487,plain,(
  sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8650,f8648,f590,f15697,f473])).
fof(f19665,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16492,f11359])).
fof(f11359,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f639,f580,f490])).
fof(f16492,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0','F_0') | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f11342,f590,f15697,f473])).
fof(f11342,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f580,f481])).
fof(f19664,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16495,f11302])).
fof(f11302,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f580,f582,f580,f472])).
fof(f16495,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11344,f11342,f590,f15697,f473])).
fof(f19663,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f16499,f12421])).
fof(f12421,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f639,f11256,f490])).
fof(f16499,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','F_0') | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f12404,f590,f15697,f473])).
fof(f12404,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f11256,f481])).
fof(f19662,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f16503,f12353])).
fof(f12353,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f11256,f582,f11256,f472])).
fof(f16503,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12406,f12404,f590,f15697,f473])).
fof(f19661,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16506,f11993])).
fof(f11993,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f639,f8562,f490])).
fof(f16506,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0','F_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f11976,f590,f15697,f473])).
fof(f11976,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f8562,f481])).
fof(f19660,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16511,f11925])).
fof(f11925,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f8562,f582,f8562,f472])).
fof(f16511,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11978,f11976,f590,f15697,f473])).
fof(f19659,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16530,f16997])).
fof(f16530,plain,(
  ( ! [X0] : (sP6(X0,'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f1781,f1774,f15697,f473])).
fof(f1774,plain,(
  ( ! [X0,X1] : (sP5(X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f1707,f478])).
fof(f19658,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16537,f1744])).
fof(f1744,plain,(
  ~sP6('T_0','T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl13),
  inference(unit_resulting_resolution,[],[f327,f327,f1707,f472])).
fof(f16537,plain,(
  sP6('T_0','T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f594,f1774,f15697,f473])).
fof(f19657,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16544,f16997])).
fof(f16544,plain,(
  sP6('const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f629,f1774,f15697,f473])).
fof(f19656,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16551,f8664])).
fof(f8664,plain,(
  ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl13)),
  inference(unit_resulting_resolution,[],[f1948,f573,f490])).
fof(f16551,plain,(
  sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f8648,f1774,f15697,f473])).
fof(f19655,plain,(
  $false | ($spl6 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16558,f11358])).
fof(f11358,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl6 | $spl13)),
  inference(unit_resulting_resolution,[],[f1948,f580,f490])).
fof(f16558,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl6 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f11342,f1774,f15697,f473])).
fof(f19654,plain,(
  $false | ($spl2 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16565,f12420])).
fof(f12420,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl2 | $spl13)),
  inference(unit_resulting_resolution,[],[f1948,f11256,f490])).
fof(f16565,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl2 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f12404,f1774,f15697,f473])).
fof(f19653,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16572,f11992])).
fof(f11992,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl13)),
  inference(unit_resulting_resolution,[],[f1948,f8562,f490])).
fof(f16572,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f11976,f1774,f15697,f473])).
fof(f19652,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16620,f16997])).
fof(f16620,plain,(
  sP6('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f1781,f1775,f15697,f473])).
fof(f1775,plain,(
  ( ! [X0,X1] : (sP5('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,X1)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f1707,f479])).
fof(f19651,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16638,f16801])).
fof(f16801,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f477])).
fof(f16638,plain,(
  ~sP5('T_0','F_0','T_0') | $spl82),
  inference(unit_resulting_resolution,[],[f602,f596,f1017,f15697,f473])).
fof(f19650,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16639,f16855])).
fof(f16855,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f481])).
fof(f16639,plain,(
  ~sP4('T_0','F_0','T_0') | $spl82),
  inference(unit_resulting_resolution,[],[f602,f590,f1017,f15697,f473])).
fof(f19649,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16640,f16861])).
fof(f16861,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f487])).
fof(f16640,plain,(
  ~sP3('F_0','T_0','T_0') | $spl82),
  inference(unit_resulting_resolution,[],[f596,f590,f1017,f15697,f473])).
fof(f19648,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16642,f16801])).
fof(f16642,plain,(
  ~sP5('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0') | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f1781,f1744,f15697,f473])).
fof(f19647,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16643,f16855])).
fof(f16643,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0') | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f1774,f1744,f15697,f473])).
fof(f19646,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16644,f16861])).
fof(f16644,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','T_0') | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1781,f1774,f1744,f15697,f473])).
fof(f19645,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16651,f16801])).
fof(f16651,plain,(
  ~sP5('T_0','F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8650,f596,f8607,f15697,f473])).
fof(f19644,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16652,f16855])).
fof(f16652,plain,(
  ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','T_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8650,f590,f8607,f15697,f473])).
fof(f19643,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16653,f16861])).
fof(f16653,plain,(
  ~sP3('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f596,f590,f8607,f15697,f473])).
fof(f19642,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16655,f16801])).
fof(f16655,plain,(
  ~sP5('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11344,f596,f11301,f15697,f473])).
fof(f19641,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16656,f16855])).
fof(f16656,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0','T_0') | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11344,f590,f11301,f15697,f473])).
fof(f19640,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16657,f16861])).
fof(f16657,plain,(
  ~sP3('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0') | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f596,f590,f11301,f15697,f473])).
fof(f19639,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f16659,f16801])).
fof(f16659,plain,(
  ~sP5('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12406,f596,f12352,f15697,f473])).
fof(f19638,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f16660,f16855])).
fof(f16660,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0') | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12406,f590,f12352,f15697,f473])).
fof(f19637,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f16661,f16861])).
fof(f16661,plain,(
  ~sP3('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f596,f590,f12352,f15697,f473])).
fof(f19636,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16663,f16801])).
fof(f16663,plain,(
  ~sP5('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11978,f596,f11924,f15697,f473])).
fof(f19635,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16664,f16855])).
fof(f16664,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0','T_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11978,f590,f11924,f15697,f473])).
fof(f19634,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16665,f16861])).
fof(f16665,plain,(
  ~sP3('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f596,f590,f11924,f15697,f473])).
fof(f19633,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16675,f16861])).
fof(f16675,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))),s(X0,X1)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1782,f592,f1448,f15697,f473])).
fof(f1782,plain,(
  ( ! [X0,X1] : (sP4('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,X1)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f1707,f483])).
fof(f19632,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16694,f16855])).
fof(f16694,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))))),s(X0,X1)),'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)),'F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f632,f592,f1445,f15697,f473])).
fof(f19631,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16695,f16855])).
fof(f16695,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))),s(X0,X1)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8651,f592,f1445,f15697,f473])).
fof(f8651,plain,(
  ( ! [X0,X1] : (sP3('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f487])).
fof(f19630,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16696,f16855])).
fof(f16696,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))))),s(X0,X1)),happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0')) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11345,f592,f1445,f15697,f473])).
fof(f11345,plain,(
  ( ! [X0,X1] : (sP3(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0,X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f580,f487])).
fof(f19629,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f16697,f16855])).
fof(f16697,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))))),s(X0,X1)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0')) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12407,f592,f1445,f15697,f473])).
fof(f12407,plain,(
  ( ! [X0,X1] : (sP3(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f11256,f487])).
fof(f19628,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16698,f16855])).
fof(f16698,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))))),s(X0,X1)),happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0')) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11979,f592,f1445,f15697,f473])).
fof(f11979,plain,(
  ( ! [X0,X1] : (sP3(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f8562,f487])).
fof(f19627,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16700,f16861])).
fof(f16700,plain,(
  ( ! [X0,X1] : (~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))),s(X0,X1)),'F_0')) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1781,f592,f1445,f15697,f473])).
fof(f19626,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16702,f16861])).
fof(f16702,plain,(
  ( ! [X0,X1] : (~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))),s(X0,X1)),'F_0')) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1781,f1774,f1445,f15697,f473])).
fof(f19625,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16711,f16801])).
fof(f16711,plain,(
  ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','T_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f596,f8665,f15697,f473])).
fof(f19624,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16712,f16855])).
fof(f16712,plain,(
  ~sP4('T_0','F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f590,f8665,f15697,f473])).
fof(f19623,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16713,f16861])).
fof(f16713,plain,(
  ~sP3('F_0','T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f596,f590,f8665,f15697,f473])).
fof(f19622,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16715,f16801])).
fof(f16715,plain,(
  ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0') | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f1781,f8664,f15697,f473])).
fof(f19621,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16716,f16855])).
fof(f16716,plain,(
  ~sP4('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f1774,f8664,f15697,f473])).
fof(f19620,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16717,f16861])).
fof(f16717,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1781,f1774,f8664,f15697,f473])).
fof(f19619,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16721,f16801])).
fof(f16721,plain,(
  ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8650,f596,f8608,f15697,f473])).
fof(f19618,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16722,f16855])).
fof(f16722,plain,(
  ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8650,f590,f8608,f15697,f473])).
fof(f19617,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16723,f16861])).
fof(f16723,plain,(
  ~sP3('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f596,f590,f8608,f15697,f473])).
fof(f19616,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16726,f16801])).
fof(f16726,plain,(
  ~sP5(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0','T_0') | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f596,f11359,f15697,f473])).
fof(f19615,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16727,f16855])).
fof(f16727,plain,(
  ~sP4('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f590,f11359,f15697,f473])).
fof(f19614,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16728,f16861])).
fof(f16728,plain,(
  ~sP3('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f596,f590,f11359,f15697,f473])).
fof(f19613,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16731,f16801])).
fof(f16731,plain,(
  ~sP5(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11344,f596,f11302,f15697,f473])).
fof(f19612,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16732,f16855])).
fof(f16732,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11344,f590,f11302,f15697,f473])).
fof(f19611,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16733,f16861])).
fof(f16733,plain,(
  ~sP3('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f596,f590,f11302,f15697,f473])).
fof(f19610,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f16735,f16801])).
fof(f16735,plain,(
  ~sP5(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0') | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f596,f12421,f15697,f473])).
fof(f19609,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f16736,f16855])).
fof(f16736,plain,(
  ~sP4('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f590,f12421,f15697,f473])).
fof(f19608,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f16737,f16861])).
fof(f16737,plain,(
  ~sP3('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f596,f590,f12421,f15697,f473])).
fof(f19607,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16740,f16801])).
fof(f16740,plain,(
  ~sP5(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0','T_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f596,f11993,f15697,f473])).
fof(f19606,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16741,f16855])).
fof(f16741,plain,(
  ~sP4('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f590,f11993,f15697,f473])).
fof(f19605,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16742,f16861])).
fof(f16742,plain,(
  ~sP3('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f596,f590,f11993,f15697,f473])).
fof(f19604,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16750,f15697])).
fof(f16750,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1774,f582,f15697,f476])).
fof(f19603,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16751,f15697])).
fof(f16751,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f592,f582,f15697,f476])).
fof(f19602,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16752,f15697])).
fof(f16752,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1775,f582,f15697,f476])).
fof(f19601,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16761,f16801])).
fof(f16761,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f327,f15697,f476])).
fof(f19600,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16773,f16801])).
fof(f16773,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f15697,f476])).
fof(f19599,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16775,f15697])).
fof(f16775,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f590,f582,f15697,f476])).
fof(f19598,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16776,f15697])).
fof(f16776,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1774,f582,f15697,f476])).
fof(f19597,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16778,f15697])).
fof(f16778,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1775,f582,f15697,f476])).
fof(f19596,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16787,f16801])).
fof(f16787,plain,(
  ( ! [X0] : (~sP5('T_0',X0,'F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f327,f15697,f476])).
fof(f19595,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16799,f16801])).
fof(f16799,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f15697,f476])).
fof(f19594,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16804,f15697])).
fof(f16804,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1781,f582,f15697,f480])).
fof(f19593,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16805,f15697])).
fof(f16805,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f598,f582,f15697,f480])).
fof(f19592,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16806,f15697])).
fof(f16806,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1782,f582,f15697,f480])).
fof(f19591,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16815,f16855])).
fof(f16815,plain,(
  ( ! [X0] : (~sP4(X0,'T_0','F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f327,f15697,f480])).
fof(f19590,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16827,f16855])).
fof(f16827,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f15697,f480])).
fof(f19589,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16829,f15697])).
fof(f16829,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f596,f582,f15697,f480])).
fof(f19588,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16830,f15697])).
fof(f16830,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1781,f582,f15697,f480])).
fof(f19587,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16832,f15697])).
fof(f16832,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1782,f582,f15697,f480])).
fof(f19586,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16841,f16855])).
fof(f16841,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f327,f15697,f480])).
fof(f19585,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16853,f16855])).
fof(f16853,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f15697,f480])).
fof(f19584,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16866,f18999])).
fof(f16866,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f9693,f15697,f489])).
fof(f9693,plain,(
  ~sP6('T_0','F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f8626,f471])).
fof(f8626,plain,(
  ~sP5('T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f573,f476])).
fof(f19583,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16867,f18999])).
fof(f16867,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11650,f15697,f489])).
fof(f11650,plain,(
  ~sP6('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl6),
  inference(unit_resulting_resolution,[],[f11320,f471])).
fof(f11320,plain,(
  ~sP5('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f327,f582,f580,f476])).
fof(f19582,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f16868,f18999])).
fof(f16868,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12876,f15697,f489])).
fof(f12876,plain,(
  ~sP6('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f12379,f471])).
fof(f12379,plain,(
  ~sP5('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f327,f582,f11256,f476])).
fof(f19581,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16869,f18999])).
fof(f16869,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f13774,f15697,f489])).
fof(f13774,plain,(
  ~sP6('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f11951,f471])).
fof(f11951,plain,(
  ~sP5('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f8562,f476])).
fof(f19580,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16870,f18999])).
fof(f16870,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8607,f15697,f489])).
fof(f19579,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16871,f18999])).
fof(f16871,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11301,f15697,f489])).
fof(f19578,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f16872,f18999])).
fof(f16872,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12352,f15697,f489])).
fof(f19577,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16873,f18999])).
fof(f16873,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11924,f15697,f489])).
fof(f19576,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16875,f428])).
fof(f16875,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f1448,f15697,f489])).
fof(f19575,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16877,f644])).
fof(f16877,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | $spl82),
  inference(unit_resulting_resolution,[],[f1030,f15697,f489])).
fof(f19574,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16878,f8567])).
fof(f16878,plain,(
  s(bool,'T_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f9723,f15697,f489])).
fof(f9723,plain,(
  ~sP6('F_0','T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f8641,f470])).
fof(f8641,plain,(
  ~sP4('T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f573,f480])).
fof(f19573,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16879,f11261])).
fof(f16879,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11672,f15697,f489])).
fof(f11672,plain,(
  ~sP6('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl6),
  inference(unit_resulting_resolution,[],[f11335,f470])).
fof(f11335,plain,(
  ~sP4('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f327,f582,f580,f480])).
fof(f19572,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f16880,f12312])).
fof(f16880,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f13013,f15697,f489])).
fof(f13013,plain,(
  ~sP6('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f12396,f470])).
fof(f12396,plain,(
  ~sP4('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f327,f582,f11256,f480])).
fof(f19571,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16881,f11884])).
fof(f11884,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f327,f8562,f394])).
fof(f16881,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f13821,f15697,f489])).
fof(f13821,plain,(
  ~sP6('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f11968,f470])).
fof(f11968,plain,(
  ~sP4('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f8562,f480])).
fof(f19570,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16882,f1720])).
fof(f16882,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f2208,f15697,f489])).
fof(f2208,plain,(
  ~sP6('F_0','F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl13),
  inference(unit_resulting_resolution,[],[f1783,f469])).
fof(f1783,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','F_0') | $spl13),
  inference(unit_resulting_resolution,[],[f582,f582,f1707,f484])).
fof(f19569,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16884,f644])).
fof(f16884,plain,(
  s(bool,'T_0') != s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16))) | $spl82),
  inference(unit_resulting_resolution,[],[f1027,f15697,f489])).
fof(f19568,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16885,f8567])).
fof(f16885,plain,(
  s(bool,'T_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f9718,f15697,f489])).
fof(f9718,plain,(
  ~sP6('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f8634,f470])).
fof(f8634,plain,(
  ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f573,f480])).
fof(f19567,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16887,f428])).
fof(f16887,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,happ(s(fun(X1,bool),'_dst_x0_1'(s(bool,X0))),s(X1,X2)))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f1445,f15697,f489])).
fof(f19566,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16888,f11261])).
fof(f16888,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11667,f15697,f489])).
fof(f11667,plain,(
  ~sP6('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0') | $spl6),
  inference(unit_resulting_resolution,[],[f11328,f470])).
fof(f11328,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f327,f582,f580,f480])).
fof(f19565,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f16889,f12312])).
fof(f16889,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f13007,f15697,f489])).
fof(f13007,plain,(
  ~sP6('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f12388,f470])).
fof(f12388,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f327,f582,f11256,f480])).
fof(f19564,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16890,f11884])).
fof(f16890,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f13805,f15697,f489])).
fof(f13805,plain,(
  ~sP6('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f11960,f470])).
fof(f11960,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f8562,f480])).
fof(f19563,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16891,f369])).
fof(f16891,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) != s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f1664,f15697,f489])).
fof(f19562,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16892,f1720])).
fof(f16892,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f2215,f15697,f489])).
fof(f2215,plain,(
  ~sP6('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl13),
  inference(unit_resulting_resolution,[],[f1793,f469])).
fof(f1793,plain,(
  ~sP3('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl13),
  inference(unit_resulting_resolution,[],[f582,f582,f1707,f484])).
fof(f19561,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f16893,f369])).
fof(f16893,plain,(
  ( ! [X0,X1] : (s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))) != s(bool,happ(s(fun(X0,bool),X1),s(X0,'const.min.@_1'(s(fun(X0,bool),X1)))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f1666,f15697,f489])).
fof(f19560,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16898,f18999])).
fof(f16898,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12738,f15697,f489])).
fof(f12738,plain,(
  ~sP6('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f8627,f471])).
fof(f8627,plain,(
  ~sP5('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f573,f582,f573,f476])).
fof(f19559,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f16899,f18999])).
fof(f16899,plain,(
  s(bool,'F_0') != s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8608,f15697,f489])).
fof(f19558,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f16903,f18999])).
fof(f16903,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11302,f15697,f489])).
fof(f19557,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16910,f1720])).
fof(f16910,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f5851,f15697,f489])).
fof(f5851,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl13),
  inference(unit_resulting_resolution,[],[f1804,f469])).
fof(f1804,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl13),
  inference(unit_resulting_resolution,[],[f1707,f582,f1707,f484])).
fof(f19556,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f16911,f1720])).
fof(f16911,plain,(
  s(bool,'F_0') != s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f5865,f15697,f489])).
fof(f5865,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl13),
  inference(unit_resulting_resolution,[],[f1805,f469])).
fof(f1805,plain,(
  ~sP3('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl13),
  inference(unit_resulting_resolution,[],[f582,f1707,f1707,f484])).
fof(f19555,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17035,f17549])).
fof(f17549,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f15697,f15697,f505])).
fof(f17035,plain,(
  ( ! [X0,X1] : (~sP8(X0,X1,'F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f493])).
fof(f19554,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17067,f17549])).
fof(f17067,plain,(
  ( ! [X0,X1] : (~sP8(X0,'F_0',X1)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f494])).
fof(f19553,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17070,f15697])).
fof(f17070,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1838,f1827,f15697,f495])).
fof(f1827,plain,(
  ( ! [X0] : (~sP8('T_0',X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl13),
  inference(unit_resulting_resolution,[],[f327,f1707,f493])).
fof(f1838,plain,(
  ( ! [X0,X1] : (sP7(X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f1707,f501])).
fof(f19552,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17071,f15697])).
fof(f17071,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f610,f762,f15697,f495])).
fof(f19551,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17072,f15697])).
fof(f17072,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f606,f762,f15697,f495])).
fof(f19550,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17074,f15697])).
fof(f17074,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1838,f762,f15697,f495])).
fof(f19549,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17075,f15697])).
fof(f17075,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1839,f1830,f15697,f495])).
fof(f1830,plain,(
  ( ! [X0] : (~sP8('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f327,f1707,f494])).
fof(f1839,plain,(
  ( ! [X0,X1] : (sP7('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,X1)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f1707,f502])).
fof(f19548,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17076,f15697])).
fof(f17076,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f606,f1830,f15697,f495])).
fof(f19547,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17078,f15697])).
fof(f17078,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1838,f1830,f15697,f495])).
fof(f19546,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17080,f15697])).
fof(f17080,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f8669,f15697,f495])).
fof(f8669,plain,(
  ( ! [X0] : (~sP8('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f582,f573,f494])).
fof(f19545,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17081,f15697])).
fof(f17081,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8700,f8669,f15697,f495])).
fof(f8700,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f500])).
fof(f19544,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17083,f15697])).
fof(f17083,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1838,f8669,f15697,f495])).
fof(f19543,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17085,f15697])).
fof(f17085,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f11363,f15697,f495])).
fof(f11363,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0',X0)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f582,f580,f494])).
fof(f19542,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17086,f15697])).
fof(f17086,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11396,f11363,f15697,f495])).
fof(f11396,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f580,f500])).
fof(f19541,plain,(
  $false | ($spl6 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17088,f15697])).
fof(f17088,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl6 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1838,f11363,f15697,f495])).
fof(f19540,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17090,f15697])).
fof(f17090,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f12425,f15697,f495])).
fof(f12425,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f11256,f494])).
fof(f19539,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17091,f15697])).
fof(f17091,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12466,f12425,f15697,f495])).
fof(f12466,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f11256,f500])).
fof(f19538,plain,(
  $false | ($spl2 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17093,f15697])).
fof(f17093,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl2 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1838,f12425,f15697,f495])).
fof(f19537,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17095,f15697])).
fof(f17095,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f11997,f15697,f495])).
fof(f11997,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f582,f8562,f494])).
fof(f19536,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17096,f15697])).
fof(f17096,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12038,f11997,f15697,f495])).
fof(f12038,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f8562,f500])).
fof(f19535,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17098,f15697])).
fof(f17098,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1838,f11997,f15697,f495])).
fof(f19534,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17107,f17325])).
fof(f17325,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f500])).
fof(f17107,plain,(
  ~sP7('F_0','T_0','T_0') | $spl82),
  inference(unit_resulting_resolution,[],[f762,f327,f15697,f495])).
fof(f19533,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17109,f17325])).
fof(f17109,plain,(
  ~sP7('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','T_0') | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1830,f327,f15697,f495])).
fof(f19532,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17116,f17325])).
fof(f17116,plain,(
  ~sP7('F_0','T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8669,f327,f15697,f495])).
fof(f19531,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17118,f17325])).
fof(f17118,plain,(
  ~sP7('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11363,f327,f15697,f495])).
fof(f19530,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17120,f17325])).
fof(f17120,plain,(
  ~sP7('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12425,f327,f15697,f495])).
fof(f19529,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17122,f17325])).
fof(f17122,plain,(
  ~sP7('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11997,f327,f15697,f495])).
fof(f19528,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17135,f17325])).
fof(f17135,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f759,f15697,f15697,f495])).
fof(f19527,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17137,f17325])).
fof(f17137,plain,(
  ( ! [X0] : (~sP7(X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0')) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1827,f15697,f15697,f495])).
fof(f19526,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17139,f17325])).
fof(f17139,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f762,f15697,f15697,f495])).
fof(f19525,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17141,f17325])).
fof(f17141,plain,(
  ( ! [X0] : (~sP7('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'T_0')) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1830,f15697,f15697,f495])).
fof(f19524,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17154,f17325])).
fof(f17154,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8668,f15697,f15697,f495])).
fof(f8668,plain,(
  ( ! [X0] : (~sP8('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f582,f573,f493])).
fof(f19523,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17156,f17325])).
fof(f17156,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8669,f15697,f15697,f495])).
fof(f19522,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17158,f17325])).
fof(f17158,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11362,f15697,f15697,f495])).
fof(f11362,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0,'F_0')) ) | $spl6),
  inference(unit_resulting_resolution,[],[f582,f580,f493])).
fof(f19521,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17160,f17325])).
fof(f17160,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11363,f15697,f15697,f495])).
fof(f19520,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17162,f17325])).
fof(f17162,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12424,f15697,f15697,f495])).
fof(f12424,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f11256,f493])).
fof(f19519,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17164,f17325])).
fof(f17164,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12425,f15697,f15697,f495])).
fof(f19518,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17166,f17325])).
fof(f17166,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11996,f15697,f15697,f495])).
fof(f11996,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f582,f8562,f493])).
fof(f19517,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17168,f17325])).
fof(f17168,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11997,f15697,f15697,f495])).
fof(f19516,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17171,f15697])).
fof(f17171,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f608,f759,f15697,f495])).
fof(f19515,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17172,f15697])).
fof(f17172,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f606,f759,f15697,f495])).
fof(f19514,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17174,f15697])).
fof(f17174,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1839,f759,f15697,f495])).
fof(f19513,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17175,f15697])).
fof(f17175,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1838,f1827,f15697,f495])).
fof(f19512,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17176,f15697])).
fof(f17176,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f606,f1827,f15697,f495])).
fof(f19511,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17178,f15697])).
fof(f17178,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1839,f1827,f15697,f495])).
fof(f19510,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17180,f15697])).
fof(f17180,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1839,f1830,f15697,f495])).
fof(f19509,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17181,f15697])).
fof(f17181,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f8668,f15697,f495])).
fof(f19508,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17182,f15697])).
fof(f17182,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8700,f8668,f15697,f495])).
fof(f19507,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17184,f15697])).
fof(f17184,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1839,f8668,f15697,f495])).
fof(f19506,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17186,f15697])).
fof(f17186,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f11362,f15697,f495])).
fof(f19505,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17187,f15697])).
fof(f17187,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11396,f11362,f15697,f495])).
fof(f19504,plain,(
  $false | ($spl6 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17189,f15697])).
fof(f17189,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl6 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1839,f11362,f15697,f495])).
fof(f19503,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17191,f15697])).
fof(f17191,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f12424,f15697,f495])).
fof(f19502,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17192,f15697])).
fof(f17192,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12466,f12424,f15697,f495])).
fof(f19501,plain,(
  $false | ($spl2 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17194,f15697])).
fof(f17194,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl2 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1839,f12424,f15697,f495])).
fof(f19500,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17196,f15697])).
fof(f17196,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f11996,f15697,f495])).
fof(f19499,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17197,f15697])).
fof(f17197,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12038,f11996,f15697,f495])).
fof(f19498,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17199,f15697])).
fof(f17199,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1839,f11996,f15697,f495])).
fof(f19497,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17209,f17325])).
fof(f17209,plain,(
  ~sP7('T_0','F_0','T_0') | $spl82),
  inference(unit_resulting_resolution,[],[f759,f327,f15697,f495])).
fof(f19496,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17211,f17325])).
fof(f17211,plain,(
  ~sP7('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0') | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1827,f327,f15697,f495])).
fof(f19495,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17218,f17325])).
fof(f17218,plain,(
  ~sP7('T_0','F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8668,f327,f15697,f495])).
fof(f19494,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17220,f17325])).
fof(f17220,plain,(
  ~sP7('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11362,f327,f15697,f495])).
fof(f19493,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17222,f17325])).
fof(f17222,plain,(
  ~sP7('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12424,f327,f15697,f495])).
fof(f19492,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17224,f17325])).
fof(f17224,plain,(
  ~sP7('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11996,f327,f15697,f495])).
fof(f19491,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17237,f17325])).
fof(f17237,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','T_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f759,f15697,f15697,f495])).
fof(f19490,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17239,f17325])).
fof(f17239,plain,(
  ( ! [X0] : (~sP7(X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0')) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1827,f15697,f15697,f495])).
fof(f19489,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17241,f17325])).
fof(f17241,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f762,f15697,f15697,f495])).
fof(f19488,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17243,f17325])).
fof(f17243,plain,(
  ( ! [X0] : (~sP7('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'T_0')) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1830,f15697,f15697,f495])).
fof(f19487,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17256,f17325])).
fof(f17256,plain,(
  ( ! [X0] : (~sP7(X0,'F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8668,f15697,f15697,f495])).
fof(f19486,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17258,f17325])).
fof(f17258,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8669,f15697,f15697,f495])).
fof(f19485,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17260,f17325])).
fof(f17260,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11362,f15697,f15697,f495])).
fof(f19484,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17262,f17325])).
fof(f17262,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11363,f15697,f15697,f495])).
fof(f19483,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17264,f17325])).
fof(f17264,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12424,f15697,f15697,f495])).
fof(f19482,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17266,f17325])).
fof(f17266,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12425,f15697,f15697,f495])).
fof(f19481,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17268,f17325])).
fof(f17268,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11996,f15697,f15697,f495])).
fof(f19480,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17270,f17325])).
fof(f17270,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11997,f15697,f15697,f495])).
fof(f19479,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17274,f15697])).
fof(f17274,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1838,f582,f15697,f499])).
fof(f19478,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17275,f15697])).
fof(f17275,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f610,f582,f15697,f499])).
fof(f19477,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17276,f15697])).
fof(f17276,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1839,f582,f15697,f499])).
fof(f19476,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17285,f17325])).
fof(f17285,plain,(
  ( ! [X0] : (~sP7(X0,'T_0','F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f327,f15697,f499])).
fof(f19475,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17297,f17325])).
fof(f17297,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f15697,f499])).
fof(f19474,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17299,f15697])).
fof(f17299,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f608,f582,f15697,f499])).
fof(f19473,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17300,f15697])).
fof(f17300,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1838,f582,f15697,f499])).
fof(f19472,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17302,f15697])).
fof(f17302,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1839,f582,f15697,f499])).
fof(f19471,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17311,f17325])).
fof(f17311,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f327,f15697,f499])).
fof(f19470,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17323,f17325])).
fof(f17323,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f15697,f499])).
fof(f19469,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17328,f15697])).
fof(f17328,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1827,f327,f15697,f505])).
fof(f19468,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17329,f15697])).
fof(f17329,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f762,f327,f15697,f505])).
fof(f19467,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17330,f15697])).
fof(f17330,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1830,f327,f15697,f505])).
fof(f19466,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17332,f15697])).
fof(f17332,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1827,f15697,f15697,f505])).
fof(f19465,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17333,f15697])).
fof(f17333,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f762,f15697,f15697,f505])).
fof(f19464,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17334,f15697])).
fof(f17334,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1830,f15697,f15697,f505])).
fof(f19463,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17335,f327])).
fof(f17335,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f1181,f15697,f15697,f505])).
fof(f19462,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17336,f327])).
fof(f17336,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8717,f15697,f15697,f505])).
fof(f8717,plain,(
  ~sP8('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f573,f508])).
fof(f19461,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17337,f327])).
fof(f17337,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11413,f15697,f15697,f505])).
fof(f11413,plain,(
  ~sP8('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0') | $spl6),
  inference(unit_resulting_resolution,[],[f327,f582,f580,f508])).
fof(f19460,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17338,f327])).
fof(f17338,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12483,f15697,f15697,f505])).
fof(f12483,plain,(
  ~sP8('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f327,f582,f11256,f508])).
fof(f19459,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17339,f327])).
fof(f17339,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12055,f15697,f15697,f505])).
fof(f12055,plain,(
  ~sP8('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f8562,f508])).
fof(f19458,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17340,f15697])).
fof(f17340,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8721,f15697,f15697,f505])).
fof(f8721,plain,(
  ~sP8('F_0','T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f573,f508])).
fof(f19457,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17341,f15697])).
fof(f17341,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8722,f15697,f15697,f505])).
fof(f8722,plain,(
  ~sP8('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f582,f573,f508])).
fof(f19456,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17342,f15697])).
fof(f17342,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11417,f15697,f15697,f505])).
fof(f11417,plain,(
  ~sP8('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl6),
  inference(unit_resulting_resolution,[],[f327,f582,f580,f508])).
fof(f19455,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17343,f15697])).
fof(f17343,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12487,f15697,f15697,f505])).
fof(f12487,plain,(
  ~sP8('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f327,f582,f11256,f508])).
fof(f19454,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17344,f15697])).
fof(f17344,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12059,f15697,f15697,f505])).
fof(f12059,plain,(
  ~sP8('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f8562,f508])).
fof(f19453,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17345,f15697])).
fof(f17345,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11418,f15697,f15697,f505])).
fof(f11418,plain,(
  ~sP8('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl6),
  inference(unit_resulting_resolution,[],[f580,f582,f580,f508])).
fof(f19452,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17347,f15697])).
fof(f17347,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8669,f15697,f15697,f505])).
fof(f19451,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17349,f15697])).
fof(f17349,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11363,f15697,f15697,f505])).
fof(f19450,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17351,f15697])).
fof(f17351,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12425,f15697,f15697,f505])).
fof(f19449,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17353,f15697])).
fof(f17353,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11997,f15697,f15697,f505])).
fof(f19448,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17354,f327])).
fof(f17354,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1850,f15697,f15697,f505])).
fof(f1850,plain,(
  ~sP8('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f327,f327,f1707,f508])).
fof(f19447,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17355,f327])).
fof(f17355,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f762,f327,f15697,f505])).
fof(f19446,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17356,f327])).
fof(f17356,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1830,f327,f15697,f505])).
fof(f19445,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17362,f15697])).
fof(f17362,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8669,f327,f15697,f505])).
fof(f19444,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17363,f15697])).
fof(f17363,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11363,f327,f15697,f505])).
fof(f19443,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17364,f15697])).
fof(f17364,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12425,f327,f15697,f505])).
fof(f19442,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17365,f15697])).
fof(f17365,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11997,f327,f15697,f505])).
fof(f19441,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17366,f15697])).
fof(f17366,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1850,f327,f15697,f505])).
fof(f19440,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17371,f327])).
fof(f17371,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f759,f15697,f15697,f505])).
fof(f19439,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17372,f327])).
fof(f17372,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1827,f15697,f15697,f505])).
fof(f19438,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17373,f327])).
fof(f17373,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f762,f15697,f15697,f505])).
fof(f19437,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17374,f327])).
fof(f17374,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1830,f15697,f15697,f505])).
fof(f19436,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17386,f15697])).
fof(f17386,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8668,f15697,f15697,f505])).
fof(f19435,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17387,f15697])).
fof(f17387,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8669,f15697,f15697,f505])).
fof(f19434,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17388,f15697])).
fof(f17388,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11362,f15697,f15697,f505])).
fof(f19433,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17389,f15697])).
fof(f17389,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11363,f15697,f15697,f505])).
fof(f19432,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17390,f15697])).
fof(f17390,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12424,f15697,f15697,f505])).
fof(f19431,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17391,f15697])).
fof(f17391,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12425,f15697,f15697,f505])).
fof(f19430,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17392,f15697])).
fof(f17392,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11996,f15697,f15697,f505])).
fof(f19429,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17393,f15697])).
fof(f17393,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11997,f15697,f15697,f505])).
fof(f19428,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17394,f15697])).
fof(f17394,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1850,f15697,f15697,f505])).
fof(f19427,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17399,f15697])).
fof(f17399,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f759,f327,f15697,f505])).
fof(f19426,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17400,f15697])).
fof(f17400,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1827,f327,f15697,f505])).
fof(f19425,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17402,f15697])).
fof(f17402,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1830,f327,f15697,f505])).
fof(f19424,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17403,f15697])).
fof(f17403,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f759,f15697,f15697,f505])).
fof(f19423,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17404,f15697])).
fof(f17404,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1827,f15697,f15697,f505])).
fof(f19422,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17406,f15697])).
fof(f17406,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1830,f15697,f15697,f505])).
fof(f19421,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17407,f327])).
fof(f17407,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f1181,f15697,f15697,f505])).
fof(f19420,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17408,f15697])).
fof(f17408,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8717,f15697,f15697,f505])).
fof(f19419,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17409,f15697])).
fof(f17409,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11413,f15697,f15697,f505])).
fof(f19418,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17410,f15697])).
fof(f17410,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12483,f15697,f15697,f505])).
fof(f19417,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17411,f15697])).
fof(f17411,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12055,f15697,f15697,f505])).
fof(f19416,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17412,f327])).
fof(f17412,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8721,f15697,f15697,f505])).
fof(f19415,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17413,f15697])).
fof(f17413,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8722,f15697,f15697,f505])).
fof(f19414,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17414,f327])).
fof(f17414,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11417,f15697,f15697,f505])).
fof(f19413,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17415,f327])).
fof(f17415,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12487,f15697,f15697,f505])).
fof(f19412,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17416,f327])).
fof(f17416,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12059,f15697,f15697,f505])).
fof(f19411,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17417,f15697])).
fof(f17417,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11418,f15697,f15697,f505])).
fof(f19410,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17418,f15697])).
fof(f17418,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8668,f15697,f15697,f505])).
fof(f19409,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17420,f15697])).
fof(f17420,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11362,f15697,f15697,f505])).
fof(f19408,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17422,f15697])).
fof(f17422,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12424,f15697,f15697,f505])).
fof(f19407,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17424,f15697])).
fof(f17424,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11996,f15697,f15697,f505])).
fof(f19406,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17426,f327])).
fof(f17426,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1850,f15697,f15697,f505])).
fof(f19405,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17427,f327])).
fof(f17427,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f759,f327,f15697,f505])).
fof(f19404,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17428,f327])).
fof(f17428,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1827,f327,f15697,f505])).
fof(f19403,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17434,f15697])).
fof(f17434,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8668,f327,f15697,f505])).
fof(f19402,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17435,f15697])).
fof(f17435,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11362,f327,f15697,f505])).
fof(f19401,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17436,f15697])).
fof(f17436,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12424,f327,f15697,f505])).
fof(f19400,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17437,f15697])).
fof(f17437,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11996,f327,f15697,f505])).
fof(f19399,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17438,f15697])).
fof(f17438,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1850,f327,f15697,f505])).
fof(f19398,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17443,f327])).
fof(f17443,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f759,f15697,f15697,f505])).
fof(f19397,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17444,f327])).
fof(f17444,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1827,f15697,f15697,f505])).
fof(f19396,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17445,f327])).
fof(f17445,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f762,f15697,f15697,f505])).
fof(f19395,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17446,f327])).
fof(f17446,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1830,f15697,f15697,f505])).
fof(f19394,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17458,f15697])).
fof(f17458,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8668,f15697,f15697,f505])).
fof(f19393,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17459,f15697])).
fof(f17459,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8669,f15697,f15697,f505])).
fof(f19392,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17460,f15697])).
fof(f17460,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11362,f15697,f15697,f505])).
fof(f19391,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17461,f15697])).
fof(f17461,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11363,f15697,f15697,f505])).
fof(f19390,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17462,f15697])).
fof(f17462,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12424,f15697,f15697,f505])).
fof(f19389,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17463,f15697])).
fof(f17463,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12425,f15697,f15697,f505])).
fof(f19388,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17464,f15697])).
fof(f17464,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11996,f15697,f15697,f505])).
fof(f19387,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17465,f15697])).
fof(f17465,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11997,f15697,f15697,f505])).
fof(f19386,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17466,f15697])).
fof(f17466,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1850,f15697,f15697,f505])).
fof(f19385,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17472,f15697])).
fof(f17472,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1830,f327,f15697,f505])).
fof(f19384,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17473,f327])).
fof(f17473,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f1181,f327,f15697,f505])).
fof(f19383,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17474,f15697])).
fof(f17474,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8717,f327,f15697,f505])).
fof(f19382,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17475,f15697])).
fof(f17475,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11413,f327,f15697,f505])).
fof(f19381,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17476,f15697])).
fof(f17476,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12483,f327,f15697,f505])).
fof(f19380,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17477,f15697])).
fof(f17477,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12055,f327,f15697,f505])).
fof(f19379,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17482,f327])).
fof(f17482,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1850,f327,f15697,f505])).
fof(f19378,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17483,f15697])).
fof(f17483,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f759,f15697,f15697,f505])).
fof(f19377,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17484,f15697])).
fof(f17484,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1827,f15697,f15697,f505])).
fof(f19376,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17486,f15697])).
fof(f17486,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1830,f15697,f15697,f505])).
fof(f19375,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17487,f327])).
fof(f17487,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f1181,f15697,f15697,f505])).
fof(f19374,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17488,f15697])).
fof(f17488,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8717,f15697,f15697,f505])).
fof(f19373,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17489,f15697])).
fof(f17489,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11413,f15697,f15697,f505])).
fof(f19372,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17490,f15697])).
fof(f17490,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12483,f15697,f15697,f505])).
fof(f19371,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17491,f15697])).
fof(f17491,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12055,f15697,f15697,f505])).
fof(f19370,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17492,f327])).
fof(f17492,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8721,f15697,f15697,f505])).
fof(f19369,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17493,f15697])).
fof(f17493,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8722,f15697,f15697,f505])).
fof(f19368,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17494,f327])).
fof(f17494,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11417,f15697,f15697,f505])).
fof(f19367,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17495,f327])).
fof(f17495,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12487,f15697,f15697,f505])).
fof(f19366,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17496,f327])).
fof(f17496,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12059,f15697,f15697,f505])).
fof(f19365,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17497,f15697])).
fof(f17497,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11418,f15697,f15697,f505])).
fof(f19364,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17498,f15697])).
fof(f17498,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8668,f15697,f15697,f505])).
fof(f19363,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17500,f15697])).
fof(f17500,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11362,f15697,f15697,f505])).
fof(f19362,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17502,f15697])).
fof(f17502,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12424,f15697,f15697,f505])).
fof(f19361,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17504,f15697])).
fof(f17504,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11996,f15697,f15697,f505])).
fof(f19360,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17506,f327])).
fof(f17506,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1850,f15697,f15697,f505])).
fof(f19359,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17508,f15697])).
fof(f17508,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1827,f327,f15697,f505])).
fof(f19358,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17509,f327])).
fof(f17509,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f1181,f327,f15697,f505])).
fof(f19357,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17510,f15697])).
fof(f17510,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8721,f327,f15697,f505])).
fof(f19356,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17511,f15697])).
fof(f17511,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11417,f327,f15697,f505])).
fof(f19355,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17512,f15697])).
fof(f17512,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12487,f327,f15697,f505])).
fof(f19354,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17513,f15697])).
fof(f17513,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12059,f327,f15697,f505])).
fof(f19353,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17518,f327])).
fof(f17518,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1850,f327,f15697,f505])).
fof(f19352,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17524,f15697])).
fof(f17524,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1827,f15697,f15697,f505])).
fof(f19351,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17525,f15697])).
fof(f17525,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f762,f15697,f15697,f505])).
fof(f19350,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17526,f15697])).
fof(f17526,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1830,f15697,f15697,f505])).
fof(f19349,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17527,f327])).
fof(f17527,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f1181,f15697,f15697,f505])).
fof(f19348,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17528,f327])).
fof(f17528,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8717,f15697,f15697,f505])).
fof(f19347,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17529,f327])).
fof(f17529,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11413,f15697,f15697,f505])).
fof(f19346,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17530,f327])).
fof(f17530,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12483,f15697,f15697,f505])).
fof(f19345,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17531,f327])).
fof(f17531,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12055,f15697,f15697,f505])).
fof(f19344,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17532,f15697])).
fof(f17532,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8721,f15697,f15697,f505])).
fof(f19343,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17533,f15697])).
fof(f17533,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8722,f15697,f15697,f505])).
fof(f19342,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17534,f15697])).
fof(f17534,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11417,f15697,f15697,f505])).
fof(f19341,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17535,f15697])).
fof(f17535,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12487,f15697,f15697,f505])).
fof(f19340,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17536,f15697])).
fof(f17536,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12059,f15697,f15697,f505])).
fof(f19339,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17537,f15697])).
fof(f17537,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11418,f15697,f15697,f505])).
fof(f19338,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17539,f15697])).
fof(f17539,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8669,f15697,f15697,f505])).
fof(f19337,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17541,f15697])).
fof(f17541,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11363,f15697,f15697,f505])).
fof(f19336,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17543,f15697])).
fof(f17543,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12425,f15697,f15697,f505])).
fof(f19335,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17545,f15697])).
fof(f17545,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11997,f15697,f15697,f505])).
fof(f19334,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17546,f327])).
fof(f17546,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1850,f15697,f15697,f505])).
fof(f19333,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17552,f15697])).
fof(f17552,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1840,f582,f15697,f508])).
fof(f1840,plain,(
  ( ! [X0] : (sP8('F_0',X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl13),
  inference(unit_resulting_resolution,[],[f582,f1707,f503])).
fof(f19332,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17553,f15697])).
fof(f17553,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f768,f582,f15697,f508])).
fof(f19331,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17554,f15697])).
fof(f17554,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1845,f582,f15697,f508])).
fof(f1845,plain,(
  ( ! [X0] : (sP8('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f582,f1707,f504])).
fof(f19330,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17605,f15697])).
fof(f17605,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f765,f582,f15697,f508])).
fof(f19329,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17606,f15697])).
fof(f17606,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1840,f582,f15697,f508])).
fof(f19328,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17608,f15697])).
fof(f17608,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1845,f582,f15697,f508])).
fof(f19327,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17691,f18763])).
fof(f18763,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f616,f15697,f564])).
fof(f17691,plain,(
  ( ! [X0,X1] : (~sP10('F_0',X0,X1)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f509])).
fof(f19326,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17725,f18763])).
fof(f17725,plain,(
  ( ! [X0,X1] : (~sP10('F_0',X0,X1)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f510])).
fof(f19325,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17791,f17914])).
fof(f17914,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f518])).
fof(f17791,plain,(
  ~sP9('F_0','T_0','F_0') | $spl82),
  inference(unit_resulting_resolution,[],[f776,f582,f15697,f513])).
fof(f19324,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17793,f17914])).
fof(f17793,plain,(
  ~sP9('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8731,f582,f15697,f513])).
fof(f8731,plain,(
  ( ! [X0] : (~sP10('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f582,f573,f510])).
fof(f19323,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17795,f17914])).
fof(f17795,plain,(
  ~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0') | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11427,f582,f15697,f513])).
fof(f11427,plain,(
  ( ! [X0] : (~sP10('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f582,f580,f510])).
fof(f19322,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17797,f17914])).
fof(f17797,plain,(
  ~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12498,f582,f15697,f513])).
fof(f12498,plain,(
  ( ! [X0] : (~sP10('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f11256,f510])).
fof(f19321,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17799,f17914])).
fof(f17799,plain,(
  ~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12070,f582,f15697,f513])).
fof(f12070,plain,(
  ( ! [X0] : (~sP10('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f582,f8562,f510])).
fof(f19320,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17805,f17914])).
fof(f17805,plain,(
  ~sP9('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','F_0') | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1852,f582,f15697,f513])).
fof(f1852,plain,(
  ( ! [X0] : (~sP10('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0',X0)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f327,f1707,f510])).
fof(f19319,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17869,f17914])).
fof(f17869,plain,(
  ~sP9('F_0','F_0','T_0') | $spl82),
  inference(unit_resulting_resolution,[],[f775,f582,f15697,f514])).
fof(f19318,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17870,f17914])).
fof(f17870,plain,(
  ~sP9('F_0','F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8725,f582,f15697,f514])).
fof(f8725,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f582,f573,f509])).
fof(f19317,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17871,f17914])).
fof(f17871,plain,(
  ~sP9('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11421,f582,f15697,f514])).
fof(f11421,plain,(
  ( ! [X0] : (~sP10('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f582,f580,f509])).
fof(f19316,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17872,f17914])).
fof(f17872,plain,(
  ~sP9('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12491,f582,f15697,f514])).
fof(f12491,plain,(
  ( ! [X0] : (~sP10('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f11256,f509])).
fof(f19315,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17873,f17914])).
fof(f17873,plain,(
  ~sP9('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12063,f582,f15697,f514])).
fof(f12063,plain,(
  ( ! [X0] : (~sP10('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f582,f8562,f509])).
fof(f19314,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17878,f17914])).
fof(f17878,plain,(
  ~sP9('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','T_0') | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1851,f582,f15697,f514])).
fof(f1851,plain,(
  ( ! [X0] : (~sP10('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'T_0')) ) | $spl13),
  inference(unit_resulting_resolution,[],[f327,f1707,f509])).
fof(f19313,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17910,f17914])).
fof(f17910,plain,(
  ( ! [X0] : (~sP9(X0,'F_0','F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f582,f15697,f516])).
fof(f19312,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17916,f327])).
fof(f17916,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f1271,f15697,f521])).
fof(f19311,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17917,f327])).
fof(f17917,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1904,f15697,f521])).
fof(f1904,plain,(
  ~sP10('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl13),
  inference(unit_resulting_resolution,[],[f582,f327,f1707,f523])).
fof(f19310,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17918,f327])).
fof(f17918,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1909,f15697,f521])).
fof(f1909,plain,(
  ~sP10('T_0','F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl13),
  inference(unit_resulting_resolution,[],[f327,f582,f1707,f523])).
fof(f19309,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17919,f327])).
fof(f17919,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1910,f15697,f521])).
fof(f1910,plain,(
  ~sP10('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl13),
  inference(unit_resulting_resolution,[],[f327,f1707,f1707,f523])).
fof(f19308,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17930,f15697])).
fof(f17930,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8762,f15697,f521])).
fof(f8762,plain,(
  ~sP10('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f582,f582,f573,f523])).
fof(f19307,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17931,f15697])).
fof(f17931,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11461,f15697,f521])).
fof(f11461,plain,(
  ~sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0','F_0') | $spl6),
  inference(unit_resulting_resolution,[],[f582,f582,f580,f523])).
fof(f19306,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17932,f15697])).
fof(f17932,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12539,f15697,f521])).
fof(f12539,plain,(
  ~sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f582,f582,f11256,f523])).
fof(f19305,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17933,f15697])).
fof(f17933,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12111,f15697,f521])).
fof(f12111,plain,(
  ~sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0','F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f582,f582,f8562,f523])).
fof(f19304,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17934,f15697])).
fof(f17934,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1851,f15697,f521])).
fof(f19303,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17935,f15697])).
fof(f17935,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1852,f15697,f521])).
fof(f19302,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17942,f15697])).
fof(f17942,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1909,f15697,f521])).
fof(f19301,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17943,f15697])).
fof(f17943,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1910,f15697,f521])).
fof(f19300,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17944,f327])).
fof(f17944,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f775,f15697,f521])).
fof(f19299,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17945,f15697])).
fof(f17945,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8725,f15697,f521])).
fof(f19298,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17946,f15697])).
fof(f17946,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11421,f15697,f521])).
fof(f19297,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17947,f15697])).
fof(f17947,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12491,f15697,f521])).
fof(f19296,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17948,f15697])).
fof(f17948,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12063,f15697,f521])).
fof(f19295,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17949,f15697])).
fof(f17949,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f776,f15697,f521])).
fof(f19294,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17950,f15697])).
fof(f17950,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8731,f15697,f521])).
fof(f19293,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17951,f15697])).
fof(f17951,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11427,f15697,f521])).
fof(f19292,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17952,f15697])).
fof(f17952,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12498,f15697,f521])).
fof(f19291,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17953,f15697])).
fof(f17953,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12070,f15697,f521])).
fof(f19290,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17958,f327])).
fof(f17958,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1851,f15697,f521])).
fof(f19289,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17959,f15697])).
fof(f17959,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1852,f15697,f521])).
fof(f19288,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17964,f327])).
fof(f17964,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f1271,f15697,f522])).
fof(f19287,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17965,f327])).
fof(f17965,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1904,f15697,f522])).
fof(f19286,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17966,f327])).
fof(f17966,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1909,f15697,f522])).
fof(f19285,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17967,f327])).
fof(f17967,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1910,f15697,f522])).
fof(f19284,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17978,f15697])).
fof(f17978,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8762,f15697,f522])).
fof(f19283,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17979,f15697])).
fof(f17979,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11461,f15697,f522])).
fof(f19282,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17980,f15697])).
fof(f17980,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12539,f15697,f522])).
fof(f19281,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17981,f15697])).
fof(f17981,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12111,f15697,f522])).
fof(f19280,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17982,f15697])).
fof(f17982,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1851,f15697,f522])).
fof(f19279,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17983,f15697])).
fof(f17983,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1852,f15697,f522])).
fof(f19278,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17989,f15697])).
fof(f17989,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1904,f15697,f522])).
fof(f19277,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f17991,f15697])).
fof(f17991,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1910,f15697,f522])).
fof(f19276,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17992,f15697])).
fof(f17992,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f775,f15697,f522])).
fof(f19275,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17993,f15697])).
fof(f17993,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8725,f15697,f522])).
fof(f19274,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17994,f15697])).
fof(f17994,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11421,f15697,f522])).
fof(f19273,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f17995,f15697])).
fof(f17995,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12491,f15697,f522])).
fof(f19272,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17996,f15697])).
fof(f17996,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12063,f15697,f522])).
fof(f19271,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f17997,f327])).
fof(f17997,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f776,f15697,f522])).
fof(f19270,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f17998,f15697])).
fof(f17998,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8731,f15697,f522])).
fof(f19269,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f17999,f15697])).
fof(f17999,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11427,f15697,f522])).
fof(f19268,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f18000,f15697])).
fof(f18000,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12498,f15697,f522])).
fof(f19267,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18001,f15697])).
fof(f18001,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12070,f15697,f522])).
fof(f19266,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18006,f15697])).
fof(f18006,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1851,f15697,f522])).
fof(f19265,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18007,f327])).
fof(f18007,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1852,f15697,f522])).
fof(f19264,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18075,f18715])).
fof(f18715,plain,(
  ( ! [X0,X1] : (sP12('F_0',X0,X1)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f622,f15697,f563])).
fof(f18075,plain,(
  ( ! [X0,X1] : (~sP12('F_0',X0,X1)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f527])).
fof(f19263,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18078,f15697])).
fof(f18078,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f622,f1916,f15697,f530])).
fof(f1916,plain,(
  ( ! [X0] : (~sP12('F_0',X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl13),
  inference(unit_resulting_resolution,[],[f582,f1707,f526])).
fof(f19262,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18079,f15697])).
fof(f18079,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f622,f801,f15697,f530])).
fof(f19261,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18080,f15697])).
fof(f18080,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f620,f801,f15697,f530])).
fof(f19260,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18082,f15697])).
fof(f18082,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1936,f801,f15697,f530])).
fof(f1936,plain,(
  ( ! [X0,X1] : (sP11(X0,X1,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl13),
  inference(unit_resulting_resolution,[],[f1707,f534])).
fof(f19259,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18083,f15697])).
fof(f18083,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f622,f8763,f15697,f530])).
fof(f8763,plain,(
  ( ! [X0] : (~sP12('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f582,f573,f527])).
fof(f19258,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18084,f15697])).
fof(f18084,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8804,f8763,f15697,f530])).
fof(f8804,plain,(
  ( ! [X0,X1] : (sP11(X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f535])).
fof(f19257,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18086,f15697])).
fof(f18086,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1936,f8763,f15697,f530])).
fof(f19256,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f18087,f15697])).
fof(f18087,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f622,f11462,f15697,f530])).
fof(f11462,plain,(
  ( ! [X0] : (~sP12('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f582,f580,f527])).
fof(f19255,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f18088,f15697])).
fof(f18088,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11506,f11462,f15697,f530])).
fof(f11506,plain,(
  ( ! [X0,X1] : (sP11(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f580,f535])).
fof(f19254,plain,(
  $false | ($spl6 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18090,f15697])).
fof(f18090,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl6 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1936,f11462,f15697,f530])).
fof(f19253,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f18091,f15697])).
fof(f18091,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f622,f12540,f15697,f530])).
fof(f12540,plain,(
  ( ! [X0] : (~sP12('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f11256,f527])).
fof(f19252,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f18092,f15697])).
fof(f18092,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12593,f12540,f15697,f530])).
fof(f12593,plain,(
  ( ! [X0,X1] : (sP11(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f11256,f535])).
fof(f19251,plain,(
  $false | ($spl2 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18094,f15697])).
fof(f18094,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl2 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1936,f12540,f15697,f530])).
fof(f19250,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18095,f15697])).
fof(f18095,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f622,f12112,f15697,f530])).
fof(f12112,plain,(
  ( ! [X0] : (~sP12('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f582,f8562,f527])).
fof(f19249,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18096,f15697])).
fof(f18096,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12165,f12112,f15697,f530])).
fof(f12165,plain,(
  ( ! [X0,X1] : (sP11(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f8562,f535])).
fof(f19248,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18098,f15697])).
fof(f18098,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1936,f12112,f15697,f530])).
fof(f19247,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18100,f15697])).
fof(f18100,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1937,f1917,f15697,f530])).
fof(f1917,plain,(
  ( ! [X0] : (~sP12('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl13),
  inference(unit_resulting_resolution,[],[f1707,f1707,f526])).
fof(f1937,plain,(
  ( ! [X0,X1] : (sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,X1)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f1707,f536])).
fof(f19246,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18101,f15697])).
fof(f18101,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1937,f1922,f15697,f530])).
fof(f1922,plain,(
  ( ! [X0] : (~sP12('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0',X0)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f327,f1707,f527])).
fof(f19245,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18102,f15697])).
fof(f18102,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f620,f1922,f15697,f530])).
fof(f19244,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18104,f15697])).
fof(f18104,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1936,f1922,f15697,f530])).
fof(f19243,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18115,f18439])).
fof(f18439,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f535])).
fof(f18115,plain,(
  ~sP11('F_0','T_0','T_0') | $spl82),
  inference(unit_resulting_resolution,[],[f801,f327,f15697,f530])).
fof(f19242,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18117,f18439])).
fof(f18117,plain,(
  ~sP11('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'T_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8763,f327,f15697,f530])).
fof(f19241,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f18119,f18439])).
fof(f18119,plain,(
  ~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'T_0') | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11462,f327,f15697,f530])).
fof(f19240,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f18121,f18439])).
fof(f18121,plain,(
  ~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12540,f327,f15697,f530])).
fof(f19239,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18123,f18439])).
fof(f18123,plain,(
  ~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'T_0') | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12112,f327,f15697,f530])).
fof(f19238,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18129,f18439])).
fof(f18129,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','T_0') | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1922,f327,f15697,f530])).
fof(f19237,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18147,f18439])).
fof(f18147,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f795,f15697,f15697,f530])).
fof(f19236,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18149,f18439])).
fof(f18149,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1916,f15697,f15697,f530])).
fof(f19235,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18151,f18439])).
fof(f18151,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f801,f15697,f15697,f530])).
fof(f19234,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18153,f18439])).
fof(f18153,plain,(
  ( ! [X0] : (~sP11('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0)) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8763,f15697,f15697,f530])).
fof(f19233,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f18155,f18439])).
fof(f18155,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0)) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11462,f15697,f15697,f530])).
fof(f19232,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f18157,f18439])).
fof(f18157,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12540,f15697,f15697,f530])).
fof(f19231,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18159,f18439])).
fof(f18159,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0)) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12112,f15697,f15697,f530])).
fof(f19230,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18167,f18439])).
fof(f18167,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'F_0')) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f15697,f15697,f530])).
fof(f1914,plain,(
  ( ! [X0] : (~sP12('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'F_0')) ) | $spl13),
  inference(unit_resulting_resolution,[],[f582,f1707,f526])).
fof(f19229,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18169,f18439])).
fof(f18169,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1917,f15697,f15697,f530])).
fof(f19228,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18171,f18439])).
fof(f18171,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0',X0)) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1922,f15697,f15697,f530])).
fof(f19227,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18179,f15697])).
fof(f18179,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f618,f15697,f530])).
fof(f19226,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18180,f15697])).
fof(f18180,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1922,f618,f15697,f530])).
fof(f19225,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18187,f15697])).
fof(f18187,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1917,f1936,f15697,f530])).
fof(f19224,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18188,f15697])).
fof(f18188,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1922,f1936,f15697,f530])).
fof(f19223,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18192,f15697])).
fof(f18192,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1922,f620,f15697,f530])).
fof(f19222,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18193,f15697])).
fof(f18193,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f620,f15697,f530])).
fof(f19221,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18194,f15697])).
fof(f18194,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1917,f620,f15697,f530])).
fof(f19220,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18197,f15697])).
fof(f18197,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f636,f15697,f530])).
fof(f19219,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18198,f15697])).
fof(f18198,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1917,f636,f15697,f530])).
fof(f19218,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18202,f15697])).
fof(f18202,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f8804,f15697,f530])).
fof(f19217,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18203,f15697])).
fof(f18203,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1917,f8804,f15697,f530])).
fof(f19216,plain,(
  $false | ($spl6 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18207,f15697])).
fof(f18207,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl6 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f11506,f15697,f530])).
fof(f19215,plain,(
  $false | ($spl6 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18208,f15697])).
fof(f18208,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl6 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1917,f11506,f15697,f530])).
fof(f19214,plain,(
  $false | ($spl2 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18212,f15697])).
fof(f18212,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl2 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f12593,f15697,f530])).
fof(f19213,plain,(
  $false | ($spl2 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18213,f15697])).
fof(f18213,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl2 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1917,f12593,f15697,f530])).
fof(f19212,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18217,f15697])).
fof(f18217,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f12165,f15697,f530])).
fof(f19211,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18218,f15697])).
fof(f18218,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1917,f12165,f15697,f530])).
fof(f19210,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18220,f15697])).
fof(f18220,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f1937,f15697,f530])).
fof(f19209,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18241,f18439])).
fof(f18241,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f795,f15697,f15697,f530])).
fof(f19208,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18242,f18439])).
fof(f18242,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1916,f15697,f15697,f530])).
fof(f19207,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18243,f18439])).
fof(f18243,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f801,f15697,f15697,f530])).
fof(f19206,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18244,f18439])).
fof(f18244,plain,(
  ( ! [X0] : (~sP11('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0)) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8763,f15697,f15697,f530])).
fof(f19205,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f18245,f18439])).
fof(f18245,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0)) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11462,f15697,f15697,f530])).
fof(f19204,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f18246,f18439])).
fof(f18246,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12540,f15697,f15697,f530])).
fof(f19203,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18247,f18439])).
fof(f18247,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0)) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12112,f15697,f15697,f530])).
fof(f19202,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18254,f18439])).
fof(f18254,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'F_0')) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f15697,f15697,f530])).
fof(f19201,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18255,f18439])).
fof(f18255,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1917,f15697,f15697,f530])).
fof(f19200,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18256,f18439])).
fof(f18256,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0',X0)) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1922,f15697,f15697,f530])).
fof(f19199,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18314,f18439])).
fof(f18314,plain,(
  ~sP11('F_0','F_0','F_0') | $spl82),
  inference(unit_resulting_resolution,[],[f795,f582,f15697,f531])).
fof(f19198,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18315,f18439])).
fof(f18315,plain,(
  ~sP11('F_0','F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1916,f582,f15697,f531])).
fof(f19197,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18322,f18439])).
fof(f18322,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','F_0') | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f582,f15697,f531])).
fof(f19196,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18323,f18439])).
fof(f18323,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1917,f582,f15697,f531])).
fof(f19195,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18369,f18439])).
fof(f18369,plain,(
  ~sP11('F_0','F_0','F_0') | $spl82),
  inference(unit_resulting_resolution,[],[f795,f582,f15697,f532])).
fof(f19194,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18370,f18439])).
fof(f18370,plain,(
  ~sP11('F_0','F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1916,f582,f15697,f532])).
fof(f19193,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18377,f18439])).
fof(f18377,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','F_0') | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f582,f15697,f532])).
fof(f19192,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18378,f18439])).
fof(f18378,plain,(
  ~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1917,f582,f15697,f532])).
fof(f19191,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18406,f15697])).
fof(f18406,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1936,f582,f15697,f533])).
fof(f19190,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18407,f15697])).
fof(f18407,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f622,f582,f15697,f533])).
fof(f19189,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18408,f15697])).
fof(f18408,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1937,f582,f15697,f533])).
fof(f19188,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18409,f18439])).
fof(f18409,plain,(
  ( ! [X0] : (~sP11(X0,'F_0','T_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f327,f582,f15697,f533])).
fof(f19187,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18411,f18439])).
fof(f18411,plain,(
  ( ! [X0,X1] : (~sP11(X0,'F_0',X1)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f582,f15697,f533])).
fof(f19186,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18431,f15697])).
fof(f18431,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f618,f582,f15697,f533])).
fof(f19185,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18432,f15697])).
fof(f18432,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1936,f582,f15697,f533])).
fof(f19184,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18434,f15697])).
fof(f18434,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1937,f582,f15697,f533])).
fof(f19183,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18435,f18439])).
fof(f18435,plain,(
  ( ! [X0] : (~sP11('T_0','F_0',X0)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f327,f582,f15697,f533])).
fof(f19182,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18437,f18439])).
fof(f18437,plain,(
  ( ! [X0,X1] : (~sP11(X0,'F_0',X1)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f582,f15697,f533])).
fof(f19181,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18489,f327])).
fof(f18489,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f1388,f15697,f539])).
fof(f19180,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18490,f327])).
fof(f18490,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1945,f15697,f539])).
fof(f1945,plain,(
  ~sP12('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f327,f327,f1707,f540])).
fof(f19179,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18491,f327])).
fof(f18491,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8816,f15697,f539])).
fof(f8816,plain,(
  ~sP12('T_0','F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f573,f540])).
fof(f19178,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f18492,f327])).
fof(f18492,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11519,f15697,f539])).
fof(f11519,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl6),
  inference(unit_resulting_resolution,[],[f327,f582,f580,f540])).
fof(f19177,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f18493,f327])).
fof(f18493,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12607,f15697,f539])).
fof(f12607,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f327,f582,f11256,f540])).
fof(f19176,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18494,f327])).
fof(f18494,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12179,f15697,f539])).
fof(f12179,plain,(
  ~sP12('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f8562,f540])).
fof(f19175,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18502,f15697])).
fof(f18502,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8821,f15697,f539])).
fof(f8821,plain,(
  ~sP12('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f573,f540])).
fof(f19174,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18503,f15697])).
fof(f18503,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8822,f15697,f539])).
fof(f8822,plain,(
  ~sP12('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f582,f573,f540])).
fof(f19173,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f18504,f15697])).
fof(f18504,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11524,f15697,f539])).
fof(f11524,plain,(
  ~sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0','T_0') | $spl6),
  inference(unit_resulting_resolution,[],[f327,f582,f580,f540])).
fof(f19172,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f18505,f15697])).
fof(f18505,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12613,f15697,f539])).
fof(f12613,plain,(
  ~sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f327,f582,f11256,f540])).
fof(f19171,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18506,f15697])).
fof(f18506,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12185,f15697,f539])).
fof(f12185,plain,(
  ~sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),'F_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f327,f582,f8562,f540])).
fof(f19170,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f18507,f15697])).
fof(f18507,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11525,f15697,f539])).
fof(f11525,plain,(
  ~sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl6),
  inference(unit_resulting_resolution,[],[f580,f582,f580,f540])).
fof(f19169,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18508,f15697])).
fof(f18508,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f15697,f539])).
fof(f19168,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18509,f15697])).
fof(f18509,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1917,f15697,f539])).
fof(f19167,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18510,f15697])).
fof(f18510,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1922,f15697,f539])).
fof(f19166,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18516,f15697])).
fof(f18516,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1945,f15697,f539])).
fof(f19165,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18521,f15697])).
fof(f18521,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f795,f15697,f539])).
fof(f19164,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18522,f15697])).
fof(f18522,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1916,f15697,f539])).
fof(f19163,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18523,f327])).
fof(f18523,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f801,f15697,f539])).
fof(f19162,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18524,f15697])).
fof(f18524,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8763,f15697,f539])).
fof(f19161,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f18525,f15697])).
fof(f18525,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11462,f15697,f539])).
fof(f19160,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f18526,f15697])).
fof(f18526,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12540,f15697,f539])).
fof(f19159,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18527,f15697])).
fof(f18527,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12112,f15697,f539])).
fof(f19158,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18534,f15697])).
fof(f18534,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f15697,f539])).
fof(f19157,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18535,f15697])).
fof(f18535,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1917,f15697,f539])).
fof(f19156,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18536,f327])).
fof(f18536,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1922,f15697,f539])).
fof(f19155,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18541,f327])).
fof(f18541,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f805,f582,f15697,f540])).
fof(f19154,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18542,f327])).
fof(f18542,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1943,f582,f15697,f540])).
fof(f1943,plain,(
  ( ! [X0] : (sP12('T_0',X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl13),
  inference(unit_resulting_resolution,[],[f327,f1707,f538])).
fof(f19153,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18548,f15697])).
fof(f18548,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8826,f582,f15697,f540])).
fof(f8826,plain,(
  ( ! [X0] : (sP12('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f618,f573,f563])).
fof(f19152,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18549,f15697])).
fof(f18549,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f8827,f582,f15697,f540])).
fof(f8827,plain,(
  ( ! [X0] : (sP12('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl0 | $spl13)),
  inference(unit_resulting_resolution,[],[f1936,f573,f563])).
fof(f19151,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f18550,f15697])).
fof(f18550,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11529,f582,f15697,f540])).
fof(f11529,plain,(
  ( ! [X0] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0,'F_0')) ) | $spl6),
  inference(unit_resulting_resolution,[],[f618,f580,f563])).
fof(f19150,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f18551,f15697])).
fof(f18551,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12618,f582,f15697,f540])).
fof(f12618,plain,(
  ( ! [X0] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f618,f11256,f563])).
fof(f19149,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18552,f15697])).
fof(f18552,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12190,f582,f15697,f540])).
fof(f12190,plain,(
  ( ! [X0] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f618,f8562,f563])).
fof(f19148,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18553,f15697])).
fof(f18553,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1941,f582,f15697,f540])).
fof(f1941,plain,(
  sP12('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f582,f327,f1707,f537])).
fof(f19147,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18554,f15697])).
fof(f18554,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f8788,f582,f15697,f540])).
fof(f8788,plain,(
  sP12('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl13)),
  inference(unit_resulting_resolution,[],[f582,f1937,f573,f532])).
fof(f19146,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18596,f15697])).
fof(f18596,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1943,f582,f15697,f540])).
fof(f19145,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18597,f327])).
fof(f18597,plain,(
  ~p(s(bool,'T_0')) | $spl82),
  inference(unit_resulting_resolution,[],[f1360,f582,f15697,f540])).
fof(f19144,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18598,f15697])).
fof(f18598,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8805,f582,f15697,f540])).
fof(f8805,plain,(
  sP12('F_0','F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f582,f582,f573,f537])).
fof(f19143,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f18599,f15697])).
fof(f18599,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11507,f582,f15697,f540])).
fof(f11507,plain,(
  sP12('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl6),
  inference(unit_resulting_resolution,[],[f582,f582,f580,f537])).
fof(f19142,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f18600,f15697])).
fof(f18600,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12594,f582,f15697,f540])).
fof(f12594,plain,(
  sP12('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | $spl2),
  inference(unit_resulting_resolution,[],[f582,f582,f11256,f537])).
fof(f19141,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18601,f15697])).
fof(f18601,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12166,f582,f15697,f540])).
fof(f12166,plain,(
  sP12('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14)))) | $spl0),
  inference(unit_resulting_resolution,[],[f582,f582,f8562,f537])).
fof(f19140,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18603,f15697])).
fof(f18603,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f8827,f582,f15697,f540])).
fof(f19139,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18607,f327])).
fof(f18607,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1941,f582,f15697,f540])).
fof(f19138,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18608,f15697])).
fof(f18608,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f8788,f582,f15697,f540])).
fof(f19137,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18657,f18999])).
fof(f18657,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f562])).
fof(f19136,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18683,f618])).
fof(f18683,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f795,f15697,f563])).
fof(f19135,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18684,f1936])).
fof(f18684,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1916,f15697,f563])).
fof(f19134,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18685,f620])).
fof(f18685,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f801,f15697,f563])).
fof(f19133,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18686,f8804])).
fof(f18686,plain,(
  ( ! [X0] : (~sP11('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0)) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8763,f15697,f563])).
fof(f19132,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f18687,f11506])).
fof(f18687,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0)) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11462,f15697,f563])).
fof(f19131,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f18688,f12593])).
fof(f18688,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12540,f15697,f563])).
fof(f19130,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18689,f12165])).
fof(f18689,plain,(
  ( ! [X0] : (~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0)) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12112,f15697,f563])).
fof(f19129,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18696,f618])).
fof(f18696,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'F_0')) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f15697,f563])).
fof(f19128,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18697,f1936])).
fof(f18697,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1917,f15697,f563])).
fof(f19127,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18698,f620])).
fof(f18698,plain,(
  ( ! [X0] : (~sP11('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0',X0)) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1922,f15697,f563])).
fof(f19126,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18723,f612])).
fof(f18723,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) ) | $spl82),
  inference(unit_resulting_resolution,[],[f775,f15697,f564])).
fof(f19125,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18724,f8748])).
fof(f8748,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f517])).
fof(f18724,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8725,f15697,f564])).
fof(f19124,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f18725,f11445])).
fof(f11445,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f580,f517])).
fof(f18725,plain,(
  ( ! [X0] : (~sP9('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11421,f15697,f564])).
fof(f19123,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f18726,f12521])).
fof(f12521,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f11256,f517])).
fof(f18726,plain,(
  ( ! [X0] : (~sP9('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12491,f15697,f564])).
fof(f19122,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18727,f12093])).
fof(f12093,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f8562,f517])).
fof(f18727,plain,(
  ( ! [X0] : (~sP9('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))))) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12063,f15697,f564])).
fof(f19121,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18728,f614])).
fof(f18728,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) ) | $spl82),
  inference(unit_resulting_resolution,[],[f776,f15697,f564])).
fof(f19120,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18729,f8749])).
fof(f8749,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f573,f518])).
fof(f18729,plain,(
  ( ! [X0] : (~sP9('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0)) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8731,f15697,f564])).
fof(f19119,plain,(
  $false | ($spl6 | $spl82)),
  inference(subsumption_resolution,[],[f18730,f11446])).
fof(f11446,plain,(
  ( ! [X0,X1] : (sP9(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X1)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f580,f518])).
fof(f18730,plain,(
  ( ! [X0] : (~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X0)) ) | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11427,f15697,f564])).
fof(f19118,plain,(
  $false | ($spl2 | $spl82)),
  inference(subsumption_resolution,[],[f18731,f12522])).
fof(f12522,plain,(
  ( ! [X0,X1] : (sP9(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f11256,f518])).
fof(f18731,plain,(
  ( ! [X0] : (~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12498,f15697,f564])).
fof(f19117,plain,(
  $false | ($spl0 | $spl82)),
  inference(subsumption_resolution,[],[f18732,f12094])).
fof(f12094,plain,(
  ( ! [X0,X1] : (sP9(X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f8562,f518])).
fof(f18732,plain,(
  ( ! [X0] : (~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK18(sK16,sK14))),X0)) ) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12070,f15697,f564])).
fof(f19116,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18737,f612])).
fof(f18737,plain,(
  ( ! [X0] : (~sP9('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,'T_0')) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1851,f15697,f564])).
fof(f19115,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18738,f614])).
fof(f18738,plain,(
  ( ! [X0] : (~sP9('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0',X0)) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1852,f15697,f564])).
fof(f19114,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18785,f15697])).
fof(f18785,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f1775,f15697,f1476])).
fof(f19113,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18798,f15697])).
fof(f18798,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f592,f15697,f1476])).
fof(f19112,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18802,f15697])).
fof(f18802,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f1775,f15697,f1476])).
fof(f19111,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18815,f15697])).
fof(f18815,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f1782,f15697,f1477])).
fof(f19110,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18828,f15697])).
fof(f18828,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f598,f15697,f1477])).
fof(f19109,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18832,f15697])).
fof(f18832,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f1782,f15697,f1477])).
fof(f19108,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18845,f15697])).
fof(f18845,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f1839,f15697,f1478])).
fof(f19107,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18858,f15697])).
fof(f18858,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f610,f15697,f1478])).
fof(f19106,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18862,f15697])).
fof(f18862,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f1839,f15697,f1478])).
fof(f19105,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18869,f15697])).
fof(f18869,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f1845,f15697,f1479])).
fof(f19104,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18877,f15697])).
fof(f18877,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f768,f15697,f1479])).
fof(f19103,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18880,f15697])).
fof(f18880,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f1845,f15697,f1479])).
fof(f19102,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18913,f15697])).
fof(f18913,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f1937,f15697,f1483])).
fof(f19101,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18924,f15697])).
fof(f18924,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f582,f622,f15697,f1483])).
fof(f19100,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18929,f15697])).
fof(f18929,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f1937,f15697,f1483])).
fof(f19099,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18958,f15697])).
fof(f18958,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f15697,f373])).
fof(f19098,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18959,f15697])).
fof(f18959,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f639,f15697,f15697,f373])).
fof(f19097,plain,(
  $false | ($spl5 | $spl82)),
  inference(subsumption_resolution,[],[f18960,f15697])).
fof(f18960,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK15))))) ) | ($spl5 | $spl82)),
  inference(unit_resulting_resolution,[],[f578,f15697,f15697,f373])).
fof(f19096,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18961,f15697])).
fof(f18961,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f15697,f373])).
fof(f19095,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18962,f15697])).
fof(f18962,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f639,f15697,f15697,f373])).
fof(f19094,plain,(
  $false | ($spl5 | $spl82)),
  inference(subsumption_resolution,[],[f18963,f15697])).
fof(f18963,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK17))))) ) | ($spl5 | $spl82)),
  inference(unit_resulting_resolution,[],[f578,f15697,f15697,f373])).
fof(f19093,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18964,f15697])).
fof(f18964,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f15697,f373])).
fof(f19092,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18965,f15697])).
fof(f18965,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f639,f15697,f15697,f373])).
fof(f19091,plain,(
  $false | ($spl5 | $spl82)),
  inference(subsumption_resolution,[],[f18966,f15697])).
fof(f18966,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK15))))) ) | ($spl5 | $spl82)),
  inference(unit_resulting_resolution,[],[f578,f15697,f15697,f373])).
fof(f19090,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18967,f15697])).
fof(f18967,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f15697,f373])).
fof(f19089,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18968,f15697])).
fof(f18968,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f639,f15697,f15697,f373])).
fof(f19088,plain,(
  $false | ($spl5 | $spl82)),
  inference(subsumption_resolution,[],[f18969,f15697])).
fof(f18969,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK17))))) ) | ($spl5 | $spl82)),
  inference(unit_resulting_resolution,[],[f578,f15697,f15697,f373])).
fof(f19087,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18973,f15697])).
fof(f18973,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f15697,f373])).
fof(f19086,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18974,f15697])).
fof(f18974,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f639,f15697,f15697,f373])).
fof(f19085,plain,(
  $false | ($spl5 | $spl82)),
  inference(subsumption_resolution,[],[f18975,f15697])).
fof(f18975,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK15))))) ) | ($spl5 | $spl82)),
  inference(unit_resulting_resolution,[],[f578,f15697,f15697,f373])).
fof(f19084,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18976,f15697])).
fof(f18976,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f15697,f373])).
fof(f19083,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18977,f15697])).
fof(f18977,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f639,f15697,f15697,f373])).
fof(f19082,plain,(
  $false | ($spl5 | $spl82)),
  inference(subsumption_resolution,[],[f18978,f15697])).
fof(f18978,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK17))))) ) | ($spl5 | $spl82)),
  inference(unit_resulting_resolution,[],[f578,f15697,f15697,f373])).
fof(f19081,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18979,f15697])).
fof(f18979,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f15697,f373])).
fof(f19080,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18980,f15697])).
fof(f18980,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f639,f15697,f15697,f373])).
fof(f19079,plain,(
  $false | ($spl5 | $spl82)),
  inference(subsumption_resolution,[],[f18981,f15697])).
fof(f18981,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),X0))))) ) | ($spl5 | $spl82)),
  inference(unit_resulting_resolution,[],[f578,f15697,f15697,f373])).
fof(f19078,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18982,f15697])).
fof(f18982,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f15697,f373])).
fof(f19077,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18983,f15697])).
fof(f18983,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f639,f15697,f15697,f373])).
fof(f19076,plain,(
  $false | ($spl5 | $spl82)),
  inference(subsumption_resolution,[],[f18984,f15697])).
fof(f18984,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),X0))))) ) | ($spl5 | $spl82)),
  inference(unit_resulting_resolution,[],[f578,f15697,f15697,f373])).
fof(f19075,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18987,f15697])).
fof(f18987,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f15697,f373])).
fof(f19074,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18988,f15697])).
fof(f18988,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f639,f15697,f15697,f373])).
fof(f19073,plain,(
  $false | ($spl5 | $spl82)),
  inference(subsumption_resolution,[],[f18989,f15697])).
fof(f18989,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK15))))) ) | ($spl5 | $spl82)),
  inference(unit_resulting_resolution,[],[f578,f15697,f15697,f373])).
fof(f19072,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18990,f15697])).
fof(f18990,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f15697,f373])).
fof(f19071,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18991,f15697])).
fof(f18991,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0'))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f639,f15697,f15697,f373])).
fof(f19070,plain,(
  $false | ($spl5 | $spl82)),
  inference(subsumption_resolution,[],[f18992,f15697])).
fof(f18992,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),X0),s(sK14,sK17))))) ) | ($spl5 | $spl82)),
  inference(unit_resulting_resolution,[],[f578,f15697,f15697,f373])).
fof(f19069,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18993,f15697])).
fof(f18993,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f15697,f373])).
fof(f19068,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18994,f15697])).
fof(f18994,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f639,f15697,f15697,f373])).
fof(f19067,plain,(
  $false | ($spl5 | $spl82)),
  inference(subsumption_resolution,[],[f18995,f15697])).
fof(f18995,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),X0))))) ) | ($spl5 | $spl82)),
  inference(unit_resulting_resolution,[],[f578,f15697,f15697,f373])).
fof(f19066,plain,(
  $false | ($spl13 | $spl82)),
  inference(subsumption_resolution,[],[f18996,f15697])).
fof(f18996,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f15697,f373])).
fof(f19065,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f18997,f15697])).
fof(f18997,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(bool,bool),X0))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f639,f15697,f15697,f373])).
fof(f19064,plain,(
  $false | ($spl5 | $spl82)),
  inference(subsumption_resolution,[],[f18998,f15697])).
fof(f18998,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),X0))))) ) | ($spl5 | $spl82)),
  inference(unit_resulting_resolution,[],[f578,f15697,f15697,f373])).
fof(f19063,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f19005,f15697])).
fof(f19005,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f445])).
fof(f19062,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f19030,f15697])).
fof(f19030,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))) ) | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f445])).
fof(f19031,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f15697,f445])).
fof(f19006,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f15697,f445])).
fof(f19000,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f15697,f15697,f373])).
fof(f18986,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f15697,f15697,f373])).
fof(f18971,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f15697,f15697,f373])).
fof(f18956,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11524,f15697,f12980])).
fof(f12980,plain,(
  ( ! [X30,X29] : (sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X29,X30) | ~p(s(bool,X29))) ) | $spl6),
  inference(subsumption_resolution,[],[f12916,f327])).
fof(f12916,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,X29)) | sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X29,X30)) ) | $spl6),
  inference(superposition,[],[f539,f11261])).
fof(f18954,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11461,f15697,f12976])).
fof(f12976,plain,(
  ( ! [X21,X22] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X21,X22) | ~p(s(bool,X21))) ) | $spl6),
  inference(subsumption_resolution,[],[f12912,f327])).
fof(f12912,plain,(
  ( ! [X21,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X21)) | sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK17)),X21,X22)) ) | $spl6),
  inference(superposition,[],[f522,f11261])).
fof(f18952,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8821,f15697,f10086])).
fof(f10086,plain,(
  ( ! [X31,X32] : (sP12('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X31,X32) | ~p(s(bool,X31))) ) | $spl0),
  inference(subsumption_resolution,[],[f10030,f327])).
fof(f10030,plain,(
  ( ! [X31,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X31)) | sP12('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X31,X32)) ) | $spl0),
  inference(superposition,[],[f539,f8567])).
fof(f18950,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8762,f15697,f10082])).
fof(f10082,plain,(
  ( ! [X24,X23] : (sP10('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X23,X24) | ~p(s(bool,X23))) ) | $spl0),
  inference(subsumption_resolution,[],[f10026,f327])).
fof(f10026,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | sP10('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X23,X24)) ) | $spl0),
  inference(superposition,[],[f522,f8567])).
fof(f18948,plain,(
  $false | ($spl32 | $spl82)),
  inference(unit_resulting_resolution,[],[f1271,f15697,f4604])).
fof(f18933,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f15697,f1937,f15697,f1483])).
fof(f18931,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f327,f1937,f15697,f1483])).
fof(f18928,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f622,f15697,f1483])).
fof(f18926,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f582,f327,f622,f15697,f1483])).
fof(f18915,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f15697,f1937,f15697,f1483])).
fof(f18912,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f622,f15697,f1483])).
fof(f18881,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f327,f582,f1845,f15697,f1479])).
fof(f18878,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f327,f582,f768,f15697,f1479])).
fof(f18871,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f15697,f1845,f15697,f1479])).
fof(f18868,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f582,f768,f15697,f1479])).
fof(f18863,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f327,f582,f1839,f15697,f1478])).
fof(f18859,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f327,f582,f610,f15697,f1478])).
fof(f18847,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f15697,f1839,f15697,f1478])).
fof(f18844,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f610,f15697,f1478])).
fof(f18833,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f327,f582,f1782,f15697,f1477])).
fof(f18829,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f327,f582,f598,f15697,f1477])).
fof(f18817,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f15697,f1782,f15697,f1477])).
fof(f18814,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f598,f15697,f1477])).
fof(f18803,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f327,f582,f1775,f15697,f1476])).
fof(f18799,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f327,f582,f592,f15697,f1476])).
fof(f18787,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f15697,f1775,f15697,f1476])).
fof(f18784,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697,f592,f15697,f1476])).
fof(f18768,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f713,f15697,f566])).
fof(f18766,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1851,f1889,f15697,f564])).
fof(f1889,plain,(
  ( ! [X0,X1] : (sP9('|dollar|exists_1'(s(fun(sK14,bool),sK16)),X0,X1)) ) | $spl13),
  inference(unit_resulting_resolution,[],[f1707,f519])).
fof(f18764,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f775,f616,f15697,f564])).
fof(f18762,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12070,f12094,f15697,f564])).
fof(f18760,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12498,f12522,f15697,f564])).
fof(f18758,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11427,f11446,f15697,f564])).
fof(f18756,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8731,f8749,f15697,f564])).
fof(f18754,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f775,f635,f15697,f564])).
fof(f18752,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f776,f614,f15697,f564])).
fof(f18750,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12063,f12093,f15697,f564])).
fof(f18748,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12491,f12521,f15697,f564])).
fof(f18746,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11421,f11445,f15697,f564])).
fof(f18744,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8725,f8748,f15697,f564])).
fof(f18742,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f776,f634,f15697,f564])).
fof(f18740,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f775,f612,f15697,f564])).
fof(f18718,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f1937,f15697,f563])).
fof(f18716,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f795,f622,f15697,f563])).
fof(f18714,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12112,f12165,f15697,f563])).
fof(f18712,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12540,f12593,f15697,f563])).
fof(f18710,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11462,f11506,f15697,f563])).
fof(f18708,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8763,f8804,f15697,f563])).
fof(f18706,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f795,f636,f15697,f563])).
fof(f18704,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f801,f620,f15697,f563])).
fof(f18702,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1916,f1936,f15697,f563])).
fof(f18700,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f795,f618,f15697,f563])).
fof(f18676,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f582,f428,f15697,f562])).
fof(f18673,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f582,f1720,f15697,f562])).
fof(f18667,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f582,f1671,f15697,f562])).
fof(f18656,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f805,f582,f15697,f15697,f540])).
fof(f18626,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1360,f582,f327,f15697,f540])).
fof(f18594,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f805,f582,f15697,f15697,f540])).
fof(f18564,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f805,f582,f327,f15697,f540])).
fof(f18540,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1388,f15697,f15697,f539])).
fof(f18538,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f801,f327,f15697,f539])).
fof(f18514,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1388,f15697,f15697,f539])).
fof(f18512,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1388,f327,f15697,f539])).
fof(f18488,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f795,f582,f15697,f538])).
fof(f18464,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f795,f582,f582,f15697,f537])).
fof(f18440,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1373,f15697,f535])).
fof(f18438,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f618,f15697,f582,f15697,f533])).
fof(f18436,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f618,f327,f582,f15697,f533])).
fof(f18412,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f618,f15697,f582,f15697,f533])).
fof(f18410,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f622,f327,f582,f15697,f533])).
fof(f18386,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f1937,f582,f15697,f532])).
fof(f18384,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f795,f622,f582,f15697,f532])).
fof(f18382,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1916,f1936,f582,f15697,f532])).
fof(f18380,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f795,f618,f582,f15697,f532])).
fof(f18331,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f1937,f582,f15697,f531])).
fof(f18329,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f795,f622,f582,f15697,f531])).
fof(f18327,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1916,f1936,f582,f15697,f531])).
fof(f18325,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f795,f618,f582,f15697,f531])).
fof(f18276,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1914,f1937,f15697,f15697,f530])).
fof(f18274,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f795,f622,f15697,f15697,f530])).
fof(f18272,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12112,f12165,f15697,f15697,f530])).
fof(f18270,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12540,f12593,f15697,f15697,f530])).
fof(f18268,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11462,f11506,f15697,f15697,f530])).
fof(f18266,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8763,f8804,f15697,f15697,f530])).
fof(f18264,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f795,f636,f15697,f15697,f530])).
fof(f18262,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f801,f620,f15697,f15697,f530])).
fof(f18260,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1916,f1936,f15697,f15697,f530])).
fof(f18258,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f795,f618,f15697,f15697,f530])).
fof(f18172,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1937,f1922,f15697,f15697,f530])).
fof(f18170,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1937,f1917,f15697,f15697,f530])).
fof(f18168,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1937,f1914,f15697,f15697,f530])).
fof(f18160,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f622,f12112,f15697,f15697,f530])).
fof(f18158,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f622,f12540,f15697,f15697,f530])).
fof(f18156,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f622,f11462,f15697,f15697,f530])).
fof(f18154,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f622,f8763,f15697,f15697,f530])).
fof(f18152,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f622,f801,f15697,f15697,f530])).
fof(f18150,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f622,f1916,f15697,f15697,f530])).
fof(f18148,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f622,f795,f15697,f15697,f530])).
fof(f18130,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1937,f1922,f327,f15697,f530])).
fof(f18124,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f622,f12112,f327,f15697,f530])).
fof(f18122,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f622,f12540,f327,f15697,f530])).
fof(f18120,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f622,f11462,f327,f15697,f530])).
fof(f18118,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f622,f8763,f327,f15697,f530])).
fof(f18116,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f622,f801,f327,f15697,f530])).
fof(f18076,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1360,f582,f15697,f527])).
fof(f18046,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1241,f582,f582,f15697,f523])).
fof(f18011,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1271,f15697,f15697,f522])).
fof(f18009,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f776,f327,f15697,f522])).
fof(f17987,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1271,f15697,f15697,f522])).
fof(f17985,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1271,f327,f15697,f522])).
fof(f17963,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1271,f15697,f15697,f521])).
fof(f17961,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f775,f327,f15697,f521])).
fof(f17939,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1271,f15697,f15697,f521])).
fof(f17937,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1271,f327,f15697,f521])).
fof(f17915,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1255,f15697,f518])).
fof(f17913,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1255,f15697,f517])).
fof(f17911,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f616,f582,f582,f15697,f516])).
fof(f17893,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1851,f1889,f582,f15697,f514])).
fof(f17891,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f775,f616,f582,f15697,f514])).
fof(f17889,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12063,f12093,f582,f15697,f514])).
fof(f17887,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12491,f12521,f582,f15697,f514])).
fof(f17885,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11421,f11445,f582,f15697,f514])).
fof(f17883,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8725,f8748,f582,f15697,f514])).
fof(f17880,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f775,f612,f582,f15697,f514])).
fof(f17806,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1889,f1852,f582,f15697,f513])).
fof(f17800,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f616,f12070,f582,f15697,f513])).
fof(f17798,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f616,f12498,f582,f15697,f513])).
fof(f17796,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f616,f11427,f582,f15697,f513])).
fof(f17794,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f616,f8731,f582,f15697,f513])).
fof(f17792,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f616,f776,f582,f15697,f513])).
fof(f17726,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1241,f582,f15697,f510])).
fof(f17692,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1241,f582,f15697,f509])).
fof(f17658,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f765,f582,f15697,f15697,f508])).
fof(f17626,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f765,f582,f327,f15697,f508])).
fof(f17604,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f765,f582,f15697,f15697,f508])).
fof(f17572,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f768,f582,f327,f15697,f508])).
fof(f17550,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f759,f15697,f15697,f15697,f505])).
fof(f17548,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f762,f327,f15697,f15697,f505])).
fof(f17522,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f759,f15697,f327,f15697,f505])).
fof(f17520,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1181,f327,f327,f15697,f505])).
fof(f17470,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f759,f15697,f15697,f15697,f505])).
fof(f17468,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f759,f327,f15697,f15697,f505])).
fof(f17442,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f759,f15697,f327,f15697,f505])).
fof(f17440,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f759,f327,f327,f15697,f505])).
fof(f17398,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f759,f15697,f15697,f15697,f505])).
fof(f17396,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f759,f327,f15697,f15697,f505])).
fof(f17370,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f762,f15697,f327,f15697,f505])).
fof(f17368,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f762,f327,f327,f15697,f505])).
fof(f17326,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1166,f15697,f500])).
fof(f17324,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f608,f582,f15697,f15697,f499])).
fof(f17312,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f608,f582,f327,f15697,f499])).
fof(f17298,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f608,f582,f15697,f15697,f499])).
fof(f17286,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f610,f582,f327,f15697,f499])).
fof(f17271,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f11997,f15697,f15697,f495])).
fof(f17269,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f11996,f15697,f15697,f495])).
fof(f17267,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f12425,f15697,f15697,f495])).
fof(f17265,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f12424,f15697,f15697,f495])).
fof(f17263,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f11363,f15697,f15697,f495])).
fof(f17261,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f11362,f15697,f15697,f495])).
fof(f17259,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f8669,f15697,f15697,f495])).
fof(f17257,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f8668,f15697,f15697,f495])).
fof(f17244,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1839,f1830,f15697,f15697,f495])).
fof(f17242,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f610,f762,f15697,f15697,f495])).
fof(f17240,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1838,f1827,f15697,f15697,f495])).
fof(f17238,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f608,f759,f15697,f15697,f495])).
fof(f17225,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f11996,f327,f15697,f495])).
fof(f17223,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f12424,f327,f15697,f495])).
fof(f17221,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f11362,f327,f15697,f495])).
fof(f17219,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f8668,f327,f15697,f495])).
fof(f17212,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1838,f1827,f327,f15697,f495])).
fof(f17210,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f608,f759,f327,f15697,f495])).
fof(f17169,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f11997,f15697,f15697,f495])).
fof(f17167,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f11996,f15697,f15697,f495])).
fof(f17165,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f12425,f15697,f15697,f495])).
fof(f17163,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f12424,f15697,f15697,f495])).
fof(f17161,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f11363,f15697,f15697,f495])).
fof(f17159,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f11362,f15697,f15697,f495])).
fof(f17157,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f8669,f15697,f15697,f495])).
fof(f17155,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f608,f8668,f15697,f15697,f495])).
fof(f17142,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1839,f1830,f15697,f15697,f495])).
fof(f17140,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f610,f762,f15697,f15697,f495])).
fof(f17138,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1838,f1827,f15697,f15697,f495])).
fof(f17136,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f608,f759,f15697,f15697,f495])).
fof(f17123,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f11997,f327,f15697,f495])).
fof(f17121,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f12425,f327,f15697,f495])).
fof(f17119,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f11363,f327,f15697,f495])).
fof(f17117,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f610,f8669,f327,f15697,f495])).
fof(f17110,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1839,f1830,f327,f15697,f495])).
fof(f17108,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f610,f762,f327,f15697,f495])).
fof(f17068,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f768,f582,f15697,f494])).
fof(f17036,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f765,f582,f15697,f493])).
fof(f17004,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f989,f639,f15697,f490])).
fof(f17002,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f2198,f1948,f15697,f490])).
fof(f2198,plain,(
  sP6('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f327,f604,f592,f1782,f473])).
fof(f17000,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f991,f639,f15697,f490])).
fof(f16998,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f2180,f1948,f15697,f490])).
fof(f2180,plain,(
  sP6('F_0','T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl13),
  inference(unit_resulting_resolution,[],[f582,f602,f592,f1781,f475])).
fof(f16941,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1448,f428,f15697,f489])).
fof(f16939,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1666,f369,f15697,f489])).
fof(f16937,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f2208,f1720,f15697,f489])).
fof(f16935,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f13013,f12312,f15697,f489])).
fof(f16933,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11672,f11261,f15697,f489])).
fof(f16931,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f9723,f8567,f15697,f489])).
fof(f16929,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1030,f644,f15697,f489])).
fof(f16926,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1445,f428,f15697,f489])).
fof(f16924,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1664,f369,f15697,f489])).
fof(f16922,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f2215,f1720,f15697,f489])).
fof(f16920,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f13007,f12312,f15697,f489])).
fof(f16918,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11667,f11261,f15697,f489])).
fof(f16916,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f9718,f8567,f15697,f489])).
fof(f16914,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1027,f644,f15697,f489])).
fof(f19061,plain,(
  $false | ($spl13 | $spl82)),
  inference(trivial_inequality_removal,[],[f16909])).
fof(f16909,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f2220,f15697,f489])).
fof(f2220,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0','F_0') | $spl13),
  inference(unit_resulting_resolution,[],[f1803,f469])).
fof(f1803,plain,(
  ~sP3('F_0','F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl13),
  inference(unit_resulting_resolution,[],[f582,f582,f1707,f484])).
fof(f19060,plain,(
  $false | ($spl13 | $spl82)),
  inference(trivial_inequality_removal,[],[f16908])).
fof(f16908,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f2172,f15697,f489])).
fof(f2172,plain,(
  ~sP6('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'T_0','T_0') | $spl13),
  inference(unit_resulting_resolution,[],[f1776,f470])).
fof(f1776,plain,(
  ~sP4('T_0','T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl13),
  inference(unit_resulting_resolution,[],[f327,f327,f1707,f480])).
fof(f19059,plain,(
  $false | ($spl13 | $spl82)),
  inference(trivial_inequality_removal,[],[f16894])).
fof(f16894,plain,(
  s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) != s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f5845,f15697,f489])).
fof(f5845,plain,(
  ~sP6('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16))) | $spl13),
  inference(unit_resulting_resolution,[],[f1794,f469])).
fof(f1794,plain,(
  ~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl13),
  inference(unit_resulting_resolution,[],[f1707,f582,f1707,f484])).
fof(f19058,plain,(
  $false | ($spl0 | $spl82)),
  inference(trivial_inequality_removal,[],[f16886])).
fof(f16886,plain,(
  s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) != s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f12748,f15697,f489])).
fof(f12748,plain,(
  ~sP6('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl0),
  inference(unit_resulting_resolution,[],[f8642,f470])).
fof(f8642,plain,(
  ~sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f573,f582,f573,f480])).
fof(f19057,plain,(
  $false | $spl82),
  inference(trivial_inequality_removal,[],[f16883])).
fof(f16883,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl82),
  inference(unit_resulting_resolution,[],[f1142,f15697,f489])).
fof(f19056,plain,(
  $false | $spl82),
  inference(trivial_inequality_removal,[],[f16876])).
fof(f16876,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl82),
  inference(unit_resulting_resolution,[],[f1121,f15697,f489])).
fof(f16862,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1129,f15697,f487])).
fof(f16860,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1129,f15697,f486])).
fof(f16858,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1129,f15697,f485])).
fof(f16856,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1106,f15697,f481])).
fof(f16854,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f596,f582,f15697,f15697,f480])).
fof(f16842,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f596,f582,f327,f15697,f480])).
fof(f16828,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f596,f582,f15697,f15697,f480])).
fof(f16816,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f598,f582,f327,f15697,f480])).
fof(f16802,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1078,f15697,f477])).
fof(f16800,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f590,f582,f15697,f15697,f476])).
fof(f16788,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f590,f582,f327,f15697,f476])).
fof(f16774,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f590,f582,f15697,f15697,f476])).
fof(f16762,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f592,f582,f327,f15697,f476])).
fof(f16743,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f596,f590,f11993,f15697,f473])).
fof(f16738,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f596,f590,f12421,f15697,f473])).
fof(f16734,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11344,f596,f590,f11302,f15697,f473])).
fof(f16729,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f596,f590,f11359,f15697,f473])).
fof(f16724,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8650,f596,f590,f8608,f15697,f473])).
fof(f16718,plain,(
  $false | ($spl0 | $spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f1781,f1774,f8664,f15697,f473])).
fof(f16714,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f596,f590,f8665,f15697,f473])).
fof(f16666,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f11978,f596,f590,f11924,f15697,f473])).
fof(f16662,plain,(
  $false | ($spl2 | $spl82)),
  inference(unit_resulting_resolution,[],[f12406,f596,f590,f12352,f15697,f473])).
fof(f16658,plain,(
  $false | ($spl6 | $spl82)),
  inference(unit_resulting_resolution,[],[f11344,f596,f590,f11301,f15697,f473])).
fof(f16654,plain,(
  $false | ($spl0 | $spl82)),
  inference(unit_resulting_resolution,[],[f8650,f596,f590,f8607,f15697,f473])).
fof(f16645,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f602,f1781,f1774,f1744,f15697,f473])).
fof(f16641,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f602,f596,f590,f1017,f15697,f473])).
fof(f16085,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1568,f15697,f582,f15697,f472])).
fof(f16083,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1568,f327,f582,f15697,f472])).
fof(f15992,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f1568,f582,f15697,f15697,f472])).
fof(f15935,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f991,f582,f327,f15697,f472])).
fof(f15901,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f713,f15697,f15697,f458])).
fof(f15899,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f713,f327,f15697,f458])).
fof(f15889,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f713,f15697,f15697,f458])).
fof(f15887,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f720,f327,f15697,f458])).
fof(f15877,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f713,f15697,f15697,f457])).
fof(f15875,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f713,f327,f15697,f457])).
fof(f15865,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f713,f15697,f15697,f457])).
fof(f15854,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f713,f15697,f15697,f456])).
fof(f15843,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f713,f15697,f15697,f456])).
fof(f15841,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f720,f327,f15697,f456])).
fof(f15831,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f676,f582,f15697,f419])).
fof(f15814,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f676,f582,f15697,f416])).
fof(f15797,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f680,f15697,f15697,f415])).
fof(f15795,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f680,f327,f15697,f415])).
fof(f15779,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f680,f15697,f15697,f415])).
fof(f15777,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f689,f327,f15697,f415])).
fof(f15761,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f703,f15697,f401])).
fof(f15759,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f703,f15697,f400])).
fof(f15757,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f639,f15697,f394])).
fof(f15755,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f15697,f1948,f15697,f394])).
fof(f15753,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f327,f639,f15697,f394])).
fof(f15751,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f327,f1948,f15697,f394])).
fof(f15747,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f15697,f639,f15697,f394])).
fof(f15745,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f15697,f1948,f15697,f394])).
fof(f15743,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f327,f639,f15697,f394])).
fof(f15741,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f327,f1948,f15697,f394])).
fof(f15735,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1720,f15697,f345])).
fof(f15733,plain,(
  $false | ($spl13 | $spl82)),
  inference(unit_resulting_resolution,[],[f1948,f15697,f342])).
fof(f15731,plain,(
  $false | $spl82),
  inference(unit_resulting_resolution,[],[f582,f15697])).
fof(f15728,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK30(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) ) | $spl84),
  inference(cnf_transformation,[],[f15728_D])).
fof(f15728_D,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK30(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f15729,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X18),s(X16,sK30(X18,X17,X15,X16)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X15,X16),X17),s(fun(X16,bool),X18)))) )),
  inference(subsumption_resolution,[],[f15724,f341])).
fof(f15724,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X18),s(X16,sK30(X18,X17,X15,X16)))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X15,X16),X17),s(fun(X16,bool),X18)))) )),
  inference(superposition,[],[f2067,f890])).
fof(f15727,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK30(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) )),
  inference(subsumption_resolution,[],[f15723,f341])).
fof(f15723,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK30(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f2067,f890])).
fof(f15696,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK28(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) ) | $spl80),
  inference(cnf_transformation,[],[f15696_D])).
fof(f15696_D,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK28(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f15698,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X18),s(X16,sK28(X18,X17,X15,X16)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X15,X16),X17),s(fun(X16,bool),X18)))) )),
  inference(subsumption_resolution,[],[f15692,f341])).
fof(f15692,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X18),s(X16,sK28(X18,X17,X15,X16)))) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X15,X16),X17),s(fun(X16,bool),X18)))) )),
  inference(superposition,[],[f2057,f890])).
fof(f15695,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK28(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13)))) )),
  inference(subsumption_resolution,[],[f15691,f341])).
fof(f15691,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,happ(s(fun(X11,bool),X13),s(X11,sK28(X13,X12,X10,X11)))) | s(bool,'T_0') = s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X10,X11),X12),s(fun(X11,bool),X13))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f2057,f890])).
fof(f14677,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP12(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2,'T_0') | p(s(bool,X2))) ) | $spl78),
  inference(cnf_transformation,[],[f14677_D])).
fof(f14677_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP12(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2,'T_0') | p(s(bool,X2))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f14679,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2,'T_0')) )),
  inference(condensation,[],[f14678])).
fof(f14678,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP12(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37,'T_0') | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f14664,f954])).
fof(f14664,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP12(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37,'T_0') | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1486,f890])).
fof(f14676,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2,'T_0')) )),
  inference(condensation,[],[f14675])).
fof(f14675,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP12(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31,'T_0') | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f14663,f954])).
fof(f14663,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP12(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31,'T_0') | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1486,f890])).
fof(f14641,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | ~sP11('T_0',X1,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X2))),s(X3,X4))) | ~p(s(bool,X2))) ) | $spl76),
  inference(cnf_transformation,[],[f14641_D])).
fof(f14641_D,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | ~sP11('T_0',X1,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X2))),s(X3,X4))) | ~p(s(bool,X2))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f14643,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11('T_0',X1,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X2))),s(X3,X4)))) )),
  inference(condensation,[],[f14642])).
fof(f14642,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (p(s(bool,X36)) | ~p(s(bool,X37)) | ~sP11('T_0',X36,happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X37))),s(X38,X39))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f14622,f954])).
fof(f14622,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | p(s(bool,X36)) | ~p(s(bool,X37)) | ~sP11('T_0',X36,happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X37))),s(X38,X39))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1483,f890])).
fof(f14640,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11('T_0',X1,happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X2))),s(X3,X4)))) )),
  inference(condensation,[],[f14639])).
fof(f14639,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~sP11('T_0',X30,happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X31))),s(X32,X33))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f14621,f954])).
fof(f14621,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | ~p(s(bool,X31)) | ~sP11('T_0',X30,happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X31))),s(X32,X33))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1483,f890])).
fof(f14575,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP10(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X2,X1)) ) | $spl74),
  inference(cnf_transformation,[],[f14575_D])).
fof(f14575_D,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | p(s(bool,X2)) | ~sP10(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X2,X1)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f14577,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X2,X1)) )),
  inference(condensation,[],[f14576])).
fof(f14576,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (p(s(bool,X36)) | p(s(bool,X37)) | ~sP10(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X38,X39)),X37,X36) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f14562,f954])).
fof(f14562,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | p(s(bool,X36)) | p(s(bool,X37)) | ~sP10(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X38,X39)),X37,X36) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1482,f890])).
fof(f14574,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X2,X1)) )),
  inference(condensation,[],[f14573])).
fof(f14573,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | ~sP10(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X32,X33)),X31,X30) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f14561,f954])).
fof(f14561,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | p(s(bool,X31)) | ~sP10(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X32,X33)),X31,X30) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1482,f890])).
fof(f14531,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | ~sP9(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X1,X2) | p(s(bool,X2))) ) | $spl72),
  inference(cnf_transformation,[],[f14531_D])).
fof(f14531_D,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | ~sP9(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X1,X2) | p(s(bool,X2))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f14533,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X1,X2)) )),
  inference(condensation,[],[f14532])).
fof(f14532,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (p(s(bool,X36)) | p(s(bool,X37)) | ~sP9(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X38,X39)),X36,X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f14508,f954])).
fof(f14508,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | p(s(bool,X36)) | p(s(bool,X37)) | ~sP9(happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X38,X39)),X36,X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1480,f890])).
fof(f14530,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X3,X4)),X1,X2)) )),
  inference(condensation,[],[f14529])).
fof(f14529,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | ~sP9(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X32,X33)),X30,X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f14507,f954])).
fof(f14507,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | p(s(bool,X31)) | ~sP9(happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X32,X33)),X30,X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1480,f890])).
fof(f14472,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,'T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) ) | $spl70),
  inference(cnf_transformation,[],[f14472_D])).
fof(f14472_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,'T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f14474,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,'T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) )),
  inference(condensation,[],[f14473])).
fof(f14473,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP8(X37,'T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f14458,f954])).
fof(f14458,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP8(X37,'T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1479,f890])).
fof(f14471,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,'T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)))) )),
  inference(condensation,[],[f14470])).
fof(f14470,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP8(X31,'T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f14457,f954])).
fof(f14457,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP8(X31,'T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1479,f890])).
fof(f14434,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP7('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) | $spl68),
  inference(cnf_transformation,[],[f14434_D])).
fof(f14434_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP7('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f14436,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f14435])).
fof(f14435,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP7('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f14415,f954])).
fof(f14415,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP7('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1478,f890])).
fof(f14433,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f14432])).
fof(f14432,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP7('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f14414,f954])).
fof(f14414,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP7('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1478,f890])).
fof(f14363,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP4('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) | $spl66),
  inference(cnf_transformation,[],[f14363_D])).
fof(f14363_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP4('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f14365,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f14364])).
fof(f14364,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f14344,f954])).
fof(f14344,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1477,f890])).
fof(f14362,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f14361])).
fof(f14361,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP4('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f14343,f954])).
fof(f14343,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP4('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1477,f890])).
fof(f14310,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP5('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) | $spl64),
  inference(cnf_transformation,[],[f14310_D])).
fof(f14310_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | ~sP5('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2) | p(s(bool,X2))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f14312,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f14311])).
fof(f14311,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP5('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f14291,f954])).
fof(f14291,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP5('T_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,X36))),s(X38,X39)),X37) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1476,f890])).
fof(f14309,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5('T_0',happ(s(fun(X3,bool),'_dst_x0_1'(s(bool,X1))),s(X3,X4)),X2)) )),
  inference(condensation,[],[f14308])).
fof(f14308,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | p(s(bool,X31)) | ~sP5('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f14290,f954])).
fof(f14290,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | ~sP5('T_0',happ(s(fun(X32,bool),'_dst_x0_1'(s(bool,X30))),s(X32,X33)),X31) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1476,f890])).
fof(f14004,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) ) | $spl62),
  inference(cnf_transformation,[],[f14004_D])).
fof(f14004_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f14006,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) )),
  inference(condensation,[],[f14005])).
fof(f14005,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | sP12(happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,X36))),s(X37,X38)),'T_0',X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f13992,f954])).
fof(f13992,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | sP12(happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,X36))),s(X37,X38)),'T_0',X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1485,f890])).
fof(f14003,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) )),
  inference(condensation,[],[f14002])).
fof(f14002,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | sP12(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,X30))),s(X31,X32)),'T_0',X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f13991,f954])).
fof(f13991,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP12(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,X30))),s(X31,X32)),'T_0',X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1485,f890])).
fof(f13961,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X4,X1)) ) | $spl60),
  inference(cnf_transformation,[],[f13961_D])).
fof(f13961_D,plain,(
  ( ! [X4,X2,X3,X1] : (p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X4,X1)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f13963,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X4,X1)) )),
  inference(condensation,[],[f13962])).
fof(f13962,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (p(s(bool,X36)) | sP12(happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X37,X38)),X39,X36) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f13949,f954])).
fof(f13949,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | p(s(bool,X36)) | sP12(happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X37,X38)),X39,X36) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1484,f890])).
fof(f13960,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X4,X1)) )),
  inference(condensation,[],[f13959])).
fof(f13959,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (p(s(bool,X30)) | sP12(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X33,X30) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f13948,f954])).
fof(f13948,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP12(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X33,X30) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1484,f890])).
fof(f13925,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | sP10(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) ) | $spl58),
  inference(cnf_transformation,[],[f13925_D])).
fof(f13925_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,X1)) | sP10(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f13927,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) )),
  inference(condensation,[],[f13926])).
fof(f13926,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X36)) | sP10(happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,X36))),s(X37,X38)),'T_0',X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f13913,f954])).
fof(f13913,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | sP10(happ(s(fun(X37,bool),'_dst_x0_1'(s(bool,X36))),s(X37,X38)),'T_0',X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1481,f890])).
fof(f13924,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,X1))),s(X2,X3)),'T_0',X4)) )),
  inference(condensation,[],[f13923])).
fof(f13923,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X30)) | sP10(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,X30))),s(X31,X32)),'T_0',X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f13912,f954])).
fof(f13912,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | sP10(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,X30))),s(X31,X32)),'T_0',X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f1481,f890])).
fof(f13754,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP1(X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) ) | $spl56),
  inference(cnf_transformation,[],[f13754_D])).
fof(f13754_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP1(X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f13756,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) )),
  inference(condensation,[],[f13755])).
fof(f13755,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | sP1(X30,happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f13742,f954])).
fof(f13742,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1(X30,happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32))) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1474,f890])).
fof(f13753,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)))) )),
  inference(condensation,[],[f13752])).
fof(f13752,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | sP1(X25,happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f13741,f954])).
fof(f13741,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP1(X25,happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27))) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1474,f890])).
fof(f13711,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP1(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) | $spl54),
  inference(cnf_transformation,[],[f13711_D])).
fof(f13711_D,plain,(
  ( ! [X2,X3,X1] : (p(s(bool,X1)) | sP1(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f13713,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) )),
  inference(condensation,[],[f13712])).
fof(f13712,plain,(
  ( ! [X30,X33,X31,X29,X32] : (p(s(bool,X30)) | sP1(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f13699,f954])).
fof(f13699,plain,(
  ( ! [X30,X33,X31,X29,X32] : (~p(s(bool,X29)) | p(s(bool,X30)) | sP1(happ(s(fun(X31,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X31,X32)),X30) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1473,f890])).
fof(f13710,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(happ(s(fun(X2,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X2,X3)),X1)) )),
  inference(condensation,[],[f13709])).
fof(f13709,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,X25)) | sP1(happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27)),X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(subsumption_resolution,[],[f13698,f954])).
fof(f13698,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,X24)) | p(s(bool,X25)) | sP1(happ(s(fun(X26,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X26,X27)),X25) | s(bool,X24) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f1473,f890])).
fof(f11243,plain,(
  $false | ($spl0 | $spl52)),
  inference(subsumption_resolution,[],[f11242,f327])).
fof(f11242,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl52)),
  inference(forward_demodulation,[],[f10706,f8567])).
fof(f10706,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | $spl52),
  inference(unit_resulting_resolution,[],[f9334,f372])).
fof(f11141,plain,(
  $false | ($spl0 | $spl52)),
  inference(subsumption_resolution,[],[f11140,f582])).
fof(f11140,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl52)),
  inference(backward_demodulation,[],[f11130,f8562])).
fof(f11130,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0)))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f639,f9334,f954])).
fof(f9334,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0))))) ) | $spl52),
  inference(cnf_transformation,[],[f9334_D])).
fof(f9334_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f575,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl2),
  inference(cnf_transformation,[],[f575_D])).
fof(f575_D,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f2549,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0))))) )),
  inference(forward_demodulation,[],[f2548,f428])).
fof(f2548,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0)))) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,happ(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(sK13,sK29(X0,sK16,'_dst_x0_1'(s(sK14,sK15)),sK13,sK14)))))))) )),
  inference(subsumption_resolution,[],[f2543,f327])).
fof(f2543,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0)))) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,happ(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(sK13,sK29(X0,sK16,'_dst_x0_1'(s(sK14,sK15)),sK13,sK14)))))))) )),
  inference(superposition,[],[f436,f644])).
fof(f9331,plain,(
  $false | ($spl5 | $spl6 | $spl8)),
  inference(subsumption_resolution,[],[f9330,f327])).
fof(f9330,plain,(
  ~p(s(bool,'T_0')) | ($spl5 | $spl6 | $spl8)),
  inference(forward_demodulation,[],[f9321,f8946])).
fof(f8946,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17))) | $spl6),
  inference(unit_resulting_resolution,[],[f327,f580,f394])).
fof(f9321,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ($spl5 | $spl8)),
  inference(unit_resulting_resolution,[],[f578,f581])).
fof(f573,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | $spl0),
  inference(cnf_transformation,[],[f573_D])).
fof(f573_D,plain,(
  p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f8561,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) | $spl50),
  inference(cnf_transformation,[],[f8561_D])).
fof(f8561_D,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f2389,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | $spl2),
  inference(subsumption_resolution,[],[f2383,f327])).
fof(f2383,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | $spl2),
  inference(superposition,[],[f374,f643])).
fof(f8559,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK20(sK16,sK14))))) | $spl48),
  inference(cnf_transformation,[],[f8559_D])).
fof(f8559_D,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK20(sK16,sK14))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2417,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK20(sK16,sK14))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | $spl2),
  inference(subsumption_resolution,[],[f2411,f327])).
fof(f2411,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK20(sK16,sK14))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | $spl2),
  inference(superposition,[],[f375,f643])).
fof(f6932,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | $spl2),
  inference(subsumption_resolution,[],[f6918,f327])).
fof(f6918,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | $spl2),
  inference(superposition,[],[f374,f643])).
fof(f8557,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK20(sK16,sK14))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | $spl2),
  inference(subsumption_resolution,[],[f6967,f327])).
fof(f6967,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK20(sK16,sK14))))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | $spl2),
  inference(superposition,[],[f375,f643])).
fof(f8556,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7649,f8402])).
fof(f8402,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f621,f7633,f1384])).
fof(f7633,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0))))) ) | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7632,f582])).
fof(f7632,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0))))) ) | ($spl1 | $spl2 | $spl8)),
  inference(forward_demodulation,[],[f7631,f663])).
fof(f663,plain,(
  s(bool,'F_0') = s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))) | $spl1),
  inference(unit_resulting_resolution,[],[f574,f346])).
fof(f7631,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0))))) ) | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7618,f6128])).
fof(f6128,plain,(
  s(sK14,sK15) = s(sK14,sK19(sK16,sK14)) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f2391,f581])).
fof(f2391,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2390,f582])).
fof(f2390,plain,(
  p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2389,f663])).
fof(f7618,plain,(
  ( ! [X0] : (s(sK14,sK15) != s(sK14,sK19(sK16,sK14)) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X0))))) ) | ($spl1 | $spl2 | $spl8)),
  inference(superposition,[],[f376,f7077])).
fof(f7077,plain,(
  s(sK14,sK15) = s(sK14,sK20(sK16,sK14)) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f2419,f581])).
fof(f2419,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK20(sK16,sK14))))) | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f2418,f582])).
fof(f2418,plain,(
  p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK20(sK16,sK14))))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f2417,f663])).
fof(f7649,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_2'(s(fun(X0,sK14),X1),s(fun(sK14,bool),sK16))))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f444])).
fof(f8555,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7651,f8402])).
fof(f7651,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X0,sK14),X1),s(fun(sK14,bool),sK16))))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f446])).
fof(f8554,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7653,f8402])).
fof(f7653,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,sK14),X1),s(fun(sK14,bool),sK16))))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f434])).
fof(f8553,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7654,f8402])).
fof(f7654,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(X0,sK14),X1),s(fun(sK14,bool),sK16))))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f437])).
fof(f8550,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7666,f8402])).
fof(f7666,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f686,f7633,f414])).
fof(f8549,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7670,f8402])).
fof(f7670,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f710,f7633,f414])).
fof(f8548,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7680,f8402])).
fof(f7680,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f625,f7633,f453])).
fof(f8547,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7683,f8402])).
fof(f7683,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6537,f7633,f453])).
fof(f6537,plain,(
  ( ! [X0] : (sP2('T_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6165])).
fof(f6165,plain,(
  ( ! [X0] : (sP2('T_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f327,f2391,f458])).
fof(f8546,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7689,f8402])).
fof(f7689,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f625,f7633,f453])).
fof(f8545,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7691,f8402])).
fof(f7691,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f637,f7633,f453])).
fof(f8544,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7692,f8402])).
fof(f7692,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6537,f7633,f453])).
fof(f8543,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7696,f8402])).
fof(f7696,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6536,f7633,f454])).
fof(f6536,plain,(
  ( ! [X0] : (sP2(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,'T_0')) ) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6159])).
fof(f6159,plain,(
  ( ! [X0] : (sP2(happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),X0,'T_0')) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f327,f2391,f458])).
fof(f8542,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7699,f8402])).
fof(f7699,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f625,f7633,f454])).
fof(f8541,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7701,f8402])).
fof(f7701,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f637,f7633,f454])).
fof(f8540,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7704,f8402])).
fof(f7704,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f625,f7633,f454])).
fof(f8539,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7708,f8402])).
fof(f7708,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6536,f7633,f454])).
fof(f8538,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7715,f589])).
fof(f7715,plain,(
  ~sP5('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f603,f596,f6539,f7633,f474])).
fof(f6539,plain,(
  ~sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6173])).
fof(f6173,plain,(
  ~sP6('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),'F_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f327,f582,f2391,f472])).
fof(f8537,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7743,f594])).
fof(f7743,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0') | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f603,f6539,f589,f7633,f474])).
fof(f8536,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7744,f600])).
fof(f7744,plain,(
  ~sP3('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f596,f6539,f589,f7633,f474])).
fof(f8535,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7747,f7892])).
fof(f7892,plain,(
  ( ! [X2,X0,X1] : (sP4(happ(s(fun(sK14,bool),sK16),s(sK14,X0)),X1,X2)) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f483])).
fof(f7747,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))))),s(X0,X1)),'F_0')) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f603,f1448,f589,f7633,f474])).
fof(f8534,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7758,f594])).
fof(f7758,plain,(
  ~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0') | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f603,f6539,f590,f7633,f474])).
fof(f8533,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7759,f600])).
fof(f7759,plain,(
  ~sP3('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f596,f6539,f590,f7633,f474])).
fof(f8532,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7785,f7892])).
fof(f7785,plain,(
  ( ! [X0,X1] : (~sP4(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15))))),s(X0,X1)),'F_0')) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f603,f1448,f592,f7633,f474])).
fof(f8531,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7792,f588])).
fof(f7792,plain,(
  ~sP5(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0') | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f602,f596,f6576,f7633,f475])).
fof(f6576,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','F_0') | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6237])).
fof(f6237,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),'T_0','F_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f639,f2391,f490])).
fof(f8530,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7805,f595])).
fof(f7805,plain,(
  ~sP4('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f602,f6576,f588,f7633,f475])).
fof(f8529,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7806,f601])).
fof(f7806,plain,(
  ~sP3('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f596,f6576,f588,f7633,f475])).
fof(f8528,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7837,f6540])).
fof(f6540,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6174])).
fof(f6174,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),'F_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f2391,f582,f2391,f472])).
fof(f7837,plain,(
  sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f603,f595,f590,f7633,f475])).
fof(f8527,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7841,f595])).
fof(f7841,plain,(
  ~sP4('T_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f602,f6576,f590,f7633,f475])).
fof(f8526,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7842,f601])).
fof(f7842,plain,(
  ~sP3('F_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f596,f6576,f590,f7633,f475])).
fof(f8525,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7898,f8402])).
fof(f7898,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f605,f7633,f484])).
fof(f8524,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7899,f8402])).
fof(f7899,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f605,f7633,f484])).
fof(f8523,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7900,f8402])).
fof(f7900,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f605,f7633,f484])).
fof(f8522,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7902,f8402])).
fof(f7902,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f605,f7633,f484])).
fof(f8521,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7905,f8402])).
fof(f7905,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f601,f7633,f484])).
fof(f8520,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7906,f8402])).
fof(f7906,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f601,f7633,f484])).
fof(f8519,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7908,f8402])).
fof(f7908,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f630,f7633,f484])).
fof(f8518,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7910,f8402])).
fof(f7910,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f631,f7633,f484])).
fof(f8517,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7911,f8402])).
fof(f7911,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f603,f7633,f484])).
fof(f8516,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7912,f8402])).
fof(f7912,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f603,f7633,f484])).
fof(f8515,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7918,f8402])).
fof(f7918,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f603,f7633,f484])).
fof(f8514,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7919,f8402])).
fof(f7919,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f603,f7633,f484])).
fof(f8513,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7920,f8402])).
fof(f7920,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f603,f7633,f484])).
fof(f8512,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7922,f8402])).
fof(f7922,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f603,f7633,f484])).
fof(f8511,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7925,f8402])).
fof(f7925,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f601,f7633,f484])).
fof(f8510,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7926,f8402])).
fof(f7926,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f601,f7633,f484])).
fof(f8509,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7928,f8402])).
fof(f7928,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f630,f7633,f484])).
fof(f8508,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7930,f8402])).
fof(f7930,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f632,f7633,f484])).
fof(f8507,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7931,f8402])).
fof(f7931,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f605,f7633,f484])).
fof(f8506,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7932,f8402])).
fof(f7932,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f605,f7633,f484])).
fof(f8505,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7938,f8402])).
fof(f7938,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f601,f7633,f484])).
fof(f8504,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7939,f8402])).
fof(f7939,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f601,f7633,f484])).
fof(f8503,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7940,f8402])).
fof(f7940,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f601,f7633,f484])).
fof(f8502,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7942,f8402])).
fof(f7942,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f601,f7633,f484])).
fof(f8501,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7945,f8402])).
fof(f7945,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f631,f7633,f484])).
fof(f8500,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7946,f8402])).
fof(f7946,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f603,f7633,f484])).
fof(f8499,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7947,f8402])).
fof(f7947,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f603,f7633,f484])).
fof(f8498,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7950,f8402])).
fof(f7950,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f632,f7633,f484])).
fof(f8497,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7951,f8402])).
fof(f7951,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f605,f7633,f484])).
fof(f8496,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7952,f8402])).
fof(f7952,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f605,f7633,f484])).
fof(f8495,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7960,f639])).
fof(f7960,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6576,f7633,f488])).
fof(f8494,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f7962,f639])).
fof(f7962,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6579,f7633,f488])).
fof(f6579,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0') | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6240])).
fof(f6240,plain,(
  ~sP6(happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),'F_0','T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f639,f2391,f490])).
fof(f8493,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8016,f8402])).
fof(f8016,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f1352,f7633,f503])).
fof(f8492,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8018,f8402])).
fof(f8018,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6582,f7633,f503])).
fof(f6582,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0',X0)) ) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6243])).
fof(f6243,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),'F_0',X0)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f582,f2391,f494])).
fof(f8491,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8022,f8402])).
fof(f8022,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6581,f7633,f503])).
fof(f6581,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,'F_0')) ) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6242])).
fof(f6242,plain,(
  ( ! [X0] : (~sP8(happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),X0,'F_0')) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f582,f2391,f493])).
fof(f8490,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8024,f8402])).
fof(f8024,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6582,f7633,f503])).
fof(f8489,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8030,f8402])).
fof(f8030,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6600,f7633,f504])).
fof(f6600,plain,(
  ~sP8('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6293])).
fof(f6293,plain,(
  ~sP8('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f327,f582,f2391,f508])).
fof(f8488,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8032,f8402])).
fof(f8032,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6581,f7633,f504])).
fof(f8487,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8036,f8402])).
fof(f8036,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6581,f7633,f504])).
fof(f8486,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8038,f8402])).
fof(f8038,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6582,f7633,f504])).
fof(f8485,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8086,f615])).
fof(f8086,plain,(
  ~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0') | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6604,f7633,f515])).
fof(f6604,plain,(
  ( ! [X0] : (~sP10('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6309])).
fof(f6309,plain,(
  ( ! [X0] : (~sP10('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),X0)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f582,f2391,f510])).
fof(f8484,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8087,f615])).
fof(f8087,plain,(
  ( ! [X0] : (~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,X0)))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6604,f7633,f515])).
fof(f8483,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8088,f613])).
fof(f8088,plain,(
  ( ! [X0] : (~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,X0)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6603,f7633,f515])).
fof(f6603,plain,(
  ( ! [X0] : (~sP10('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6303])).
fof(f6303,plain,(
  ( ! [X0] : (~sP10('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))))) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f582,f2391,f509])).
fof(f8482,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8091,f8402])).
fof(f8091,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6604,f615,f7633,f515])).
fof(f8481,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8094,f8402])).
fof(f8094,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6603,f615,f7633,f515])).
fof(f8480,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8099,f8402])).
fof(f8099,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6604,f613,f7633,f515])).
fof(f8479,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8101,f8402])).
fof(f8101,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6604,f634,f7633,f515])).
fof(f8478,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8104,f8402])).
fof(f8104,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6604,f616,f7633,f515])).
fof(f8477,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8107,f8402])).
fof(f8107,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6603,f616,f7633,f515])).
fof(f8476,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8114,f613])).
fof(f8114,plain,(
  ~sP9('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6603,f7633,f515])).
fof(f8475,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8115,f613])).
fof(f8115,plain,(
  ( ! [X0] : (~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,X0)),happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6603,f7633,f515])).
fof(f8474,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8116,f615])).
fof(f8116,plain,(
  ( ! [X0] : (~sP9('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),happ(s(fun(sK14,bool),sK16),s(sK14,X0)))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6604,f7633,f515])).
fof(f8473,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8119,f8402])).
fof(f8119,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6603,f613,f7633,f515])).
fof(f8472,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8122,f8402])).
fof(f8122,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6604,f613,f7633,f515])).
fof(f8471,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8126,f8402])).
fof(f8126,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6603,f635,f7633,f515])).
fof(f8470,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8128,f8402])).
fof(f8128,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6604,f615,f7633,f515])).
fof(f8469,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8132,f8402])).
fof(f8132,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6603,f616,f7633,f515])).
fof(f8468,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8135,f8402])).
fof(f8135,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6604,f616,f7633,f515])).
fof(f8467,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8144,f8402])).
fof(f8144,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f615,f7633,f516])).
fof(f8466,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8150,f8402])).
fof(f8150,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f613,f7633,f516])).
fof(f8465,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8153,f8402])).
fof(f8153,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f634,f7633,f516])).
fof(f8464,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8161,f8402])).
fof(f8161,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f613,f7633,f516])).
fof(f8463,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8164,f8402])).
fof(f8164,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f635,f7633,f516])).
fof(f8462,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8165,f8402])).
fof(f8165,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f615,f7633,f516])).
fof(f8461,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8177,f8402])).
fof(f8177,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6604,f7633,f520])).
fof(f8460,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8178,f8402])).
fof(f8178,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6604,f7633,f520])).
fof(f8459,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8181,f8402])).
fof(f8181,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6603,f7633,f520])).
fof(f8458,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8182,f8402])).
fof(f8182,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6603,f7633,f520])).
fof(f8457,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8188,f8402])).
fof(f8188,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6603,f7633,f520])).
fof(f8456,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8189,f8402])).
fof(f8189,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6603,f7633,f520])).
fof(f8455,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8192,f8402])).
fof(f8192,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6604,f7633,f520])).
fof(f8454,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8193,f8402])).
fof(f8193,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6604,f7633,f520])).
fof(f8453,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8199,f8402])).
fof(f8199,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f1436,f7633,f520])).
fof(f8452,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8200,f8402])).
fof(f8200,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f1436,f7633,f520])).
fof(f8451,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8205,f8402])).
fof(f8205,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f867,f7633,f523])).
fof(f8450,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8208,f8402])).
fof(f8208,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f863,f7633,f523])).
fof(f8449,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8211,f8402])).
fof(f8211,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f864,f7633,f523])).
fof(f8448,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8215,f8402])).
fof(f8215,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f863,f7633,f523])).
fof(f8447,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8218,f8402])).
fof(f8218,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f866,f7633,f523])).
fof(f8446,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8219,f8402])).
fof(f8219,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f867,f7633,f523])).
fof(f8445,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8224,f8402])).
fof(f8224,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f814,f7633,f526])).
fof(f8444,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8226,f8402])).
fof(f8226,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f934,f7633,f526])).
fof(f8443,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8233,f8402])).
fof(f8233,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6620,f7633,f526])).
fof(f6620,plain,(
  sP12('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6388])).
fof(f6388,plain,(
  sP12('F_0','F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14)))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f582,f582,f2391,f537])).
fof(f8442,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8235,f8402])).
fof(f8235,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f934,f7633,f526])).
fof(f8441,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8260,f621])).
fof(f8260,plain,(
  ~sP11('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f6613,f7633,f532])).
fof(f6613,plain,(
  ( ! [X0] : (~sP12('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0)) ) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6341])).
fof(f6341,plain,(
  ( ! [X0] : (~sP12('F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),X0)) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f582,f2391,f527])).
fof(f8440,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8296,f8402])).
fof(f8296,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f6613,f7633,f537])).
fof(f8439,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8304,f8402])).
fof(f8304,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f6624,f7633,f537])).
fof(f6624,plain,(
  ~sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0') | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6404])).
fof(f6404,plain,(
  ~sP12(happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),'F_0','T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f327,f582,f2391,f540])).
fof(f8438,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8332,f607])).
fof(f8332,plain,(
  ( ! [X0] : (~sP7(X0,'F_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6581,f7633,f565])).
fof(f8437,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8334,f607])).
fof(f8334,plain,(
  ( ! [X0] : (~sP7('F_0',X0,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6582,f7633,f565])).
fof(f8436,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8338,f8402])).
fof(f8338,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f587,f7633,f583])).
fof(f8435,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8341,f8402])).
fof(f8341,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f585,f7633,f583])).
fof(f8434,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8343,f8402])).
fof(f8343,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f626,f7633,f583])).
fof(f8433,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8346,f8402])).
fof(f8346,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f585,f7633,f583])).
fof(f8432,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8349,f8402])).
fof(f8349,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f627,f7633,f583])).
fof(f8431,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8350,f8402])).
fof(f8350,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f587,f7633,f583])).
fof(f8426,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8381,f8402])).
fof(f8381,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f613,f7633,f1267])).
fof(f8425,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8384,f8402])).
fof(f8384,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f635,f7633,f1267])).
fof(f8424,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8385,f8402])).
fof(f8385,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f615,f7633,f1267])).
fof(f8423,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8389,f8402])).
fof(f8389,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f615,f7633,f1267])).
fof(f8422,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8392,f8402])).
fof(f8392,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f613,f7633,f1267])).
fof(f8421,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f8394,f8402])).
fof(f8394,plain,(
  p(s(bool,'const.quantHeuristics.GUESS_EXISTS_2'(s(fun(sK13,sK14),'_dst_x0_1'(s(sK14,sK15))),s(fun(sK14,bool),sK16)))) | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f634,f7633,f1267])).
fof(f8403,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f621,f7633,f1384])).
fof(f8393,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f613,f7633,f1267])).
fof(f8390,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f615,f7633,f1267])).
fof(f8386,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f615,f7633,f1267])).
fof(f8382,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f613,f7633,f1267])).
fof(f8374,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f607,f7633,f1177])).
fof(f8368,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f595,f7633,f1117])).
fof(f8362,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f589,f7633,f1089])).
fof(f8351,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f587,f7633,f583])).
fof(f8347,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f585,f7633,f583])).
fof(f8342,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f585,f7633,f583])).
fof(f8339,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f587,f7633,f583])).
fof(f8335,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f610,f6582,f7633,f565])).
fof(f8333,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f608,f6581,f7633,f565])).
fof(f8321,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f643,f7633,f562])).
fof(f8315,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f327,f843,f7633,f540])).
fof(f8308,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f3710,f7633,f538])).
fof(f8305,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f327,f6624,f7633,f537])).
fof(f8298,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f582,f6613,f7633,f537])).
fof(f8293,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6616,f7633,f536])).
fof(f6616,plain,(
  ~sP11(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'F_0','T_0') | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6371])).
fof(f6371,plain,(
  ~sP11(happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),'F_0','T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f582,f327,f2391,f533])).
fof(f8291,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f1506,f7633,f534])).
fof(f8283,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f327,f621,f7633,f533])).
fof(f8273,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f6613,f622,f7633,f532])).
fof(f8263,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f6613,f621,f7633,f532])).
fof(f8241,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f934,f7633,f527])).
fof(f8239,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f814,f7633,f527])).
fof(f8236,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f934,f7633,f526])).
fof(f8234,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6620,f7633,f526])).
fof(f8227,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f934,f7633,f526])).
fof(f8225,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f814,f7633,f526])).
fof(f8221,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f7633,f867,f7633,f523])).
fof(f8216,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f327,f863,f7633,f523])).
fof(f8210,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f7633,f863,f7633,f523])).
fof(f8206,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f327,f867,f7633,f523])).
fof(f8201,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f582,f1436,f7633,f520])).
fof(f8194,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f7633,f6604,f7633,f520])).
fof(f8190,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f582,f6603,f7633,f520])).
fof(f8183,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f7633,f6603,f7633,f520])).
fof(f8179,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f582,f6604,f7633,f520])).
fof(f8172,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f1403,f7633,f519])).
fof(f8167,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f7633,f615,f7633,f516])).
fof(f8162,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f327,f613,f7633,f516])).
fof(f8152,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f7633,f613,f7633,f516])).
fof(f8148,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f7633,f615,f7633,f516])).
fof(f8146,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f582,f615,f7633,f516])).
fof(f8136,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6604,f616,f7633,f515])).
fof(f8133,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6603,f616,f7633,f515])).
fof(f8129,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6604,f615,f7633,f515])).
fof(f8123,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6604,f613,f7633,f515])).
fof(f8120,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6603,f613,f7633,f515])).
fof(f8108,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6603,f616,f7633,f515])).
fof(f8105,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6604,f616,f7633,f515])).
fof(f8100,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6604,f613,f7633,f515])).
fof(f8095,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6603,f615,f7633,f515])).
fof(f8092,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6604,f615,f7633,f515])).
fof(f8053,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f6632,f7633,f510])).
fof(f6632,plain,(
  ( ! [X0] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),X0,'T_0')) ) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6415])).
fof(f6415,plain,(
  ( ! [X0] : (sP10(happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),X0,'T_0')) ) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f612,f2391,f564])).
fof(f8051,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f794,f7633,f510])).
fof(f8048,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f6632,f7633,f509])).
fof(f8046,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f794,f7633,f509])).
fof(f8043,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f327,f1147,f7633,f508])).
fof(f8039,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6582,f7633,f504])).
fof(f8037,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6581,f7633,f504])).
fof(f8033,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6581,f7633,f504])).
fof(f8031,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6600,f7633,f504])).
fof(f8025,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6582,f7633,f503])).
fof(f8023,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6581,f7633,f503])).
fof(f8019,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6582,f7633,f503])).
fof(f8017,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f1352,f7633,f503])).
fof(f8011,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6589,f7633,f502])).
fof(f6589,plain,(
  ~sP7(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','F_0') | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6260])).
fof(f6260,plain,(
  ~sP7(happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),'T_0','F_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f327,f582,f2391,f499])).
fof(f8009,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f1337,f7633,f501])).
fof(f8001,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f327,f607,f7633,f499])).
fof(f7991,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f6593,f7633,f494])).
fof(f6593,plain,(
  sP8('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6277])).
fof(f6277,plain,(
  sP8('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f327,f327,f2391,f505])).
fof(f7982,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f6594,f7633,f493])).
fof(f6594,plain,(
  sP8('T_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6283])).
fof(f6283,plain,(
  sP8('T_0','T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14)))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f327,f327,f2391,f505])).
fof(f8418,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(trivial_inequality_removal,[],[f7979])).
fof(f7979,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f1569,f7633,f491])).
fof(f8417,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(trivial_inequality_removal,[],[f7978])).
fof(f7978,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f1518,f7633,f491])).
fof(f7963,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f639,f6579,f7633,f488])).
fof(f7961,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f639,f6576,f7633,f488])).
fof(f7953,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f7633,f605,f7633,f484])).
fof(f7948,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f7633,f603,f7633,f484])).
fof(f7943,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f7633,f601,f7633,f484])).
fof(f7941,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f582,f601,f7633,f484])).
fof(f7933,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f7633,f605,f7633,f484])).
fof(f7927,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f7633,f601,f7633,f484])).
fof(f7923,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f7633,f603,f7633,f484])).
fof(f7921,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f582,f603,f7633,f484])).
fof(f7913,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f7633,f603,f7633,f484])).
fof(f7907,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f7633,f601,f7633,f484])).
fof(f7903,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f7633,f605,f7633,f484])).
fof(f7901,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f582,f605,f7633,f484])).
fof(f7893,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f1534,f7633,f483])).
fof(f7891,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f1319,f7633,f482])).
fof(f7883,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f327,f595,f7633,f480])).
fof(f7879,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f6550,f7633,f479])).
fof(f6550,plain,(
  ~sP5(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0','F_0') | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6185])).
fof(f6185,plain,(
  ~sP5(happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),'T_0','F_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f327,f582,f2391,f476])).
fof(f7877,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f1282,f7633,f478])).
fof(f7869,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f327,f589,f7633,f476])).
fof(f7843,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f602,f596,f6576,f590,f7633,f475])).
fof(f7807,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f602,f596,f6576,f588,f7633,f475])).
fof(f7760,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f603,f596,f6539,f590,f7633,f474])).
fof(f7745,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f603,f596,f6539,f589,f7633,f474])).
fof(f7713,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f327,f1001,f7633,f472])).
fof(f7709,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6536,f7633,f454])).
fof(f7705,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f625,f7633,f454])).
fof(f7700,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f625,f7633,f454])).
fof(f7697,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6536,f7633,f454])).
fof(f7693,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f6537,f7633,f453])).
fof(f7690,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f625,f7633,f453])).
fof(f7684,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f7633,f6537,f7633,f453])).
fof(f7681,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f625,f7633,f453])).
fof(f7677,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f6531,f7633,f419])).
fof(f6531,plain,(
  ~sP1(happ(s(fun(sK14,bool),sK16),s(sK14,sK15)),'T_0') | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6137])).
fof(f6137,plain,(
  ~sP1(happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14))),'T_0') | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f327,f2391,f415])).
fof(f7674,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f327,f6532,f7633,f416])).
fof(f6532,plain,(
  ~sP1('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK15))) | ($spl1 | $spl2 | $spl8)),
  inference(backward_demodulation,[],[f6128,f6139])).
fof(f6139,plain,(
  ~sP1('T_0',happ(s(fun(sK14,bool),sK16),s(sK14,sK19(sK16,sK14)))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f327,f2391,f415])).
fof(f7671,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f710,f7633,f414])).
fof(f7667,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f582,f686,f7633,f414])).
fof(f7656,plain,(
  $false | ($spl1 | $spl2 | $spl8)),
  inference(unit_resulting_resolution,[],[f643,f7633,f341])).
fof(f5044,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4888,f4889,f4890,f4891])).
fof(f4891,plain,(
  ( ! [X14,X12,X15,X13] : (sP6('T_0',happ(s(fun(X12,bool),X13),s(X12,'const.min.@_1'(s(fun(X12,bool),X13)))),happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X14,X15)))) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f2690])).
fof(f4890,plain,(
  ( ! [X14,X12,X15,X13] : (~sP6('F_0',happ(s(fun(X12,bool),X13),s(X12,'const.min.@_1'(s(fun(X12,bool),X13)))),happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X14,X15)))) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f2659])).
fof(f4889,plain,(
  ( ! [X14,X12,X15,X13] : (sP6('T_0',happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X14,X15)),happ(s(fun(X12,bool),X13),s(X12,'const.min.@_1'(s(fun(X12,bool),X13)))))) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f2626])).
fof(f4888,plain,(
  ( ! [X14,X12,X15,X13] : (~sP6('F_0',happ(s(fun(X14,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X14,X15)),happ(s(fun(X12,bool),X13),s(X12,'const.min.@_1'(s(fun(X12,bool),X13)))))) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f2590])).
fof(f5035,plain,(
  ( ! [X92,X90,X93,X91] : (~sP9(happ(s(fun(X90,bool),X91),s(X90,'const.min.@_1'(s(fun(X90,bool),X91)))),X92,X93) | sP10(happ(s(fun(X90,bool),X91),s(X90,'const.min.@_1'(s(fun(X90,bool),X91)))),X92,X93)) ) | $spl16),
  inference(subsumption_resolution,[],[f4885,f327])).
fof(f4885,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,'T_0')) | ~sP9(happ(s(fun(X90,bool),X91),s(X90,'const.min.@_1'(s(fun(X90,bool),X91)))),X92,X93) | sP10(happ(s(fun(X90,bool),X91),s(X90,'const.min.@_1'(s(fun(X90,bool),X91)))),X92,X93)) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1700])).
fof(f5033,plain,(
  ( ! [X88,X87,X89,X86] : (~sP11(happ(s(fun(X86,bool),X87),s(X86,'const.min.@_1'(s(fun(X86,bool),X87)))),X88,X89) | sP12(happ(s(fun(X86,bool),X87),s(X86,'const.min.@_1'(s(fun(X86,bool),X87)))),X88,X89)) ) | $spl16),
  inference(subsumption_resolution,[],[f4884,f327])).
fof(f4884,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,'T_0')) | ~sP11(happ(s(fun(X86,bool),X87),s(X86,'const.min.@_1'(s(fun(X86,bool),X87)))),X88,X89) | sP12(happ(s(fun(X86,bool),X87),s(X86,'const.min.@_1'(s(fun(X86,bool),X87)))),X88,X89)) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1699])).
fof(f5031,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X81)) | p(s(bool,X82)) | ~sP12(happ(s(fun(X79,bool),X80),s(X79,'const.min.@_1'(s(fun(X79,bool),X80)))),X82,X81)) ) | $spl16),
  inference(subsumption_resolution,[],[f4883,f327])).
fof(f4883,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~sP12(happ(s(fun(X79,bool),X80),s(X79,'const.min.@_1'(s(fun(X79,bool),X80)))),X82,X81)) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1697])).
fof(f5029,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X77)) | sP12(happ(s(fun(X75,bool),X76),s(X75,'const.min.@_1'(s(fun(X75,bool),X76)))),X77,X78)) ) | $spl16),
  inference(subsumption_resolution,[],[f4882,f327])).
fof(f4882,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X77)) | sP12(happ(s(fun(X75,bool),X76),s(X75,'const.min.@_1'(s(fun(X75,bool),X76)))),X77,X78)) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1696])).
fof(f5027,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,X73)) | sP12(happ(s(fun(X71,bool),X72),s(X71,'const.min.@_1'(s(fun(X71,bool),X72)))),X74,X73)) ) | $spl16),
  inference(subsumption_resolution,[],[f4881,f327])).
fof(f4881,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | p(s(bool,X73)) | sP12(happ(s(fun(X71,bool),X72),s(X71,'const.min.@_1'(s(fun(X71,bool),X72)))),X74,X73)) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1695])).
fof(f5025,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X69)) | ~p(s(bool,X70)) | ~sP11(X70,X69,happ(s(fun(X67,bool),X68),s(X67,'const.min.@_1'(s(fun(X67,bool),X68)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f4880,f327])).
fof(f4880,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X69)) | ~p(s(bool,X70)) | ~sP11(X70,X69,happ(s(fun(X67,bool),X68),s(X67,'const.min.@_1'(s(fun(X67,bool),X68)))))) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1694])).
fof(f5023,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X65)) | p(s(bool,X66)) | ~sP10(happ(s(fun(X63,bool),X64),s(X63,'const.min.@_1'(s(fun(X63,bool),X64)))),X66,X65)) ) | $spl16),
  inference(subsumption_resolution,[],[f4879,f327])).
fof(f4879,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X65)) | p(s(bool,X66)) | ~sP10(happ(s(fun(X63,bool),X64),s(X63,'const.min.@_1'(s(fun(X63,bool),X64)))),X66,X65)) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1693])).
fof(f5021,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X61)) | sP10(happ(s(fun(X59,bool),X60),s(X59,'const.min.@_1'(s(fun(X59,bool),X60)))),X61,X62)) ) | $spl16),
  inference(subsumption_resolution,[],[f4878,f327])).
fof(f4878,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | sP10(happ(s(fun(X59,bool),X60),s(X59,'const.min.@_1'(s(fun(X59,bool),X60)))),X61,X62)) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1692])).
fof(f5019,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X57)) | p(s(bool,X58)) | ~sP9(happ(s(fun(X55,bool),X56),s(X55,'const.min.@_1'(s(fun(X55,bool),X56)))),X57,X58)) ) | $spl16),
  inference(subsumption_resolution,[],[f4877,f327])).
fof(f4877,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X57)) | p(s(bool,X58)) | ~sP9(happ(s(fun(X55,bool),X56),s(X55,'const.min.@_1'(s(fun(X55,bool),X56)))),X57,X58)) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1691])).
fof(f5017,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X53)) | p(s(bool,X54)) | ~sP8(X54,X53,happ(s(fun(X51,bool),X52),s(X51,'const.min.@_1'(s(fun(X51,bool),X52)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f4876,f327])).
fof(f4876,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X53)) | p(s(bool,X54)) | ~sP8(X54,X53,happ(s(fun(X51,bool),X52),s(X51,'const.min.@_1'(s(fun(X51,bool),X52)))))) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1690])).
fof(f5015,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X49,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48)))),X50)) ) | $spl16),
  inference(subsumption_resolution,[],[f4875,f327])).
fof(f4875,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | p(s(bool,X50)) | ~sP7(X49,happ(s(fun(X47,bool),X48),s(X47,'const.min.@_1'(s(fun(X47,bool),X48)))),X50)) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1689])).
fof(f5013,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,happ(s(fun(X43,bool),X44),s(X43,'const.min.@_1'(s(fun(X43,bool),X44)))),X46)) ) | $spl16),
  inference(subsumption_resolution,[],[f4874,f327])).
fof(f4874,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | p(s(bool,X46)) | ~sP4(X45,happ(s(fun(X43,bool),X44),s(X43,'const.min.@_1'(s(fun(X43,bool),X44)))),X46)) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1688])).
fof(f5011,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X41)) | p(s(bool,X42)) | ~sP5(X41,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))),X42)) ) | $spl16),
  inference(subsumption_resolution,[],[f4873,f327])).
fof(f4873,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X41)) | p(s(bool,X42)) | ~sP5(X41,happ(s(fun(X39,bool),X40),s(X39,'const.min.@_1'(s(fun(X39,bool),X40)))),X42)) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1687])).
fof(f5009,plain,(
  ( ! [X37,X35,X38,X36] : (sP6(X37,X38,happ(s(fun(X35,bool),X36),s(X35,'const.min.@_1'(s(fun(X35,bool),X36))))) | ~sP5(X37,happ(s(fun(X35,bool),X36),s(X35,'const.min.@_1'(s(fun(X35,bool),X36)))),X38) | ~sP4(X38,happ(s(fun(X35,bool),X36),s(X35,'const.min.@_1'(s(fun(X35,bool),X36)))),X37) | ~sP3(happ(s(fun(X35,bool),X36),s(X35,'const.min.@_1'(s(fun(X35,bool),X36)))),X38,X37)) ) | $spl16),
  inference(subsumption_resolution,[],[f4872,f327])).
fof(f4872,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | sP6(X37,X38,happ(s(fun(X35,bool),X36),s(X35,'const.min.@_1'(s(fun(X35,bool),X36))))) | ~sP5(X37,happ(s(fun(X35,bool),X36),s(X35,'const.min.@_1'(s(fun(X35,bool),X36)))),X38) | ~sP4(X38,happ(s(fun(X35,bool),X36),s(X35,'const.min.@_1'(s(fun(X35,bool),X36)))),X37) | ~sP3(happ(s(fun(X35,bool),X36),s(X35,'const.min.@_1'(s(fun(X35,bool),X36)))),X38,X37)) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1686])).
fof(f5007,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X30)) | sP1(X30,happ(s(fun(X28,bool),X29),s(X28,'const.min.@_1'(s(fun(X28,bool),X29)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f4871,f327])).
fof(f4871,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X30)) | sP1(X30,happ(s(fun(X28,bool),X29),s(X28,'const.min.@_1'(s(fun(X28,bool),X29)))))) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1684])).
fof(f5005,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X27)) | sP1(happ(s(fun(X25,bool),X26),s(X25,'const.min.@_1'(s(fun(X25,bool),X26)))),X27)) ) | $spl16),
  inference(subsumption_resolution,[],[f4870,f327])).
fof(f4870,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X27)) | sP1(happ(s(fun(X25,bool),X26),s(X25,'const.min.@_1'(s(fun(X25,bool),X26)))),X27)) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1683])).
fof(f5003,plain,(
  ( ! [X24,X23,X22] : (sP0(happ(s(fun(X22,bool),X23),s(X22,'const.min.@_1'(s(fun(X22,bool),X23)))),X24)) ) | $spl16),
  inference(subsumption_resolution,[],[f4869,f327])).
fof(f4869,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X22,bool),X23),s(X22,'const.min.@_1'(s(fun(X22,bool),X23)))),X24)) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1682])).
fof(f5001,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,happ(s(fun(X19,bool),X20),s(X19,'const.min.@_1'(s(fun(X19,bool),X20)))))) ) | $spl16),
  inference(subsumption_resolution,[],[f4868,f327])).
fof(f4868,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X21,happ(s(fun(X19,bool),X20),s(X19,'const.min.@_1'(s(fun(X19,bool),X20)))))) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f1681])).
fof(f4980,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK27(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) ) | $spl16),
  inference(subsumption_resolution,[],[f4979,f327])).
fof(f4979,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK27(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) ) | $spl16),
  inference(forward_demodulation,[],[f4978,f4579])).
fof(f4978,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK27(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) ) | $spl16),
  inference(subsumption_resolution,[],[f4801,f327])).
fof(f4801,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X4,X0),X5),s(fun(X0,bool),X2)))) | s(X0,X6) = s(X0,happ(s(fun(X4,X0),X5),s(X4,sK27(X6,happ(s(fun(X1,fun(X0,bool)),'_dst_x0_1'(s(fun(X0,bool),X2))),s(X1,X3)),X5,X4,X0))))) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f2497])).
fof(f4976,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,sK27(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))) ) | $spl16),
  inference(subsumption_resolution,[],[f4975,f327])).
fof(f4975,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,sK27(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1))))) ) | $spl16),
  inference(forward_demodulation,[],[f4974,f4579])).
fof(f4974,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,sK27(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)))) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5))))) ) | $spl16),
  inference(subsumption_resolution,[],[f4800,f327])).
fof(f4800,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X6) = s(X1,happ(s(fun(X0,X1),X3),s(X0,sK27(X6,X5,happ(s(fun(X2,fun(X0,X1)),'_dst_x0_1'(s(fun(X0,X1),X3))),s(X2,X4)),X0,X1)))) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X0,X1),X3),s(fun(X1,bool),X5))))) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f2496])).
fof(f4970,plain,(
  ( ! [X6,X2,X0,X3,X1] : (s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X6,X3,X2,X1,X0))))) ) | $spl16),
  inference(subsumption_resolution,[],[f4969,f327])).
fof(f4969,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X6,X3,X2,X1,X0))))) ) | $spl16),
  inference(forward_demodulation,[],[f4968,f4579])).
fof(f4968,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X6,X3,X2,X1,X0))))) ) | $spl16),
  inference(subsumption_resolution,[],[f4790,f327])).
fof(f4790,plain,(
  ( ! [X6,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.quantHeuristics.GUESS_EXISTS_GAP_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3)))) | s(X0,X6) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK27(X6,X3,X2,X1,X0))))) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f433])).
fof(f4964,plain,(
  ( ! [X6,X0,X5] : (s(X0,X5) = s(X0,X6)) ) | $spl16),
  inference(subsumption_resolution,[],[f4963,f327])).
fof(f4963,plain,(
  ( ! [X6,X0,X5] : (~p(s(bool,'T_0')) | s(X0,X5) = s(X0,X6)) ) | $spl16),
  inference(forward_demodulation,[],[f4962,f4579])).
fof(f4962,plain,(
  ( ! [X6,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | s(X0,X5) = s(X0,X6)) ) | $spl16),
  inference(subsumption_resolution,[],[f4961,f327])).
fof(f4961,plain,(
  ( ! [X6,X0,X5,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | s(X0,X5) = s(X0,X6)) ) | $spl16),
  inference(forward_demodulation,[],[f4960,f4579])).
fof(f4960,plain,(
  ( ! [X6,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | s(X0,X5) = s(X0,X6)) ) | $spl16),
  inference(subsumption_resolution,[],[f4771,f327])).
fof(f4771,plain,(
  ( ! [X6,X0,X5,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | s(X0,X5) = s(X0,X6)) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f373])).
fof(f4740,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0','F_0')) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f2651])).
fof(f4733,plain,(
  ( ! [X0,X1] : (~sP3('F_0',happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f2579])).
fof(f4933,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl16),
  inference(subsumption_resolution,[],[f4704,f4579])).
fof(f4704,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f345])).
fof(f5043,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4888,f4889,f4890])).
fof(f5042,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4888,f4889])).
fof(f5041,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035,f4888])).
fof(f5036,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033,f5035])).
fof(f5034,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031,f5033])).
fof(f5032,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029,f5031])).
fof(f5030,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027,f5029])).
fof(f5028,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025,f5027])).
fof(f5026,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023,f5025])).
fof(f5024,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021,f5023])).
fof(f5022,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019,f5021])).
fof(f5020,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017,f5019])).
fof(f5018,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013,f5015,f5017])).
fof(f5016,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013,f5015])).
fof(f5014,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011,f5013])).
fof(f5012,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009,f5011])).
fof(f5010,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007,f5009])).
fof(f5008,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005,f5007])).
fof(f5006,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003,f5005])).
fof(f5004,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001,f5003])).
fof(f5002,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980,f5001])).
fof(f5000,plain,(
  $false | ($spl8 | $spl16)),
  inference(global_subsumption,[],[f4999,f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980])).
fof(f4999,plain,(
  ( ! [X5] : (s(sK14,sK15) = s(sK14,X5)) ) | ($spl8 | $spl16)),
  inference(subsumption_resolution,[],[f4847,f327])).
fof(f4847,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(sK14,sK15) = s(sK14,X5)) ) | ($spl8 | $spl16)),
  inference(backward_demodulation,[],[f4579,f581])).
fof(f4984,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f4983,f327])).
fof(f4983,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(forward_demodulation,[],[f4982,f4579])).
fof(f4982,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) ) | $spl16),
  inference(subsumption_resolution,[],[f4822,f327])).
fof(f4822,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.quantHeuristics.GUESS_FORALL_POINT_2'(s(fun(X1,X0),X2),s(fun(X0,bool),X3))))) ) | $spl16),
  inference(backward_demodulation,[],[f4579,f445])).
fof(f4981,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976,f4980])).
fof(f4977,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970,f4976])).
fof(f4971,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964,f4970])).
fof(f4965,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740,f4964])).
fof(f4949,plain,(
  $false | ($spl13 | $spl16)),
  inference(global_subsumption,[],[f4751,f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740])).
fof(f4751,plain,(
  ( ! [X45,X44] : (sP6('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X44,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X44,X45)))) ) | ($spl13 | $spl16)),
  inference(backward_demodulation,[],[f4579,f3648])).
fof(f3648,plain,(
  ( ! [X45,X44] : (sP6('T_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X44,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X44,X45)))) ) | $spl13),
  inference(superposition,[],[f1449,f1720])).
fof(f4948,plain,(
  $false | ($spl13 | $spl16)),
  inference(global_subsumption,[],[f4750,f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740])).
fof(f4750,plain,(
  ( ! [X43,X42] : (~sP6('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X42,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X42,X43)))) ) | ($spl13 | $spl16)),
  inference(backward_demodulation,[],[f4579,f3647])).
fof(f3647,plain,(
  ( ! [X43,X42] : (~sP6('F_0','|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X42,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X42,X43)))) ) | $spl13),
  inference(superposition,[],[f1448,f1720])).
fof(f4947,plain,(
  $false | ($spl13 | $spl16)),
  inference(global_subsumption,[],[f4749,f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740])).
fof(f4749,plain,(
  ( ! [X41,X40] : (sP6('T_0',happ(s(fun(X40,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X40,X41)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl13 | $spl16)),
  inference(backward_demodulation,[],[f4579,f3646])).
fof(f3646,plain,(
  ( ! [X41,X40] : (sP6('T_0',happ(s(fun(X40,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X40,X41)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl13),
  inference(superposition,[],[f1446,f1720])).
fof(f4946,plain,(
  $false | ($spl13 | $spl16)),
  inference(global_subsumption,[],[f4748,f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740])).
fof(f4748,plain,(
  ( ! [X39,X38] : (~sP6('F_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X38,X39)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | ($spl13 | $spl16)),
  inference(backward_demodulation,[],[f4579,f3645])).
fof(f3645,plain,(
  ( ! [X39,X38] : (~sP6('F_0',happ(s(fun(X38,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X38,X39)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) ) | $spl13),
  inference(superposition,[],[f1445,f1720])).
fof(f4945,plain,(
  $false | ($spl1 | $spl16)),
  inference(global_subsumption,[],[f4746,f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740])).
fof(f4746,plain,(
  ( ! [X8,X9] : (sP6('T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X8,X9)))) ) | ($spl1 | $spl16)),
  inference(backward_demodulation,[],[f4579,f2688])).
fof(f2688,plain,(
  ( ! [X8,X9] : (sP6('T_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X8,X9)))) ) | $spl1),
  inference(superposition,[],[f1449,f663])).
fof(f4944,plain,(
  $false | ($spl1 | $spl16)),
  inference(global_subsumption,[],[f4744,f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740])).
fof(f4744,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | ($spl1 | $spl16)),
  inference(backward_demodulation,[],[f4579,f2671])).
fof(f2671,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | $spl1),
  inference(forward_demodulation,[],[f2646,f663])).
fof(f2646,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))),s(X0,X1)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f582,f599,f592,f1448,f475])).
fof(f599,plain,(
  ( ! [X0,X1] : (sP4('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f574,f483])).
fof(f4943,plain,(
  $false | ($spl13 | $spl16)),
  inference(global_subsumption,[],[f4743,f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740])).
fof(f4743,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | ($spl13 | $spl16)),
  inference(backward_demodulation,[],[f4579,f2670])).
fof(f2670,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | $spl13),
  inference(forward_demodulation,[],[f2647,f1720])).
fof(f2647,plain,(
  ( ! [X0,X1] : (~sP3(happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))),s(X0,X1)),'|dollar|exists_1'(s(fun(sK14,bool),sK16)),'F_0')) ) | $spl13),
  inference(unit_resulting_resolution,[],[f582,f1782,f592,f1448,f475])).
fof(f4942,plain,(
  $false | ($spl1 | $spl16)),
  inference(global_subsumption,[],[f4742,f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740])).
fof(f4742,plain,(
  ( ! [X8,X9] : (~sP6('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X8,X9)))) ) | ($spl1 | $spl16)),
  inference(backward_demodulation,[],[f4579,f2657])).
fof(f2657,plain,(
  ( ! [X8,X9] : (~sP6('F_0','const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X8,X9)))) ) | $spl1),
  inference(superposition,[],[f1448,f663])).
fof(f4941,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733,f4740])).
fof(f4940,plain,(
  $false | ($spl1 | $spl16)),
  inference(global_subsumption,[],[f4739,f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733])).
fof(f4739,plain,(
  ( ! [X8,X9] : (sP6('T_0',happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X8,X9)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | ($spl1 | $spl16)),
  inference(backward_demodulation,[],[f4579,f2624])).
fof(f2624,plain,(
  ( ! [X8,X9] : (sP6('T_0',happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X8,X9)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl1),
  inference(superposition,[],[f1446,f663])).
fof(f4939,plain,(
  $false | ($spl1 | $spl16)),
  inference(global_subsumption,[],[f4737,f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733])).
fof(f4737,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) ) | ($spl1 | $spl16)),
  inference(backward_demodulation,[],[f4579,f2600])).
fof(f2600,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'F_0')) ) | $spl1),
  inference(forward_demodulation,[],[f2580,f663])).
fof(f2580,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))),s(X0,X1)),'F_0')) ) | $spl1),
  inference(unit_resulting_resolution,[],[f582,f597,f591,f1445,f475])).
fof(f591,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f574,f478])).
fof(f597,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f574,f482])).
fof(f4938,plain,(
  $false | ($spl13 | $spl16)),
  inference(global_subsumption,[],[f4736,f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733])).
fof(f4736,plain,(
  ( ! [X0,X1] : (~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X0,X1)),'F_0')) ) | ($spl13 | $spl16)),
  inference(backward_demodulation,[],[f4579,f2599])).
fof(f2599,plain,(
  ( ! [X0,X1] : (~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X0,X1)),'F_0')) ) | $spl13),
  inference(forward_demodulation,[],[f2581,f1720])).
fof(f2581,plain,(
  ( ! [X0,X1] : (~sP3('|dollar|exists_1'(s(fun(sK14,bool),sK16)),happ(s(fun(X0,bool),'_dst_x0_1'(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16))))),s(X0,X1)),'F_0')) ) | $spl13),
  inference(unit_resulting_resolution,[],[f582,f1781,f1774,f1445,f475])).
fof(f4937,plain,(
  $false | ($spl1 | $spl16)),
  inference(global_subsumption,[],[f4735,f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733])).
fof(f4735,plain,(
  ( ! [X8,X9] : (~sP6('F_0',happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'T_0'))),s(X8,X9)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | ($spl1 | $spl16)),
  inference(backward_demodulation,[],[f4579,f2588])).
fof(f2588,plain,(
  ( ! [X8,X9] : (~sP6('F_0',happ(s(fun(X8,bool),'_dst_x0_1'(s(bool,'F_0'))),s(X8,X9)),'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) ) | $spl1),
  inference(superposition,[],[f1445,f663])).
fof(f4936,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933,f4733])).
fof(f4935,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f4706,f327])).
fof(f4706,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(backward_demodulation,[],[f4579,f582])).
fof(f4934,plain,(
  $false | $spl16),
  inference(global_subsumption,[],[f326,f325,f324,f323,f322,f327,f582,f334,f342,f341,f346,f345,f369,f376,f375,f374,f373,f372,f386,f384,f389,f387,f394,f391,f401,f400,f583,f419,f416,f415,f414,f428,f430,f432,f431,f435,f434,f433,f438,f437,f436,f441,f440,f439,f444,f443,f442,f446,f445,f458,f457,f456,f566,f454,f453,f475,f474,f473,f472,f471,f470,f469,f479,f478,f477,f476,f483,f482,f481,f480,f487,f486,f485,f484,f491,f490,f489,f488,f565,f495,f494,f493,f492,f502,f501,f500,f499,f508,f505,f504,f503,f515,f514,f513,f564,f511,f510,f509,f519,f518,f517,f516,f523,f522,f521,f520,f532,f531,f530,f563,f528,f527,f526,f536,f535,f534,f533,f540,f539,f538,f537,f562,f584,f586,f588,f590,f592,f594,f596,f598,f600,f602,f604,f606,f608,f610,f612,f614,f616,f618,f620,f622,f624,f637,f636,f635,f634,f633,f632,f631,f630,f629,f628,f627,f626,f639,f644,f676,f677,f680,f689,f703,f713,f720,f746,f747,f750,f755,f758,f759,f762,f765,f768,f775,f776,f778,f785,f790,f795,f801,f802,f805,f810,f817,f822,f842,f848,f866,f864,f872,f890,f954,f991,f989,f1002,f999,f1017,f1013,f1030,f1027,f1078,f1088,f1106,f1116,f1121,f1129,f1142,f1148,f1146,f1166,f1176,f1181,f1207,f1241,f1255,f1266,f1271,f1315,f1360,f1373,f1383,f1388,f1431,f1449,f1448,f1446,f1445,f1462,f1471,f1472,f1473,f1474,f1476,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1488,f1489,f1533,f1525,f1519,f1517,f1543,f1540,f1583,f1570,f1568,f1636,f1628,f1657,f1667,f1666,f1665,f1664,f1669,f1670,f1671,f1681,f1682,f1683,f1684,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1699,f1700,f2056,f2053,f2057,f2055,f2066,f2063,f2067,f2065,f2077,f2074,f2076,f2090,f2087,f2089,f2104,f2100,f2121,f2117,f2139,f2135,f2167,f2163,f2194,f2210,f2211,f2222,f2235,f2232,f2233,f2234,f2248,f2261,f2258,f2259,f2260,f2302,f2303,f2304,f2309,f2334,f2335,f2336,f2341,f2342,f2388,f2384,f2416,f2412,f2443,f2444,f2438,f2439,f2440,f2442,f2459,f2460,f2454,f2455,f2456,f2461,f2458,f2496,f2497,f2515,f2516,f2545,f2547,f2561,f2563,f2582,f2579,f2604,f2585,f2586,f2587,f2590,f2591,f2618,f2621,f2622,f2623,f2626,f2627,f2651,f2668,f2648,f2654,f2655,f2656,f2659,f2682,f2685,f2686,f2687,f2690,f2714,f2713,f2711,f2710,f2723,f2732,f2733,f2734,f2735,f2737,f2738,f2739,f2740,f2741,f2742,f2743,f2744,f2745,f2746,f2747,f2748,f2750,f2751,f2752,f2753,f2754,f2755,f3077,f3075,f3073,f3072,f3071,f3069,f3067,f3066,f3063,f3060,f3059,f3058,f3053,f3052,f3046,f3043,f3042,f3035,f3034,f3026,f2996,f2995,f2994,f2989,f2988,f2983,f2982,f2981,f2974,f2964,f2958,f2954,f2953,f2948,f2947,f2946,f2945,f2944,f2940,f2938,f2932,f2931,f2924,f2923,f2915,f2912,f2911,f2903,f2898,f2897,f2896,f2894,f2892,f2891,f2889,f2887,f2886,f2885,f2883,f2881,f2880,f2879,f2878,f2877,f2870,f2869,f2861,f2860,f2853,f2852,f2844,f2843,f2839,f2831,f2827,f2824,f2823,f2817,f2816,f2811,f2796,f2795,f2794,f2793,f2791,f2790,f2789,f2788,f2787,f3495,f3493,f3491,f3490,f3489,f3487,f3485,f3484,f3481,f3478,f3477,f3476,f3471,f3470,f3464,f3461,f3460,f3453,f3452,f3444,f3414,f3413,f3412,f3407,f3406,f3401,f3400,f3399,f3392,f3382,f3376,f3372,f3371,f3366,f3365,f3364,f3363,f3362,f3358,f3356,f3350,f3349,f3342,f3341,f3333,f3330,f3329,f3321,f3316,f3315,f3314,f3312,f3310,f3309,f3307,f3305,f3304,f3303,f3301,f3299,f3298,f3297,f3296,f3295,f3288,f3287,f3279,f3278,f3271,f3270,f3262,f3261,f3257,f3249,f3245,f3242,f3241,f3235,f3234,f3229,f3214,f3213,f3212,f3211,f3209,f3208,f3207,f3206,f3205,f3717,f3739,f3791,f3792,f3794,f3810,f3821,f3822,f3824,f3869,f3870,f3871,f3872,f3924,f3926,f3927,f3928,f3929,f3930,f3931,f3932,f3933,f3934,f3935,f3936,f3937,f3938,f3939,f3940,f3941,f3942,f3943,f3944,f3945,f3946,f4055,f4577,f4581,f4585,f4588,f4591,f4594,f4597,f4600,f4603,f4606,f4609,f4612,f4615,f4618,f4621,f4624,f4084,f4085,f4086,f4087,f4626,f4627,f4628,f4629,f4630,f4631,f4632,f4633,f4634,f4635,f4636,f4637,f4638,f4639,f4640,f4641,f4642,f4643,f4644,f4645,f4132,f4133,f4134,f4135,f4646,f4647,f4648,f4649,f4650,f4652,f4653,f4654,f4655,f4656,f4657,f4658,f4659,f4660,f4661,f4662,f4663,f4664,f4665,f4666,f4185,f4187,f4188,f4189,f4190,f4191,f4192,f4193,f4194,f4195,f4196,f4197,f4198,f4199,f4200,f4201,f4202,f4203,f4204,f4205,f4206,f4207,f4672,f4674,f4677,f4679,f4681,f4683,f4685,f4687,f4689,f4691,f4693,f4695,f4697,f4699,f4701,f4703,f4579,f4933])).
fof(f4625,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~sP9('T_0',X1,X2)) ) | $spl46),
  inference(cnf_transformation,[],[f4625_D])).
fof(f4625_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~sP9('T_0',X1,X2)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4703,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f4702])).
fof(f4702,plain,(
  ( ! [X99,X97,X98,X96] : (~sP9('T_0',X97,X98) | sP10('T_0',X97,X98) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f4236,f954])).
fof(f4236,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X96)) | ~sP9('T_0',X97,X98) | sP10('T_0',X97,X98) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f564,f890])).
fof(f4622,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~sP11('T_0',X1,X2)) ) | $spl44),
  inference(cnf_transformation,[],[f4622_D])).
fof(f4622_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~sP11('T_0',X1,X2)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4701,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f4700])).
fof(f4700,plain,(
  ( ! [X94,X92,X95,X93] : (~sP11('T_0',X93,X94) | sP12('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f4235,f954])).
fof(f4235,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~sP11('T_0',X93,X94) | sP12('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f563,f890])).
fof(f4619,plain,(
  ( ! [X2,X1] : (~sP12('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl42),
  inference(cnf_transformation,[],[f4619_D])).
fof(f4619_D,plain,(
  ( ! [X2,X1] : (~sP12('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4699,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f4698])).
fof(f4698,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X86)) | p(s(bool,X87)) | ~sP12('T_0',X87,X86) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(subsumption_resolution,[],[f4233,f954])).
fof(f4233,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X85)) | ~p(s(bool,X86)) | p(s(bool,X87)) | ~sP12('T_0',X87,X86) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f540,f890])).
fof(f4616,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl40),
  inference(cnf_transformation,[],[f4616_D])).
fof(f4616_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4697,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f4696])).
fof(f4696,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X82)) | sP12('T_0',X82,X83) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(subsumption_resolution,[],[f4232,f954])).
fof(f4232,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X81)) | ~p(s(bool,X82)) | sP12('T_0',X82,X83) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f539,f890])).
fof(f4613,plain,(
  ( ! [X2,X1] : (sP12('T_0',X2,X1) | p(s(bool,X1))) ) | $spl38),
  inference(cnf_transformation,[],[f4613_D])).
fof(f4613_D,plain,(
  ( ! [X2,X1] : (sP12('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f4695,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f4694])).
fof(f4694,plain,(
  ( ! [X80,X78,X79,X77] : (p(s(bool,X78)) | sP12('T_0',X79,X78) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(subsumption_resolution,[],[f4231,f954])).
fof(f4231,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,X77)) | p(s(bool,X78)) | sP12('T_0',X79,X78) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f538,f890])).
fof(f4610,plain,(
  ( ! [X2,X1] : (~sP11(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl36),
  inference(cnf_transformation,[],[f4610_D])).
fof(f4610_D,plain,(
  ( ! [X2,X1] : (~sP11(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4693,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) )),
  inference(condensation,[],[f4692])).
fof(f4692,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X74)) | ~p(s(bool,X75)) | ~sP11(X75,X74,'T_0') | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(subsumption_resolution,[],[f4230,f954])).
fof(f4230,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,X73)) | p(s(bool,X74)) | ~p(s(bool,X75)) | ~sP11(X75,X74,'T_0') | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f533,f890])).
fof(f4607,plain,(
  ( ! [X2,X1] : (~sP10('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f4607_D])).
fof(f4607_D,plain,(
  ( ! [X2,X1] : (~sP10('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4691,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f4690])).
fof(f4690,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X70)) | p(s(bool,X71)) | ~sP10('T_0',X71,X70) | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f4229,f954])).
fof(f4229,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,X69)) | p(s(bool,X70)) | p(s(bool,X71)) | ~sP10('T_0',X71,X70) | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f523,f890])).
fof(f4604,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f4604_D])).
fof(f4604_D,plain,(
  ( ! [X2,X1] : (sP10('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f4689,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f4688])).
fof(f4688,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,X66)) | sP10('T_0',X66,X67) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f4228,f954])).
fof(f4228,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,X65)) | ~p(s(bool,X66)) | sP10('T_0',X66,X67) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f522,f890])).
fof(f4601,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl30),
  inference(cnf_transformation,[],[f4601_D])).
fof(f4601_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4687,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4686])).
fof(f4686,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,X62)) | p(s(bool,X63)) | ~sP9('T_0',X62,X63) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f4227,f954])).
fof(f4227,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | p(s(bool,X63)) | ~sP9('T_0',X62,X63) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f516,f890])).
fof(f4598,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f4598_D])).
fof(f4598_D,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f4685,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f4684])).
fof(f4684,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X58)) | p(s(bool,X59)) | ~sP8(X59,X58,'T_0') | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f4226,f954])).
fof(f4226,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X57)) | ~p(s(bool,X58)) | p(s(bool,X59)) | ~sP8(X59,X58,'T_0') | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f508,f890])).
fof(f4595,plain,(
  ( ! [X2,X1] : (~sP7(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl26),
  inference(cnf_transformation,[],[f4595_D])).
fof(f4595_D,plain,(
  ( ! [X2,X1] : (~sP7(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4683,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X1,'T_0',X2)) )),
  inference(condensation,[],[f4682])).
fof(f4682,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | p(s(bool,X55)) | ~sP7(X54,'T_0',X55) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f4225,f954])).
fof(f4225,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | p(s(bool,X55)) | ~sP7(X54,'T_0',X55) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f499,f890])).
fof(f4592,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl24),
  inference(cnf_transformation,[],[f4592_D])).
fof(f4592_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4681,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f4680])).
fof(f4680,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP4(X50,'T_0',X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f4224,f954])).
fof(f4224,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP4(X50,'T_0',X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f480,f890])).
fof(f4589,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl22),
  inference(cnf_transformation,[],[f4589_D])).
fof(f4589_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f4679,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f4678])).
fof(f4678,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X46)) | p(s(bool,X47)) | ~sP5(X46,'T_0',X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f4223,f954])).
fof(f4223,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP5(X46,'T_0',X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f476,f890])).
fof(f4586,plain,(
  ( ! [X2,X1] : (~sP4(X2,'T_0',X1) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2)) ) | $spl20),
  inference(cnf_transformation,[],[f4586_D])).
fof(f4586_D,plain,(
  ( ! [X2,X1] : (~sP4(X2,'T_0',X1) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f4677,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f4676,f604])).
fof(f4676,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1) | ~sP3('T_0',X2,X1)) )),
  inference(condensation,[],[f4675])).
fof(f4675,plain,(
  ( ! [X43,X41,X44,X42] : (sP6(X42,X43,'T_0') | ~sP5(X42,'T_0',X43) | ~sP4(X43,'T_0',X42) | ~sP3('T_0',X43,X42) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f4222,f954])).
fof(f4222,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X41)) | sP6(X42,X43,'T_0') | ~sP5(X42,'T_0',X43) | ~sP4(X43,'T_0',X42) | ~sP3('T_0',X43,X42) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f473,f890])).
fof(f4582,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl18),
  inference(cnf_transformation,[],[f4582_D])).
fof(f4582_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4674,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4673])).
fof(f4673,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X35)) | sP1(X35,'T_0') | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f4220,f954])).
fof(f4220,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,X34)) | p(s(bool,X35)) | sP1(X35,'T_0') | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f419,f890])).
fof(f4578,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl14),
  inference(cnf_transformation,[],[f4578_D])).
fof(f4578_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f4672,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4671])).
fof(f4671,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X32)) | sP1('T_0',X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4219,f954])).
fof(f4219,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | sP1('T_0',X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f416,f890])).
fof(f4624,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP9('T_0',X1,X2) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f4623])).
fof(f4623,plain,(
  ( ! [X99,X97,X98,X96] : (~sP9('T_0',X97,X98) | sP10('T_0',X97,X98) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(subsumption_resolution,[],[f4082,f954])).
fof(f4082,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,X96)) | ~sP9('T_0',X97,X98) | sP10('T_0',X97,X98) | s(bool,X96) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f564,f890])).
fof(f4621,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP11('T_0',X1,X2) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f4620])).
fof(f4620,plain,(
  ( ! [X94,X92,X95,X93] : (~sP11('T_0',X93,X94) | sP12('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(subsumption_resolution,[],[f4081,f954])).
fof(f4081,plain,(
  ( ! [X94,X92,X95,X93] : (~p(s(bool,X92)) | ~sP11('T_0',X93,X94) | sP12('T_0',X93,X94) | s(bool,X92) = s(bool,X95) | s(bool,'T_0') = s(bool,X95)) )),
  inference(superposition,[],[f563,f890])).
fof(f4618,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f4617])).
fof(f4617,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X86)) | p(s(bool,X87)) | ~sP12('T_0',X87,X86) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(subsumption_resolution,[],[f4079,f954])).
fof(f4079,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,X85)) | ~p(s(bool,X86)) | p(s(bool,X87)) | ~sP12('T_0',X87,X86) | s(bool,X85) = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f540,f890])).
fof(f4615,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP12('T_0',X1,X2)) )),
  inference(condensation,[],[f4614])).
fof(f4614,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X82)) | sP12('T_0',X82,X83) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(subsumption_resolution,[],[f4078,f954])).
fof(f4078,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,X81)) | ~p(s(bool,X82)) | sP12('T_0',X82,X83) | s(bool,X81) = s(bool,X84) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f539,f890])).
fof(f4612,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP12('T_0',X2,X1)) )),
  inference(condensation,[],[f4611])).
fof(f4611,plain,(
  ( ! [X80,X78,X79,X77] : (p(s(bool,X78)) | sP12('T_0',X79,X78) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(subsumption_resolution,[],[f4077,f954])).
fof(f4077,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,X77)) | p(s(bool,X78)) | sP12('T_0',X79,X78) | s(bool,X77) = s(bool,X80) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f538,f890])).
fof(f4609,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP11(X2,X1,'T_0')) )),
  inference(condensation,[],[f4608])).
fof(f4608,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X74)) | ~p(s(bool,X75)) | ~sP11(X75,X74,'T_0') | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(subsumption_resolution,[],[f4076,f954])).
fof(f4076,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,X73)) | p(s(bool,X74)) | ~p(s(bool,X75)) | ~sP11(X75,X74,'T_0') | s(bool,X73) = s(bool,X76) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f533,f890])).
fof(f4606,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP10('T_0',X2,X1)) )),
  inference(condensation,[],[f4605])).
fof(f4605,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,X70)) | p(s(bool,X71)) | ~sP10('T_0',X71,X70) | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f4075,f954])).
fof(f4075,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,X69)) | p(s(bool,X70)) | p(s(bool,X71)) | ~sP10('T_0',X71,X70) | s(bool,X69) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f523,f890])).
fof(f4603,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP10('T_0',X1,X2)) )),
  inference(condensation,[],[f4602])).
fof(f4602,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,X66)) | sP10('T_0',X66,X67) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(subsumption_resolution,[],[f4074,f954])).
fof(f4074,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,X65)) | ~p(s(bool,X66)) | sP10('T_0',X66,X67) | s(bool,X65) = s(bool,X68) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f522,f890])).
fof(f4600,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4599])).
fof(f4599,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,X62)) | p(s(bool,X63)) | ~sP9('T_0',X62,X63) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f4073,f954])).
fof(f4073,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,X61)) | p(s(bool,X62)) | p(s(bool,X63)) | ~sP9('T_0',X62,X63) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f516,f890])).
fof(f4597,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f4596])).
fof(f4596,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X58)) | p(s(bool,X59)) | ~sP8(X59,X58,'T_0') | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f4072,f954])).
fof(f4072,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,X57)) | ~p(s(bool,X58)) | p(s(bool,X59)) | ~sP8(X59,X58,'T_0') | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f508,f890])).
fof(f4594,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X1,'T_0',X2)) )),
  inference(condensation,[],[f4593])).
fof(f4593,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X54)) | p(s(bool,X55)) | ~sP7(X54,'T_0',X55) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(subsumption_resolution,[],[f4071,f954])).
fof(f4071,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,X54)) | p(s(bool,X55)) | ~sP7(X54,'T_0',X55) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f499,f890])).
fof(f4591,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f4590])).
fof(f4590,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X50)) | p(s(bool,X51)) | ~sP4(X50,'T_0',X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f4070,f954])).
fof(f4070,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | p(s(bool,X51)) | ~sP4(X50,'T_0',X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f480,f890])).
fof(f4588,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f4587])).
fof(f4587,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X46)) | p(s(bool,X47)) | ~sP5(X46,'T_0',X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f4069,f954])).
fof(f4069,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP5(X46,'T_0',X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f476,f890])).
fof(f4585,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f4584,f604])).
fof(f4584,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP6(X1,X2,'T_0') | ~sP5(X1,'T_0',X2) | ~sP4(X2,'T_0',X1) | ~sP3('T_0',X2,X1)) )),
  inference(condensation,[],[f4583])).
fof(f4583,plain,(
  ( ! [X43,X41,X44,X42] : (sP6(X42,X43,'T_0') | ~sP5(X42,'T_0',X43) | ~sP4(X43,'T_0',X42) | ~sP3('T_0',X43,X42) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f4068,f954])).
fof(f4068,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X41)) | sP6(X42,X43,'T_0') | ~sP5(X42,'T_0',X43) | ~sP4(X43,'T_0',X42) | ~sP3('T_0',X43,X42) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f473,f890])).
fof(f4581,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4580])).
fof(f4580,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X35)) | sP1(X35,'T_0') | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(subsumption_resolution,[],[f4066,f954])).
fof(f4066,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,X34)) | p(s(bool,X35)) | sP1(X35,'T_0') | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f419,f890])).
fof(f4577,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4576])).
fof(f4576,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X32)) | sP1('T_0',X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f4065,f954])).
fof(f4065,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,X31)) | p(s(bool,X32)) | sP1('T_0',X32) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f416,f890])).
fof(f1705,plain,(
  s(sK14,sK15) = s(sK14,'const.min.@_1'(s(fun(sK14,bool),sK16))) | $spl10),
  inference(cnf_transformation,[],[f1705_D])).
fof(f1705_D,plain,(
  s(sK14,sK15) = s(sK14,'const.min.@_1'(s(fun(sK14,bool),sK16))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1707,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | $spl13),
  inference(cnf_transformation,[],[f1707_D])).
fof(f1707_D,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f1672,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK14,bool),sK16)))) | s(sK14,sK15) = s(sK14,'const.min.@_1'(s(fun(sK14,bool),sK16))) | $spl8),
  inference(superposition,[],[f581,f369])).
fof(f323,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))),
  inference(cnf_transformation,[],[f208])).
fof(f581,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X5)))) | s(sK14,sK15) = s(sK14,X5)) ) | $spl8),
  inference(cnf_transformation,[],[f581_D])).
fof(f581_D,plain,(
  ( ! [X5] : (~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X5)))) | s(sK14,sK15) = s(sK14,X5)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f324,plain,(
  ( ! [X5] : (s(sK14,sK15) = s(sK14,X5) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,X5)))) | p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))) )),
  inference(cnf_transformation,[],[f208])).
fof(f574,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) | $spl1),
  inference(cnf_transformation,[],[f574_D])).
fof(f574_D,plain,(
  ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f576,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | $spl3),
  inference(cnf_transformation,[],[f576_D])).
fof(f576_D,plain,(
  ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f580,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | $spl6),
  inference(cnf_transformation,[],[f580_D])).
fof(f580_D,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f325,plain,(
  p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK17)))) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))),
  inference(cnf_transformation,[],[f208])).
fof(f578,plain,(
  s(sK14,sK15) != s(sK14,sK17) | $spl5),
  inference(cnf_transformation,[],[f578_D])).
fof(f578_D,plain,(
  s(sK14,sK15) != s(sK14,sK17) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f326,plain,(
  s(sK14,sK15) != s(sK14,sK17) | ~p(s(bool,happ(s(fun(sK14,bool),sK16),s(sK14,sK15)))) | ~p(s(bool,'const.bool.?|33|_1'(s(fun(sK14,bool),sK16))))),
  inference(cnf_transformation,[],[f208])).
% SZS output end Proof for thm.quantHeuristics
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 2094
% Passive clauses: 9264
% Generated clauses: 58036
% Final active clauses: 1419
% Final passive clauses: 3131
% Input formulas: 56
% Initial clauses: 229
% 
% Duplicate literals: 412
% Trivial inequalities: 49
% Fw subsumption resolutions: 1925
% Fw demodulations: 9352
% Bw demodulations: 3128
% Condensations: 1033
% Global subsumptions: 45
% 
% Simple tautologies: 117
% Equational tautologies: 90
% Forward subsumptions: 34377
% Bw demodulations to eq. taut.: 12
% 
% Unit resulting resolution: 30614
% Factoring: 1
% Forward superposition: 7605
% Backward superposition: 3292
% Self superposition: 29
% Equality factoring: 28
% Equality resolution: 1
% 
% Split clauses: 109
% Split components: 225
% SAT solver clauses: 10792
% SAT solver unit clauses: 687
% SAT solver binary clauses: 3974
% 
% Sat splits: 109
% Sat splitting refutations: 1263
% 
% Memory used [KB]: 31214
% Time elapsed: 3.037 s
% ------------------------------
% ------------------------------
% Success in time 4.366 s
