Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.rich_list
% (60902)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 112
% Passive clauses: 3245
% Generated clauses: 12867
% Final active clauses: 106
% Final passive clauses: 2978
% Input formulas: 20
% Initial clauses: 46
% 
% Duplicate literals: 606
% Trivial inequalities: 233
% Fw subsumption resolutions: 2282
% Fw demodulations: 621
% Bw demodulations: 4
% 
% Simple tautologies: 362
% Equational tautologies: 79
% Forward subsumptions: 4128
% 
% Binary resolution: 588
% Factoring: 2
% Forward superposition: 3607
% Backward superposition: 4288
% Self superposition: 222
% Equality factoring: 192
% Equality resolution: 28
% 
% Split clauses: 567
% Split components: 1165
% SAT solver clauses: 1414
% SAT solver unit clauses: 89
% SAT solver binary clauses: 448
% 
% Sat splits: 567
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1411
% TWLsolver calls for satisfiability: 193
% 
% Memory used [KB]: 5500
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rich_list
% (60915)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 172
% Passive clauses: 1439
% Generated clauses: 18973
% Final active clauses: 77
% Final passive clauses: 331
% Input formulas: 20
% Initial clauses: 46
% 
% Duplicate literals: 422
% Trivial inequalities: 178
% Fw subsumption resolutions: 1921
% Fw demodulations: 2373
% Bw demodulations: 838
% Condensations: 656
% 
% Simple tautologies: 182
% Equational tautologies: 31
% Forward subsumptions: 7209
% Bw demodulations to eq. taut.: 136
% 
% Binary resolution: 470
% Factoring: 2
% Forward superposition: 2835
% Backward superposition: 2791
% Self superposition: 87
% Equality factoring: 46
% Equality resolution: 14
% 
% Split clauses: 427
% Split components: 870
% SAT solver clauses: 2196
% SAT solver unit clauses: 40
% SAT solver binary clauses: 46
% SAT solver learnt clauses: 6
% SAT solver learnt literals: 16
% 
% Sat splits: 427
% Sat splitting refutations: 101
% 
% TWLsolver clauses: 2186
% TWLsolver calls for satisfiability: 445
% 
% Memory used [KB]: 10106
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rich_list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rich_list
% SZS output start Proof for thm.rich_list
fof(f17367,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f152,f162,f160,f1301,f1304,f1303,f1455,f1457,f2529,f2531,f2530,f2521,f2532,f2540,f2541,f2507,f2542,f2503,f2544,f2556,f2558,f2557,f2782,f2783,f2784,f2785,f2786,f2789,f2793,f2797,f2798,f2799,f2800,f2801,f2802,f2803,f2804,f2805,f2807,f2809,f2810,f2811,f2812,f2813,f2814,f2815,f2816,f3100,f3101,f3102,f3103,f3104,f3173,f3176,f3786,f3789,f3787,f4297,f4292,f4299,f4300,f4298,f4279,f4310,f4278,f4311,f4277,f4312,f4276,f4313,f4315,f4316,f4319,f4320,f4322,f4324,f4323,f7032,f7033,f7925,f7908,f7931,f7907,f7933,f7930,f7939,f7940,f7881,f7947,f7880,f7948,f7879,f7949,f7878,f7950,f7858,f7932,f7952,f7953,f7957,f7958,f7962,f7964,f7963,f8888,f8028,f8915,f8913,f8051,f8862,f8845,f8844,f8827,f8826,f8825,f8824,f8793,f8912,f8791,f8926,f8790,f8789,f8955,f7034,f8956,f8959,f8961,f8960,f8972,f8973,f8976,f8977,f8752,f8923,f8745,f8984,f8744,f8985,f8743,f8986,f8742,f8987,f8722,f8721,f8719,f8988,f8685,f8924,f8684,f8925,f7786,f9028,f5585,f5435,f5432,f5330,f5327,f5222,f3112,f9029,f5072,f9030,f5069,f9031,f3106,f4967,f9032,f4964,f3113,f9034,f4859,f4754,f9035,f3107,f3110,f4221,f9036,f9037,f3175,f3174,f9033,f3172,f3109,f9027,f4757,f4862,f9038,f3183,f5588,f5225,f3185,f9039,f9113,f9114,f9115,f9116,f9117,f9118,f9119,f9120,f9121,f9122,f9123,f9124,f9125,f9126,f9127,f9128,f9129,f9130,f9131,f9132,f9133,f9135,f9138,f9141,f9148,f9153,f9778,f9780,f9786,f9787,f11320,f11321,f11334,f11344,f11343,f11335,f11346,f11345,f13543,f13544,f13547,f8919,f13553,f13555,f13558,f13561,f13564,f13567,f13569,f13571,f17068,f17051,f17050,f17074,f17071,f17044,f17075,f17020,f17085,f17019,f17086,f17018,f17087,f17017,f17088,f16997,f17072,f16996,f17073,f17090,f17092,f17091,f17098,f17099,f17101,f17103,f17102,f17146,f17147,f8918,f17253,f17341,f17342,f17343,f1302])).
fof(f1302,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl4),
  inference(cnf_transformation,[],[f1302_D])).
fof(f1302_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f17343,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl178),
  inference(cnf_transformation,[],[f17343_D])).
fof(f17343_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f17342,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f17312,f138])).
fof(f138,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f88,f83])).
fof(f83,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.rich_list.MEM_COUNT_LIST',aHL_BOOL_CASES)).
fof(f88,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.rich_list.MEM_COUNT_LIST','thm.bool.EQ_CLAUSES|split|3')).
fof(f17312,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s(bool,X2) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f3521,f913])).
fof(f913,plain,(
  ( ! [X59,X60,X58] : (s('type.num.num',X58) = s('type.num.num','const.list.EL_2'(s('type.num.num',X58),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X59))))) | s(bool,X60) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X58),s('type.num.num',X59))) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f691,f79])).
fof(f79,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rich_list.MEM_COUNT_LIST',aHL_TRUTH)).
fof(f691,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'T_0')) | s('type.num.num',X58) = s('type.num.num','const.list.EL_2'(s('type.num.num',X58),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X59))))) | s(bool,X60) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X58),s('type.num.num',X59))) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f104,f136])).
fof(f136,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f83,f83])).
fof(f104,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num','const.list.EL_2'(s('type.num.num',X0),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s('type.num.num',X0) = s('type.num.num','const.list.EL_2'(s('type.num.num',X0),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X1))))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) => s('type.num.num',X0) = s('type.num.num','const.list.EL_2'(s('type.num.num',X0),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X1))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7)))) => s('type.num.num',X6) = s('type.num.num','const.list.EL_2'(s('type.num.num',X6),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X7))))))),
  file('thm.rich_list.MEM_COUNT_LIST','thm.rich_list.EL_COUNT_LIST')).
fof(f3521,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s('type.num.num','const.list.EL_2'(s('type.num.num',X0),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X1))))),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(equality_resolution,[],[f490])).
fof(f490,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X2) != s('type.num.num','const.list.EL_2'(s('type.num.num',X1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X0))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.bool.IN_2'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X0))))))))) )),
  inference(superposition,[],[f100,f82])).
fof(f82,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X7)))))),
  file('thm.rich_list.MEM_COUNT_LIST','thm.rich_list.LENGTH_COUNT_LIST')).
fof(f100,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) | s(X0,X2) != s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) & s(X0,X2) = s(X0,'const.list.EL_2'(s('type.num.num',sK3(X2,X1,X0)),s('type.list.list'(X0),X1))))) & (! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) | s(X0,X2) != s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | ? [X4] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) & s(X0,X2) = s(X0,'const.list.EL_2'(s('type.num.num',X4),s('type.list.list'(X0),X1))))) & (! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) | s(X0,X2) != s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))))),
  inference(rectify,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | ? [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) & s(X0,X2) = s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1))))) & (! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) | s(X0,X2) != s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) <=> ? [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) & s(X0,X2) = s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1)))))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) <=> ? [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1)))))) & s(X0,X2) = s(X0,'const.list.EL_2'(s('type.num.num',X3),s('type.list.list'(X0),X1)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X10] : ! [X20,X12] : (p(s(bool,'const.bool.IN_2'(s(X10,X12),s(fun(X10,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X10),X20)))))) <=> ? [X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X10),X20)))))) & s(X10,X12) = s(X10,'const.list.EL_2'(s('type.num.num',X7),s('type.list.list'(X10),X20)))))),
  file('thm.rich_list.MEM_COUNT_LIST','thm.list.MEM_EL')).
fof(f17341,plain,(
  $false | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f17340,f79])).
fof(f17340,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f17339,f2830])).
fof(f2830,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | $spl3),
  inference(unit_resulting_resolution,[],[f81,f162,f203])).
fof(f203,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f195,f81])).
fof(f195,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f105,f83])).
fof(f105,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.rich_list.MEM_COUNT_LIST','thm.bool.IMP_ANTISYM_AX')).
fof(f81,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f23])).
fof(f23,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.rich_list.MEM_COUNT_LIST',aHL_FALSITY)).
fof(f17339,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)))) | ($spl1 | $spl3)),
  inference(subsumption_resolution,[],[f17338,f81])).
fof(f17338,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)))) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f17311,f2913])).
fof(f2913,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f81,f159,f203])).
fof(f159,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | $spl1),
  inference(cnf_transformation,[],[f159_D])).
fof(f159_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f17311,plain,(
  p(s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)))) | $spl3),
  inference(superposition,[],[f3521,f2898])).
fof(f2898,plain,(
  s('type.num.num',sK1) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))) | $spl3),
  inference(subsumption_resolution,[],[f2869,f79])).
fof(f2869,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK1) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))) | $spl3),
  inference(superposition,[],[f104,f2830])).
fof(f17253,plain,(
  ( ! [X14,X15,X18] : (s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,X18) | s(bool,X18) = s(bool,sK5(X14,X15,bool,bool)) | s(bool,'T_0') = s(bool,f(s(bool,sK5(X14,X15,bool,bool))))) ) | $spl176),
  inference(cnf_transformation,[],[f17253_D])).
fof(f17253_D,plain,(
  ( ! [X14,X15,X18] : (s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,X18) | s(bool,X18) = s(bool,sK5(X14,X15,bool,bool)) | s(bool,'T_0') = s(bool,f(s(bool,sK5(X14,X15,bool,bool))))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f8918,plain,(
  ( ! [X70,X68] : (s(bool,X68) = s(bool,X70) | s(bool,'T_0') = s(bool,g(s(bool,X68))) | s(bool,'T_0') = s(bool,X70)) ) | $spl90),
  inference(cnf_transformation,[],[f8918_D])).
fof(f8918_D,plain,(
  ( ! [X70,X68] : (s(bool,X68) = s(bool,X70) | s(bool,'T_0') = s(bool,g(s(bool,X68))) | s(bool,'T_0') = s(bool,X70)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f17147,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,g(s(bool,X21))) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,'T_0') = s(bool,f(s(bool,sK5(X19,X20,bool,bool)))) | s(bool,X22) = s(bool,sK5(X19,X20,bool,bool)) | s(bool,'T_0') = s(bool,X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f3182,f535])).
fof(f535,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f136,f136])).
fof(f3182,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK5(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(X0,sK5(X1,X2,bool,X0))))) )),
  inference(trivial_inequality_removal,[],[f3177])).
fof(f3177,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK5(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,f(s(X0,sK5(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f356,f83])).
fof(f356,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK5(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK5(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f109,f83])).
fof(f109,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK5(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK5(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.rich_list.MEM_COUNT_LIST',aHL_EXT)).
fof(f17146,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,g(s(bool,X16))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,'T_0') = s(bool,f(s(bool,sK5(X14,X15,bool,bool)))) | s(bool,X16) = s(bool,X17) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,sK5(X14,X15,bool,bool)) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f3182,f535])).
fof(f17102,plain,(
  ( ! [X39,X41,X38] : (p(s(bool,X41)) | s('type.num.num',X39) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X39)))),X38,'type.num.num')) | s(bool,X41) = s(bool,happ(s(fun('type.num.num',bool),X38),s('type.num.num',X39)))) ) | $spl174),
  inference(cnf_transformation,[],[f17102_D])).
fof(f17102_D,plain,(
  ( ! [X39,X41,X38] : (p(s(bool,X41)) | s('type.num.num',X39) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X39)))),X38,'type.num.num')) | s(bool,X41) = s(bool,happ(s(fun('type.num.num',bool),X38),s('type.num.num',X39)))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f17103,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,X45)) | s(bool,X45) = s(bool,happ(s(fun('type.num.num',bool),X42),s('type.num.num',X43))) | s('type.num.num',X43) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X43)))),X42,'type.num.num')) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f16925,f85])).
fof(f85,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.rich_list.MEM_COUNT_LIST','thm.bool.EQ_CLAUSES|split|1')).
fof(f16925,plain,(
  ( ! [X45,X43,X44,X42] : (p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X45) = s(bool,happ(s(fun('type.num.num',bool),X42),s('type.num.num',X43))) | s('type.num.num',X43) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X43)))),X42,'type.num.num')) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f105,f2940])).
fof(f2940,plain,(
  ( ! [X4,X2,X3] : (s('type.num.num',X3) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X3)))),X2,'type.num.num')) | s(bool,X4) = s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f2935,f79])).
fof(f2935,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s('type.num.num',X3) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X3)))),X2,'type.num.num')) | s(bool,X4) = s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f297,f136])).
fof(f297,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X0)))) | s('type.num.num',X0) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X0)))),X1,'type.num.num'))) )),
  inference(forward_demodulation,[],[f291,f82])).
fof(f291,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X0)))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X0))))) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X0)))),X1,'type.num.num'))) )),
  inference(superposition,[],[f102,f82])).
fof(f102,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) = s(X0,sK4(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | (s(X0,X2) = s(X0,sK4(X2,X1,X0)) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2] : ((! [X4] : (s(X0,X2) != s(X0,X4) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : ((! [X3] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (s(X0,X2) = s(X0,X3) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : ! [X15,X19] : (? [X12] : (s(X10,X12) = s(X10,X19) & p(s(bool,happ(s(fun(X10,bool),X15),s(X10,X12))))) <=> p(s(bool,happ(s(fun(X10,bool),X15),s(X10,X19)))))),
  file('thm.rich_list.MEM_COUNT_LIST','thm.bool.UNWIND_THM2')).
fof(f17101,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X41)) | s(bool,X41) = s(bool,happ(s(fun('type.num.num',bool),X38),s('type.num.num',X39))) | s('type.num.num',X39) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X39)))),X38,'type.num.num')) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f16924,f85])).
fof(f16924,plain,(
  ( ! [X39,X41,X38,X40] : (p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X41) = s(bool,happ(s(fun('type.num.num',bool),X38),s('type.num.num',X39))) | s('type.num.num',X39) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X39)))),X38,'type.num.num')) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f105,f2940])).
fof(f17099,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun('type.num.num',bool),X20),s('type.num.num',X21)))) | s('type.num.num',sK4(X21,X20,'type.num.num')) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK4(X21,X20,'type.num.num'))))),X20,'type.num.num'))) ) | $spl172),
  inference(cnf_transformation,[],[f17099_D])).
fof(f17099_D,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun('type.num.num',bool),X20),s('type.num.num',X21)))) | s('type.num.num',sK4(X21,X20,'type.num.num')) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK4(X21,X20,'type.num.num'))))),X20,'type.num.num'))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f17098,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,happ(s(fun('type.num.num',bool),X20),s('type.num.num',X21)))) | s('type.num.num',sK4(X21,X20,'type.num.num')) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK4(X21,X20,'type.num.num'))))),X20,'type.num.num')) | s(bool,'T_0') = s(bool,X22)) )),
  inference(subsumption_resolution,[],[f16918,f85])).
fof(f16918,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X22)) | ~p(s(bool,happ(s(fun('type.num.num',bool),X20),s('type.num.num',X21)))) | s('type.num.num',sK4(X21,X20,'type.num.num')) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK4(X21,X20,'type.num.num'))))),X20,'type.num.num')) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f103,f2940])).
fof(f103,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f17091,plain,(
  ( ! [X175,X176] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X175),s('type.num.num',X176))) | s('type.num.num',X176) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X176)))),X175,'type.num.num'))) ) | $spl170),
  inference(cnf_transformation,[],[f17091_D])).
fof(f17091_D,plain,(
  ( ! [X175,X176] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X175),s('type.num.num',X176))) | s('type.num.num',X176) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X176)))),X175,'type.num.num'))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f17092,plain,(
  ( ! [X182,X183,X181] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X182),s('type.num.num',X183))) | s(bool,'T_0') = s(bool,X181) | s('type.num.num',X183) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X183)))),X182,'type.num.num'))) ) | $spl30),
  inference(forward_demodulation,[],[f16985,f13832])).
fof(f13832,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl30),
  inference(condensation,[],[f13831])).
fof(f13831,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl30),
  inference(subsumption_resolution,[],[f13795,f83])).
fof(f13795,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl30),
  inference(superposition,[],[f3113,f83])).
fof(f16985,plain,(
  ( ! [X182,X183,X181] : (s(bool,f(s(bool,'F_0'))) != s(bool,happ(s(fun('type.num.num',bool),X182),s('type.num.num',X183))) | s(bool,'T_0') = s(bool,X181) | s('type.num.num',X183) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X183)))),X182,'type.num.num'))) ) | $spl30),
  inference(duplicate_literal_removal,[],[f16866])).
fof(f16866,plain,(
  ( ! [X182,X183,X181] : (s(bool,f(s(bool,'F_0'))) != s(bool,happ(s(fun('type.num.num',bool),X182),s('type.num.num',X183))) | s(bool,'T_0') = s(bool,X181) | s('type.num.num',X183) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X183)))),X182,'type.num.num')) | s(bool,'T_0') = s(bool,X181)) ) | $spl30),
  inference(superposition,[],[f3113,f2940])).
fof(f17090,plain,(
  ( ! [X175,X176,X174] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X175),s('type.num.num',X176))) | s(bool,'T_0') = s(bool,X174) | s('type.num.num',X176) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X176)))),X175,'type.num.num'))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f16988,f2830])).
fof(f16988,plain,(
  ( ! [X175,X176,X174] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s(bool,happ(s(fun('type.num.num',bool),X175),s('type.num.num',X176))) | s(bool,'T_0') = s(bool,X174) | s('type.num.num',X176) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X176)))),X175,'type.num.num'))) ) | $spl1),
  inference(duplicate_literal_removal,[],[f16864])).
fof(f16864,plain,(
  ( ! [X175,X176,X174] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s(bool,happ(s(fun('type.num.num',bool),X175),s('type.num.num',X176))) | s(bool,'T_0') = s(bool,X174) | s('type.num.num',X176) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X176)))),X175,'type.num.num')) | s(bool,'T_0') = s(bool,X174)) ) | $spl1),
  inference(superposition,[],[f2823,f2940])).
fof(f2823,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | s(bool,'T_0') = s(bool,X1)) ) | $spl1),
  inference(subsumption_resolution,[],[f2819,f159])).
fof(f2819,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))))) )),
  inference(superposition,[],[f78,f136])).
fof(f78,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))))),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2])],[f46])).
fof(f46,plain,(
  ? [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) != s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X1)))))))),
  inference(ennf_transformation,[],[f22])).
fof(f22,plain,(
  ~! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X1)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7))) = s(bool,'const.bool.IN_2'(s('type.num.num',X6),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X7)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7))) = s(bool,'const.bool.IN_2'(s('type.num.num',X6),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X7)))))))),
  file('thm.rich_list.MEM_COUNT_LIST','thm.rich_list.MEM_COUNT_LIST')).
fof(f17073,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s(bool,happ(s(fun('type.num.num',bool),X76),s('type.num.num',X77))) = s(bool,happ(s(fun(X73,bool),X74),s(X73,X75))) | s(X73,X75) = s(X73,sK4(X75,X74,X73)) | s('type.num.num',X77) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X77)))),X76,'type.num.num'))) ) | $spl158),
  inference(cnf_transformation,[],[f17073_D])).
fof(f17073_D,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s(bool,happ(s(fun('type.num.num',bool),X76),s('type.num.num',X77))) = s(bool,happ(s(fun(X73,bool),X74),s(X73,X75))) | s(X73,X75) = s(X73,sK4(X75,X74,X73)) | s('type.num.num',X77) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X77)))),X76,'type.num.num'))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f16996,plain,(
  ( ! [X154,X152,X151,X155,X153,X150] : (s(bool,happ(s(fun('type.num.num',bool),X151),s('type.num.num',X152))) = s(bool,happ(s(fun(X153,bool),X154),s(X153,X155))) | s(X153,X155) = s(X153,sK4(X155,X154,X153)) | s(bool,'T_0') = s(bool,X150) | s('type.num.num',X152) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X152)))),X151,'type.num.num'))) )),
  inference(duplicate_literal_removal,[],[f16859])).
fof(f16859,plain,(
  ( ! [X154,X152,X151,X155,X153,X150] : (s(bool,happ(s(fun('type.num.num',bool),X151),s('type.num.num',X152))) = s(bool,happ(s(fun(X153,bool),X154),s(X153,X155))) | s(X153,X155) = s(X153,sK4(X155,X154,X153)) | s(bool,'T_0') = s(bool,X150) | s('type.num.num',X152) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X152)))),X151,'type.num.num')) | s(bool,'T_0') = s(bool,X150)) )),
  inference(superposition,[],[f916,f2940])).
fof(f916,plain,(
  ( ! [X76,X74,X75,X73] : (s(bool,X76) = s(bool,happ(s(fun(X73,bool),X74),s(X73,X75))) | s(X73,X75) = s(X73,sK4(X75,X74,X73)) | s(bool,'T_0') = s(bool,X76)) )),
  inference(subsumption_resolution,[],[f695,f79])).
fof(f695,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,'T_0')) | s(X73,X75) = s(X73,sK4(X75,X74,X73)) | s(bool,X76) = s(bool,happ(s(fun(X73,bool),X74),s(X73,X75))) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f102,f136])).
fof(f17072,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X68),s('type.num.num',X69))) = s(bool,happ(s(fun('type.num.num',bool),X70),s('type.num.num',X71))) | s('type.num.num',X68) = s('type.num.num','const.list.EL_2'(s('type.num.num',X68),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X69))))) | s('type.num.num',X71) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X71)))),X70,'type.num.num'))) ) | $spl156),
  inference(cnf_transformation,[],[f17072_D])).
fof(f17072_D,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X68),s('type.num.num',X69))) = s(bool,happ(s(fun('type.num.num',bool),X70),s('type.num.num',X71))) | s('type.num.num',X68) = s('type.num.num','const.list.EL_2'(s('type.num.num',X68),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X69))))) | s('type.num.num',X71) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X71)))),X70,'type.num.num'))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f16997,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X148),s('type.num.num',X149))) = s(bool,happ(s(fun('type.num.num',bool),X146),s('type.num.num',X147))) | s('type.num.num',X148) = s('type.num.num','const.list.EL_2'(s('type.num.num',X148),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X149))))) | s(bool,'T_0') = s(bool,X145) | s('type.num.num',X147) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X147)))),X146,'type.num.num'))) )),
  inference(duplicate_literal_removal,[],[f16858])).
fof(f16858,plain,(
  ( ! [X146,X149,X147,X145,X148] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X148),s('type.num.num',X149))) = s(bool,happ(s(fun('type.num.num',bool),X146),s('type.num.num',X147))) | s('type.num.num',X148) = s('type.num.num','const.list.EL_2'(s('type.num.num',X148),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X149))))) | s(bool,'T_0') = s(bool,X145) | s('type.num.num',X147) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X147)))),X146,'type.num.num')) | s(bool,'T_0') = s(bool,X145)) )),
  inference(superposition,[],[f913,f2940])).
fof(f17088,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun('type.num.num',bool),X69),s('type.num.num',X70))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X71)))))) | s('type.num.num',X70) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X70)))),X69,'type.num.num')) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X71,bool)),s('type.list.list'(bool),X71)))) ) | $spl168),
  inference(cnf_transformation,[],[f17088_D])).
fof(f17088_D,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun('type.num.num',bool),X69),s('type.num.num',X70))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X71)))))) | s('type.num.num',X70) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X70)))),X69,'type.num.num')) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X71,bool)),s('type.list.list'(bool),X71)))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f17017,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun('type.num.num',bool),X69),s('type.num.num',X70))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X71)))))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X71,bool)),s('type.list.list'(bool),X71))) | s(bool,'T_0') = s(bool,X68) | s('type.num.num',X70) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X70)))),X69,'type.num.num'))) )),
  inference(duplicate_literal_removal,[],[f16844])).
fof(f16844,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun('type.num.num',bool),X69),s('type.num.num',X70))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X71)))))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X71,bool)),s('type.list.list'(bool),X71))) | s(bool,'T_0') = s(bool,X68) | s('type.num.num',X70) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X70)))),X69,'type.num.num')) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f380,f2940])).
fof(f380,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X6)))))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X6,bool)),s('type.list.list'(bool),X6))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f99,f83])).
fof(f99,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | s(X0,X2) = s(X0,'const.list.EL_2'(s('type.num.num',sK3(X2,X1,X0)),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f68])).
fof(f17087,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun('type.num.num',bool),X65),s('type.num.num',X66))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X67)))))) | s('type.num.num',X66) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X66)))),X65,'type.num.num')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X67,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X67))))))) ) | $spl166),
  inference(cnf_transformation,[],[f17087_D])).
fof(f17087_D,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun('type.num.num',bool),X65),s('type.num.num',X66))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X67)))))) | s('type.num.num',X66) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X66)))),X65,'type.num.num')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X67,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X67))))))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f17018,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun('type.num.num',bool),X65),s('type.num.num',X66))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X67)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X67,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X67)))))) | s(bool,'T_0') = s(bool,X64) | s('type.num.num',X66) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X66)))),X65,'type.num.num'))) )),
  inference(duplicate_literal_removal,[],[f16843])).
fof(f16843,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun('type.num.num',bool),X65),s('type.num.num',X66))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X67)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X67,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X67)))))) | s(bool,'T_0') = s(bool,X64) | s('type.num.num',X66) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X66)))),X65,'type.num.num')) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f362,f2940])).
fof(f362,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X6)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X6,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X6)))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f98,f83])).
fof(f98,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(X2,X1,X0)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X0),X1))))))) )),
  inference(cnf_transformation,[],[f68])).
fof(f17086,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,happ(s(fun(bool,bool),X62),s(bool,happ(s(fun('type.num.num',bool),X60),s('type.num.num',X61)))))) | s('type.num.num',X61) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X61)))),X60,'type.num.num')) | s(bool,'F_0') != s(bool,X63) | p(s(bool,happ(s(fun(bool,bool),X62),s(bool,X63))))) ) | $spl164),
  inference(cnf_transformation,[],[f17086_D])).
fof(f17086_D,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,happ(s(fun(bool,bool),X62),s(bool,happ(s(fun('type.num.num',bool),X60),s('type.num.num',X61)))))) | s('type.num.num',X61) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X61)))),X60,'type.num.num')) | s(bool,'F_0') != s(bool,X63) | p(s(bool,happ(s(fun(bool,bool),X62),s(bool,X63))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f17019,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun(bool,bool),X62),s(bool,happ(s(fun('type.num.num',bool),X60),s('type.num.num',X61)))))) | p(s(bool,happ(s(fun(bool,bool),X62),s(bool,X63)))) | s(bool,'F_0') != s(bool,X63) | s(bool,'T_0') = s(bool,X59) | s('type.num.num',X61) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X61)))),X60,'type.num.num'))) )),
  inference(duplicate_literal_removal,[],[f16842])).
fof(f16842,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun(bool,bool),X62),s(bool,happ(s(fun('type.num.num',bool),X60),s('type.num.num',X61)))))) | p(s(bool,happ(s(fun(bool,bool),X62),s(bool,X63)))) | s(bool,'F_0') != s(bool,X63) | s(bool,'T_0') = s(bool,X59) | s('type.num.num',X61) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X61)))),X60,'type.num.num')) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f345,f2940])).
fof(f345,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X8)))) | p(s(bool,happ(s(fun(bool,bool),X9),s(bool,X10)))) | s(bool,'F_0') != s(bool,X10) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f101,f83])).
fof(f101,plain,(
  ( ! [X4,X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) != s(X0,X4)) )),
  inference(cnf_transformation,[],[f71])).
fof(f17085,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,happ(s(fun(bool,bool),X54),s(bool,happ(s(fun('type.num.num',bool),X52),s('type.num.num',X53)))))) | s('type.num.num',X53) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X53)))),X52,'type.num.num')) | s(bool,'F_0') = s(bool,sK4('F_0',X54,bool))) ) | $spl162),
  inference(cnf_transformation,[],[f17085_D])).
fof(f17085_D,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,happ(s(fun(bool,bool),X54),s(bool,happ(s(fun('type.num.num',bool),X52),s('type.num.num',X53)))))) | s('type.num.num',X53) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X53)))),X52,'type.num.num')) | s(bool,'F_0') = s(bool,sK4('F_0',X54,bool))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f17020,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,happ(s(fun(bool,bool),X54),s(bool,happ(s(fun('type.num.num',bool),X52),s('type.num.num',X53)))))) | s(bool,'F_0') = s(bool,sK4('F_0',X54,bool)) | s(bool,'T_0') = s(bool,X51) | s('type.num.num',X53) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X53)))),X52,'type.num.num'))) )),
  inference(duplicate_literal_removal,[],[f16840])).
fof(f16840,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,happ(s(fun(bool,bool),X54),s(bool,happ(s(fun('type.num.num',bool),X52),s('type.num.num',X53)))))) | s(bool,'F_0') = s(bool,sK4('F_0',X54,bool)) | s(bool,'T_0') = s(bool,X51) | s('type.num.num',X53) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X53)))),X52,'type.num.num')) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f293,f2940])).
fof(f293,plain,(
  ( ! [X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X4)))) | s(bool,'F_0') = s(bool,sK4('F_0',X5,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f102,f83])).
fof(f17075,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,happ(s(fun('type.num.num',bool),X92),s('type.num.num',X93))) = s(bool,happ(s(fun('type.num.num',bool),X94),s('type.num.num',X95))) | s('type.num.num',X93) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X93)))),X92,'type.num.num')) | s('type.num.num',X95) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X95)))),X94,'type.num.num'))) ) | $spl160),
  inference(cnf_transformation,[],[f17075_D])).
fof(f17075_D,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,happ(s(fun('type.num.num',bool),X92),s('type.num.num',X93))) = s(bool,happ(s(fun('type.num.num',bool),X94),s('type.num.num',X95))) | s('type.num.num',X93) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X93)))),X92,'type.num.num')) | s('type.num.num',X95) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X95)))),X94,'type.num.num'))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f17044,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,happ(s(fun('type.num.num',bool),X92),s('type.num.num',X93))) = s(bool,happ(s(fun('type.num.num',bool),X94),s('type.num.num',X95))) | s('type.num.num',X95) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X95)))),X94,'type.num.num')) | s(bool,'T_0') = s(bool,X91) | s('type.num.num',X93) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X93)))),X92,'type.num.num'))) )),
  inference(duplicate_literal_removal,[],[f16770])).
fof(f16770,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s(bool,happ(s(fun('type.num.num',bool),X92),s('type.num.num',X93))) = s(bool,happ(s(fun('type.num.num',bool),X94),s('type.num.num',X95))) | s('type.num.num',X95) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X95)))),X94,'type.num.num')) | s(bool,'T_0') = s(bool,X91) | s('type.num.num',X93) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X93)))),X92,'type.num.num')) | s(bool,'T_0') = s(bool,X91)) )),
  inference(superposition,[],[f2940,f2940])).
fof(f17071,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X2))) | s('type.num.num',X2) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X2)))),X1,'type.num.num'))) ) | $spl154),
  inference(cnf_transformation,[],[f17071_D])).
fof(f17071_D,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X2))) | s('type.num.num',X2) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X2)))),X1,'type.num.num'))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f17074,plain,(
  ( ! [X90,X88,X89] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X89),s('type.num.num',X90))) | s('type.num.num',X90) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X90)))),X89,'type.num.num')) | s(bool,'T_0') = s(bool,X88)) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f17045,f2913])).
fof(f17045,plain,(
  ( ! [X90,X88,X89] : (s(bool,happ(s(fun('type.num.num',bool),X89),s('type.num.num',X90))) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s('type.num.num',X90) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X90)))),X89,'type.num.num')) | s(bool,'T_0') = s(bool,X88)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f16769])).
fof(f16769,plain,(
  ( ! [X90,X88,X89] : (s(bool,happ(s(fun('type.num.num',bool),X89),s('type.num.num',X90))) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s('type.num.num',X90) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X90)))),X89,'type.num.num')) | s(bool,'T_0') = s(bool,X88) | s(bool,'T_0') = s(bool,X88)) ) | $spl3),
  inference(superposition,[],[f2940,f2822])).
fof(f2822,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(global_subsumption,[],[f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f164,f163,f169,f171,f203,f222,f223,f227,f234,f238,f251,f266,f286,f285,f283,f281,f280,f278,f272,f271,f269,f332,f376,f535,f604,f907,f908,f910,f706,f939,f1155,f1301,f1304,f916,f1099,f1091,f1084,f1083,f1069,f1037,f1009,f1006,f971,f917,f703,f299,f1032,f1031,f1008,f1005,f970,f705,f702,f663,f662,f650,f603,f601,f600,f339,f919,f297,f915,f102,f101,f103,f306,f104,f913,f1003,f1004,f969,f968,f912,f701,f700,f499,f388,f364,f1030,f1029,f661,f660,f651,f599,f598,f441,f362,f490,f359,f385,f360,f386,f100,f99,f98,f911,f914,f602,f704,f1007,f1305,f1295,f1270,f1269,f1258,f1240,f1232,f1153,f1139,f1138,f1129,f1128,f1127,f1126,f1077,f918,f482,f480,f387,f380,f379,f363,f361,f356,f355,f345,f344,f304,f298,f293,f292,f170,f89,f88,f81,f2216,f2217,f2218,f2219,f2220,f2259,f2274,f2278,f2295,f2529,f2531,f2523,f2521,f2533,f2401,f2402,f2403,f2406,f2407,f2408,f2409,f2540,f2507,f2503,f2545,f2459,f2460,f2461,f2462,f2463,f2464,f2465,f2556,f2558,f2491,f2492,f2493,f2494,f162,f78,f2818])).
fof(f2818,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f78,f136])).
fof(f2494,plain,(
  ( ! [X78,X79,X77] : (~p(s(bool,happ(s(fun(bool,bool),X79),s(bool,X78)))) | s(bool,happ(s(fun('type.num.num',bool),X77),s('type.num.num',sK1))) = s(bool,sK4(happ(s(fun('type.num.num',bool),X77),s('type.num.num',sK1)),X79,bool)) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X77,'type.num.num')) | s(bool,'T_0') = s(bool,X78)) ) | $spl3),
  inference(superposition,[],[f102,f919])).
fof(f2493,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,happ(s(fun(bool,bool),X75),s(bool,X74)))) | p(s(bool,happ(s(fun(bool,bool),X75),s(bool,X76)))) | s(bool,X76) != s(bool,happ(s(fun('type.num.num',bool),X73),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X73,'type.num.num')) | s(bool,'T_0') = s(bool,X74)) ) | $spl3),
  inference(superposition,[],[f101,f919])).
fof(f2492,plain,(
  ( ! [X70,X72,X71] : (~p(s(bool,'const.bool.IN_2'(s(bool,X71),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X72)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(happ(s(fun('type.num.num',bool),X70),s('type.num.num',sK1)),X72,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X72)))))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X70,'type.num.num')) | s(bool,'T_0') = s(bool,X71)) ) | $spl3),
  inference(superposition,[],[f98,f919])).
fof(f2491,plain,(
  ( ! [X68,X69,X67] : (~p(s(bool,'const.bool.IN_2'(s(bool,X68),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X69)))))) | s(bool,happ(s(fun('type.num.num',bool),X67),s('type.num.num',sK1))) = s(bool,'const.list.EL_2'(s('type.num.num',sK3(happ(s(fun('type.num.num',bool),X67),s('type.num.num',sK1)),X69,bool)),s('type.list.list'(bool),X69))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X67,'type.num.num')) | s(bool,'T_0') = s(bool,X68)) ) | $spl3),
  inference(superposition,[],[f99,f919])).
fof(f2465,plain,(
  ( ! [X180,X178,X181,X179] : (s(X180,f(s(bool,happ(s(fun('type.num.num',bool),X181),s('type.num.num',sK1))))) != s(X180,g(s(bool,happ(s(fun('type.num.num',bool),X181),s('type.num.num',sK1))))) | s(fun(bool,X180),X178) = s(fun(bool,X180),X179) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X181,'type.num.num')) | s(bool,'T_0') = s(bool,sK5(X178,X179,X180,bool))) ) | $spl3),
  inference(superposition,[],[f109,f919])).
fof(f2464,plain,(
  ( ! [X177,X175,X176] : (p(s(bool,happ(s(fun(bool,bool),X176),s(bool,happ(s(fun('type.num.num',bool),X177),s('type.num.num',sK1)))))) | ~p(s(bool,happ(s(fun(bool,bool),X176),s(bool,X175)))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X177,'type.num.num')) | s(bool,'T_0') = s(bool,sK4(X175,X176,bool))) ) | $spl3),
  inference(superposition,[],[f103,f919])).
fof(f2463,plain,(
  ( ! [X173,X171,X174,X172] : (s(bool,f(s(X171,sK5(X172,X173,bool,X171)))) != s(bool,happ(s(fun('type.num.num',bool),X174),s('type.num.num',sK1))) | s(fun(X171,bool),X172) = s(fun(X171,bool),X173) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X174,'type.num.num')) | s(bool,'T_0') = s(bool,g(s(X171,sK5(X172,X173,bool,X171))))) ) | $spl3),
  inference(superposition,[],[f109,f919])).
fof(f2462,plain,(
  ( ! [X165,X166,X164] : (~p(s(bool,happ(s(fun(bool,bool),X166),s(bool,happ(s(fun('type.num.num',bool),X165),s('type.num.num',sK1)))))) | s(bool,X164) = s(bool,sK4(X164,X166,bool)) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X165,'type.num.num')) | s(bool,'T_0') = s(bool,X164)) ) | $spl3),
  inference(superposition,[],[f102,f919])).
fof(f2461,plain,(
  ( ! [X163,X161,X162,X160] : (~p(s(bool,happ(s(fun(bool,bool),X162),s(bool,happ(s(fun('type.num.num',bool),X161),s('type.num.num',sK1)))))) | p(s(bool,happ(s(fun(bool,bool),X162),s(bool,X163)))) | s(bool,X160) != s(bool,X163) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X161,'type.num.num')) | s(bool,'T_0') = s(bool,X160)) ) | $spl3),
  inference(superposition,[],[f101,f919])).
fof(f2460,plain,(
  ( ! [X158,X159,X157] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun('type.num.num',bool),X158),s('type.num.num',sK1))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X159)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(X157,X159,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X159)))))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X158,'type.num.num')) | s(bool,'T_0') = s(bool,X157)) ) | $spl3),
  inference(superposition,[],[f98,f919])).
fof(f2459,plain,(
  ( ! [X156,X154,X155] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun('type.num.num',bool),X155),s('type.num.num',sK1))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X156)))))) | s(bool,X154) = s(bool,'const.list.EL_2'(s('type.num.num',sK3(X154,X156,bool)),s('type.list.list'(bool),X156))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X155,'type.num.num')) | s(bool,'T_0') = s(bool,X154)) ) | $spl3),
  inference(superposition,[],[f99,f919])).
fof(f2545,plain,(
  ( ! [X81] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X81),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X81,'type.num.num'))) ) | $spl3),
  inference(subsumption_resolution,[],[f2434,f169])).
fof(f2434,plain,(
  ( ! [X81] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X81),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X81,'type.num.num')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))))) ) | $spl3),
  inference(superposition,[],[f169,f919])).
fof(f2409,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,happ(s(fun('type.num.num',bool),X26),s('type.num.num',sK1)))) | s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X25,'type.num.num')) = s('type.num.num',sK4(sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X25,'type.num.num'),X26,'type.num.num')) | s(bool,X27) = s(bool,happ(s(fun('type.num.num',bool),X25),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X27)) ) | $spl3),
  inference(superposition,[],[f102,f919])).
fof(f2408,plain,(
  ( ! [X24,X23,X21,X22] : (~p(s(bool,happ(s(fun('type.num.num',bool),X22),s('type.num.num',sK1)))) | p(s(bool,happ(s(fun('type.num.num',bool),X22),s('type.num.num',X23)))) | s('type.num.num',X23) != s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X21,'type.num.num')) | s(bool,X24) = s(bool,happ(s(fun('type.num.num',bool),X21),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X24)) ) | $spl3),
  inference(superposition,[],[f101,f919])).
fof(f2407,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X19)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X18,'type.num.num'),X19,'type.num.num')),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),X19)))))) | s(bool,X20) = s(bool,happ(s(fun('type.num.num',bool),X18),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X20)) ) | $spl3),
  inference(superposition,[],[f98,f919])).
fof(f2406,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X16)))))) | s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X15,'type.num.num')) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK3(sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X15,'type.num.num'),X16,'type.num.num')),s('type.list.list'('type.num.num'),X16))) | s(bool,X17) = s(bool,happ(s(fun('type.num.num',bool),X15),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X17)) ) | $spl3),
  inference(superposition,[],[f99,f919])).
fof(f2403,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',X9)))) | s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X8,'type.num.num')) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X8,'type.num.num')),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X9))))) | s(bool,X10) = s(bool,happ(s(fun('type.num.num',bool),X8),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X10)) ) | $spl3),
  inference(superposition,[],[f104,f919])).
fof(f2402,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',sK1)))) | s('type.num.num',X6) = s('type.num.num','const.list.EL_2'(s('type.num.num',X6),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X5,'type.num.num')))))) | s(bool,X7) = s(bool,happ(s(fun('type.num.num',bool),X5),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X7)) ) | $spl3),
  inference(superposition,[],[f104,f919])).
fof(f2401,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X1),X2)))))) | s(X1,X3) != s(X1,'const.list.EL_2'(s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X0,'type.num.num')),s('type.list.list'(X1),X2))) | p(s(bool,'const.bool.IN_2'(s(X1,X3),s(fun(X1,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X1),X2)))))) | s(bool,X4) = s(bool,happ(s(fun('type.num.num',bool),X0),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(superposition,[],[f100,f919])).
fof(f2533,plain,(
  ( ! [X28] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X28),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X28,'type.num.num'))) ) | $spl3),
  inference(subsumption_resolution,[],[f2385,f169])).
fof(f2385,plain,(
  ( ! [X28] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X28),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X28,'type.num.num')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))))) ) | $spl3),
  inference(superposition,[],[f919,f170])).
fof(f2523,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool)) | s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) = s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X1,'type.num.num'))) ) | $spl3),
  inference(condensation,[],[f2522])).
fof(f2522,plain,(
  ( ! [X17,X18,X16] : (s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) = s(bool,happ(s(fun('type.num.num',bool),X18),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X18,'type.num.num')) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X17,bool))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2383])).
fof(f2383,plain,(
  ( ! [X17,X18,X16] : (s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) = s(bool,happ(s(fun('type.num.num',bool),X18),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X18,'type.num.num')) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X17,bool)) | s(bool,'T_0') = s(bool,X16)) ) | $spl3),
  inference(superposition,[],[f919,f917])).
fof(f2295,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X1,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X1))))) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X1))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(subsumption_resolution,[],[f2285,f79])).
fof(f2285,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X1,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X1))))) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X1))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f499,f136])).
fof(f2278,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X1)))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X1,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X1))))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f499,f136])).
fof(f2274,plain,(
  ( ! [X4,X3] : (s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X3,bool)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(trivial_inequality_removal,[],[f2267])).
fof(f2267,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X3,bool)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(superposition,[],[f2220,f136])).
fof(f2259,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0'))) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X1,bool)) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f2220,f136])).
fof(f2220,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool))) ) | $spl3),
  inference(equality_factoring,[],[f1037])).
fof(f2219,plain,(
  ( ! [X47,X48] : (~p(s(bool,happ(s(fun(bool,bool),X48),s(bool,happ(s(fun(bool,bool),X47),s(bool,'T_0')))))) | s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X47,bool)) = s(bool,sK4(sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X47,bool),X48,bool)) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X47,bool))) ) | $spl3),
  inference(superposition,[],[f102,f1037])).
fof(f2218,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,happ(s(fun(bool,bool),X45),s(bool,happ(s(fun(bool,bool),X44),s(bool,'T_0')))))) | p(s(bool,happ(s(fun(bool,bool),X45),s(bool,X46)))) | s(bool,X46) != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X44,bool)) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X44,bool))) ) | $spl3),
  inference(superposition,[],[f101,f1037])).
fof(f2217,plain,(
  ( ! [X43,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(bool,bool),X42),s(bool,'T_0'))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X43)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X42,bool),X43,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X43)))))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X42,bool))) ) | $spl3),
  inference(superposition,[],[f98,f1037])).
fof(f2216,plain,(
  ( ! [X41,X40] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(bool,bool),X40),s(bool,'T_0'))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X41)))))) | s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X40,bool)) = s(bool,'const.list.EL_2'(s('type.num.num',sK3(sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X40,bool),X41,bool)),s('type.list.list'(bool),X41))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X40,bool))) ) | $spl3),
  inference(superposition,[],[f99,f1037])).
fof(f89,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f60])).
fof(f170,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f169,f83])).
fof(f292,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) | s(bool,X2) = s(bool,sK4(X2,X3,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f102,f83])).
fof(f298,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0')))) | s(bool,'F_0') = s(bool,sK4('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X6,bool))) ) | $spl3),
  inference(forward_demodulation,[],[f294,f170])).
fof(f294,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0')))) | s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) = s(bool,sK4('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X6,bool))) ) | $spl3),
  inference(superposition,[],[f102,f170])).
fof(f304,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool))) )),
  inference(superposition,[],[f103,f83])).
fof(f344,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X6),s(bool,X7)))) | s(bool,X5) != s(bool,X7) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f101,f83])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,X2,bool))) )),
  inference(superposition,[],[f109,f83])).
fof(f361,plain,(
  ( ! [X4,X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(X3,X4,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X4)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4)))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f98,f83])).
fof(f363,plain,(
  ( ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X7)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X7,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X7))))))) ) | $spl3),
  inference(superposition,[],[f98,f170])).
fof(f379,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X4)))))) | s(bool,X3) = s(bool,'const.list.EL_2'(s('type.num.num',sK3(X3,X4,bool)),s('type.list.list'(bool),X4))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f99,f83])).
fof(f387,plain,(
  ( ! [X7] : (s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X7,bool)),s('type.list.list'(bool),X7))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X7))))))) ) | $spl3),
  inference(forward_demodulation,[],[f381,f170])).
fof(f381,plain,(
  ( ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X7)))))) | s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) = s(bool,'const.list.EL_2'(s('type.num.num',sK3('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X7,bool)),s('type.list.list'(bool),X7)))) ) | $spl3),
  inference(superposition,[],[f99,f170])).
fof(f480,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X21,bool)),s('type.list.list'(bool),X21))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X20),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X21)))))) | s(bool,'T_0') = s(bool,X20)) ) | $spl3),
  inference(forward_demodulation,[],[f440,f170])).
fof(f440,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,X20),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X21)))))) | s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) = s(bool,'const.list.EL_2'(s('type.num.num',sK3('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X21,bool)),s('type.list.list'(bool),X21))) | s(bool,'T_0') = s(bool,X20)) ) | $spl3),
  inference(superposition,[],[f99,f376])).
fof(f482,plain,(
  ( ! [X28,X27] : (s(bool,'F_0') = s(bool,sK4('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X28,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X28),s(bool,X27)))) | s(bool,'T_0') = s(bool,X27)) ) | $spl3),
  inference(forward_demodulation,[],[f443,f170])).
fof(f443,plain,(
  ( ! [X28,X27] : (~p(s(bool,happ(s(fun(bool,bool),X28),s(bool,X27)))) | s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) = s(bool,sK4('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X28,bool)) | s(bool,'T_0') = s(bool,X27)) ) | $spl3),
  inference(superposition,[],[f102,f376])).
fof(f918,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') = s(bool,sK4('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X83,bool)) | s(bool,X84) = s(bool,happ(s(fun(bool,bool),X83),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X84)) ) | $spl3),
  inference(subsumption_resolution,[],[f698,f79])).
fof(f698,plain,(
  ( ! [X83,X84] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,sK4('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X83,bool)) | s(bool,X84) = s(bool,happ(s(fun(bool,bool),X83),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X84)) ) | $spl3),
  inference(superposition,[],[f298,f136])).
fof(f1077,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) ) | $spl3),
  inference(condensation,[],[f1076])).
fof(f1076,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X1,bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f926])).
fof(f926,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X1,bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f917,f83])).
fof(f1126,plain,(
  ( ! [X28,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X28)))))) | s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X27,bool)) = s(bool,'const.list.EL_2'(s('type.num.num',sK3(sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X27,bool),X28,bool)),s('type.list.list'(bool),X28))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X27),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f99,f1077])).
fof(f1127,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X30)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X29,bool),X30,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X30)))))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X29),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f98,f1077])).
fof(f1128,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,happ(s(fun(bool,bool),X32),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X32),s(bool,X33)))) | s(bool,X33) != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X31,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X31),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f101,f1077])).
fof(f1129,plain,(
  ( ! [X35,X34] : (~p(s(bool,happ(s(fun(bool,bool),X35),s(bool,'T_0')))) | s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X34,bool)) = s(bool,sK4(sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X34,bool),X35,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X34),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f102,f1077])).
fof(f1138,plain,(
  ( ! [X10,X11] : (~p(s(bool,X11)) | s(bool,X11) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X10,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X10),s(bool,'T_0')))) ) | $spl3),
  inference(subsumption_resolution,[],[f1117,f79])).
fof(f1117,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X11)) | s(bool,X11) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X10,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X10),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f108,f1077])).
fof(f1139,plain,(
  ( ! [X19,X18] : (s(bool,X19) != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X18,bool)) | s(bool,'T_0') = s(bool,X19) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X18),s(bool,'T_0')))) ) | $spl3),
  inference(subsumption_resolution,[],[f1120,f79])).
fof(f1120,plain,(
  ( ! [X19,X18] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X18,bool)) | s(bool,'T_0') = s(bool,X19) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X18),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f138,f1077])).
fof(f1153,plain,(
  ( ! [X4,X2,X3] : (s(bool,X3) = s(bool,sK4(X3,X2,bool)) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1147,f79])).
fof(f1147,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s(bool,X3) = s(bool,sK4(X3,X2,bool)) | s(bool,'T_0') = s(bool,X3) | s(bool,X4) = s(bool,happ(s(fun(bool,bool),X2),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f292,f136])).
fof(f1232,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))))) | s(bool,'F_0') = s(bool,sK4('F_0',X1,bool))) ) | $spl3),
  inference(condensation,[],[f1231])).
fof(f1231,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X18),s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0')))))) | s(bool,'F_0') = s(bool,sK4('F_0',X18,bool)) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X17,bool))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f1213])).
fof(f1213,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X18),s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0')))))) | s(bool,'F_0') = s(bool,sK4('F_0',X18,bool)) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X17,bool)) | s(bool,'T_0') = s(bool,X16)) ) | $spl3),
  inference(superposition,[],[f293,f917])).
fof(f1240,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) = s(bool,happ(s(fun(bool,bool),X4),s(bool,X5))) | s(bool,'T_0') = s(bool,X5) | s(bool,'F_0') = s(bool,sK4('F_0',X4,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1225,f79])).
fof(f1225,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,sK4('F_0',X4,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,happ(s(fun(bool,bool),X4),s(bool,X5))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f293,f136])).
fof(f1258,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X3,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(subsumption_resolution,[],[f1243,f79])).
fof(f1243,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | s(bool,X2) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X3,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'T_0'))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(superposition,[],[f1138,f136])).
fof(f1269,plain,(
  ( ! [X35,X34] : (s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X34,bool)) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X35,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X35),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X34),s(bool,'T_0')))) ) | $spl3),
  inference(subsumption_resolution,[],[f1256,f79])).
fof(f1256,plain,(
  ( ! [X35,X34] : (~p(s(bool,'T_0')) | s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X34,bool)) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X35,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X35),s(bool,'T_0'))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X34),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f1138,f1077])).
fof(f1270,plain,(
  ( ! [X37,X38,X36] : (s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X36,bool)) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X37,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X37),s(bool,'T_0'))) | s(bool,X38) = s(bool,happ(s(fun(bool,bool),X36),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X38)) ) | $spl3),
  inference(subsumption_resolution,[],[f1257,f79])).
fof(f1257,plain,(
  ( ! [X37,X38,X36] : (~p(s(bool,'T_0')) | s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X36,bool)) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X37,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X37),s(bool,'T_0'))) | s(bool,X38) = s(bool,happ(s(fun(bool,bool),X36),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X38)) ) | $spl3),
  inference(superposition,[],[f1138,f917])).
fof(f1295,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool)) | s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X1,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) ) | $spl3),
  inference(condensation,[],[f1294])).
fof(f1294,plain,(
  ( ! [X17,X18,X16] : (s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X18,bool)) | s(bool,'T_0') = s(bool,X16) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X18),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X17,bool))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f1277])).
fof(f1277,plain,(
  ( ! [X17,X18,X16] : (s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X18,bool)) | s(bool,'T_0') = s(bool,X16) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X18),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X17,bool)) | s(bool,'T_0') = s(bool,X16)) ) | $spl3),
  inference(superposition,[],[f1139,f917])).
fof(f1305,plain,(
  ( ! [X27] : (s(bool,'F_0') != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X27,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X27),s(bool,'T_0')))) ) | $spl3),
  inference(subsumption_resolution,[],[f1279,f169])).
fof(f1279,plain,(
  ( ! [X27] : (s(bool,'F_0') != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X27,bool)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X27),s(bool,'T_0')))) ) | $spl3),
  inference(superposition,[],[f1139,f170])).
fof(f1007,plain,(
  ( ! [X118,X116,X117,X115] : (s(bool,f(s(X115,sK5(X116,X117,bool,X115)))) != s(bool,happ(s(fun(bool,bool),X118),s(bool,'T_0'))) | s(fun(X115,bool),X116) = s(fun(X115,bool),X117) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X118,bool)) | s(bool,'T_0') = s(bool,g(s(X115,sK5(X116,X117,bool,X115))))) ) | $spl3),
  inference(superposition,[],[f109,f917])).
fof(f704,plain,(
  ( ! [X103,X105,X106,X104] : (s(bool,'T_0') != s(bool,f(s(X103,sK5(X104,X105,bool,X103)))) | s(fun(X103,bool),X104) = s(fun(X103,bool),X105) | s(bool,X106) = s(bool,g(s(X103,sK5(X104,X105,bool,X103)))) | s(bool,'T_0') = s(bool,X106)) )),
  inference(superposition,[],[f109,f136])).
fof(f602,plain,(
  ( ! [X109,X107,X110,X108] : (s(bool,X110) != s(bool,f(s(X107,sK5(X108,X109,bool,X107)))) | s(fun(X107,bool),X108) = s(fun(X107,bool),X109) | s(bool,'T_0') = s(bool,X110) | s(bool,'T_0') = s(bool,g(s(X107,sK5(X108,X109,bool,X107))))) )),
  inference(superposition,[],[f109,f136])).
fof(f914,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(X62,X64) != s(X62,'const.list.EL_2'(s('type.num.num',X61),s('type.list.list'(X62),X63))) | p(s(bool,'const.bool.IN_2'(s(X62,X64),s(fun(X62,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X62),X63)))))) | s(bool,X65) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X61),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X62),X63))))) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f692,f79])).
fof(f692,plain,(
  ( ! [X61,X64,X62,X65,X63] : (~p(s(bool,'T_0')) | s(X62,X64) != s(X62,'const.list.EL_2'(s('type.num.num',X61),s('type.list.list'(X62),X63))) | p(s(bool,'const.bool.IN_2'(s(X62,X64),s(fun(X62,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X62),X63)))))) | s(bool,X65) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X61),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X62),X63))))) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f100,f136])).
fof(f911,plain,(
  ( ! [X52,X50,X53,X51] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(X51,X52,X50)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X50),X52)))))) | s(bool,X53) = s(bool,'const.bool.IN_2'(s(X50,X51),s(fun(X50,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X50),X52))))) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f688,f79])).
fof(f688,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(X51,X52,X50)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(X50),X52)))))) | s(bool,X53) = s(bool,'const.bool.IN_2'(s(X50,X51),s(fun(X50,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X50),X52))))) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f98,f136])).
fof(f386,plain,(
  ( ! [X2] : (s('type.num.num',sK1) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK3('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X2,'type.num.num')),s('type.list.list'('type.num.num'),X2))) | ~p(s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X2))))))) ) | $spl3),
  inference(forward_demodulation,[],[f378,f332])).
fof(f378,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X2)))))) | s('type.num.num','const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK3('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X2,'type.num.num')),s('type.list.list'('type.num.num'),X2)))) ) | $spl3),
  inference(superposition,[],[f99,f332])).
fof(f360,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X2,'type.num.num')),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),X2))))))) ) | $spl3),
  inference(superposition,[],[f98,f332])).
fof(f385,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X1)))))) | s('type.num.num',X0) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK3('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X0)))),X1,'type.num.num')),s('type.list.list'('type.num.num'),X1)))) )),
  inference(forward_demodulation,[],[f377,f82])).
fof(f377,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X1)))))) | s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X0))))) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK3('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X0)))),X1,'type.num.num')),s('type.list.list'('type.num.num'),X1)))) )),
  inference(superposition,[],[f99,f82])).
fof(f359,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X0)))),X1,'type.num.num')),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X1))))))) )),
  inference(superposition,[],[f98,f82])).
fof(f441,plain,(
  ( ! [X23,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X23)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X23,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X23)))))) | s(bool,'T_0') = s(bool,X22)) ) | $spl3),
  inference(superposition,[],[f98,f376])).
fof(f598,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,'const.bool.IN_2'(s(bool,X91),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X92)))))) | s(bool,X90) = s(bool,'const.list.EL_2'(s('type.num.num',sK3(X90,X92,bool)),s('type.list.list'(bool),X92))) | s(bool,'T_0') = s(bool,X91) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f99,f136])).
fof(f599,plain,(
  ( ! [X94,X95,X93] : (~p(s(bool,'const.bool.IN_2'(s(bool,X94),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X95)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(X93,X95,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X95)))))) | s(bool,'T_0') = s(bool,X94) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f98,f136])).
fof(f651,plain,(
  ( ! [X8,X7,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X8)))))) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X8,bool)),s('type.list.list'(bool),X8))) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl3),
  inference(superposition,[],[f388,f136])).
fof(f660,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'const.bool.IN_2'(s(bool,X34),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X35)))))) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('T_0',X35,bool)),s('type.list.list'(bool),X35))) | s(bool,X34) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f99,f136])).
fof(f661,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'const.bool.IN_2'(s(bool,X37),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X38)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('T_0',X38,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X38)))))) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f98,f136])).
fof(f1029,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,X49),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X50)))))) | s(bool,happ(s(fun(bool,bool),X48),s(bool,'T_0'))) = s(bool,'const.list.EL_2'(s('type.num.num',sK3(happ(s(fun(bool,bool),X48),s(bool,'T_0')),X50,bool)),s('type.list.list'(bool),X50))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X48,bool)) | s(bool,'T_0') = s(bool,X49)) ) | $spl3),
  inference(superposition,[],[f99,f917])).
fof(f1030,plain,(
  ( ! [X52,X53,X51] : (~p(s(bool,'const.bool.IN_2'(s(bool,X52),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X53)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(happ(s(fun(bool,bool),X51),s(bool,'T_0')),X53,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X53)))))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X51,bool)) | s(bool,'T_0') = s(bool,X52)) ) | $spl3),
  inference(superposition,[],[f98,f917])).
fof(f364,plain,(
  ( ! [X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X8,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X8)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X8))))))) ) | $spl3),
  inference(superposition,[],[f98,f163])).
fof(f388,plain,(
  ( ! [X8] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X8)))))) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X8,bool)),s('type.list.list'(bool),X8)))) ) | $spl3),
  inference(forward_demodulation,[],[f382,f163])).
fof(f382,plain,(
  ( ! [X8] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X8)))))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) = s(bool,'const.list.EL_2'(s('type.num.num',sK3('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X8,bool)),s('type.list.list'(bool),X8)))) ) | $spl3),
  inference(superposition,[],[f99,f163])).
fof(f499,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X0)))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X0)))))) ) | $spl3),
  inference(subsumption_resolution,[],[f498,f81])).
fof(f498,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X0)))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X0)))))) ) | $spl3),
  inference(forward_demodulation,[],[f496,f170])).
fof(f496,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X0)))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X0)))))) ) | $spl3),
  inference(superposition,[],[f364,f376])).
fof(f700,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X88)))))) | s(bool,X87) = s(bool,'const.list.EL_2'(s('type.num.num',sK3(X87,X88,bool)),s('type.list.list'(bool),X88))) | s(bool,X87) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f99,f136])).
fof(f701,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X91)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(X90,X91,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X91)))))) | s(bool,X90) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f98,f136])).
fof(f912,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X54,bool)),s('type.list.list'(bool),X54))) | s(bool,X55) = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X54))))) | s(bool,'T_0') = s(bool,X55)) ) | $spl3),
  inference(subsumption_resolution,[],[f689,f79])).
fof(f689,plain,(
  ( ! [X54,X55] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X54,bool)),s('type.list.list'(bool),X54))) | s(bool,X55) = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X54))))) | s(bool,'T_0') = s(bool,X55)) ) | $spl3),
  inference(superposition,[],[f388,f136])).
fof(f968,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X40)))))) | s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X39,bool)) = s(bool,'const.list.EL_2'(s('type.num.num',sK3(sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X39,bool),X40,bool)),s('type.list.list'(bool),X40))) | s(bool,X41) = s(bool,happ(s(fun(bool,bool),X39),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X41)) ) | $spl3),
  inference(superposition,[],[f99,f917])).
fof(f969,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X43)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X42,bool),X43,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X43)))))) | s(bool,X44) = s(bool,happ(s(fun(bool,bool),X42),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X44)) ) | $spl3),
  inference(superposition,[],[f98,f917])).
fof(f1004,plain,(
  ( ! [X103,X101,X102] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(bool,bool),X102),s(bool,'T_0'))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X103)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(X101,X103,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X103)))))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X102,bool)) | s(bool,'T_0') = s(bool,X101)) ) | $spl3),
  inference(superposition,[],[f98,f917])).
fof(f1003,plain,(
  ( ! [X99,X100,X98] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(bool,bool),X99),s(bool,'T_0'))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X100)))))) | s(bool,X98) = s(bool,'const.list.EL_2'(s('type.num.num',sK3(X98,X100,bool)),s('type.list.list'(bool),X100))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X99,bool)) | s(bool,'T_0') = s(bool,X98)) ) | $spl3),
  inference(superposition,[],[f99,f917])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) )),
  inference(subsumption_resolution,[],[f305,f81])).
fof(f305,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(bool,'T_0') = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0))))) )),
  inference(superposition,[],[f103,f83])).
fof(f915,plain,(
  ( ! [X70,X68,X72,X71,X69] : (p(s(bool,happ(s(fun(X68,bool),X69),s(X68,X71)))) | s(X68,X70) != s(X68,X71) | s(bool,X72) = s(bool,happ(s(fun(X68,bool),X69),s(X68,X70))) | s(bool,'T_0') = s(bool,X72)) )),
  inference(subsumption_resolution,[],[f694,f79])).
fof(f694,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X68,bool),X69),s(X68,X71)))) | s(X68,X70) != s(X68,X71) | s(bool,X72) = s(bool,happ(s(fun(X68,bool),X69),s(X68,X70))) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f101,f136])).
fof(f919,plain,(
  ( ! [X85,X86] : (s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X85,'type.num.num')) | s(bool,X86) = s(bool,happ(s(fun('type.num.num',bool),X85),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X86)) ) | $spl3),
  inference(subsumption_resolution,[],[f699,f79])).
fof(f699,plain,(
  ( ! [X85,X86] : (~p(s(bool,'T_0')) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X85,'type.num.num')) | s(bool,X86) = s(bool,happ(s(fun('type.num.num',bool),X85),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X86)) ) | $spl3),
  inference(superposition,[],[f339,f136])).
fof(f339,plain,(
  ( ! [X2] : (~p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK1)))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X2,'type.num.num'))) ) | $spl3),
  inference(forward_demodulation,[],[f336,f332])).
fof(f336,plain,(
  ( ! [X2] : (~p(s(bool,happ(s(fun('type.num.num',bool),X2),s('type.num.num',sK1)))) | s('type.num.num','const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X2,'type.num.num'))) ) | $spl3),
  inference(superposition,[],[f102,f332])).
fof(f600,plain,(
  ( ! [X99,X97,X98,X96] : (~p(s(bool,happ(s(fun(bool,bool),X98),s(bool,X97)))) | p(s(bool,happ(s(fun(bool,bool),X98),s(bool,X99)))) | s(bool,X96) != s(bool,X99) | s(bool,'T_0') = s(bool,X97) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f101,f136])).
fof(f601,plain,(
  ( ! [X101,X102,X100] : (~p(s(bool,happ(s(fun(bool,bool),X102),s(bool,X101)))) | s(bool,X100) = s(bool,sK4(X100,X102,bool)) | s(bool,'T_0') = s(bool,X101) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f102,f136])).
fof(f603,plain,(
  ( ! [X111,X112,X113] : (p(s(bool,happ(s(fun(bool,bool),X112),s(bool,X113)))) | ~p(s(bool,happ(s(fun(bool,bool),X112),s(bool,X111)))) | s(bool,'T_0') = s(bool,X113) | s(bool,'T_0') = s(bool,sK4(X111,X112,bool))) )),
  inference(superposition,[],[f103,f136])).
fof(f650,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(bool,bool),X5),s(bool,X4)))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X5,bool)) | s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X6)) ) | $spl3),
  inference(superposition,[],[f299,f136])).
fof(f662,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,happ(s(fun(bool,bool),X41),s(bool,X40)))) | p(s(bool,happ(s(fun(bool,bool),X41),s(bool,X42)))) | s(bool,'T_0') != s(bool,X42) | s(bool,X40) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f101,f136])).
fof(f663,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,happ(s(fun(bool,bool),X45),s(bool,X44)))) | s(bool,'T_0') = s(bool,sK4('T_0',X45,bool)) | s(bool,X44) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f102,f136])).
fof(f702,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,happ(s(fun(bool,bool),X94),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X94),s(bool,X95)))) | s(bool,X93) != s(bool,X95) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f101,f136])).
fof(f705,plain,(
  ( ! [X109,X107,X108] : (p(s(bool,happ(s(fun(bool,bool),X108),s(bool,'T_0')))) | ~p(s(bool,happ(s(fun(bool,bool),X108),s(bool,X107)))) | s(bool,X109) = s(bool,sK4(X107,X108,bool)) | s(bool,'T_0') = s(bool,X109)) )),
  inference(superposition,[],[f103,f136])).
fof(f970,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,happ(s(fun(bool,bool),X46),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X46),s(bool,X47)))) | s(bool,X47) != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X45,bool)) | s(bool,X48) = s(bool,happ(s(fun(bool,bool),X45),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X48)) ) | $spl3),
  inference(superposition,[],[f101,f917])).
fof(f1005,plain,(
  ( ! [X107,X105,X106,X104] : (~p(s(bool,happ(s(fun(bool,bool),X106),s(bool,happ(s(fun(bool,bool),X105),s(bool,'T_0')))))) | p(s(bool,happ(s(fun(bool,bool),X106),s(bool,X107)))) | s(bool,X104) != s(bool,X107) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X105,bool)) | s(bool,'T_0') = s(bool,X104)) ) | $spl3),
  inference(superposition,[],[f101,f917])).
fof(f1008,plain,(
  ( ! [X121,X120,X119] : (p(s(bool,happ(s(fun(bool,bool),X120),s(bool,happ(s(fun(bool,bool),X121),s(bool,'T_0')))))) | ~p(s(bool,happ(s(fun(bool,bool),X120),s(bool,X119)))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X121,bool)) | s(bool,'T_0') = s(bool,sK4(X119,X120,bool))) ) | $spl3),
  inference(superposition,[],[f103,f917])).
fof(f1031,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,happ(s(fun(bool,bool),X56),s(bool,X55)))) | p(s(bool,happ(s(fun(bool,bool),X56),s(bool,X57)))) | s(bool,X57) != s(bool,happ(s(fun(bool,bool),X54),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X54,bool)) | s(bool,'T_0') = s(bool,X55)) ) | $spl3),
  inference(superposition,[],[f101,f917])).
fof(f1032,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,happ(s(fun(bool,bool),X60),s(bool,X59)))) | s(bool,happ(s(fun(bool,bool),X58),s(bool,'T_0'))) = s(bool,sK4(happ(s(fun(bool,bool),X58),s(bool,'T_0')),X60,bool)) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X58,bool)) | s(bool,'T_0') = s(bool,X59)) ) | $spl3),
  inference(superposition,[],[f102,f917])).
fof(f299,plain,(
  ( ! [X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X7,bool))) ) | $spl3),
  inference(forward_demodulation,[],[f295,f163])).
fof(f295,plain,(
  ( ! [X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,'T_0')))) | s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X7,bool))) ) | $spl3),
  inference(superposition,[],[f102,f163])).
fof(f703,plain,(
  ( ! [X99,X97,X98] : (~p(s(bool,happ(s(fun(bool,bool),X98),s(bool,'T_0')))) | s(bool,X97) = s(bool,sK4(X97,X98,bool)) | s(bool,X97) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) )),
  inference(superposition,[],[f102,f136])).
fof(f917,plain,(
  ( ! [X81,X82] : (s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X81,bool)) | s(bool,X82) = s(bool,happ(s(fun(bool,bool),X81),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X82)) ) | $spl3),
  inference(subsumption_resolution,[],[f697,f79])).
fof(f697,plain,(
  ( ! [X81,X82] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X81,bool)) | s(bool,X82) = s(bool,happ(s(fun(bool,bool),X81),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X82)) ) | $spl3),
  inference(superposition,[],[f299,f136])).
fof(f971,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,happ(s(fun(bool,bool),X50),s(bool,'T_0')))) | s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X49,bool)) = s(bool,sK4(sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X49,bool),X50,bool)) | s(bool,X51) = s(bool,happ(s(fun(bool,bool),X49),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X51)) ) | $spl3),
  inference(superposition,[],[f102,f917])).
fof(f1006,plain,(
  ( ! [X109,X110,X108] : (~p(s(bool,happ(s(fun(bool,bool),X110),s(bool,happ(s(fun(bool,bool),X109),s(bool,'T_0')))))) | s(bool,X108) = s(bool,sK4(X108,X110,bool)) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X109,bool)) | s(bool,'T_0') = s(bool,X108)) ) | $spl3),
  inference(superposition,[],[f102,f917])).
fof(f1009,plain,(
  ( ! [X125,X123,X124,X122] : (s(X124,f(s(bool,happ(s(fun(bool,bool),X125),s(bool,'T_0'))))) != s(X124,g(s(bool,happ(s(fun(bool,bool),X125),s(bool,'T_0'))))) | s(fun(bool,X124),X122) = s(fun(bool,X124),X123) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X125,bool)) | s(bool,'T_0') = s(bool,sK5(X122,X123,X124,bool))) ) | $spl3),
  inference(superposition,[],[f109,f917])).
fof(f1037,plain,(
  ( ! [X0] : (s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool)) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f1033])).
fof(f1033,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool)) | s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool))) ) | $spl3),
  inference(equality_factoring,[],[f917])).
fof(f1069,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool)) | s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X1,bool))) ) | $spl3),
  inference(condensation,[],[f1068])).
fof(f1068,plain,(
  ( ! [X17,X18,X16] : (s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) = s(bool,happ(s(fun(bool,bool),X18),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X18,bool)) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X17,bool))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f932])).
fof(f932,plain,(
  ( ! [X17,X18,X16] : (s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) = s(bool,happ(s(fun(bool,bool),X18),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X18,bool)) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X17,bool)) | s(bool,'T_0') = s(bool,X16)) ) | $spl3),
  inference(superposition,[],[f917,f917])).
fof(f1083,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X17) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X16,bool)) | s(bool,X18) = s(bool,happ(s(fun(bool,bool),X16),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X18)) ) | $spl3),
  inference(subsumption_resolution,[],[f960,f79])).
fof(f960,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X17) = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X16,bool)) | s(bool,X18) = s(bool,happ(s(fun(bool,bool),X16),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X18)) ) | $spl3),
  inference(superposition,[],[f108,f917])).
fof(f1084,plain,(
  ( ! [X28,X29,X27] : (s(bool,X28) != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X27,bool)) | s(bool,'T_0') = s(bool,X28) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X27),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X29)) ) | $spl3),
  inference(subsumption_resolution,[],[f963,f79])).
fof(f963,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'T_0')) | s(bool,X28) != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X27,bool)) | s(bool,'T_0') = s(bool,X28) | s(bool,X29) = s(bool,happ(s(fun(bool,bool),X27),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X29)) ) | $spl3),
  inference(superposition,[],[f138,f917])).
fof(f1091,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool))) ) | $spl3),
  inference(forward_demodulation,[],[f1046,f163])).
fof(f1046,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool)) | s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) ) | $spl3),
  inference(condensation,[],[f1045])).
fof(f1045,plain,(
  ( ! [X37,X38] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s(bool,happ(s(fun(bool,bool),X38),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X38,bool))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f986])).
fof(f986,plain,(
  ( ! [X37,X38] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s(bool,happ(s(fun(bool,bool),X38),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X38,bool)) | s(bool,'T_0') = s(bool,X37)) ) | $spl3),
  inference(superposition,[],[f266,f917])).
fof(f1099,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,bool)) | s(bool,X1) = s(bool,happ(s(fun(bool,bool),X0),s(bool,'T_0')))) ) | $spl3),
  inference(condensation,[],[f1098])).
fof(f1098,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X19)) | s(bool,X19) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X17,bool)) | s(bool,'T_0') = s(bool,X18)) ) | $spl3),
  inference(subsumption_resolution,[],[f1018,f85])).
fof(f1018,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) = s(bool,happ(s(fun(bool,bool),X17),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X17,bool)) | s(bool,'T_0') = s(bool,X18)) ) | $spl3),
  inference(superposition,[],[f105,f917])).
fof(f1155,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X1,bool)) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f1091,f136])).
fof(f939,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) = s(bool,happ(s(fun(bool,bool),X2),s(bool,X0))) | s(bool,'T_0') = s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X2,bool)) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl3),
  inference(superposition,[],[f917,f136])).
fof(f706,plain,(
  ( ! [X111,X112,X110,X113] : (s(X112,f(s(bool,'T_0'))) != s(X112,g(s(bool,'T_0'))) | s(fun(bool,X112),X110) = s(fun(bool,X112),X111) | s(bool,X113) = s(bool,sK5(X110,X111,X112,bool)) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f109,f136])).
fof(f910,plain,(
  ( ! [X47,X48,X46,X49] : (s(X46,X47) = s(X46,'const.list.EL_2'(s('type.num.num',sK3(X47,X48,X46)),s('type.list.list'(X46),X48))) | s(bool,X49) = s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X46),X48))))) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f687,f79])).
fof(f687,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'T_0')) | s(X46,X47) = s(X46,'const.list.EL_2'(s('type.num.num',sK3(X47,X48,X46)),s('type.list.list'(X46),X48))) | s(bool,X49) = s(bool,'const.bool.IN_2'(s(X46,X47),s(fun(X46,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X46),X48))))) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f99,f136])).
fof(f908,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f672,f79])).
fof(f672,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,X20) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f138,f136])).
fof(f907,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f671,f79])).
fof(f671,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f108,f136])).
fof(f604,plain,(
  ( ! [X116,X114,X117,X115] : (s(X116,f(s(bool,X117))) != s(X116,g(s(bool,X117))) | s(fun(bool,X116),X114) = s(fun(bool,X116),X115) | s(bool,'T_0') = s(bool,X117) | s(bool,'T_0') = s(bool,sK5(X114,X115,X116,bool))) )),
  inference(superposition,[],[f109,f136])).
fof(f376,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(subsumption_resolution,[],[f374,f81])).
fof(f374,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(superposition,[],[f203,f170])).
fof(f332,plain,(
  s('type.num.num',sK1) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))) | $spl3),
  inference(subsumption_resolution,[],[f328,f79])).
fof(f328,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK1) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))) | $spl3),
  inference(superposition,[],[f104,f163])).
fof(f269,plain,(
  sP0('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'T_0','T_0','const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | $spl3),
  inference(unit_resulting_resolution,[],[f163,f163,f123])).
fof(f271,plain,(
  sP0('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'T_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f163,f163,f123])).
fof(f272,plain,(
  sP0('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),'T_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f170,f163,f123])).
fof(f278,plain,(
  sP0('T_0','T_0','const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | $spl3),
  inference(unit_resulting_resolution,[],[f163,f163,f123])).
fof(f280,plain,(
  sP0('T_0','const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f163,f163,f123])).
fof(f281,plain,(
  sP0('T_0','const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f170,f163,f123])).
fof(f283,plain,(
  sP0('F_0','T_0','const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | $spl3),
  inference(unit_resulting_resolution,[],[f163,f170,f123])).
fof(f285,plain,(
  sP0('F_0','const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),'T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f163,f170,f123])).
fof(f286,plain,(
  sP0('F_0','const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),'F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f170,f170,f123])).
fof(f266,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | s(bool,'T_0') = s(bool,X5)) ) | $spl3),
  inference(subsumption_resolution,[],[f264,f79])).
fof(f264,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | s(bool,'T_0') = s(bool,X5)) ) | $spl3),
  inference(superposition,[],[f138,f163])).
fof(f251,plain,(
  ( ! [X0] : (~sP0('T_0','T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f81,f79,f79,f119])).
fof(f238,plain,(
  ( ! [X0] : (~sP0('T_0','T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f81,f79,f79,f118])).
fof(f234,plain,(
  ( ! [X0] : (~sP0('F_0',X0,'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f81,f79,f79,f117])).
fof(f227,plain,(
  ( ! [X0] : (~sP0(X0,'F_0','T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f81,f79,f79,f116])).
fof(f223,plain,(
  ( ! [X0,X1] : (sP0('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X0,'F_0',X1)) ) | $spl3),
  inference(unit_resulting_resolution,[],[f81,f170,f122])).
fof(f222,plain,(
  ( ! [X0,X1] : (sP0(X0,'F_0',X1,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f81,f170,f121])).
fof(f171,plain,(
  ~p(s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))))) | $spl3),
  inference(unit_resulting_resolution,[],[f169,f85])).
fof(f169,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | $spl3),
  inference(backward_demodulation,[],[f163,f78])).
fof(f163,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | $spl3),
  inference(unit_resulting_resolution,[],[f162,f83])).
fof(f164,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)))) | $spl3),
  inference(unit_resulting_resolution,[],[f162,f89])).
fof(f132,plain,(
  ( ! [X0,X1] : (sP0(X0,'F_0','F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f81,f81,f120])).
fof(f125,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f81,f84])).
fof(f120,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X3,X1,X2,X0) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,X2)) & s(bool,X0) != s(bool,X1)) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | sP0(X3,X1,X2,X0))),
  inference(definition_folding,[],[f54,f55])).
fof(f55,plain,(
  ! [X3,X1,X2,X0] : (~sP0(X3,X1,X2,X0) | ((p(s(bool,X0)) & p(s(bool,X2))) <=> (p(s(bool,X1)) & p(s(bool,X3)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f54,plain,(
  ! [X0,X1,X2,X3] : ((p(s(bool,X2)) & s(bool,X0) != s(bool,X1)) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((p(s(bool,X0)) & p(s(bool,X2))) <=> (p(s(bool,X1)) & p(s(bool,X3)))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3] : (((p(s(bool,X2)) & s(bool,X0) != s(bool,X1)) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((p(s(bool,X0)) & p(s(bool,X2))) <=> (p(s(bool,X1)) & p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3] : (((p(s(bool,X2)) => s(bool,X0) = s(bool,X1)) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) & p(s(bool,X2))) <=> (p(s(bool,X1)) & p(s(bool,X3)))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X15,X16,X17,X18] : (((p(s(bool,X17)) => s(bool,X15) = s(bool,X16)) & (p(s(bool,X16)) => s(bool,X17) = s(bool,X18))) => ((p(s(bool,X15)) & p(s(bool,X17))) <=> (p(s(bool,X16)) & p(s(bool,X18)))))),
  file('thm.rich_list.MEM_COUNT_LIST','thm.bool.AND_CONG')).
fof(f121,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X3,X1,X2,X0) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f56])).
fof(f122,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X3,X1,X2,X0) | s(bool,X2) != s(bool,X3) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f56])).
fof(f123,plain,(
  ( ! [X2,X0,X3,X1] : (sP0(X3,X1,X2,X0) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f56])).
fof(f116,plain,(
  ( ! [X2,X0,X3,X1] : (~sP0(X0,X1,X2,X3) | ~p(s(bool,X2)) | ~p(s(bool,X3)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : (~sP0(X0,X1,X2,X3) | ((~p(s(bool,X3)) | ~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | (p(s(bool,X3)) & p(s(bool,X2))))))),
  inference(rectify,[],[f76])).
fof(f76,plain,(
  ! [X3,X1,X2,X0] : (~sP0(X3,X1,X2,X0) | ((~p(s(bool,X0)) | ~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X3)))) & (~p(s(bool,X1)) | ~p(s(bool,X3)) | (p(s(bool,X0)) & p(s(bool,X2))))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X3,X1,X2,X0] : (~sP0(X3,X1,X2,X0) | (((~p(s(bool,X0)) | ~p(s(bool,X2))) | (p(s(bool,X1)) & p(s(bool,X3)))) & ((~p(s(bool,X1)) | ~p(s(bool,X3))) | (p(s(bool,X0)) & p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f55])).
fof(f117,plain,(
  ( ! [X2,X0,X3,X1] : (~sP0(X0,X1,X2,X3) | ~p(s(bool,X2)) | ~p(s(bool,X3)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f77])).
fof(f118,plain,(
  ( ! [X2,X0,X3,X1] : (~sP0(X0,X1,X2,X3) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X3))) )),
  inference(cnf_transformation,[],[f77])).
fof(f119,plain,(
  ( ! [X2,X0,X3,X1] : (~sP0(X0,X1,X2,X3) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f77])).
fof(f108,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f49])).
fof(f84,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f58])).
fof(f17050,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (s(bool,happ(s(fun('type.num.num',bool),X76),s('type.num.num',X77))) = s(bool,happ(s(fun(X73,bool),X74),s(X73,X75))) | s('type.num.num',X77) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X77)))),X76,'type.num.num')) | s(bool,'T_0') = s(bool,X72) | s(X73,X75) = s(X73,sK4(X75,X74,X73))) )),
  inference(duplicate_literal_removal,[],[f16766])).
fof(f16766,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (s(bool,happ(s(fun('type.num.num',bool),X76),s('type.num.num',X77))) = s(bool,happ(s(fun(X73,bool),X74),s(X73,X75))) | s('type.num.num',X77) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X77)))),X76,'type.num.num')) | s(bool,'T_0') = s(bool,X72) | s(X73,X75) = s(X73,sK4(X75,X74,X73)) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f2940,f916])).
fof(f17051,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X68),s('type.num.num',X69))) = s(bool,happ(s(fun('type.num.num',bool),X70),s('type.num.num',X71))) | s('type.num.num',X71) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X71)))),X70,'type.num.num')) | s(bool,'T_0') = s(bool,X67) | s('type.num.num',X68) = s('type.num.num','const.list.EL_2'(s('type.num.num',X68),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X69)))))) )),
  inference(duplicate_literal_removal,[],[f16765])).
fof(f16765,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X68),s('type.num.num',X69))) = s(bool,happ(s(fun('type.num.num',bool),X70),s('type.num.num',X71))) | s('type.num.num',X71) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X71)))),X70,'type.num.num')) | s(bool,'T_0') = s(bool,X67) | s('type.num.num',X68) = s('type.num.num','const.list.EL_2'(s('type.num.num',X68),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X69))))) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f2940,f913])).
fof(f17068,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X2))) | s('type.num.num',X2) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X2)))),X1,'type.num.num')) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f16752])).
fof(f16752,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',X2))) | s('type.num.num',X2) = s('type.num.num',sK4('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X2)))),X1,'type.num.num')) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2940,f83])).
fof(f13571,plain,(
  $false | ($spl38 | $spl102)),
  inference(subsumption_resolution,[],[f13570,f125])).
fof(f13570,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl38 | $spl102)),
  inference(forward_demodulation,[],[f12773,f3788])).
fof(f3788,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl38),
  inference(cnf_transformation,[],[f3788_D])).
fof(f3788_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f12773,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | $spl102),
  inference(equality_resolution,[],[f8956])).
fof(f13569,plain,(
  $false | ($spl3 | $spl38 | $spl102)),
  inference(subsumption_resolution,[],[f13568,f125])).
fof(f13568,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl38 | $spl102)),
  inference(forward_demodulation,[],[f12815,f3788])).
fof(f12815,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl3 | $spl102)),
  inference(forward_demodulation,[],[f12774,f2830])).
fof(f12774,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))))) | ($spl3 | $spl102)),
  inference(trivial_inequality_removal,[],[f12772])).
fof(f12772,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))))) | ($spl3 | $spl102)),
  inference(superposition,[],[f8956,f2830])).
fof(f13567,plain,(
  $false | ($spl38 | $spl102)),
  inference(subsumption_resolution,[],[f13566,f125])).
fof(f13566,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl38 | $spl102)),
  inference(condensation,[],[f13565])).
fof(f13565,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl38 | $spl102)),
  inference(forward_demodulation,[],[f12777,f3788])).
fof(f12777,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,g(s(bool,'T_0')))) ) | $spl102),
  inference(condensation,[],[f12776])).
fof(f12776,plain,(
  ( ! [X52,X53] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X52) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) ) | $spl102),
  inference(trivial_inequality_removal,[],[f12765])).
fof(f12765,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X52) | s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X52) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) ) | $spl102),
  inference(superposition,[],[f8956,f136])).
fof(f13564,plain,(
  $false | ($spl38 | $spl102)),
  inference(subsumption_resolution,[],[f13563,f125])).
fof(f13563,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl38 | $spl102)),
  inference(condensation,[],[f13562])).
fof(f13562,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl38 | $spl102)),
  inference(forward_demodulation,[],[f12779,f3788])).
fof(f12779,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,g(s(bool,'T_0')))) ) | $spl102),
  inference(condensation,[],[f12778])).
fof(f12778,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) ) | $spl102),
  inference(trivial_inequality_removal,[],[f12764])).
fof(f12764,plain,(
  ( ! [X50,X51] : (s(bool,X50) != s(bool,X50) | s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X50) = s(bool,X51) | s(bool,'T_0') = s(bool,X51)) ) | $spl102),
  inference(superposition,[],[f8956,f136])).
fof(f13561,plain,(
  $false | ($spl38 | $spl102)),
  inference(subsumption_resolution,[],[f13560,f125])).
fof(f13560,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl38 | $spl102)),
  inference(condensation,[],[f13559])).
fof(f13559,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl38 | $spl102)),
  inference(forward_demodulation,[],[f12783,f3788])).
fof(f12783,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,g(s(bool,'T_0')))) ) | $spl102),
  inference(condensation,[],[f12782])).
fof(f12782,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl102),
  inference(condensation,[],[f12781])).
fof(f12781,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl102),
  inference(condensation,[],[f12780])).
fof(f12780,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X47) = s(bool,X48) | s(bool,X46) = s(bool,X48) | s(bool,X47) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) ) | $spl102),
  inference(trivial_inequality_removal,[],[f12763])).
fof(f12763,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,X46) != s(bool,X46) | s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X47) = s(bool,X48) | s(bool,X46) = s(bool,X48) | s(bool,X47) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) ) | $spl102),
  inference(superposition,[],[f8956,f535])).
fof(f13558,plain,(
  $false | ($spl38 | $spl102)),
  inference(subsumption_resolution,[],[f13557,f125])).
fof(f13557,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl38 | $spl102)),
  inference(condensation,[],[f13556])).
fof(f13556,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl38 | $spl102)),
  inference(forward_demodulation,[],[f12787,f3788])).
fof(f12787,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,g(s(bool,'T_0')))) ) | $spl102),
  inference(condensation,[],[f12786])).
fof(f12786,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X0) = s(bool,X1)) ) | $spl102),
  inference(condensation,[],[f12785])).
fof(f12785,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,X1) = s(bool,X3)) ) | $spl102),
  inference(condensation,[],[f12784])).
fof(f12784,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X43) = s(bool,X44) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,X42) = s(bool,X43)) ) | $spl102),
  inference(trivial_inequality_removal,[],[f12762])).
fof(f12762,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X42) != s(bool,X42) | s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X43) = s(bool,X44) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,X42) = s(bool,X43)) ) | $spl102),
  inference(superposition,[],[f8956,f535])).
fof(f13555,plain,(
  $false | ($spl3 | $spl38 | $spl102)),
  inference(subsumption_resolution,[],[f13554,f125])).
fof(f13554,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl38 | $spl102)),
  inference(forward_demodulation,[],[f12805,f3788])).
fof(f12805,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl3 | $spl102)),
  inference(forward_demodulation,[],[f12730,f2830])).
fof(f12730,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))))) | ($spl3 | $spl102)),
  inference(unit_resulting_resolution,[],[f2830,f8956])).
fof(f13553,plain,(
  $false | ($spl38 | $spl104)),
  inference(subsumption_resolution,[],[f13552,f125])).
fof(f13552,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl38 | $spl104)),
  inference(backward_demodulation,[],[f3788,f13508])).
fof(f13508,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | $spl104),
  inference(unit_resulting_resolution,[],[f81,f125,f8960])).
fof(f8919,plain,(
  ( ! [X69] : (s(bool,'F_0') = s(bool,g(s(bool,X69))) | s(bool,'F_0') = s(bool,X69)) ) | $spl92),
  inference(cnf_transformation,[],[f8919_D])).
fof(f8919_D,plain,(
  ( ! [X69] : (s(bool,'F_0') = s(bool,g(s(bool,X69))) | s(bool,'F_0') = s(bool,X69)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f13547,plain,(
  ( ! [X74,X75] : (s(bool,'F_0') = s(bool,g(s(bool,X75))) | s(bool,'F_0') = s(bool,X75) | s(bool,'T_0') = s(bool,X74)) ) | $spl104),
  inference(subsumption_resolution,[],[f13528,f85])).
fof(f13528,plain,(
  ( ! [X74,X75] : (p(s(bool,X74)) | s(bool,'F_0') = s(bool,g(s(bool,X75))) | s(bool,'F_0') = s(bool,X75) | s(bool,'T_0') = s(bool,X74)) ) | $spl104),
  inference(superposition,[],[f8960,f83])).
fof(f13544,plain,(
  ( ! [X80] : (s(bool,'F_0') = s(bool,g(s(bool,X80))) | s(bool,X80) = s(bool,g(s(bool,'T_0')))) ) | $spl152),
  inference(cnf_transformation,[],[f13544_D])).
fof(f13544_D,plain,(
  ( ! [X80] : (s(bool,'F_0') = s(bool,g(s(bool,X80))) | s(bool,X80) = s(bool,g(s(bool,'T_0')))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f13543,plain,(
  ( ! [X80,X79] : (s(bool,'F_0') = s(bool,g(s(bool,X80))) | s(bool,X80) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X79)) ) | ($spl92 | $spl104)),
  inference(subsumption_resolution,[],[f13531,f85])).
fof(f13531,plain,(
  ( ! [X80,X79] : (p(s(bool,X79)) | s(bool,'F_0') = s(bool,g(s(bool,X80))) | s(bool,X80) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X79)) ) | ($spl92 | $spl104)),
  inference(superposition,[],[f8960,f11570])).
fof(f11570,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl92),
  inference(condensation,[],[f11569])).
fof(f11569,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl92),
  inference(condensation,[],[f11568])).
fof(f11568,plain,(
  ( ! [X6,X4,X5] : (s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl92),
  inference(subsumption_resolution,[],[f11514,f4079])).
fof(f4079,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,X1) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f4078])).
fof(f4078,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f4056])).
fof(f4056,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f908,f83])).
fof(f11514,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl92),
  inference(superposition,[],[f11348,f535])).
fof(f11348,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | $spl92),
  inference(unit_resulting_resolution,[],[f125,f8919])).
fof(f11345,plain,(
  ( ! [X21,X19] : (s(bool,X21) = s(bool,sK4(X21,X19,bool)) | s(bool,'T_0') = s(bool,X21)) ) | $spl148),
  inference(cnf_transformation,[],[f11345_D])).
fof(f11345_D,plain,(
  ( ! [X21,X19] : (s(bool,X21) = s(bool,sK4(X21,X19,bool)) | s(bool,'T_0') = s(bool,X21)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f11346,plain,(
  ( ! [X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'F_0'))))) ) | $spl150),
  inference(cnf_transformation,[],[f11346_D])).
fof(f11346_D,plain,(
  ( ! [X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'F_0'))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f11335,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'F_0')))) | s(bool,X21) = s(bool,sK4(X21,X19,bool)) | s(bool,'T_0') = s(bool,X21)) ) | $spl36),
  inference(superposition,[],[f292,f3787])).
fof(f11343,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,sK4('F_0',X16,bool))) ) | $spl144),
  inference(cnf_transformation,[],[f11343_D])).
fof(f11343_D,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,sK4('F_0',X16,bool))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f11344,plain,(
  ( ! [X17,X18] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))) | s(bool,'T_0') = s(bool,X18)) ) | $spl146),
  inference(cnf_transformation,[],[f11344_D])).
fof(f11344_D,plain,(
  ( ! [X17,X18] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))) | s(bool,'T_0') = s(bool,X18)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f11334,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))) | s(bool,'F_0') = s(bool,sK4('F_0',X16,bool)) | s(bool,'T_0') = s(bool,X18)) ) | $spl36),
  inference(superposition,[],[f293,f3787])).
fof(f11321,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,'F_0')))) | s(bool,X21) = s(bool,sK4(X21,X19,bool)) | s(bool,'T_0') = s(bool,X21)) ) | $spl36),
  inference(superposition,[],[f292,f3787])).
fof(f11320,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))) | s(bool,'F_0') = s(bool,sK4('F_0',X16,bool)) | s(bool,'T_0') = s(bool,X18)) ) | $spl36),
  inference(superposition,[],[f293,f3787])).
fof(f9787,plain,(
  $false | ($spl29 | $spl30 | $spl35)),
  inference(subsumption_resolution,[],[f9417,f9383])).
fof(f9383,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl29 | $spl30)),
  inference(subsumption_resolution,[],[f9344,f3113])).
fof(f9344,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl29),
  inference(trivial_inequality_removal,[],[f9329])).
fof(f9329,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl29),
  inference(superposition,[],[f3112,f136])).
fof(f9417,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl29 | $spl30 | $spl35)),
  inference(backward_demodulation,[],[f9383,f3185])).
fof(f9786,plain,(
  $false | ($spl29 | $spl30)),
  inference(subsumption_resolution,[],[f9413,f9383])).
fof(f9413,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl29 | $spl30)),
  inference(backward_demodulation,[],[f9383,f3112])).
fof(f9780,plain,(
  $false | ($spl29 | $spl30)),
  inference(subsumption_resolution,[],[f9779,f9383])).
fof(f9779,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | ($spl29 | $spl30)),
  inference(subsumption_resolution,[],[f9385,f84])).
fof(f9385,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl29 | $spl30)),
  inference(backward_demodulation,[],[f9383,f88])).
fof(f9778,plain,(
  $false | ($spl29 | $spl30)),
  inference(subsumption_resolution,[],[f9384,f79])).
fof(f9384,plain,(
  ~p(s(bool,'T_0')) | ($spl29 | $spl30)),
  inference(backward_demodulation,[],[f9383,f81])).
fof(f9153,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9152,f3112])).
fof(f9152,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(condensation,[],[f9151])).
fof(f9151,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl25),
  inference(condensation,[],[f9150])).
fof(f9150,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl25),
  inference(condensation,[],[f9149])).
fof(f9149,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,X26) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl25),
  inference(subsumption_resolution,[],[f9110,f6672])).
fof(f6672,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f6671])).
fof(f6671,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(duplicate_literal_removal,[],[f4508])).
fof(f4508,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f535,f83])).
fof(f9110,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') != s(bool,X25) | s(bool,X25) = s(bool,X26) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl25),
  inference(superposition,[],[f3109,f535])).
fof(f9148,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9147,f3112])).
fof(f9147,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(condensation,[],[f9146])).
fof(f9146,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl25),
  inference(duplicate_literal_removal,[],[f9145])).
fof(f9145,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl25),
  inference(condensation,[],[f9144])).
fof(f9144,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | $spl25),
  inference(subsumption_resolution,[],[f9108,f6672])).
fof(f9108,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X19) | s(bool,X19) = s(bool,X20) | s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | $spl25),
  inference(superposition,[],[f3109,f535])).
fof(f9141,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9140,f3112])).
fof(f9140,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(condensation,[],[f9139])).
fof(f9139,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl25),
  inference(subsumption_resolution,[],[f9101,f83])).
fof(f9101,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl25),
  inference(superposition,[],[f3109,f136])).
fof(f9138,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9137,f3112])).
fof(f9137,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(condensation,[],[f9136])).
fof(f9136,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl25),
  inference(subsumption_resolution,[],[f9100,f83])).
fof(f9100,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,f(s(bool,'F_0')))) ) | $spl25),
  inference(superposition,[],[f3109,f136])).
fof(f9135,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9111,f3112])).
fof(f9111,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(trivial_inequality_removal,[],[f9098])).
fof(f9098,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(superposition,[],[f3109,f83])).
fof(f9133,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9042,f3112])).
fof(f9042,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f3109,f83])).
fof(f9132,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9044,f3112])).
fof(f9044,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f125,f3109,f136])).
fof(f9131,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9048,f3112])).
fof(f9048,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f3109,f125,f3109,f3109,f535])).
fof(f9130,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9050,f3112])).
fof(f9050,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f125,f3109,f3109,f3109,f535])).
fof(f9129,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9051,f3112])).
fof(f9051,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f125,f125,f125,f3109,f535])).
fof(f9128,plain,(
  $false | ($spl1 | $spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9052,f3112])).
fof(f9052,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl1 | $spl25)),
  inference(unit_resulting_resolution,[],[f159,f159,f125,f3109,f535])).
fof(f9127,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9053,f3112])).
fof(f9053,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f3109,f3109,f125,f3109,f535])).
fof(f9126,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9054,f3112])).
fof(f9054,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f125,f125,f125,f3109,f535])).
fof(f9125,plain,(
  $false | ($spl1 | $spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9055,f3112])).
fof(f9055,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl1 | $spl25)),
  inference(unit_resulting_resolution,[],[f125,f159,f159,f3109,f535])).
fof(f9124,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9056,f3112])).
fof(f9056,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f79,f125,f3109,f907])).
fof(f9123,plain,(
  $false | ($spl1 | $spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9057,f3112])).
fof(f9057,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl1 | $spl25)),
  inference(unit_resulting_resolution,[],[f159,f2913,f3109,f908])).
fof(f9122,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9059,f3112])).
fof(f9059,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f125,f3109,f136])).
fof(f9121,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9060,f3112])).
fof(f9060,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f81,f3109,f203])).
fof(f9120,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9063,f3112])).
fof(f9063,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f3109,f125,f3109,f3109,f535])).
fof(f9119,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9066,f3112])).
fof(f9066,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f125,f125,f125,f3109,f535])).
fof(f9118,plain,(
  $false | ($spl1 | $spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9067,f3112])).
fof(f9067,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl1 | $spl25)),
  inference(unit_resulting_resolution,[],[f159,f159,f125,f3109,f535])).
fof(f9117,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9070,f3112])).
fof(f9070,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f3109,f125,f3109,f3109,f535])).
fof(f9116,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9073,f3112])).
fof(f9073,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f125,f125,f125,f3109,f535])).
fof(f9115,plain,(
  $false | ($spl1 | $spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9074,f3112])).
fof(f9074,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl1 | $spl25)),
  inference(unit_resulting_resolution,[],[f159,f159,f125,f3109,f535])).
fof(f9114,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9077,f3112])).
fof(f9077,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f3109,f125,f3109,f3109,f535])).
fof(f9113,plain,(
  $false | ($spl25 | $spl29)),
  inference(subsumption_resolution,[],[f9080,f3112])).
fof(f9080,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(unit_resulting_resolution,[],[f79,f125,f3109,f907])).
fof(f9039,plain,(
  ( ! [X461,X459,X460,X462,X458] : (s(fun(bool,bool),X458) = s(fun(bool,bool),X459) | s(bool,X460) = s(bool,sK5(X458,X459,bool,bool)) | s(bool,X460) = s(bool,X461) | s(bool,'T_0') = s(bool,X462) | s(bool,X461) = s(bool,X462) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X458,X459,bool,bool))))) ) | $spl142),
  inference(cnf_transformation,[],[f9039_D])).
fof(f9039_D,plain,(
  ( ! [X461,X459,X460,X462,X458] : (s(fun(bool,bool),X458) = s(fun(bool,bool),X459) | s(bool,X460) = s(bool,sK5(X458,X459,bool,bool)) | s(bool,X460) = s(bool,X461) | s(bool,'T_0') = s(bool,X462) | s(bool,X461) = s(bool,X462) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X458,X459,bool,bool))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f3185,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl35),
  inference(cnf_transformation,[],[f3185_D])).
fof(f3185_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f5225,plain,(
  ( ! [X461,X459,X460,X462,X458] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X458) = s(fun(bool,bool),X459) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X458,X459,bool,bool)))) | s(bool,X460) = s(bool,X461) | s(bool,X461) = s(bool,sK5(X458,X459,bool,bool)) | s(bool,X460) = s(bool,X462) | s(bool,'T_0') = s(bool,X462)) )),
  inference(superposition,[],[f356,f535])).
fof(f5588,plain,(
  ( ! [X461,X459,X460,X462,X458] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X458) = s(fun(bool,bool),X459) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X458,X459,bool,bool)))) | s(bool,X460) = s(bool,X461) | s(bool,X461) = s(bool,X462) | s(bool,'T_0') = s(bool,X462) | s(bool,X460) = s(bool,sK5(X458,X459,bool,bool))) )),
  inference(superposition,[],[f356,f535])).
fof(f3183,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool))))) ) | $spl32),
  inference(cnf_transformation,[],[f3183_D])).
fof(f3183_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f9038,plain,(
  ( ! [X482,X483] : (s(bool,'F_0') != s(bool,f(s(bool,X482))) | s(bool,'T_0') = s(bool,X483) | s(bool,X482) = s(bool,X483)) ) | $spl140),
  inference(cnf_transformation,[],[f9038_D])).
fof(f9038_D,plain,(
  ( ! [X482,X483] : (s(bool,'F_0') != s(bool,f(s(bool,X482))) | s(bool,'T_0') = s(bool,X483) | s(bool,X482) = s(bool,X483)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f4862,plain,(
  ( ! [X484,X480,X482,X481,X483] : (s(bool,'F_0') != s(bool,f(s(bool,X482))) | s(fun(bool,bool),X480) = s(fun(bool,bool),X481) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X480,X481,bool,bool)))) | s(bool,X483) = s(bool,sK5(X480,X481,bool,bool)) | s(bool,'T_0') = s(bool,X483) | s(bool,X482) = s(bool,X484) | s(bool,'T_0') = s(bool,X484)) )),
  inference(superposition,[],[f356,f535])).
fof(f4757,plain,(
  ( ! [X484,X480,X482,X481,X483] : (s(bool,'F_0') != s(bool,f(s(bool,X482))) | s(fun(bool,bool),X480) = s(fun(bool,bool),X481) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X480,X481,bool,bool)))) | s(bool,X482) = s(bool,X483) | s(bool,'T_0') = s(bool,X483) | s(bool,X484) = s(bool,sK5(X480,X481,bool,bool)) | s(bool,'T_0') = s(bool,X484)) )),
  inference(superposition,[],[f356,f535])).
fof(f9027,plain,(
  ( ! [X459,X458] : (s(fun(bool,bool),X458) = s(fun(bool,bool),X459) | s(bool,'T_0') = s(bool,sK5(X458,X459,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X458,X459,bool,bool))))) ) | $spl120),
  inference(cnf_transformation,[],[f9027_D])).
fof(f9027_D,plain,(
  ( ! [X459,X458] : (s(fun(bool,bool),X458) = s(fun(bool,bool),X459) | s(bool,'T_0') = s(bool,sK5(X458,X459,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X458,X459,bool,bool))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f3109,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl25),
  inference(cnf_transformation,[],[f3109_D])).
fof(f3109_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f3172,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X0,X1,bool,bool)))) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool))) )),
  inference(superposition,[],[f356,f83])).
fof(f9033,plain,(
  ( ! [X482] : (s(bool,'F_0') != s(bool,f(s(bool,X482))) | s(bool,'T_0') = s(bool,X482)) ) | $spl132),
  inference(cnf_transformation,[],[f9033_D])).
fof(f9033_D,plain,(
  ( ! [X482] : (s(bool,'F_0') != s(bool,f(s(bool,X482))) | s(bool,'T_0') = s(bool,X482)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f3174,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,f(s(bool,X7))) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X5,X6,bool,bool)))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,sK5(X5,X6,bool,bool))) )),
  inference(superposition,[],[f356,f136])).
fof(f3175,plain,(
  ( ! [X10,X8,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X10))) | s(fun(bool,bool),X8) = s(fun(bool,bool),X9) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X8,X9,bool,bool)))) | s(bool,'T_0') = s(bool,sK5(X8,X9,bool,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f356,f136])).
fof(f9037,plain,(
  ( ! [X112,X113] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X112) = s(fun(bool,bool),X113) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X112,X113,bool,bool)))) | s(bool,'T_0') = s(bool,sK5(X112,X113,bool,bool))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f3618,f2913])).
fof(f3618,plain,(
  ( ! [X112,X113] : (s(bool,'F_0') != s(bool,f(s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))))) | s(fun(bool,bool),X112) = s(fun(bool,bool),X113) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X112,X113,bool,bool)))) | s(bool,'T_0') = s(bool,sK5(X112,X113,bool,bool))) ) | $spl3),
  inference(superposition,[],[f356,f2822])).
fof(f9036,plain,(
  ( ! [X441,X440,X439] : (s(bool,'F_0') != s(bool,f(s(bool,happ(s(fun(X439,bool),X440),s(X439,X441))))) | s(X439,X441) = s(X439,sK4(X441,X440,X439))) ) | $spl138),
  inference(cnf_transformation,[],[f9036_D])).
fof(f9036_D,plain,(
  ( ! [X441,X440,X439] : (s(bool,'F_0') != s(bool,f(s(bool,happ(s(fun(X439,bool),X440),s(X439,X441))))) | s(X439,X441) = s(X439,sK4(X441,X440,X439))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f4221,plain,(
  ( ! [X441,X438,X440,X437,X439] : (s(bool,'F_0') != s(bool,f(s(bool,happ(s(fun(X439,bool),X440),s(X439,X441))))) | s(fun(bool,bool),X437) = s(fun(bool,bool),X438) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X437,X438,bool,bool)))) | s(X439,X441) = s(X439,sK4(X441,X440,X439)) | s(bool,'T_0') = s(bool,sK5(X437,X438,bool,bool))) )),
  inference(superposition,[],[f356,f916])).
fof(f3110,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) | $spl26),
  inference(cnf_transformation,[],[f3110_D])).
fof(f3110_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3107,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl22),
  inference(cnf_transformation,[],[f3107_D])).
fof(f3107_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f9035,plain,(
  ( ! [X468,X464] : (s(bool,X464) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X468) | s(bool,X464) = s(bool,X468)) ) | $spl136),
  inference(cnf_transformation,[],[f9035_D])).
fof(f9035_D,plain,(
  ( ! [X468,X464] : (s(bool,X464) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X468) | s(bool,X464) = s(bool,X468)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f4754,plain,(
  ( ! [X466,X465,X467,X468,X464] : (s(bool,X464) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X465) = s(fun(bool,bool),X466) | s(bool,'T_0') = s(bool,sK5(X465,X466,bool,bool)) | s(bool,X464) = s(bool,X467) | s(bool,'T_0') = s(bool,X467) | s(bool,X468) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X468)) )),
  inference(superposition,[],[f355,f535])).
fof(f4859,plain,(
  ( ! [X466,X465,X467,X468,X464] : (s(bool,X464) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X465) = s(fun(bool,bool),X466) | s(bool,'T_0') = s(bool,sK5(X465,X466,bool,bool)) | s(bool,X467) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X467) | s(bool,X464) = s(bool,X468) | s(bool,'T_0') = s(bool,X468)) )),
  inference(superposition,[],[f355,f535])).
fof(f9034,plain,(
  ( ! [X467,X468] : (s(bool,X467) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X468) | s(bool,X467) = s(bool,X468)) ) | $spl134),
  inference(cnf_transformation,[],[f9034_D])).
fof(f9034_D,plain,(
  ( ! [X467,X468] : (s(bool,X467) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X468) | s(bool,X467) = s(bool,X468)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f3113,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl30),
  inference(cnf_transformation,[],[f3113_D])).
fof(f3113_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4964,plain,(
  ( ! [X466,X465,X467,X468,X464] : (s(bool,X464) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X465) = s(fun(bool,bool),X466) | s(bool,'T_0') = s(bool,sK5(X465,X466,bool,bool)) | s(bool,X467) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X464) | s(bool,X467) = s(bool,X468) | s(bool,'T_0') = s(bool,X468)) )),
  inference(superposition,[],[f355,f535])).
fof(f9032,plain,(
  ( ! [X484,X480,X481,X483] : (s(fun(bool,bool),X480) = s(fun(bool,bool),X481) | s(bool,'T_0') = s(bool,X484) | s(bool,X483) = s(bool,X484) | s(bool,X483) = s(bool,sK5(X480,X481,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X480,X481,bool,bool))))) ) | $spl130),
  inference(cnf_transformation,[],[f9032_D])).
fof(f9032_D,plain,(
  ( ! [X484,X480,X481,X483] : (s(fun(bool,bool),X480) = s(fun(bool,bool),X481) | s(bool,'T_0') = s(bool,X484) | s(bool,X483) = s(bool,X484) | s(bool,X483) = s(bool,sK5(X480,X481,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X480,X481,bool,bool))))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f4967,plain,(
  ( ! [X484,X480,X482,X481,X483] : (s(bool,'F_0') != s(bool,f(s(bool,X482))) | s(fun(bool,bool),X480) = s(fun(bool,bool),X481) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X480,X481,bool,bool)))) | s(bool,X483) = s(bool,sK5(X480,X481,bool,bool)) | s(bool,'T_0') = s(bool,X482) | s(bool,X483) = s(bool,X484) | s(bool,'T_0') = s(bool,X484)) )),
  inference(superposition,[],[f356,f535])).
fof(f3106,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl20),
  inference(cnf_transformation,[],[f3106_D])).
fof(f3106_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f9031,plain,(
  ( ! [X467,X468,X464] : (s(bool,X464) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X468) | s(bool,X467) = s(bool,X468) | s(bool,X464) = s(bool,X467)) ) | $spl128),
  inference(cnf_transformation,[],[f9031_D])).
fof(f9031_D,plain,(
  ( ! [X467,X468,X464] : (s(bool,X464) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X468) | s(bool,X467) = s(bool,X468) | s(bool,X464) = s(bool,X467)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f5069,plain,(
  ( ! [X436,X438,X437,X439,X435] : (s(bool,X435) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X436) = s(fun(bool,bool),X437) | s(bool,'T_0') = s(bool,sK5(X436,X437,bool,bool)) | s(bool,X435) = s(bool,X438) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X438) = s(bool,X439) | s(bool,'T_0') = s(bool,X439)) )),
  inference(superposition,[],[f355,f535])).
fof(f9030,plain,(
  ( ! [X484,X482,X483] : (s(bool,'F_0') != s(bool,f(s(bool,X482))) | s(bool,'T_0') = s(bool,X484) | s(bool,X483) = s(bool,X484) | s(bool,X482) = s(bool,X483)) ) | $spl126),
  inference(cnf_transformation,[],[f9030_D])).
fof(f9030_D,plain,(
  ( ! [X484,X482,X483] : (s(bool,'F_0') != s(bool,f(s(bool,X482))) | s(bool,'T_0') = s(bool,X484) | s(bool,X483) = s(bool,X484) | s(bool,X482) = s(bool,X483)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f5072,plain,(
  ( ! [X453,X455,X451,X452,X454] : (s(bool,'F_0') != s(bool,f(s(bool,X453))) | s(fun(bool,bool),X451) = s(fun(bool,bool),X452) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X451,X452,bool,bool)))) | s(bool,X453) = s(bool,X454) | s(bool,'T_0') = s(bool,sK5(X451,X452,bool,bool)) | s(bool,X454) = s(bool,X455) | s(bool,'T_0') = s(bool,X455)) )),
  inference(superposition,[],[f356,f535])).
fof(f9029,plain,(
  ( ! [X445,X444,X446] : (s(bool,X444) = s(bool,X445) | s(bool,X444) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X446) | s(bool,X445) = s(bool,X446)) ) | $spl124),
  inference(cnf_transformation,[],[f9029_D])).
fof(f9029_D,plain,(
  ( ! [X445,X444,X446] : (s(bool,X444) = s(bool,X445) | s(bool,X444) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X446) | s(bool,X445) = s(bool,X446)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f3112,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl29),
  inference(cnf_transformation,[],[f3112_D])).
fof(f3112_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f5222,plain,(
  ( ! [X445,X443,X444,X446,X442] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X442) = s(fun(bool,bool),X443) | s(bool,'T_0') = s(bool,sK5(X442,X443,bool,bool)) | s(bool,X444) = s(bool,X445) | s(bool,X445) = s(bool,g(s(bool,'F_0'))) | s(bool,X444) = s(bool,X446) | s(bool,'T_0') = s(bool,X446)) )),
  inference(superposition,[],[f355,f535])).
fof(f5327,plain,(
  ( ! [X466,X465,X467,X468,X464] : (s(bool,X464) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X465) = s(fun(bool,bool),X466) | s(bool,'T_0') = s(bool,sK5(X465,X466,bool,bool)) | s(bool,X467) = s(bool,g(s(bool,'F_0'))) | s(bool,X467) = s(bool,X468) | s(bool,'T_0') = s(bool,X468) | s(bool,'T_0') = s(bool,X464)) )),
  inference(superposition,[],[f355,f535])).
fof(f5330,plain,(
  ( ! [X484,X480,X482,X481,X483] : (s(bool,'F_0') != s(bool,f(s(bool,X482))) | s(fun(bool,bool),X480) = s(fun(bool,bool),X481) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X480,X481,bool,bool)))) | s(bool,X483) = s(bool,sK5(X480,X481,bool,bool)) | s(bool,X483) = s(bool,X484) | s(bool,'T_0') = s(bool,X484) | s(bool,'T_0') = s(bool,X482)) )),
  inference(superposition,[],[f356,f535])).
fof(f5432,plain,(
  ( ! [X466,X465,X467,X468,X464] : (s(bool,X464) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X465) = s(fun(bool,bool),X466) | s(bool,'T_0') = s(bool,sK5(X465,X466,bool,bool)) | s(bool,X464) = s(bool,X467) | s(bool,X467) = s(bool,X468) | s(bool,'T_0') = s(bool,X468) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f355,f535])).
fof(f5435,plain,(
  ( ! [X484,X480,X482,X481,X483] : (s(bool,'F_0') != s(bool,f(s(bool,X482))) | s(fun(bool,bool),X480) = s(fun(bool,bool),X481) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X480,X481,bool,bool)))) | s(bool,X482) = s(bool,X483) | s(bool,X483) = s(bool,X484) | s(bool,'T_0') = s(bool,X484) | s(bool,'T_0') = s(bool,sK5(X480,X481,bool,bool))) )),
  inference(superposition,[],[f356,f535])).
fof(f5585,plain,(
  ( ! [X445,X443,X444,X446,X442] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X442) = s(fun(bool,bool),X443) | s(bool,'T_0') = s(bool,sK5(X442,X443,bool,bool)) | s(bool,X444) = s(bool,X445) | s(bool,X445) = s(bool,X446) | s(bool,'T_0') = s(bool,X446) | s(bool,X444) = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f355,f535])).
fof(f9028,plain,(
  ( ! [X461,X460] : (s(bool,'F_0') != s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X460),s('type.num.num',X461))))) | s('type.num.num',X460) = s('type.num.num','const.list.EL_2'(s('type.num.num',X460),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X461)))))) ) | $spl122),
  inference(cnf_transformation,[],[f9028_D])).
fof(f9028_D,plain,(
  ( ! [X461,X460] : (s(bool,'F_0') != s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X460),s('type.num.num',X461))))) | s('type.num.num',X460) = s('type.num.num','const.list.EL_2'(s('type.num.num',X460),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X461)))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f7786,plain,(
  ( ! [X461,X459,X460,X458] : (s(bool,'F_0') != s(bool,f(s(bool,'const.prim_rec.<_2'(s('type.num.num',X460),s('type.num.num',X461))))) | s(fun(bool,bool),X458) = s(fun(bool,bool),X459) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X458,X459,bool,bool)))) | s('type.num.num',X460) = s('type.num.num','const.list.EL_2'(s('type.num.num',X460),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X461))))) | s(bool,'T_0') = s(bool,sK5(X458,X459,bool,bool))) )),
  inference(superposition,[],[f356,f913])).
fof(f8925,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,X62) = s(bool,happ(s(fun(X59,bool),X60),s(X59,X61))) | s(X59,X61) = s(X59,sK4(X61,X60,X59)) | s(bool,'F_0') = s(bool,g(s(bool,X62)))) ) | $spl98),
  inference(cnf_transformation,[],[f8925_D])).
fof(f8925_D,plain,(
  ( ! [X61,X59,X62,X60] : (s(bool,X62) = s(bool,happ(s(fun(X59,bool),X60),s(X59,X61))) | s(X59,X61) = s(X59,sK4(X61,X60,X59)) | s(bool,'F_0') = s(bool,g(s(bool,X62)))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f8684,plain,(
  ( ! [X257,X259,X260,X256,X258] : (s(bool,'T_0') = s(bool,X259) | s(X256,X258) = s(X256,sK4(X258,X257,X256)) | s(bool,'F_0') = s(bool,g(s(bool,X260))) | s(bool,X260) = s(bool,happ(s(fun(X256,bool),X257),s(X256,X258)))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8644])).
fof(f8644,plain,(
  ( ! [X257,X259,X260,X256,X258] : (s(bool,'T_0') = s(bool,X259) | s(X256,X258) = s(X256,sK4(X258,X257,X256)) | s(bool,'T_0') = s(bool,X259) | s(bool,'F_0') = s(bool,g(s(bool,X260))) | s(bool,X260) = s(bool,happ(s(fun(X256,bool),X257),s(X256,X258)))) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f916,f3868])).
fof(f3868,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f3826,f136])).
fof(f3826,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl1 | $spl3 | $spl39)),
  inference(forward_demodulation,[],[f3803,f2913])).
fof(f3803,plain,(
  s(bool,g(s(bool,'T_0'))) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | ($spl3 | $spl39)),
  inference(unit_resulting_resolution,[],[f3789,f2822])).
fof(f8924,plain,(
  ( ! [X57,X56,X55] : (s(bool,X57) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X55),s('type.num.num',X56))) | s('type.num.num',X55) = s('type.num.num','const.list.EL_2'(s('type.num.num',X55),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X56))))) | s(bool,'F_0') = s(bool,g(s(bool,X57)))) ) | $spl96),
  inference(cnf_transformation,[],[f8924_D])).
fof(f8924_D,plain,(
  ( ! [X57,X56,X55] : (s(bool,X57) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X55),s('type.num.num',X56))) | s('type.num.num',X55) = s('type.num.num','const.list.EL_2'(s('type.num.num',X55),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X56))))) | s(bool,'F_0') = s(bool,g(s(bool,X57)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f8685,plain,(
  ( ! [X239,X237,X238,X236] : (s(bool,'T_0') = s(bool,X238) | s('type.num.num',X236) = s('type.num.num','const.list.EL_2'(s('type.num.num',X236),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X237))))) | s(bool,'F_0') = s(bool,g(s(bool,X239))) | s(bool,X239) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X236),s('type.num.num',X237)))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8638])).
fof(f8638,plain,(
  ( ! [X239,X237,X238,X236] : (s(bool,'T_0') = s(bool,X238) | s('type.num.num',X236) = s('type.num.num','const.list.EL_2'(s('type.num.num',X236),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X237))))) | s(bool,'T_0') = s(bool,X238) | s(bool,'F_0') = s(bool,g(s(bool,X239))) | s(bool,X239) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X236),s('type.num.num',X237)))) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f913,f3868])).
fof(f8988,plain,(
  ( ! [X125] : (s(bool,X125) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | s(bool,'F_0') = s(bool,g(s(bool,X125)))) ) | $spl118),
  inference(cnf_transformation,[],[f8988_D])).
fof(f8988_D,plain,(
  ( ! [X125] : (s(bool,X125) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | s(bool,'F_0') = s(bool,g(s(bool,X125)))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f8719,plain,(
  ( ! [X125,X124] : (s(bool,X125) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | s(bool,'T_0') = s(bool,X124) | s(bool,'F_0') = s(bool,g(s(bool,X125)))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8375])).
fof(f8375,plain,(
  ( ! [X125,X124] : (s(bool,X125) != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | s(bool,'T_0') = s(bool,X124) | s(bool,'F_0') = s(bool,g(s(bool,X125))) | s(bool,'T_0') = s(bool,X124)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f2823,f3868])).
fof(f8721,plain,(
  ( ! [X121,X118,X120,X119,X117] : (s(bool,X118) = s(bool,happ(s(fun(X119,bool),X120),s(X119,X121))) | s(X119,X121) = s(X119,sK4(X121,X120,X119)) | s(bool,'T_0') = s(bool,X117) | s(bool,'F_0') = s(bool,g(s(bool,X118)))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8373])).
fof(f8373,plain,(
  ( ! [X121,X118,X120,X119,X117] : (s(bool,X118) = s(bool,happ(s(fun(X119,bool),X120),s(X119,X121))) | s(X119,X121) = s(X119,sK4(X121,X120,X119)) | s(bool,'T_0') = s(bool,X117) | s(bool,'F_0') = s(bool,g(s(bool,X118))) | s(bool,'T_0') = s(bool,X117)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f916,f3868])).
fof(f8722,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,X114) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X115),s('type.num.num',X116))) | s('type.num.num',X115) = s('type.num.num','const.list.EL_2'(s('type.num.num',X115),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X116))))) | s(bool,'T_0') = s(bool,X113) | s(bool,'F_0') = s(bool,g(s(bool,X114)))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8372])).
fof(f8372,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,X114) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X115),s('type.num.num',X116))) | s('type.num.num',X115) = s('type.num.num','const.list.EL_2'(s('type.num.num',X115),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X116))))) | s(bool,'T_0') = s(bool,X113) | s(bool,'F_0') = s(bool,g(s(bool,X114))) | s(bool,'T_0') = s(bool,X113)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f913,f3868])).
fof(f8987,plain,(
  ( ! [X52,X51] : (~p(s(bool,'const.bool.IN_2'(s(bool,X51),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X52)))))) | s(bool,'F_0') = s(bool,g(s(bool,X51))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X52,bool)),s('type.list.list'(bool),X52)))) ) | $spl116),
  inference(cnf_transformation,[],[f8987_D])).
fof(f8987_D,plain,(
  ( ! [X52,X51] : (~p(s(bool,'const.bool.IN_2'(s(bool,X51),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X52)))))) | s(bool,'F_0') = s(bool,g(s(bool,X51))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X52,bool)),s('type.list.list'(bool),X52)))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f8742,plain,(
  ( ! [X52,X50,X51] : (~p(s(bool,'const.bool.IN_2'(s(bool,X51),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X52)))))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X52,bool)),s('type.list.list'(bool),X52))) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,g(s(bool,X51)))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8358])).
fof(f8358,plain,(
  ( ! [X52,X50,X51] : (~p(s(bool,'const.bool.IN_2'(s(bool,X51),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X52)))))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X52,bool)),s('type.list.list'(bool),X52))) | s(bool,'T_0') = s(bool,X50) | s(bool,'F_0') = s(bool,g(s(bool,X51))) | s(bool,'T_0') = s(bool,X50)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f380,f3868])).
fof(f8986,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,X48),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X49)))))) | s(bool,'F_0') = s(bool,g(s(bool,X48))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X49,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X49))))))) ) | $spl114),
  inference(cnf_transformation,[],[f8986_D])).
fof(f8986_D,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,X48),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X49)))))) | s(bool,'F_0') = s(bool,g(s(bool,X48))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X49,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X49))))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f8743,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,X48),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X49)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X49,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X49)))))) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,g(s(bool,X48)))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8357])).
fof(f8357,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,X48),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X49)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X49,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X49)))))) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,g(s(bool,X48))) | s(bool,'T_0') = s(bool,X47)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f362,f3868])).
fof(f8985,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,happ(s(fun(bool,bool),X45),s(bool,X44)))) | s(bool,'F_0') = s(bool,g(s(bool,X44))) | s(bool,'F_0') != s(bool,X46) | p(s(bool,happ(s(fun(bool,bool),X45),s(bool,X46))))) ) | $spl112),
  inference(cnf_transformation,[],[f8985_D])).
fof(f8985_D,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,happ(s(fun(bool,bool),X45),s(bool,X44)))) | s(bool,'F_0') = s(bool,g(s(bool,X44))) | s(bool,'F_0') != s(bool,X46) | p(s(bool,happ(s(fun(bool,bool),X45),s(bool,X46))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f8744,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,happ(s(fun(bool,bool),X45),s(bool,X44)))) | p(s(bool,happ(s(fun(bool,bool),X45),s(bool,X46)))) | s(bool,'F_0') != s(bool,X46) | s(bool,'T_0') = s(bool,X43) | s(bool,'F_0') = s(bool,g(s(bool,X44)))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8356])).
fof(f8356,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,happ(s(fun(bool,bool),X45),s(bool,X44)))) | p(s(bool,happ(s(fun(bool,bool),X45),s(bool,X46)))) | s(bool,'F_0') != s(bool,X46) | s(bool,'T_0') = s(bool,X43) | s(bool,'F_0') = s(bool,g(s(bool,X44))) | s(bool,'T_0') = s(bool,X43)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f345,f3868])).
fof(f8984,plain,(
  ( ! [X39,X38] : (~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X38)))) | s(bool,'F_0') = s(bool,g(s(bool,X38))) | s(bool,'F_0') = s(bool,sK4('F_0',X39,bool))) ) | $spl110),
  inference(cnf_transformation,[],[f8984_D])).
fof(f8984_D,plain,(
  ( ! [X39,X38] : (~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X38)))) | s(bool,'F_0') = s(bool,g(s(bool,X38))) | s(bool,'F_0') = s(bool,sK4('F_0',X39,bool))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f8745,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X38)))) | s(bool,'F_0') = s(bool,sK4('F_0',X39,bool)) | s(bool,'T_0') = s(bool,X37) | s(bool,'F_0') = s(bool,g(s(bool,X38)))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8354])).
fof(f8354,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,happ(s(fun(bool,bool),X39),s(bool,X38)))) | s(bool,'F_0') = s(bool,sK4('F_0',X39,bool)) | s(bool,'T_0') = s(bool,X37) | s(bool,'F_0') = s(bool,g(s(bool,X38))) | s(bool,'T_0') = s(bool,X37)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f293,f3868])).
fof(f8923,plain,(
  ( ! [X114,X113] : (s(bool,'F_0') = s(bool,g(s(bool,sK5(X113,X114,bool,bool)))) | s(fun(bool,bool),X113) = s(fun(bool,bool),X114)) ) | $spl94),
  inference(cnf_transformation,[],[f8923_D])).
fof(f8923_D,plain,(
  ( ! [X114,X113] : (s(bool,'F_0') = s(bool,g(s(bool,sK5(X113,X114,bool,bool)))) | s(fun(bool,bool),X113) = s(fun(bool,bool),X114)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f8752,plain,(
  ( ! [X329,X331,X330] : (s(bool,'T_0') = s(bool,X331) | s(fun(bool,bool),X329) = s(fun(bool,bool),X330) | s(bool,'F_0') = s(bool,g(s(bool,sK5(X329,X330,bool,bool))))) ) | ($spl1 | $spl3 | $spl22 | $spl39)),
  inference(duplicate_literal_removal,[],[f8328])).
fof(f8328,plain,(
  ( ! [X329,X331,X330] : (s(bool,'T_0') = s(bool,X331) | s(fun(bool,bool),X329) = s(fun(bool,bool),X330) | s(bool,'F_0') = s(bool,g(s(bool,sK5(X329,X330,bool,bool)))) | s(bool,'T_0') = s(bool,X331)) ) | ($spl1 | $spl3 | $spl22 | $spl39)),
  inference(superposition,[],[f3107,f3868])).
fof(f8977,plain,(
  ( ! [X292,X293,X291] : (~p(s(bool,happ(s(fun(X291,bool),X292),s(X291,X293)))) | s(bool,'F_0') = s(bool,g(s(bool,happ(s(fun(X291,bool),X292),s(X291,sK4(X293,X292,X291))))))) ) | $spl108),
  inference(cnf_transformation,[],[f8977_D])).
fof(f8977_D,plain,(
  ( ! [X292,X293,X291] : (~p(s(bool,happ(s(fun(X291,bool),X292),s(X291,X293)))) | s(bool,'F_0') = s(bool,g(s(bool,happ(s(fun(X291,bool),X292),s(X291,sK4(X293,X292,X291))))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f8976,plain,(
  ( ! [X292,X294,X293,X291] : (~p(s(bool,happ(s(fun(X291,bool),X292),s(X291,X293)))) | s(bool,'F_0') = s(bool,g(s(bool,happ(s(fun(X291,bool),X292),s(X291,sK4(X293,X292,X291)))))) | s(bool,'T_0') = s(bool,X294)) ) | ($spl1 | $spl3 | $spl39)),
  inference(subsumption_resolution,[],[f8317,f85])).
fof(f8317,plain,(
  ( ! [X292,X294,X293,X291] : (p(s(bool,X294)) | ~p(s(bool,happ(s(fun(X291,bool),X292),s(X291,X293)))) | s(bool,'F_0') = s(bool,g(s(bool,happ(s(fun(X291,bool),X292),s(X291,sK4(X293,X292,X291)))))) | s(bool,'T_0') = s(bool,X294)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f103,f3868])).
fof(f8973,plain,(
  ( ! [X267,X268] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X267),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X268)))))) | s(bool,'F_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X267)))),X268,'type.num.num')),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),X268)))))))) ) | $spl106),
  inference(cnf_transformation,[],[f8973_D])).
fof(f8973_D,plain,(
  ( ! [X267,X268] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X267),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X268)))))) | s(bool,'F_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X267)))),X268,'type.num.num')),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),X268)))))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f8972,plain,(
  ( ! [X269,X267,X268] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X267),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X268)))))) | s(bool,'F_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X267)))),X268,'type.num.num')),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),X268))))))) | s(bool,'T_0') = s(bool,X269)) ) | ($spl1 | $spl3 | $spl39)),
  inference(subsumption_resolution,[],[f8311,f85])).
fof(f8311,plain,(
  ( ! [X269,X267,X268] : (p(s(bool,X269)) | ~p(s(bool,'const.bool.IN_2'(s('type.num.num',X267),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X268)))))) | s(bool,'F_0') = s(bool,g(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X267)))),X268,'type.num.num')),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),X268))))))) | s(bool,'T_0') = s(bool,X269)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f359,f3868])).
fof(f8960,plain,(
  ( ! [X12,X10] : (p(s(bool,X12)) | s(bool,'F_0') = s(bool,g(s(bool,X10))) | s(bool,X10) = s(bool,X12)) ) | $spl104),
  inference(cnf_transformation,[],[f8960_D])).
fof(f8960_D,plain,(
  ( ! [X12,X10] : (p(s(bool,X12)) | s(bool,'F_0') = s(bool,g(s(bool,X10))) | s(bool,X10) = s(bool,X12)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f8961,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X15)) | s(bool,X13) = s(bool,X15) | s(bool,'F_0') = s(bool,g(s(bool,X13))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl1 | $spl3 | $spl39)),
  inference(subsumption_resolution,[],[f8214,f85])).
fof(f8214,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X13) = s(bool,X15) | s(bool,'F_0') = s(bool,g(s(bool,X13))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f105,f3868])).
fof(f8959,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X12)) | s(bool,X10) = s(bool,X12) | s(bool,'F_0') = s(bool,g(s(bool,X10))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl1 | $spl3 | $spl39)),
  inference(subsumption_resolution,[],[f8213,f85])).
fof(f8213,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X10) = s(bool,X12) | s(bool,'F_0') = s(bool,g(s(bool,X10))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f105,f3868])).
fof(f8956,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') = s(bool,g(s(bool,X2)))) ) | $spl102),
  inference(cnf_transformation,[],[f8956_D])).
fof(f8956_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') = s(bool,g(s(bool,X2)))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f7034,plain,(
  ( ! [X5] : (p(s(bool,X5))) ) | $spl60),
  inference(cnf_transformation,[],[f7034_D])).
fof(f7034_D,plain,(
  ( ! [X5] : (p(s(bool,X5))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f8955,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') = s(bool,g(s(bool,X2)))) ) | ($spl1 | $spl3 | $spl39)),
  inference(subsumption_resolution,[],[f8209,f84])).
fof(f8209,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') = s(bool,g(s(bool,X2))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f84,f3868])).
fof(f8789,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,'T_0') = s(bool,X81) | s(bool,'F_0') = s(bool,g(s(bool,X82))) | s(bool,X82) = s(bool,happ(s(fun(X78,bool),X79),s(X78,X80))) | s(X78,X80) = s(X78,sK4(X80,X79,X78))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8161])).
fof(f8161,plain,(
  ( ! [X80,X78,X81,X79,X82] : (s(bool,'T_0') = s(bool,X81) | s(bool,'F_0') = s(bool,g(s(bool,X82))) | s(bool,X82) = s(bool,happ(s(fun(X78,bool),X79),s(X78,X80))) | s(X78,X80) = s(X78,sK4(X80,X79,X78)) | s(bool,'T_0') = s(bool,X81)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f3868,f916])).
fof(f8790,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,g(s(bool,X77))) | s(bool,X77) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X74),s('type.num.num',X75))) | s('type.num.num',X74) = s('type.num.num','const.list.EL_2'(s('type.num.num',X74),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X75)))))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8160])).
fof(f8160,plain,(
  ( ! [X76,X74,X77,X75] : (s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,g(s(bool,X77))) | s(bool,X77) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X74),s('type.num.num',X75))) | s('type.num.num',X74) = s('type.num.num','const.list.EL_2'(s('type.num.num',X74),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X75))))) | s(bool,'T_0') = s(bool,X76)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f3868,f913])).
fof(f8926,plain,(
  ( ! [X64] : (s(bool,X64) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(bool,'F_0') = s(bool,g(s(bool,X64)))) ) | $spl100),
  inference(cnf_transformation,[],[f8926_D])).
fof(f8926_D,plain,(
  ( ! [X64] : (s(bool,X64) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(bool,'F_0') = s(bool,g(s(bool,X64)))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f8791,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') = s(bool,X71) | s(bool,'F_0') = s(bool,g(s(bool,X72))) | s(bool,X72) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8158])).
fof(f8158,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') = s(bool,X71) | s(bool,'F_0') = s(bool,g(s(bool,X72))) | s(bool,X72) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(bool,'T_0') = s(bool,X71)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f3868,f2822])).
fof(f8912,plain,(
  ( ! [X76,X74] : (s(bool,'F_0') = s(bool,g(s(bool,X74))) | s(bool,X74) = s(bool,X76) | s(bool,'F_0') = s(bool,g(s(bool,X76)))) ) | $spl84),
  inference(cnf_transformation,[],[f8912_D])).
fof(f8912_D,plain,(
  ( ! [X76,X74] : (s(bool,'F_0') = s(bool,g(s(bool,X74))) | s(bool,X74) = s(bool,X76) | s(bool,'F_0') = s(bool,g(s(bool,X76)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f8793,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,g(s(bool,X55))) | s(bool,X53) = s(bool,X55) | s(bool,'F_0') = s(bool,g(s(bool,X53)))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8153])).
fof(f8153,plain,(
  ( ! [X54,X55,X53] : (s(bool,'T_0') = s(bool,X54) | s(bool,'F_0') = s(bool,g(s(bool,X55))) | s(bool,X53) = s(bool,X55) | s(bool,'F_0') = s(bool,g(s(bool,X53))) | s(bool,'T_0') = s(bool,X54)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f3868,f3868])).
fof(f8824,plain,(
  ( ! [X70,X68,X69] : (s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,g(s(bool,X70))) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,g(s(bool,X69)))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8104])).
fof(f8104,plain,(
  ( ! [X70,X68,X69] : (s(bool,X69) = s(bool,X70) | s(bool,'F_0') = s(bool,g(s(bool,X70))) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,g(s(bool,X69))) | s(bool,'T_0') = s(bool,X68)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f3868,f3868])).
fof(f8825,plain,(
  ( ! [X64,X63] : (s(bool,X64) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(bool,'F_0') = s(bool,g(s(bool,X64))) | s(bool,'T_0') = s(bool,X63)) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8102])).
fof(f8102,plain,(
  ( ! [X64,X63] : (s(bool,X64) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(bool,'F_0') = s(bool,g(s(bool,X64))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f3868,f2822])).
fof(f8826,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,X62) = s(bool,happ(s(fun(X59,bool),X60),s(X59,X61))) | s(bool,'F_0') = s(bool,g(s(bool,X62))) | s(bool,'T_0') = s(bool,X58) | s(X59,X61) = s(X59,sK4(X61,X60,X59))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8101])).
fof(f8101,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,X62) = s(bool,happ(s(fun(X59,bool),X60),s(X59,X61))) | s(bool,'F_0') = s(bool,g(s(bool,X62))) | s(bool,'T_0') = s(bool,X58) | s(X59,X61) = s(X59,sK4(X61,X60,X59)) | s(bool,'T_0') = s(bool,X58)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f3868,f916])).
fof(f8827,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X57) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X55),s('type.num.num',X56))) | s(bool,'F_0') = s(bool,g(s(bool,X57))) | s(bool,'T_0') = s(bool,X54) | s('type.num.num',X55) = s('type.num.num','const.list.EL_2'(s('type.num.num',X55),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X56)))))) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8100])).
fof(f8100,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X57) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X55),s('type.num.num',X56))) | s(bool,'F_0') = s(bool,g(s(bool,X57))) | s(bool,'T_0') = s(bool,X54) | s('type.num.num',X55) = s('type.num.num','const.list.EL_2'(s('type.num.num',X55),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X56))))) | s(bool,'T_0') = s(bool,X54)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f3868,f913])).
fof(f8844,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8087])).
fof(f8087,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f3868,f83])).
fof(f8845,plain,(
  ( ! [X114,X115,X113] : (s(bool,'T_0') = s(bool,X115) | s(bool,'F_0') = s(bool,g(s(bool,sK5(X113,X114,bool,bool)))) | s(fun(bool,bool),X113) = s(fun(bool,bool),X114)) ) | ($spl1 | $spl3 | $spl22 | $spl39)),
  inference(duplicate_literal_removal,[],[f8086])).
fof(f8086,plain,(
  ( ! [X114,X115,X113] : (s(bool,'T_0') = s(bool,X115) | s(bool,'F_0') = s(bool,g(s(bool,sK5(X113,X114,bool,bool)))) | s(bool,'T_0') = s(bool,X115) | s(fun(bool,bool),X113) = s(fun(bool,bool),X114)) ) | ($spl1 | $spl3 | $spl22 | $spl39)),
  inference(superposition,[],[f3868,f3107])).
fof(f8862,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') = s(bool,X72) | s(bool,'F_0') = s(bool,g(s(bool,X71))) | s(bool,'F_0') = s(bool,g(s(bool,X73))) | s(bool,X71) = s(bool,X73)) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8071])).
fof(f8071,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') = s(bool,X72) | s(bool,'F_0') = s(bool,g(s(bool,X71))) | s(bool,'T_0') = s(bool,X72) | s(bool,'F_0') = s(bool,g(s(bool,X73))) | s(bool,X71) = s(bool,X73)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f3868,f3868])).
fof(f8051,plain,(
  ( ! [X70,X68,X69] : (s(bool,'F_0') = s(bool,X69) | s(bool,X68) = s(bool,X70) | s(bool,'T_0') = s(bool,X70) | s(bool,'F_0') = s(bool,g(s(bool,X69))) | s(bool,'T_0') = s(bool,g(s(bool,X68)))) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f3868,f3868])).
fof(f8913,plain,(
  ( ! [X99,X97,X98] : (s(bool,X98) = s(bool,g(s(bool,X97))) | s(bool,'T_0') = s(bool,X99) | s(bool,X97) = s(bool,X99) | s(bool,'T_0') = s(bool,X98)) ) | $spl86),
  inference(cnf_transformation,[],[f8913_D])).
fof(f8913_D,plain,(
  ( ! [X99,X97,X98] : (s(bool,X98) = s(bool,g(s(bool,X97))) | s(bool,'T_0') = s(bool,X99) | s(bool,X97) = s(bool,X99) | s(bool,'T_0') = s(bool,X98)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f8915,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | $spl88),
  inference(cnf_transformation,[],[f8915_D])).
fof(f8915_D,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f8028,plain,(
  ( ! [X99,X97,X98] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X98) = s(bool,g(s(bool,X97))) | s(bool,'T_0') = s(bool,X98) | s(bool,X97) = s(bool,X99) | s(bool,'T_0') = s(bool,X99)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f3868,f3868])).
fof(f8888,plain,(
  ( ! [X76,X74,X75] : (s(bool,'F_0') = s(bool,g(s(bool,X74))) | s(bool,'T_0') = s(bool,X75) | s(bool,'F_0') = s(bool,g(s(bool,X76))) | s(bool,X74) = s(bool,X76)) ) | ($spl1 | $spl3 | $spl39)),
  inference(duplicate_literal_removal,[],[f8020])).
fof(f8020,plain,(
  ( ! [X76,X74,X75] : (s(bool,'F_0') = s(bool,g(s(bool,X74))) | s(bool,'T_0') = s(bool,X75) | s(bool,'T_0') = s(bool,X75) | s(bool,'F_0') = s(bool,g(s(bool,X76))) | s(bool,X74) = s(bool,X76)) ) | ($spl1 | $spl3 | $spl39)),
  inference(superposition,[],[f3868,f3868])).
fof(f7963,plain,(
  ( ! [X35,X33,X32] : (p(s(bool,X35)) | s('type.num.num',X32) = s('type.num.num','const.list.EL_2'(s('type.num.num',X32),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X33))))) | s(bool,X35) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num',X33)))) ) | $spl82),
  inference(cnf_transformation,[],[f7963_D])).
fof(f7963_D,plain,(
  ( ! [X35,X33,X32] : (p(s(bool,X35)) | s('type.num.num',X32) = s('type.num.num','const.list.EL_2'(s('type.num.num',X32),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X33))))) | s(bool,X35) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num',X33)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f7964,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,X39)) | s(bool,X39) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num',X37))) | s('type.num.num',X36) = s('type.num.num','const.list.EL_2'(s('type.num.num',X36),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X37))))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f7799,f85])).
fof(f7799,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,X38)) | p(s(bool,X39)) | s(bool,X39) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num',X37))) | s('type.num.num',X36) = s('type.num.num','const.list.EL_2'(s('type.num.num',X36),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X37))))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f105,f913])).
fof(f7962,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,X35)) | s(bool,X35) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num',X33))) | s('type.num.num',X32) = s('type.num.num','const.list.EL_2'(s('type.num.num',X32),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X33))))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f7798,f85])).
fof(f7798,plain,(
  ( ! [X35,X33,X34,X32] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num',X33))) | s('type.num.num',X32) = s('type.num.num','const.list.EL_2'(s('type.num.num',X32),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X33))))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f105,f913])).
fof(f7958,plain,(
  ( ! [X12,X11] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X11),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X12)))))) | s('type.num.num',sK3('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X11)))),X12,'type.num.num')) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK3('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X11)))),X12,'type.num.num')),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),X12)))))))) ) | $spl80),
  inference(cnf_transformation,[],[f7958_D])).
fof(f7958_D,plain,(
  ( ! [X12,X11] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X11),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X12)))))) | s('type.num.num',sK3('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X11)))),X12,'type.num.num')) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK3('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X11)))),X12,'type.num.num')),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),X12)))))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f7957,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X11),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X12)))))) | s('type.num.num',sK3('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X11)))),X12,'type.num.num')) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK3('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X11)))),X12,'type.num.num')),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),X12))))))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(subsumption_resolution,[],[f7791,f85])).
fof(f7791,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X13)) | ~p(s(bool,'const.bool.IN_2'(s('type.num.num',X11),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),X12)))))) | s('type.num.num',sK3('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X11)))),X12,'type.num.num')) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK3('const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X11)))),X12,'type.num.num')),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),X12))))))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f359,f913])).
fof(f7953,plain,(
  ( ! [X156,X155] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X155),s('type.num.num',X156))) | s('type.num.num',X155) = s('type.num.num','const.list.EL_2'(s('type.num.num',X155),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X156)))))) ) | $spl78),
  inference(cnf_transformation,[],[f7953_D])).
fof(f7953_D,plain,(
  ( ! [X156,X155] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X155),s('type.num.num',X156))) | s('type.num.num',X155) = s('type.num.num','const.list.EL_2'(s('type.num.num',X155),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X156)))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f7952,plain,(
  ( ! [X156,X154,X155] : (s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',X155),s('type.num.num',X156))) | s(bool,'T_0') = s(bool,X154) | s('type.num.num',X155) = s('type.num.num','const.list.EL_2'(s('type.num.num',X155),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X156)))))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f7856,f2830])).
fof(f7856,plain,(
  ( ! [X156,X154,X155] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X155),s('type.num.num',X156))) | s(bool,'T_0') = s(bool,X154) | s('type.num.num',X155) = s('type.num.num','const.list.EL_2'(s('type.num.num',X155),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X156)))))) ) | $spl1),
  inference(duplicate_literal_removal,[],[f7742])).
fof(f7742,plain,(
  ( ! [X156,X154,X155] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s(bool,'const.prim_rec.<_2'(s('type.num.num',X155),s('type.num.num',X156))) | s(bool,'T_0') = s(bool,X154) | s('type.num.num',X155) = s('type.num.num','const.list.EL_2'(s('type.num.num',X155),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X156))))) | s(bool,'T_0') = s(bool,X154)) ) | $spl1),
  inference(superposition,[],[f2823,f913])).
fof(f7932,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X76),s('type.num.num',X77))) = s(bool,happ(s(fun(X73,bool),X74),s(X73,X75))) | s(X73,X75) = s(X73,sK4(X75,X74,X73)) | s('type.num.num',X76) = s('type.num.num','const.list.EL_2'(s('type.num.num',X76),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X77)))))) ) | $spl66),
  inference(cnf_transformation,[],[f7932_D])).
fof(f7932_D,plain,(
  ( ! [X76,X74,X77,X75,X73] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X76),s('type.num.num',X77))) = s(bool,happ(s(fun(X73,bool),X74),s(X73,X75))) | s(X73,X75) = s(X73,sK4(X75,X74,X73)) | s('type.num.num',X76) = s('type.num.num','const.list.EL_2'(s('type.num.num',X76),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X77)))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f7858,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X146),s('type.num.num',X147))) = s(bool,happ(s(fun(X148,bool),X149),s(X148,X150))) | s(X148,X150) = s(X148,sK4(X150,X149,X148)) | s(bool,'T_0') = s(bool,X145) | s('type.num.num',X146) = s('type.num.num','const.list.EL_2'(s('type.num.num',X146),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X147)))))) )),
  inference(duplicate_literal_removal,[],[f7740])).
fof(f7740,plain,(
  ( ! [X146,X149,X147,X145,X150,X148] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X146),s('type.num.num',X147))) = s(bool,happ(s(fun(X148,bool),X149),s(X148,X150))) | s(X148,X150) = s(X148,sK4(X150,X149,X148)) | s(bool,'T_0') = s(bool,X145) | s('type.num.num',X146) = s('type.num.num','const.list.EL_2'(s('type.num.num',X146),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X147))))) | s(bool,'T_0') = s(bool,X145)) )),
  inference(superposition,[],[f916,f913])).
fof(f7950,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X69),s('type.num.num',X70))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X71)))))) | s('type.num.num',X69) = s('type.num.num','const.list.EL_2'(s('type.num.num',X69),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X70))))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X71,bool)),s('type.list.list'(bool),X71)))) ) | $spl76),
  inference(cnf_transformation,[],[f7950_D])).
fof(f7950_D,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X69),s('type.num.num',X70))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X71)))))) | s('type.num.num',X69) = s('type.num.num','const.list.EL_2'(s('type.num.num',X69),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X70))))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X71,bool)),s('type.list.list'(bool),X71)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f7878,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X69),s('type.num.num',X70))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X71)))))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X71,bool)),s('type.list.list'(bool),X71))) | s(bool,'T_0') = s(bool,X68) | s('type.num.num',X69) = s('type.num.num','const.list.EL_2'(s('type.num.num',X69),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X70)))))) )),
  inference(duplicate_literal_removal,[],[f7726])).
fof(f7726,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X69),s('type.num.num',X70))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X71)))))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X71,bool)),s('type.list.list'(bool),X71))) | s(bool,'T_0') = s(bool,X68) | s('type.num.num',X69) = s('type.num.num','const.list.EL_2'(s('type.num.num',X69),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X70))))) | s(bool,'T_0') = s(bool,X68)) )),
  inference(superposition,[],[f380,f913])).
fof(f7949,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X65),s('type.num.num',X66))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X67)))))) | s('type.num.num',X65) = s('type.num.num','const.list.EL_2'(s('type.num.num',X65),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X66))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X67,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X67))))))) ) | $spl74),
  inference(cnf_transformation,[],[f7949_D])).
fof(f7949_D,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X65),s('type.num.num',X66))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X67)))))) | s('type.num.num',X65) = s('type.num.num','const.list.EL_2'(s('type.num.num',X65),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X66))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X67,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X67))))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f7879,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X65),s('type.num.num',X66))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X67)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X67,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X67)))))) | s(bool,'T_0') = s(bool,X64) | s('type.num.num',X65) = s('type.num.num','const.list.EL_2'(s('type.num.num',X65),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X66)))))) )),
  inference(duplicate_literal_removal,[],[f7725])).
fof(f7725,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X65),s('type.num.num',X66))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X67)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X67,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X67)))))) | s(bool,'T_0') = s(bool,X64) | s('type.num.num',X65) = s('type.num.num','const.list.EL_2'(s('type.num.num',X65),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X66))))) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f362,f913])).
fof(f7948,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,happ(s(fun(bool,bool),X62),s(bool,'const.prim_rec.<_2'(s('type.num.num',X60),s('type.num.num',X61)))))) | s('type.num.num',X60) = s('type.num.num','const.list.EL_2'(s('type.num.num',X60),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X61))))) | s(bool,'F_0') != s(bool,X63) | p(s(bool,happ(s(fun(bool,bool),X62),s(bool,X63))))) ) | $spl72),
  inference(cnf_transformation,[],[f7948_D])).
fof(f7948_D,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,happ(s(fun(bool,bool),X62),s(bool,'const.prim_rec.<_2'(s('type.num.num',X60),s('type.num.num',X61)))))) | s('type.num.num',X60) = s('type.num.num','const.list.EL_2'(s('type.num.num',X60),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X61))))) | s(bool,'F_0') != s(bool,X63) | p(s(bool,happ(s(fun(bool,bool),X62),s(bool,X63))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f7880,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun(bool,bool),X62),s(bool,'const.prim_rec.<_2'(s('type.num.num',X60),s('type.num.num',X61)))))) | p(s(bool,happ(s(fun(bool,bool),X62),s(bool,X63)))) | s(bool,'F_0') != s(bool,X63) | s(bool,'T_0') = s(bool,X59) | s('type.num.num',X60) = s('type.num.num','const.list.EL_2'(s('type.num.num',X60),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X61)))))) )),
  inference(duplicate_literal_removal,[],[f7724])).
fof(f7724,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,happ(s(fun(bool,bool),X62),s(bool,'const.prim_rec.<_2'(s('type.num.num',X60),s('type.num.num',X61)))))) | p(s(bool,happ(s(fun(bool,bool),X62),s(bool,X63)))) | s(bool,'F_0') != s(bool,X63) | s(bool,'T_0') = s(bool,X59) | s('type.num.num',X60) = s('type.num.num','const.list.EL_2'(s('type.num.num',X60),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X61))))) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f345,f913])).
fof(f7947,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,happ(s(fun(bool,bool),X54),s(bool,'const.prim_rec.<_2'(s('type.num.num',X52),s('type.num.num',X53)))))) | s('type.num.num',X52) = s('type.num.num','const.list.EL_2'(s('type.num.num',X52),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X53))))) | s(bool,'F_0') = s(bool,sK4('F_0',X54,bool))) ) | $spl70),
  inference(cnf_transformation,[],[f7947_D])).
fof(f7947_D,plain,(
  ( ! [X54,X52,X53] : (~p(s(bool,happ(s(fun(bool,bool),X54),s(bool,'const.prim_rec.<_2'(s('type.num.num',X52),s('type.num.num',X53)))))) | s('type.num.num',X52) = s('type.num.num','const.list.EL_2'(s('type.num.num',X52),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X53))))) | s(bool,'F_0') = s(bool,sK4('F_0',X54,bool))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f7881,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,happ(s(fun(bool,bool),X54),s(bool,'const.prim_rec.<_2'(s('type.num.num',X52),s('type.num.num',X53)))))) | s(bool,'F_0') = s(bool,sK4('F_0',X54,bool)) | s(bool,'T_0') = s(bool,X51) | s('type.num.num',X52) = s('type.num.num','const.list.EL_2'(s('type.num.num',X52),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X53)))))) )),
  inference(duplicate_literal_removal,[],[f7722])).
fof(f7722,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,happ(s(fun(bool,bool),X54),s(bool,'const.prim_rec.<_2'(s('type.num.num',X52),s('type.num.num',X53)))))) | s(bool,'F_0') = s(bool,sK4('F_0',X54,bool)) | s(bool,'T_0') = s(bool,X51) | s('type.num.num',X52) = s('type.num.num','const.list.EL_2'(s('type.num.num',X52),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X53))))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f293,f913])).
fof(f7940,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X0) != s('type.num.num',X2) | p(s(bool,'const.bool.IN_2'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) | $spl68),
  inference(cnf_transformation,[],[f7940_D])).
fof(f7940_D,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num',X0) != s('type.num.num',X2) | p(s(bool,'const.bool.IN_2'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X1)))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f7939,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num',X0) != s('type.num.num',X2) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.bool.IN_2'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X1)))))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f7691,f138])).
fof(f7691,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.num.num',X0) != s('type.num.num',X2) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.bool.IN_2'(s('type.num.num',X2),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X1)))))))) | s(bool,X3) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f490,f913])).
fof(f7930,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2))) | s('type.num.num',X1) = s('type.num.num','const.list.EL_2'(s('type.num.num',X1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X2)))))) ) | $spl62),
  inference(cnf_transformation,[],[f7930_D])).
fof(f7930_D,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2))) | s('type.num.num',X1) = s('type.num.num','const.list.EL_2'(s('type.num.num',X1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X2)))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f7933,plain,(
  ( ! [X80,X78,X79] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X79),s('type.num.num',X80))) | s('type.num.num',X79) = s('type.num.num','const.list.EL_2'(s('type.num.num',X79),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X80))))) | s(bool,'T_0') = s(bool,X78)) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f7906,f2913])).
fof(f7906,plain,(
  ( ! [X80,X78,X79] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X79),s('type.num.num',X80))) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s('type.num.num',X79) = s('type.num.num','const.list.EL_2'(s('type.num.num',X79),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X80))))) | s(bool,'T_0') = s(bool,X78)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f7660])).
fof(f7660,plain,(
  ( ! [X80,X78,X79] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X79),s('type.num.num',X80))) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s('type.num.num',X79) = s('type.num.num','const.list.EL_2'(s('type.num.num',X79),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X80))))) | s(bool,'T_0') = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) ) | $spl3),
  inference(superposition,[],[f913,f2822])).
fof(f7907,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X76),s('type.num.num',X77))) = s(bool,happ(s(fun(X73,bool),X74),s(X73,X75))) | s('type.num.num',X76) = s('type.num.num','const.list.EL_2'(s('type.num.num',X76),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X77))))) | s(bool,'T_0') = s(bool,X72) | s(X73,X75) = s(X73,sK4(X75,X74,X73))) )),
  inference(duplicate_literal_removal,[],[f7659])).
fof(f7659,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X76),s('type.num.num',X77))) = s(bool,happ(s(fun(X73,bool),X74),s(X73,X75))) | s('type.num.num',X76) = s('type.num.num','const.list.EL_2'(s('type.num.num',X76),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X77))))) | s(bool,'T_0') = s(bool,X72) | s(X73,X75) = s(X73,sK4(X75,X74,X73)) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f913,f916])).
fof(f7931,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X68),s('type.num.num',X69))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X70),s('type.num.num',X71))) | s('type.num.num',X68) = s('type.num.num','const.list.EL_2'(s('type.num.num',X68),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X69))))) | s('type.num.num',X70) = s('type.num.num','const.list.EL_2'(s('type.num.num',X70),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X71)))))) ) | $spl64),
  inference(cnf_transformation,[],[f7931_D])).
fof(f7931_D,plain,(
  ( ! [X70,X68,X71,X69] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X68),s('type.num.num',X69))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X70),s('type.num.num',X71))) | s('type.num.num',X68) = s('type.num.num','const.list.EL_2'(s('type.num.num',X68),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X69))))) | s('type.num.num',X70) = s('type.num.num','const.list.EL_2'(s('type.num.num',X70),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X71)))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f7908,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X68),s('type.num.num',X69))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X70),s('type.num.num',X71))) | s('type.num.num',X70) = s('type.num.num','const.list.EL_2'(s('type.num.num',X70),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X71))))) | s(bool,'T_0') = s(bool,X67) | s('type.num.num',X68) = s('type.num.num','const.list.EL_2'(s('type.num.num',X68),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X69)))))) )),
  inference(duplicate_literal_removal,[],[f7658])).
fof(f7658,plain,(
  ( ! [X70,X68,X71,X69,X67] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X68),s('type.num.num',X69))) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X70),s('type.num.num',X71))) | s('type.num.num',X70) = s('type.num.num','const.list.EL_2'(s('type.num.num',X70),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X71))))) | s(bool,'T_0') = s(bool,X67) | s('type.num.num',X68) = s('type.num.num','const.list.EL_2'(s('type.num.num',X68),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X69))))) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f913,f913])).
fof(f7925,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2))) | s('type.num.num',X1) = s('type.num.num','const.list.EL_2'(s('type.num.num',X1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X2))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f7645])).
fof(f7645,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2))) | s('type.num.num',X1) = s('type.num.num','const.list.EL_2'(s('type.num.num',X1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',X2))))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f913,f83])).
fof(f7033,plain,(
  ( ! [X6,X4,X7] : (s(bool,'T_0') != s(bool,X4) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X7) | s(bool,X4) = s(bool,X6)) ) | $spl58),
  inference(cnf_transformation,[],[f7033_D])).
fof(f7033_D,plain,(
  ( ! [X6,X4,X7] : (s(bool,'T_0') != s(bool,X4) | s(bool,'T_0') = s(bool,X7) | s(bool,X6) = s(bool,X7) | s(bool,X4) = s(bool,X6)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f7032,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X5)) | s(bool,'T_0') != s(bool,X4) | s(bool,X4) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f4864,f84])).
fof(f4864,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,X5)) | s(bool,'T_0') != s(bool,X4) | s(bool,X4) = s(bool,X6) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f84,f535])).
fof(f4323,plain,(
  ( ! [X61,X59,X60,X63] : (p(s(bool,X63)) | s(X59,X61) = s(X59,sK4(X61,X60,X59)) | s(bool,X63) = s(bool,happ(s(fun(X59,bool),X60),s(X59,X61)))) ) | $spl56),
  inference(cnf_transformation,[],[f4323_D])).
fof(f4323_D,plain,(
  ( ! [X61,X59,X60,X63] : (p(s(bool,X63)) | s(X59,X61) = s(X59,sK4(X61,X60,X59)) | s(bool,X63) = s(bool,happ(s(fun(X59,bool),X60),s(X59,X61)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4324,plain,(
  ( ! [X68,X66,X64,X67,X65] : (p(s(bool,X68)) | s(bool,X68) = s(bool,happ(s(fun(X64,bool),X65),s(X64,X66))) | s(X64,X66) = s(X64,sK4(X66,X65,X64)) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f4239,f85])).
fof(f4239,plain,(
  ( ! [X68,X66,X64,X67,X65] : (p(s(bool,X67)) | p(s(bool,X68)) | s(bool,X68) = s(bool,happ(s(fun(X64,bool),X65),s(X64,X66))) | s(X64,X66) = s(X64,sK4(X66,X65,X64)) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f105,f916])).
fof(f4322,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,X63)) | s(bool,X63) = s(bool,happ(s(fun(X59,bool),X60),s(X59,X61))) | s(X59,X61) = s(X59,sK4(X61,X60,X59)) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f4238,f85])).
fof(f4238,plain,(
  ( ! [X61,X59,X62,X60,X63] : (p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X63) = s(bool,happ(s(fun(X59,bool),X60),s(X59,X61))) | s(X59,X61) = s(X59,sK4(X61,X60,X59)) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f105,f916])).
fof(f4320,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,X15)))) | s(X13,sK4(X15,X14,X13)) = s(X13,sK4(sK4(X15,X14,X13),X14,X13))) ) | $spl54),
  inference(cnf_transformation,[],[f4320_D])).
fof(f4320_D,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,X15)))) | s(X13,sK4(X15,X14,X13)) = s(X13,sK4(sK4(X15,X14,X13),X14,X13))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f4319,plain,(
  ( ! [X14,X15,X13,X16] : (~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,X15)))) | s(X13,sK4(X15,X14,X13)) = s(X13,sK4(sK4(X15,X14,X13),X14,X13)) | s(bool,'T_0') = s(bool,X16)) )),
  inference(subsumption_resolution,[],[f4225,f85])).
fof(f4225,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X16)) | ~p(s(bool,happ(s(fun(X13,bool),X14),s(X13,X15)))) | s(X13,sK4(X15,X14,X13)) = s(X13,sK4(sK4(X15,X14,X13),X14,X13)) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f103,f916])).
fof(f4316,plain,(
  ( ! [X116,X114,X115] : (s(bool,'T_0') != s(bool,happ(s(fun(X114,bool),X115),s(X114,X116))) | s(X114,X116) = s(X114,sK4(X116,X115,X114))) ) | $spl52),
  inference(cnf_transformation,[],[f4316_D])).
fof(f4316_D,plain,(
  ( ! [X116,X114,X115] : (s(bool,'T_0') != s(bool,happ(s(fun(X114,bool),X115),s(X114,X116))) | s(X114,X116) = s(X114,sK4(X116,X115,X114))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f4315,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,'T_0') != s(bool,happ(s(fun(X114,bool),X115),s(X114,X116))) | s(bool,'T_0') = s(bool,X113) | s(X114,X116) = s(X114,sK4(X116,X115,X114))) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f4269,f2830])).
fof(f4269,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s(bool,happ(s(fun(X114,bool),X115),s(X114,X116))) | s(bool,'T_0') = s(bool,X113) | s(X114,X116) = s(X114,sK4(X116,X115,X114))) ) | $spl1),
  inference(duplicate_literal_removal,[],[f4186])).
fof(f4186,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s(bool,happ(s(fun(X114,bool),X115),s(X114,X116))) | s(bool,'T_0') = s(bool,X113) | s(X114,X116) = s(X114,sK4(X116,X115,X114)) | s(bool,'T_0') = s(bool,X113)) ) | $spl1),
  inference(superposition,[],[f2823,f916])).
fof(f4313,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X87,bool),X88),s(X87,X89))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X90)))))) | s(X87,X89) = s(X87,sK4(X89,X88,X87)) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X90,bool)),s('type.list.list'(bool),X90)))) ) | $spl50),
  inference(cnf_transformation,[],[f4313_D])).
fof(f4313_D,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X87,bool),X88),s(X87,X89))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X90)))))) | s(X87,X89) = s(X87,sK4(X89,X88,X87)) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X90,bool)),s('type.list.list'(bool),X90)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4276,plain,(
  ( ! [X90,X88,X87,X89,X86] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X87,bool),X88),s(X87,X89))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X90)))))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X90,bool)),s('type.list.list'(bool),X90))) | s(bool,'T_0') = s(bool,X86) | s(X87,X89) = s(X87,sK4(X89,X88,X87))) )),
  inference(duplicate_literal_removal,[],[f4181])).
fof(f4181,plain,(
  ( ! [X90,X88,X87,X89,X86] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X87,bool),X88),s(X87,X89))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X90)))))) | s(bool,'F_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('F_0',X90,bool)),s('type.list.list'(bool),X90))) | s(bool,'T_0') = s(bool,X86) | s(X87,X89) = s(X87,sK4(X89,X88,X87)) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f380,f916])).
fof(f4312,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X82,bool),X83),s(X82,X84))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X85)))))) | s(X82,X84) = s(X82,sK4(X84,X83,X82)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X85,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X85))))))) ) | $spl48),
  inference(cnf_transformation,[],[f4312_D])).
fof(f4312_D,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X82,bool),X83),s(X82,X84))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X85)))))) | s(X82,X84) = s(X82,sK4(X84,X83,X82)) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X85,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X85))))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f4277,plain,(
  ( ! [X85,X83,X81,X84,X82] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X82,bool),X83),s(X82,X84))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X85)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X85,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X85)))))) | s(bool,'T_0') = s(bool,X81) | s(X82,X84) = s(X82,sK4(X84,X83,X82))) )),
  inference(duplicate_literal_removal,[],[f4180])).
fof(f4180,plain,(
  ( ! [X85,X83,X81,X84,X82] : (~p(s(bool,'const.bool.IN_2'(s(bool,happ(s(fun(X82,bool),X83),s(X82,X84))),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X85)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('F_0',X85,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X85)))))) | s(bool,'T_0') = s(bool,X81) | s(X82,X84) = s(X82,sK4(X84,X83,X82)) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f362,f916])).
fof(f4311,plain,(
  ( ! [X80,X78,X76,X79,X77] : (~p(s(bool,happ(s(fun(bool,bool),X79),s(bool,happ(s(fun(X76,bool),X77),s(X76,X78)))))) | s(X76,X78) = s(X76,sK4(X78,X77,X76)) | s(bool,'F_0') != s(bool,X80) | p(s(bool,happ(s(fun(bool,bool),X79),s(bool,X80))))) ) | $spl46),
  inference(cnf_transformation,[],[f4311_D])).
fof(f4311_D,plain,(
  ( ! [X80,X78,X76,X79,X77] : (~p(s(bool,happ(s(fun(bool,bool),X79),s(bool,happ(s(fun(X76,bool),X77),s(X76,X78)))))) | s(X76,X78) = s(X76,sK4(X78,X77,X76)) | s(bool,'F_0') != s(bool,X80) | p(s(bool,happ(s(fun(bool,bool),X79),s(bool,X80))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4278,plain,(
  ( ! [X80,X78,X76,X79,X77,X75] : (~p(s(bool,happ(s(fun(bool,bool),X79),s(bool,happ(s(fun(X76,bool),X77),s(X76,X78)))))) | p(s(bool,happ(s(fun(bool,bool),X79),s(bool,X80)))) | s(bool,'F_0') != s(bool,X80) | s(bool,'T_0') = s(bool,X75) | s(X76,X78) = s(X76,sK4(X78,X77,X76))) )),
  inference(duplicate_literal_removal,[],[f4179])).
fof(f4179,plain,(
  ( ! [X80,X78,X76,X79,X77,X75] : (~p(s(bool,happ(s(fun(bool,bool),X79),s(bool,happ(s(fun(X76,bool),X77),s(X76,X78)))))) | p(s(bool,happ(s(fun(bool,bool),X79),s(bool,X80)))) | s(bool,'F_0') != s(bool,X80) | s(bool,'T_0') = s(bool,X75) | s(X76,X78) = s(X76,sK4(X78,X77,X76)) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f345,f916])).
fof(f4310,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,happ(s(fun(bool,bool),X69),s(bool,happ(s(fun(X66,bool),X67),s(X66,X68)))))) | s(X66,X68) = s(X66,sK4(X68,X67,X66)) | s(bool,'F_0') = s(bool,sK4('F_0',X69,bool))) ) | $spl44),
  inference(cnf_transformation,[],[f4310_D])).
fof(f4310_D,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,happ(s(fun(bool,bool),X69),s(bool,happ(s(fun(X66,bool),X67),s(X66,X68)))))) | s(X66,X68) = s(X66,sK4(X68,X67,X66)) | s(bool,'F_0') = s(bool,sK4('F_0',X69,bool))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4279,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,happ(s(fun(bool,bool),X69),s(bool,happ(s(fun(X66,bool),X67),s(X66,X68)))))) | s(bool,'F_0') = s(bool,sK4('F_0',X69,bool)) | s(bool,'T_0') = s(bool,X65) | s(X66,X68) = s(X66,sK4(X68,X67,X66))) )),
  inference(duplicate_literal_removal,[],[f4177])).
fof(f4177,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,happ(s(fun(bool,bool),X69),s(bool,happ(s(fun(X66,bool),X67),s(X66,X68)))))) | s(bool,'F_0') = s(bool,sK4('F_0',X69,bool)) | s(bool,'T_0') = s(bool,X65) | s(X66,X68) = s(X66,sK4(X68,X67,X66)) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f293,f916])).
fof(f4298,plain,(
  ( ! [X2,X3,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(X1,X3) = s(X1,sK4(X3,X2,X1))) ) | $spl40),
  inference(cnf_transformation,[],[f4298_D])).
fof(f4298_D,plain,(
  ( ! [X2,X3,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(X1,X3) = s(X1,sK4(X3,X2,X1))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f4300,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,'F_0') = s(bool,happ(s(fun(X32,bool),X33),s(X32,X34))) | s(X32,X34) = s(X32,sK4(X34,X33,X32)) | s(bool,'T_0') = s(bool,X31)) ) | ($spl1 | $spl3)),
  inference(forward_demodulation,[],[f4291,f2913])).
fof(f4291,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,happ(s(fun(X32,bool),X33),s(X32,X34))) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(X32,X34) = s(X32,sK4(X34,X33,X32)) | s(bool,'T_0') = s(bool,X31)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f4131])).
fof(f4131,plain,(
  ( ! [X33,X31,X34,X32] : (s(bool,happ(s(fun(X32,bool),X33),s(X32,X34))) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(X32,X34) = s(X32,sK4(X34,X33,X32)) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | $spl3),
  inference(superposition,[],[f916,f2822])).
fof(f4299,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,happ(s(fun(X25,bool),X26),s(X25,X27))) = s(bool,happ(s(fun(X28,bool),X29),s(X28,X30))) | s(X25,X27) = s(X25,sK4(X27,X26,X25)) | s(X28,X30) = s(X28,sK4(X30,X29,X28))) ) | $spl42),
  inference(cnf_transformation,[],[f4299_D])).
fof(f4299_D,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,happ(s(fun(X25,bool),X26),s(X25,X27))) = s(bool,happ(s(fun(X28,bool),X29),s(X28,X30))) | s(X25,X27) = s(X25,sK4(X27,X26,X25)) | s(X28,X30) = s(X28,sK4(X30,X29,X28))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4292,plain,(
  ( ! [X30,X28,X26,X24,X29,X27,X25] : (s(bool,happ(s(fun(X25,bool),X26),s(X25,X27))) = s(bool,happ(s(fun(X28,bool),X29),s(X28,X30))) | s(X28,X30) = s(X28,sK4(X30,X29,X28)) | s(bool,'T_0') = s(bool,X24) | s(X25,X27) = s(X25,sK4(X27,X26,X25))) )),
  inference(duplicate_literal_removal,[],[f4130])).
fof(f4130,plain,(
  ( ! [X30,X28,X26,X24,X29,X27,X25] : (s(bool,happ(s(fun(X25,bool),X26),s(X25,X27))) = s(bool,happ(s(fun(X28,bool),X29),s(X28,X30))) | s(X28,X30) = s(X28,sK4(X30,X29,X28)) | s(bool,'T_0') = s(bool,X24) | s(X25,X27) = s(X25,sK4(X27,X26,X25)) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f916,f916])).
fof(f4297,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(X1,X3) = s(X1,sK4(X3,X2,X1)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f4125])).
fof(f4125,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,happ(s(fun(X1,bool),X2),s(X1,X3))) | s(X1,X3) = s(X1,sK4(X3,X2,X1)) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f916,f83])).
fof(f3787,plain,(
  ( ! [X57,X58] : (s(fun(bool,bool),X57) = s(fun(bool,bool),X58)) ) | $spl36),
  inference(cnf_transformation,[],[f3787_D])).
fof(f3787_D,plain,(
  ( ! [X57,X58] : (s(fun(bool,bool),X57) = s(fun(bool,bool),X58)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f3789,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | $spl39),
  inference(cnf_transformation,[],[f3789_D])).
fof(f3789_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f3786,plain,(
  ( ! [X57,X58] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X57) = s(fun(bool,bool),X58)) ) | ($spl22 | $spl35)),
  inference(forward_demodulation,[],[f3773,f3196])).
fof(f3196,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl35),
  inference(unit_resulting_resolution,[],[f81,f3185,f203])).
fof(f3773,plain,(
  ( ! [X57,X58] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X57) = s(fun(bool,bool),X58)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f3767])).
fof(f3767,plain,(
  ( ! [X57,X58] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X57) = s(fun(bool,bool),X58) | s(fun(bool,bool),X57) = s(fun(bool,bool),X58)) ) | $spl22),
  inference(superposition,[],[f109,f3107])).
fof(f3176,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK5(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f356,f136])).
fof(f3173,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK5(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK5(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f356,f136])).
fof(f3104,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK5(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f355,f136])).
fof(f3103,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK5(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f355,f136])).
fof(f3102,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK5(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f355,f136])).
fof(f3101,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK5(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f355,f136])).
fof(f3100,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK5(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f355,f83])).
fof(f2816,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f162,f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f203,f227,f234,f238,f251,f535,f604,f907,f908,f910,f706,f916,f703,f705,f702,f663,f662,f603,f601,f600,f297,f915,f102,f101,f103,f306,f104,f913,f701,f700,f661,f660,f599,f598,f362,f490,f359,f385,f100,f99,f98,f911,f914,f602,f704,f1240,f1153,f380,f379,f361,f356,f355,f345,f344,f304,f293,f292,f89,f88,f81,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792,f2788,f2781,f2779,f2772,f2771,f2769,f2806,f2808])).
fof(f2808,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | $spl0),
  inference(forward_demodulation,[],[f78,f160])).
fof(f2806,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | $spl0),
  inference(forward_demodulation,[],[f2752,f2767])).
fof(f2752,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2))) | $spl0),
  inference(unit_resulting_resolution,[],[f125,f2684,f138])).
fof(f2769,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | $spl0),
  inference(backward_demodulation,[],[f2767,f2724])).
fof(f2771,plain,(
  ( ! [X0] : (~sP0(X0,'F_0','T_0','const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)))) ) | $spl0),
  inference(backward_demodulation,[],[f2767,f2729])).
fof(f2772,plain,(
  ( ! [X0] : (~sP0(X0,'F_0','const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)))) ) | $spl0),
  inference(backward_demodulation,[],[f2767,f2730])).
fof(f2779,plain,(
  ( ! [X0] : (~sP0('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'T_0','F_0',X0)) ) | $spl0),
  inference(backward_demodulation,[],[f2767,f2745])).
fof(f2781,plain,(
  ( ! [X0] : (~sP0('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'F_0',X0)) ) | $spl0),
  inference(backward_demodulation,[],[f2767,f2750])).
fof(f2788,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | $spl0),
  inference(forward_demodulation,[],[f2787,f2767])).
fof(f2787,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2))) | $spl0),
  inference(subsumption_resolution,[],[f2754,f81])).
fof(f2754,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2))) | $spl0),
  inference(superposition,[],[f2684,f83])).
fof(f2792,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | $spl0),
  inference(forward_demodulation,[],[f2791,f2767])).
fof(f2791,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2))) | $spl0),
  inference(condensation,[],[f2790])).
fof(f2790,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)))) ) | $spl0),
  inference(subsumption_resolution,[],[f2756,f85])).
fof(f2756,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)))) ) | $spl0),
  inference(superposition,[],[f2684,f136])).
fof(f2796,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | $spl0),
  inference(forward_demodulation,[],[f2795,f2767])).
fof(f2795,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2))) | $spl0),
  inference(condensation,[],[f2794])).
fof(f2794,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2))) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(subsumption_resolution,[],[f2757,f85])).
fof(f2757,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2))) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(superposition,[],[f2684,f136])).
fof(f2780,plain,(
  ( ! [X0] : (~sP0('T_0','const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'F_0',X0)) ) | $spl0),
  inference(backward_demodulation,[],[f2767,f2749])).
fof(f2778,plain,(
  ( ! [X0] : (~sP0('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,'F_0')) ) | $spl0),
  inference(backward_demodulation,[],[f2767,f2743])).
fof(f2777,plain,(
  ( ! [X0] : (~sP0('T_0','const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,'F_0')) ) | $spl0),
  inference(backward_demodulation,[],[f2767,f2742])).
fof(f2776,plain,(
  ( ! [X0] : (~sP0('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'T_0',X0,'F_0')) ) | $spl0),
  inference(backward_demodulation,[],[f2767,f2739])).
fof(f2775,plain,(
  ( ! [X0] : (~sP0('F_0',X0,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)))) ) | $spl0),
  inference(backward_demodulation,[],[f2767,f2737])).
fof(f2774,plain,(
  ( ! [X0] : (~sP0('F_0',X0,'T_0','const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)))) ) | $spl0),
  inference(backward_demodulation,[],[f2767,f2736])).
fof(f2773,plain,(
  ( ! [X0] : (~sP0('F_0',X0,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'T_0')) ) | $spl0),
  inference(backward_demodulation,[],[f2767,f2733])).
fof(f2770,plain,(
  ( ! [X0] : (~sP0(X0,'F_0','const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),'T_0')) ) | $spl0),
  inference(backward_demodulation,[],[f2767,f2726])).
fof(f2768,plain,(
  s('type.num.num',sK1) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))) | $spl0),
  inference(backward_demodulation,[],[f2767,f2682])).
fof(f2767,plain,(
  s('type.num.num',sK1) = s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')) | $spl0),
  inference(forward_demodulation,[],[f2719,f2682])).
fof(f2719,plain,(
  s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))) | $spl0),
  inference(unit_resulting_resolution,[],[f2684,f104])).
fof(f2724,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2))) | $spl0),
  inference(unit_resulting_resolution,[],[f79,f2684,f108])).
fof(f2726,plain,(
  ( ! [X0] : (~sP0(X0,'F_0','const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)),'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f81,f79,f2684,f116])).
fof(f2729,plain,(
  ( ! [X0] : (~sP0(X0,'F_0','T_0','const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f81,f79,f2684,f116])).
fof(f2730,plain,(
  ( ! [X0] : (~sP0(X0,'F_0','const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)),'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f81,f2684,f2684,f116])).
fof(f2733,plain,(
  ( ! [X0] : (~sP0('F_0',X0,'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)),'T_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f81,f79,f2684,f117])).
fof(f2736,plain,(
  ( ! [X0] : (~sP0('F_0',X0,'T_0','const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f81,f79,f2684,f117])).
fof(f2737,plain,(
  ( ! [X0] : (~sP0('F_0',X0,'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)),'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f81,f2684,f2684,f117])).
fof(f2739,plain,(
  ( ! [X0] : (~sP0('const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)),'T_0',X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f81,f79,f2684,f118])).
fof(f2742,plain,(
  ( ! [X0] : (~sP0('T_0','const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f81,f79,f2684,f118])).
fof(f2743,plain,(
  ( ! [X0] : (~sP0('const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)),'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)),X0,'F_0')) ) | $spl0),
  inference(unit_resulting_resolution,[],[f81,f2684,f2684,f118])).
fof(f2745,plain,(
  ( ! [X0] : (~sP0('const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)),'T_0','F_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f81,f79,f2684,f119])).
fof(f2749,plain,(
  ( ! [X0] : (~sP0('T_0','const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)),'F_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f81,f79,f2684,f119])).
fof(f2750,plain,(
  ( ! [X0] : (~sP0('const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)),'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)),'F_0',X0)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f81,f2684,f2684,f119])).
fof(f2690,plain,(
  ( ! [X12] : (s(bool,'T_0') = s(bool,sK4('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X12,bool)) | ~p(s(bool,happ(s(fun(bool,bool),X12),s(bool,'T_0'))))) ) | $spl0),
  inference(forward_demodulation,[],[f2670,f160])).
fof(f2670,plain,(
  ( ! [X12] : (~p(s(bool,happ(s(fun(bool,bool),X12),s(bool,'T_0')))) | s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) = s(bool,sK4('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X12,bool))) ) | $spl0),
  inference(superposition,[],[f102,f160])).
fof(f2668,plain,(
  ( ! [X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X9)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X9,bool)),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'(bool),X9))))))) ) | $spl0),
  inference(superposition,[],[f98,f160])).
fof(f2688,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,'const.list.EL_2'(s('type.num.num',sK3('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X8,bool)),s('type.list.list'(bool),X8))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X8))))))) ) | $spl0),
  inference(forward_demodulation,[],[f2667,f160])).
fof(f2667,plain,(
  ( ! [X8] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(bool),X8)))))) | s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) = s(bool,'const.list.EL_2'(s('type.num.num',sK3('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),X8,bool)),s('type.list.list'(bool),X8)))) ) | $spl0),
  inference(superposition,[],[f99,f160])).
fof(f2687,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(bool,'T_0') = s(bool,X5)) ) | $spl0),
  inference(subsumption_resolution,[],[f2664,f79])).
fof(f2664,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | s(bool,X5) != s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(bool,'T_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f138,f160])).
fof(f2684,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num',sK2)))) | $spl0),
  inference(forward_demodulation,[],[f2683,f82])).
fof(f2683,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))))) | $spl0),
  inference(subsumption_resolution,[],[f2653,f79])).
fof(f2653,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.num.num','const.list.LENGTH_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))))) | $spl0),
  inference(superposition,[],[f98,f160])).
fof(f2682,plain,(
  s('type.num.num',sK1) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))) | $spl0),
  inference(subsumption_resolution,[],[f2652,f79])).
fof(f2652,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK1) = s('type.num.num','const.list.EL_2'(s('type.num.num',sK3(sK1,'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)),'type.num.num')),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))) | $spl0),
  inference(superposition,[],[f99,f160])).
fof(f2632,plain,(
  sP0('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),'T_0','T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f160,f160,f123])).
fof(f2637,plain,(
  sP0('const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),'T_0','T_0','const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f160,f160,f123])).
fof(f2642,plain,(
  sP0('T_0','T_0','const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | $spl0),
  inference(unit_resulting_resolution,[],[f160,f160,f123])).
fof(f2644,plain,(
  sP0('T_0','const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))),'T_0') | $spl0),
  inference(unit_resulting_resolution,[],[f160,f160,f123])).
fof(f2815,plain,(
  $false | ($spl0 | $spl8)),
  inference(global_subsumption,[],[f2530,f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f203,f227,f234,f238,f251,f535,f604,f907,f908,f910,f706,f916,f703,f705,f702,f663,f662,f603,f601,f600,f297,f915,f102,f101,f103,f306,f104,f913,f701,f700,f661,f660,f599,f598,f362,f490,f359,f385,f100,f99,f98,f911,f914,f602,f704,f1240,f1153,f380,f379,f361,f356,f355,f345,f344,f304,f293,f292,f89,f88,f81,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792,f2788,f2781,f2779,f2772,f2771,f2769,f2806,f2808])).
fof(f2814,plain,(
  $false | ($spl0 | $spl10)),
  inference(global_subsumption,[],[f2532,f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f203,f227,f234,f238,f251,f535,f604,f907,f908,f910,f706,f916,f703,f705,f702,f663,f662,f603,f601,f600,f297,f915,f102,f101,f103,f306,f104,f913,f701,f700,f661,f660,f599,f598,f362,f490,f359,f385,f100,f99,f98,f911,f914,f602,f704,f1240,f1153,f380,f379,f361,f356,f355,f345,f344,f304,f293,f292,f89,f88,f81,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792,f2788,f2781,f2779,f2772,f2771,f2769,f2806,f2808])).
fof(f2813,plain,(
  $false | ($spl0 | $spl12)),
  inference(global_subsumption,[],[f2541,f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f203,f227,f234,f238,f251,f535,f604,f907,f908,f910,f706,f916,f703,f705,f702,f663,f662,f603,f601,f600,f297,f915,f102,f101,f103,f306,f104,f913,f701,f700,f661,f660,f599,f598,f362,f490,f359,f385,f100,f99,f98,f911,f914,f602,f704,f1240,f1153,f380,f379,f361,f356,f355,f345,f344,f304,f293,f292,f89,f88,f81,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792,f2788,f2781,f2779,f2772,f2771,f2769,f2806,f2808])).
fof(f2812,plain,(
  $false | ($spl0 | $spl14)),
  inference(global_subsumption,[],[f2542,f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f203,f227,f234,f238,f251,f535,f604,f907,f908,f910,f706,f916,f703,f705,f702,f663,f662,f603,f601,f600,f297,f915,f102,f101,f103,f306,f104,f913,f701,f700,f661,f660,f599,f598,f362,f490,f359,f385,f100,f99,f98,f911,f914,f602,f704,f1240,f1153,f380,f379,f361,f356,f355,f345,f344,f304,f293,f292,f89,f88,f81,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792,f2788,f2781,f2779,f2772,f2771,f2769,f2806,f2808])).
fof(f2811,plain,(
  $false | ($spl0 | $spl16)),
  inference(global_subsumption,[],[f2544,f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f203,f227,f234,f238,f251,f535,f604,f907,f908,f910,f706,f916,f703,f705,f702,f663,f662,f603,f601,f600,f297,f915,f102,f101,f103,f306,f104,f913,f701,f700,f661,f660,f599,f598,f362,f490,f359,f385,f100,f99,f98,f911,f914,f602,f704,f1240,f1153,f380,f379,f361,f356,f355,f345,f344,f304,f293,f292,f89,f88,f81,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792,f2788,f2781,f2779,f2772,f2771,f2769,f2806,f2808])).
fof(f2810,plain,(
  $false | ($spl0 | $spl18)),
  inference(global_subsumption,[],[f2557,f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f203,f227,f234,f238,f251,f535,f604,f907,f908,f910,f706,f916,f703,f705,f702,f663,f662,f603,f601,f600,f297,f915,f102,f101,f103,f306,f104,f913,f701,f700,f661,f660,f599,f598,f362,f490,f359,f385,f100,f99,f98,f911,f914,f602,f704,f1240,f1153,f380,f379,f361,f356,f355,f345,f344,f304,f293,f292,f89,f88,f81,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792,f2788,f2781,f2779,f2772,f2771,f2769,f2806,f2808])).
fof(f2809,plain,(
  $false | $spl0),
  inference(global_subsumption,[],[f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f203,f227,f234,f238,f251,f535,f604,f907,f908,f910,f706,f916,f703,f705,f702,f663,f662,f603,f601,f600,f297,f915,f102,f101,f103,f306,f104,f913,f701,f700,f661,f660,f599,f598,f362,f490,f359,f385,f100,f99,f98,f911,f914,f602,f704,f1240,f1153,f380,f379,f361,f356,f355,f345,f344,f304,f293,f292,f89,f88,f81,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792,f2788,f2781,f2779,f2772,f2771,f2769,f2806,f2808])).
fof(f2807,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f162,f164,f163,f169,f171,f203,f222,f223,f227,f234,f238,f251,f266,f286,f285,f283,f281,f280,f278,f272,f271,f269,f332,f376,f535,f604,f907,f908,f910,f706,f939,f1155,f1301,f1304,f916,f1099,f1091,f1084,f1083,f1069,f1037,f1009,f1006,f971,f917,f703,f299,f1032,f1031,f1008,f1005,f970,f705,f702,f663,f662,f650,f603,f601,f600,f339,f919,f297,f915,f102,f101,f103,f306,f104,f913,f1003,f1004,f969,f968,f912,f701,f700,f499,f388,f364,f1030,f1029,f661,f660,f651,f599,f598,f441,f362,f490,f359,f385,f360,f386,f100,f99,f98,f911,f914,f602,f704,f1007,f1305,f1295,f1270,f1269,f1258,f1240,f1232,f1153,f1139,f1138,f1129,f1128,f1127,f1126,f1077,f918,f482,f480,f387,f380,f379,f363,f361,f356,f355,f345,f344,f304,f298,f293,f292,f170,f89,f88,f81,f2216,f2217,f2218,f2219,f2220,f2259,f2274,f2278,f2295,f2529,f2531,f2523,f2521,f2533,f2401,f2402,f2403,f2406,f2407,f2408,f2409,f2540,f2507,f2503,f2545,f2459,f2460,f2461,f2462,f2463,f2464,f2465,f2556,f2558,f2491,f2492,f2493,f2494,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792,f2788,f2781,f2779,f2772,f2771,f2769,f2806])).
fof(f2805,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f162,f164,f163,f169,f171,f203,f222,f223,f227,f234,f238,f251,f266,f286,f285,f283,f281,f280,f278,f272,f271,f269,f332,f376,f535,f604,f907,f908,f910,f706,f939,f1155,f1301,f1304,f916,f1099,f1091,f1084,f1083,f1069,f1037,f1009,f1006,f971,f917,f703,f299,f1032,f1031,f1008,f1005,f970,f705,f702,f663,f662,f650,f603,f601,f600,f339,f919,f297,f915,f102,f101,f103,f306,f104,f913,f1003,f1004,f969,f968,f912,f701,f700,f499,f388,f364,f1030,f1029,f661,f660,f651,f599,f598,f441,f362,f490,f359,f385,f360,f386,f100,f99,f98,f911,f914,f602,f704,f1007,f1305,f1295,f1270,f1269,f1258,f1240,f1232,f1153,f1139,f1138,f1129,f1128,f1127,f1126,f1077,f918,f482,f480,f387,f380,f379,f363,f361,f356,f355,f345,f344,f304,f298,f293,f292,f170,f89,f88,f81,f2216,f2217,f2218,f2219,f2220,f2259,f2274,f2278,f2295,f2529,f2531,f2523,f2521,f2533,f2401,f2402,f2403,f2406,f2407,f2408,f2409,f2540,f2507,f2503,f2545,f2459,f2460,f2461,f2462,f2463,f2464,f2465,f2556,f2558,f2491,f2492,f2493,f2494,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792,f2788,f2781,f2779,f2772,f2771,f2769])).
fof(f2804,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f162,f164,f163,f169,f171,f203,f222,f223,f227,f234,f238,f251,f266,f286,f285,f283,f281,f280,f278,f272,f271,f269,f332,f376,f535,f604,f907,f908,f910,f706,f939,f1155,f1301,f1304,f916,f1099,f1091,f1084,f1083,f1069,f1037,f1009,f1006,f971,f917,f703,f299,f1032,f1031,f1008,f1005,f970,f705,f702,f663,f662,f650,f603,f601,f600,f339,f919,f297,f915,f102,f101,f103,f306,f104,f913,f1003,f1004,f969,f968,f912,f701,f700,f499,f388,f364,f1030,f1029,f661,f660,f651,f599,f598,f441,f362,f490,f359,f385,f360,f386,f100,f99,f98,f911,f914,f602,f704,f1007,f1305,f1295,f1270,f1269,f1258,f1240,f1232,f1153,f1139,f1138,f1129,f1128,f1127,f1126,f1077,f918,f482,f480,f387,f380,f379,f363,f361,f356,f355,f345,f344,f304,f298,f293,f292,f170,f89,f88,f81,f2216,f2217,f2218,f2219,f2220,f2259,f2274,f2278,f2295,f2529,f2531,f2523,f2521,f2533,f2401,f2402,f2403,f2406,f2407,f2408,f2409,f2540,f2507,f2503,f2545,f2459,f2460,f2461,f2462,f2463,f2464,f2465,f2556,f2558,f2491,f2492,f2493,f2494,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792,f2788,f2781,f2779,f2772,f2771])).
fof(f2803,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f162,f164,f163,f169,f171,f203,f222,f223,f227,f234,f238,f251,f266,f286,f285,f283,f281,f280,f278,f272,f271,f269,f332,f376,f535,f604,f907,f908,f910,f706,f939,f1155,f1301,f1304,f916,f1099,f1091,f1084,f1083,f1069,f1037,f1009,f1006,f971,f917,f703,f299,f1032,f1031,f1008,f1005,f970,f705,f702,f663,f662,f650,f603,f601,f600,f339,f919,f297,f915,f102,f101,f103,f306,f104,f913,f1003,f1004,f969,f968,f912,f701,f700,f499,f388,f364,f1030,f1029,f661,f660,f651,f599,f598,f441,f362,f490,f359,f385,f360,f386,f100,f99,f98,f911,f914,f602,f704,f1007,f1305,f1295,f1270,f1269,f1258,f1240,f1232,f1153,f1139,f1138,f1129,f1128,f1127,f1126,f1077,f918,f482,f480,f387,f380,f379,f363,f361,f356,f355,f345,f344,f304,f298,f293,f292,f170,f89,f88,f81,f2216,f2217,f2218,f2219,f2220,f2259,f2274,f2278,f2295,f2529,f2531,f2523,f2521,f2533,f2401,f2402,f2403,f2406,f2407,f2408,f2409,f2540,f2507,f2503,f2545,f2459,f2460,f2461,f2462,f2463,f2464,f2465,f2556,f2558,f2491,f2492,f2493,f2494,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792,f2788,f2781,f2779,f2772])).
fof(f2802,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f162,f164,f163,f169,f171,f203,f222,f223,f227,f234,f238,f251,f266,f286,f285,f283,f281,f280,f278,f272,f271,f269,f332,f376,f535,f604,f907,f908,f910,f706,f939,f1155,f1301,f1304,f916,f1099,f1091,f1084,f1083,f1069,f1037,f1009,f1006,f971,f917,f703,f299,f1032,f1031,f1008,f1005,f970,f705,f702,f663,f662,f650,f603,f601,f600,f339,f919,f297,f915,f102,f101,f103,f306,f104,f913,f1003,f1004,f969,f968,f912,f701,f700,f499,f388,f364,f1030,f1029,f661,f660,f651,f599,f598,f441,f362,f490,f359,f385,f360,f386,f100,f99,f98,f911,f914,f602,f704,f1007,f1305,f1295,f1270,f1269,f1258,f1240,f1232,f1153,f1139,f1138,f1129,f1128,f1127,f1126,f1077,f918,f482,f480,f387,f380,f379,f363,f361,f356,f355,f345,f344,f304,f298,f293,f292,f170,f89,f88,f81,f2216,f2217,f2218,f2219,f2220,f2259,f2274,f2278,f2295,f2529,f2531,f2523,f2521,f2533,f2401,f2402,f2403,f2406,f2407,f2408,f2409,f2540,f2507,f2503,f2545,f2459,f2460,f2461,f2462,f2463,f2464,f2465,f2556,f2558,f2491,f2492,f2493,f2494,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792,f2788,f2781,f2779])).
fof(f2801,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f162,f164,f163,f169,f171,f203,f222,f223,f227,f234,f238,f251,f266,f286,f285,f283,f281,f280,f278,f272,f271,f269,f332,f376,f535,f604,f907,f908,f910,f706,f939,f1155,f1301,f1304,f916,f1099,f1091,f1084,f1083,f1069,f1037,f1009,f1006,f971,f917,f703,f299,f1032,f1031,f1008,f1005,f970,f705,f702,f663,f662,f650,f603,f601,f600,f339,f919,f297,f915,f102,f101,f103,f306,f104,f913,f1003,f1004,f969,f968,f912,f701,f700,f499,f388,f364,f1030,f1029,f661,f660,f651,f599,f598,f441,f362,f490,f359,f385,f360,f386,f100,f99,f98,f911,f914,f602,f704,f1007,f1305,f1295,f1270,f1269,f1258,f1240,f1232,f1153,f1139,f1138,f1129,f1128,f1127,f1126,f1077,f918,f482,f480,f387,f380,f379,f363,f361,f356,f355,f345,f344,f304,f298,f293,f292,f170,f89,f88,f81,f2216,f2217,f2218,f2219,f2220,f2259,f2274,f2278,f2295,f2529,f2531,f2523,f2521,f2533,f2401,f2402,f2403,f2406,f2407,f2408,f2409,f2540,f2507,f2503,f2545,f2459,f2460,f2461,f2462,f2463,f2464,f2465,f2556,f2558,f2491,f2492,f2493,f2494,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792,f2788,f2781])).
fof(f2800,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f162,f164,f163,f169,f171,f203,f222,f223,f227,f234,f238,f251,f266,f286,f285,f283,f281,f280,f278,f272,f271,f269,f332,f376,f535,f604,f907,f908,f910,f706,f939,f1155,f1301,f1304,f916,f1099,f1091,f1084,f1083,f1069,f1037,f1009,f1006,f971,f917,f703,f299,f1032,f1031,f1008,f1005,f970,f705,f702,f663,f662,f650,f603,f601,f600,f339,f919,f297,f915,f102,f101,f103,f306,f104,f913,f1003,f1004,f969,f968,f912,f701,f700,f499,f388,f364,f1030,f1029,f661,f660,f651,f599,f598,f441,f362,f490,f359,f385,f360,f386,f100,f99,f98,f911,f914,f602,f704,f1007,f1305,f1295,f1270,f1269,f1258,f1240,f1232,f1153,f1139,f1138,f1129,f1128,f1127,f1126,f1077,f918,f482,f480,f387,f380,f379,f363,f361,f356,f355,f345,f344,f304,f298,f293,f292,f170,f89,f88,f81,f2216,f2217,f2218,f2219,f2220,f2259,f2274,f2278,f2295,f2529,f2531,f2523,f2521,f2533,f2401,f2402,f2403,f2406,f2407,f2408,f2409,f2540,f2507,f2503,f2545,f2459,f2460,f2461,f2462,f2463,f2464,f2465,f2556,f2558,f2491,f2492,f2493,f2494,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792,f2788])).
fof(f2799,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f162,f164,f163,f169,f171,f203,f222,f223,f227,f234,f238,f251,f266,f286,f285,f283,f281,f280,f278,f272,f271,f269,f332,f376,f535,f604,f907,f908,f910,f706,f939,f1155,f1301,f1304,f916,f1099,f1091,f1084,f1083,f1069,f1037,f1009,f1006,f971,f917,f703,f299,f1032,f1031,f1008,f1005,f970,f705,f702,f663,f662,f650,f603,f601,f600,f339,f919,f297,f915,f102,f101,f103,f306,f104,f913,f1003,f1004,f969,f968,f912,f701,f700,f499,f388,f364,f1030,f1029,f661,f660,f651,f599,f598,f441,f362,f490,f359,f385,f360,f386,f100,f99,f98,f911,f914,f602,f704,f1007,f1305,f1295,f1270,f1269,f1258,f1240,f1232,f1153,f1139,f1138,f1129,f1128,f1127,f1126,f1077,f918,f482,f480,f387,f380,f379,f363,f361,f356,f355,f345,f344,f304,f298,f293,f292,f170,f89,f88,f81,f2216,f2217,f2218,f2219,f2220,f2259,f2274,f2278,f2295,f2529,f2531,f2523,f2521,f2533,f2401,f2402,f2403,f2406,f2407,f2408,f2409,f2540,f2507,f2503,f2545,f2459,f2460,f2461,f2462,f2463,f2464,f2465,f2556,f2558,f2491,f2492,f2493,f2494,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796,f2792])).
fof(f2798,plain,(
  $false | ($spl0 | $spl3)),
  inference(global_subsumption,[],[f79,f82,f83,f85,f84,f108,f105,f109,f119,f118,f117,f116,f123,f122,f121,f120,f125,f132,f136,f138,f152,f162,f164,f163,f169,f171,f203,f222,f223,f227,f234,f238,f251,f266,f286,f285,f283,f281,f280,f278,f272,f271,f269,f332,f376,f535,f604,f907,f908,f910,f706,f939,f1155,f1301,f1304,f916,f1099,f1091,f1084,f1083,f1069,f1037,f1009,f1006,f971,f917,f703,f299,f1032,f1031,f1008,f1005,f970,f705,f702,f663,f662,f650,f603,f601,f600,f339,f919,f297,f915,f102,f101,f103,f306,f104,f913,f1003,f1004,f969,f968,f912,f701,f700,f499,f388,f364,f1030,f1029,f661,f660,f651,f599,f598,f441,f362,f490,f359,f385,f360,f386,f100,f99,f98,f911,f914,f602,f704,f1007,f1305,f1295,f1270,f1269,f1258,f1240,f1232,f1153,f1139,f1138,f1129,f1128,f1127,f1126,f1077,f918,f482,f480,f387,f380,f379,f363,f361,f356,f355,f345,f344,f304,f298,f293,f292,f170,f89,f88,f81,f2216,f2217,f2218,f2219,f2220,f2259,f2274,f2278,f2295,f2529,f2531,f2523,f2521,f2533,f2401,f2402,f2403,f2406,f2407,f2408,f2409,f2540,f2507,f2503,f2545,f2459,f2460,f2461,f2462,f2463,f2464,f2465,f2556,f2558,f2491,f2492,f2493,f2494,f78,f160,f2644,f2642,f2637,f2632,f2682,f2684,f2687,f2688,f2668,f2690,f2750,f2749,f2745,f2743,f2742,f2739,f2737,f2736,f2733,f2730,f2729,f2726,f2724,f2767,f2768,f2770,f2773,f2774,f2775,f2776,f2777,f2778,f2780,f2796])).
fof(f2797,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f2796,f2604])).
fof(f2604,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | $spl2),
  inference(subsumption_resolution,[],[f2582,f81])).
fof(f2582,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | $spl2),
  inference(superposition,[],[f84,f161])).
fof(f161,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | $spl2),
  inference(cnf_transformation,[],[f161_D])).
fof(f161_D,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f2793,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f2792,f2604])).
fof(f2789,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f2788,f2604])).
fof(f2786,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f2781,f2571])).
fof(f2571,plain,(
  ( ! [X0,X1] : (sP0('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X0,'F_0',X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f81,f161,f122])).
fof(f2785,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f2779,f2571])).
fof(f2784,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f2772,f2576])).
fof(f2576,plain,(
  ( ! [X0,X1] : (sP0(X0,'F_0',X1,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f81,f161,f121])).
fof(f2783,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f2771,f2576])).
fof(f2782,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f2769,f2604])).
fof(f2557,plain,(
  ( ! [X19,X17] : (p(s(bool,X19)) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X17,'type.num.num')) | s(bool,X19) = s(bool,happ(s(fun('type.num.num',bool),X17),s('type.num.num',sK1)))) ) | $spl18),
  inference(cnf_transformation,[],[f2557_D])).
fof(f2557_D,plain,(
  ( ! [X19,X17] : (p(s(bool,X19)) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X17,'type.num.num')) | s(bool,X19) = s(bool,happ(s(fun('type.num.num',bool),X17),s('type.num.num',sK1)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2558,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X22)) | s(bool,X22) = s(bool,happ(s(fun('type.num.num',bool),X20),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X20,'type.num.num')) | s(bool,'T_0') = s(bool,X21)) ) | $spl3),
  inference(subsumption_resolution,[],[f2475,f85])).
fof(f2475,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) = s(bool,happ(s(fun('type.num.num',bool),X20),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X20,'type.num.num')) | s(bool,'T_0') = s(bool,X21)) ) | $spl3),
  inference(superposition,[],[f105,f919])).
fof(f2556,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X19)) | s(bool,X19) = s(bool,happ(s(fun('type.num.num',bool),X17),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X17,'type.num.num')) | s(bool,'T_0') = s(bool,X18)) ) | $spl3),
  inference(subsumption_resolution,[],[f2474,f85])).
fof(f2474,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X19) = s(bool,happ(s(fun('type.num.num',bool),X17),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X17,'type.num.num')) | s(bool,'T_0') = s(bool,X18)) ) | $spl3),
  inference(superposition,[],[f105,f919])).
fof(f2544,plain,(
  ( ! [X59,X58] : (s(bool,happ(s(fun('type.num.num',bool),X58),s('type.num.num',sK1))) != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X59,bool)) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X58,'type.num.num')) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0')))) ) | $spl16),
  inference(cnf_transformation,[],[f2544_D])).
fof(f2544_D,plain,(
  ( ! [X59,X58] : (s(bool,happ(s(fun('type.num.num',bool),X58),s('type.num.num',sK1))) != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X59,bool)) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X58,'type.num.num')) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0')))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2503,plain,(
  ( ! [X59,X57,X58] : (s(bool,happ(s(fun('type.num.num',bool),X58),s('type.num.num',sK1))) != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X59,bool)) | s(bool,'T_0') = s(bool,X57) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0'))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X58,'type.num.num'))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2431])).
fof(f2431,plain,(
  ( ! [X59,X57,X58] : (s(bool,happ(s(fun('type.num.num',bool),X58),s('type.num.num',sK1))) != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X59,bool)) | s(bool,'T_0') = s(bool,X57) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X59),s(bool,'T_0'))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X58,'type.num.num')) | s(bool,'T_0') = s(bool,X57)) ) | $spl3),
  inference(superposition,[],[f1139,f919])).
fof(f2542,plain,(
  ( ! [X41,X40] : (~p(s(bool,happ(s(fun(bool,bool),X41),s(bool,happ(s(fun('type.num.num',bool),X40),s('type.num.num',sK1)))))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X40,'type.num.num')) | s(bool,'F_0') = s(bool,sK4('F_0',X41,bool))) ) | $spl14),
  inference(cnf_transformation,[],[f2542_D])).
fof(f2542_D,plain,(
  ( ! [X41,X40] : (~p(s(bool,happ(s(fun(bool,bool),X41),s(bool,happ(s(fun('type.num.num',bool),X40),s('type.num.num',sK1)))))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X40,'type.num.num')) | s(bool,'F_0') = s(bool,sK4('F_0',X41,bool))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2507,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,happ(s(fun(bool,bool),X41),s(bool,happ(s(fun('type.num.num',bool),X40),s('type.num.num',sK1)))))) | s(bool,'F_0') = s(bool,sK4('F_0',X41,bool)) | s(bool,'T_0') = s(bool,X39) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X40,'type.num.num'))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2425])).
fof(f2425,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,happ(s(fun(bool,bool),X41),s(bool,happ(s(fun('type.num.num',bool),X40),s('type.num.num',sK1)))))) | s(bool,'F_0') = s(bool,sK4('F_0',X41,bool)) | s(bool,'T_0') = s(bool,X39) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X40,'type.num.num')) | s(bool,'T_0') = s(bool,X39)) ) | $spl3),
  inference(superposition,[],[f293,f919])).
fof(f2541,plain,(
  ( ! [X38] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X38),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X38,'type.num.num'))) ) | $spl12),
  inference(cnf_transformation,[],[f2541_D])).
fof(f2541_D,plain,(
  ( ! [X38] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X38),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X38,'type.num.num'))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2540,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),X38),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X37) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X38,'type.num.num'))) ) | $spl3),
  inference(forward_demodulation,[],[f2508,f163])).
fof(f2508,plain,(
  ( ! [X37,X38] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s(bool,happ(s(fun('type.num.num',bool),X38),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X37) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X38,'type.num.num'))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2424])).
fof(f2424,plain,(
  ( ! [X37,X38] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) != s(bool,happ(s(fun('type.num.num',bool),X38),s('type.num.num',sK1))) | s(bool,'T_0') = s(bool,X37) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X38,'type.num.num')) | s(bool,'T_0') = s(bool,X37)) ) | $spl3),
  inference(superposition,[],[f266,f919])).
fof(f2532,plain,(
  ( ! [X21,X20] : (s(bool,happ(s(fun('type.num.num',bool),X20),s('type.num.num',sK1))) = s(bool,happ(s(fun('type.num.num',bool),X21),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X20,'type.num.num')) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X21,'type.num.num'))) ) | $spl10),
  inference(cnf_transformation,[],[f2532_D])).
fof(f2532_D,plain,(
  ( ! [X21,X20] : (s(bool,happ(s(fun('type.num.num',bool),X20),s('type.num.num',sK1))) = s(bool,happ(s(fun('type.num.num',bool),X21),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X20,'type.num.num')) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X21,'type.num.num'))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2521,plain,(
  ( ! [X21,X19,X20] : (s(bool,happ(s(fun('type.num.num',bool),X20),s('type.num.num',sK1))) = s(bool,happ(s(fun('type.num.num',bool),X21),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X21,'type.num.num')) | s(bool,'T_0') = s(bool,X19) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X20,'type.num.num'))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2384])).
fof(f2384,plain,(
  ( ! [X21,X19,X20] : (s(bool,happ(s(fun('type.num.num',bool),X20),s('type.num.num',sK1))) = s(bool,happ(s(fun('type.num.num',bool),X21),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X21,'type.num.num')) | s(bool,'T_0') = s(bool,X19) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X20,'type.num.num')) | s(bool,'T_0') = s(bool,X19)) ) | $spl3),
  inference(superposition,[],[f919,f919])).
fof(f2530,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X1,'type.num.num'))) ) | $spl8),
  inference(cnf_transformation,[],[f2530_D])).
fof(f2530_D,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X1,'type.num.num'))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f2531,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X15),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X15,'type.num.num')) | s(bool,'T_0') = s(bool,X14)) ) | $spl3),
  inference(forward_demodulation,[],[f2524,f170])).
fof(f2524,plain,(
  ( ! [X14,X15] : (s(bool,happ(s(fun('type.num.num',bool),X15),s('type.num.num',sK1))) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X15,'type.num.num')) | s(bool,'T_0') = s(bool,X14)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2382])).
fof(f2382,plain,(
  ( ! [X14,X15] : (s(bool,happ(s(fun('type.num.num',bool),X15),s('type.num.num',sK1))) = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X15,'type.num.num')) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) ) | $spl3),
  inference(superposition,[],[f919,f376])).
fof(f2529,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X1,'type.num.num')) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2377])).
fof(f2377,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),X1),s('type.num.num',sK1))) | s('type.num.num',sK1) = s('type.num.num',sK4('const.list.EL_2'(s('type.num.num',sK1),s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))),X1,'type.num.num')) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f919,f83])).
fof(f1457,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f1456,f1302])).
fof(f1456,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl4),
  inference(subsumption_resolution,[],[f1320,f84])).
fof(f1320,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl4),
  inference(backward_demodulation,[],[f1302,f88])).
fof(f1455,plain,(
  $false | $spl4),
  inference(subsumption_resolution,[],[f1319,f79])).
fof(f1319,plain,(
  ~p(s(bool,'T_0')) | $spl4),
  inference(backward_demodulation,[],[f1302,f81])).
fof(f1303,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X1,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) ) | $spl6),
  inference(cnf_transformation,[],[f1303_D])).
fof(f1303_D,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X1,bool)) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1304,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X15,bool)) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X15),s(bool,'T_0')))) ) | $spl3),
  inference(forward_demodulation,[],[f1296,f170])).
fof(f1296,plain,(
  ( ! [X14,X15] : (s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X15,bool)) != s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X15),s(bool,'T_0')))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f1276])).
fof(f1276,plain,(
  ( ! [X14,X15] : (s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X15,bool)) != s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | s(bool,'T_0') = s(bool,X14) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X15),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X14)) ) | $spl3),
  inference(superposition,[],[f1139,f376])).
fof(f1301,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X1,bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f1271])).
fof(f1271,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,sK4('const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2)),X1,bool)) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f1139,f83])).
fof(f160,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) | $spl0),
  inference(cnf_transformation,[],[f160_D])).
fof(f160_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f162,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | $spl3),
  inference(cnf_transformation,[],[f162_D])).
fof(f162_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f152,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK1),s('type.num.num',sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.num.num',sK1),s(fun('type.num.num',bool),'const.list.LIST_TO_SET_1'(s('type.list.list'('type.num.num'),'const.rich_list.COUNT_LIST_1'(s('type.num.num',sK2)))))))),
  inference(superposition,[],[f78,f83])).
% SZS output end Proof for thm.rich_list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 302
% Passive clauses: 1637
% Generated clauses: 17117
% Final active clauses: 113
% Final passive clauses: 544
% Input formulas: 20
% Initial clauses: 46
% 
% Duplicate literals: 2115
% Trivial inequalities: 104
% Fw subsumption resolutions: 774
% Fw demodulations: 836
% Bw demodulations: 378
% Condensations: 2685
% Global subsumptions: 22
% 
% Simple tautologies: 190
% Equational tautologies: 413
% Forward subsumptions: 8231
% Bw demodulations to eq. taut.: 10
% 
% Unit resulting resolution: 1569
% Factoring: 1
% Forward superposition: 4177
% Backward superposition: 3596
% Self superposition: 203
% Equality factoring: 156
% Equality resolution: 12
% 
% Split clauses: 121
% Split components: 255
% SAT solver clauses: 1969
% SAT solver unit clauses: 17
% SAT solver binary clauses: 295
% 
% Sat splits: 121
% Sat splitting refutations: 66
% 
% Memory used [KB]: 8571
% Time elapsed: 1.392 s
% ------------------------------
% ------------------------------
% Success in time 2.766 s
