Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f12406,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f12363])).
fof(f12363,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')),
  inference(superposition,[],[f68,f12270])).
fof(f12270,plain,(
  ( ! [X4,X3] : (s('type.canonical.canonical_sum'(X3),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X3),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X3),X4),s('type.canonical.canonical_sum'(X3),'const.canonical.Nil_monom_0')))) )),
  inference(superposition,[],[f11882,f72])).
fof(f72,plain,(
  ( ! [X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0] : ! [X1] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X10] : s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X6),X10),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|2','thm.canonical.canonical_sum_simplify_def|split|2')).
fof(f11882,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))) )),
  inference(forward_demodulation,[],[f1373,f76])).
fof(f76,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X20,X21] : s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X6),X20),s('type.canonical.canonical_sum'(X6),X21))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X6),X20))),s('type.canonical.canonical_sum'(X6),X21)))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|2','ARITY_1_2_const.ringNorm.r_canonical_sum_simplify')).
fof(f1373,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))) )),
  inference(superposition,[],[f75,f73])).
fof(f73,plain,(
  ( ! [X0,X1] : (s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1] : s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X0),X1))) = s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X0),X1)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7] : s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),'const.ringNorm.r_canonical_sum_simplify_1'(s('type.ring.ring'(X6),X7))) = s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X6),'const.ring.semi_ring_of_1'(s('type.ring.ring'(X6),X7)))))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|2','thm.ringNorm.r_canonical_sum_simplify_def')).
fof(f75,plain,(
  ( ! [X2,X0,X1] : (s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0] : ! [X1,X2] : s('type.canonical.canonical_sum'(X0),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X0),X1),s('type.canonical.canonical_sum'(X0),X2))) = s('type.canonical.canonical_sum'(X0),happ(s(fun('type.canonical.canonical_sum'(X0),'type.canonical.canonical_sum'(X0)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.canonical.canonical_sum'(X0),X2)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X20,X21] : s('type.canonical.canonical_sum'(X6),'const.canonical.canonical_sum_simplify_2'(s('type.semi_ring.semi_ring'(X6),X20),s('type.canonical.canonical_sum'(X6),X21))) = s('type.canonical.canonical_sum'(X6),happ(s(fun('type.canonical.canonical_sum'(X6),'type.canonical.canonical_sum'(X6)),'const.canonical.canonical_sum_simplify_1'(s('type.semi_ring.semi_ring'(X6),X20))),s('type.canonical.canonical_sum'(X6),X21)))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|2','ARITY_1_2_const.canonical.canonical_sum_simplify')).
fof(f68,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(sK0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(sK0),sK1),s('type.canonical.canonical_sum'(sK0),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f60])).
fof(f60,plain,(
  ? [X0,X1] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') != s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ~! [X0] : ! [X1] : s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X0),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X0),X1),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0') = s('type.canonical.canonical_sum'(X6),'const.ringNorm.r_canonical_sum_simplify_2'(s('type.ring.ring'(X6),X7),s('type.canonical.canonical_sum'(X6),'const.canonical.Nil_monom_0')))),
  file('thm.ringNorm.canonical_sum_simplify_def|split|2','thm.ringNorm.canonical_sum_simplify_def|split|2')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 121
% Passive clauses: 5426
% Generated clauses: 12201
% Final active clauses: 116
% Final passive clauses: 5217
% Input formulas: 21
% Initial clauses: 26
% 
% Duplicate literals: 284
% Trivial inequalities: 139
% Fw subsumption resolutions: 860
% Fw demodulations: 352
% Bw demodulations: 2
% 
% Simple tautologies: 220
% Equational tautologies: 23
% Forward subsumptions: 4597
% 
% Binary resolution: 66
% Forward superposition: 5711
% Backward superposition: 4461
% Self superposition: 156
% Equality factoring: 62
% Equality resolution: 22
% 
% Split clauses: 239
% Split components: 486
% SAT solver clauses: 308
% SAT solver unit clauses: 6
% SAT solver binary clauses: 28
% 
% Sat splits: 239
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 307
% TWLsolver calls for satisfiability: 57
% 
% Memory used [KB]: 13688
% Time elapsed: 0.500 s
% ------------------------------
% ------------------------------
% Success in time 0.51 s
