Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.complex
% (92702)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 83
% Passive clauses: 2401
% Generated clauses: 14770
% Final active clauses: 75
% Final passive clauses: 2246
% Input formulas: 11
% Initial clauses: 17
% 
% Duplicate literals: 969
% Trivial inequalities: 205
% Fw subsumption resolutions: 1755
% Fw demodulations: 1007
% Bw demodulations: 5
% 
% Simple tautologies: 327
% Equational tautologies: 93
% Forward subsumptions: 6064
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 24
% Forward superposition: 6694
% Backward superposition: 3421
% Self superposition: 298
% Equality factoring: 233
% Equality resolution: 16
% 
% Split clauses: 734
% Split components: 1490
% SAT solver clauses: 1837
% SAT solver unit clauses: 259
% SAT solver binary clauses: 533
% 
% Sat splits: 734
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1833
% TWLsolver calls for satisfiability: 83
% 
% Memory used [KB]: 5756
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.complex
% (92719)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 69
% Passive clauses: 293
% Generated clauses: 8088
% Final active clauses: 68
% Final passive clauses: 201
% Input formulas: 11
% Initial clauses: 17
% 
% Duplicate literals: 1019
% Trivial inequalities: 123
% Fw subsumption resolutions: 555
% Fw demodulations: 447
% Bw demodulations: 20
% Condensations: 1503
% 
% Simple tautologies: 163
% Equational tautologies: 50
% Forward subsumptions: 3195
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 11
% Forward superposition: 2663
% Backward superposition: 1365
% Self superposition: 172
% Equality factoring: 134
% Equality resolution: 14
% 
% Split clauses: 101
% Split components: 221
% SAT solver clauses: 815
% SAT solver unit clauses: 17
% SAT solver binary clauses: 69
% 
% Sat splits: 101
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 815
% TWLsolver calls for satisfiability: 117
% 
% Memory used [KB]: 5245
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.complex
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.complex
% SZS output start Proof for thm.complex
fof(f15685,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f37,f58,f56,f54,f36,f57,f35,f55,f53,f64,f67,f1545,f1548,f1547,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1556,f1554,f1484,f1485,f1486,f1487,f1550,f1557,f1488,f1555,f1558,f1489,f1559,f1551,f2337,f2293,f2275,f2268,f2267,f2250,f2219,f2364,f2218,f2163,f2365,f2158,f2370,f2157,f2371,f2152,f2139,f2378,f2973,f2980,f3055,f3056,f2976,f3086,f5165,f5164,f5198,f5089,f5070,f4924,f4922,f5317,f3965,f3964,f3963,f3962,f5320,f3961,f5321,f3960,f5322,f3959,f3958,f5323,f3957,f3956,f3955,f5325,f3954,f5324,f3953,f5319,f5326,f5327,f5333,f5335,f3193,f5339,f3190,f5337,f5340,f5341,f5342,f2791,f2794,f2795,f5336,f2798,f2799,f2800,f2801,f5318,f5384,f5385,f5386,f5387,f5388,f5389,f5390,f5391,f5392,f5393,f5394,f5395,f5407,f5409,f5411,f5413,f5415,f5417,f5419,f5421,f5423,f5425,f5427,f5429,f5431,f5435,f3570,f3567,f5437,f5440,f5485,f5487,f5489,f3528,f5492,f5494,f3531,f3534,f5495,f3537,f5491,f5432,f5497,f5496,f2956,f2955,f2954,f2953,f5404,f2952,f2951,f2950,f2949,f5401,f2948,f2947,f2946,f5400,f2945,f5399,f2975,f5499,f4217,f4220,f4222,f4224,f4226,f4228,f5501,f3538,f3535,f5493,f5502,f3533,f3532,f5503,f5436,f3192,f3189,f5504,f5505,f5506,f5500,f5507,f5508,f5433,f5509,f2978,f5510,f6336,f6339,f5338,f6358,f6366,f6368,f6599,f6600,f6601,f6602,f6603,f6604,f6605,f6606,f6607,f6608,f6609,f6610,f6611,f6612,f5328,f6613,f6614,f6615,f6616,f6618,f6619,f6640,f6620,f6621,f6622,f6623,f5398,f5402,f6624,f6639,f5403,f6625,f6641,f1553,f6629,f5329,f2974,f6909,f6911,f6929,f6938,f7578,f7587,f7591,f7590,f7597,f7622,f7623,f7678,f7680,f8284,f8343,f8345,f8347,f8348,f8400,f8402,f9136,f9215,f9216,f9220,f9224,f9278,f9280,f11658,f11571,f11570,f11665,f5197,f11548,f11507,f5196,f11347,f5268,f11343,f5270,f11312,f11770,f11769,f11826,f11828,f11829,f11847,f11848,f11850,f11852,f11855,f11860,f11862,f12914,f12916,f12919,f12922,f12925,f12983,f14204,f65,f14205,f14206,f2363,f14663,f14737,f15148,f15147,f15680,f15682,f15684])).
fof(f15684,plain,(
  $false | ($spl0 | $spl3)),
  inference(trivial_inequality_removal,[],[f15683])).
fof(f15683,plain,(
  s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) != s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f56,f53])).
fof(f15682,plain,(
  $false | ($spl0 | $spl5)),
  inference(trivial_inequality_removal,[],[f15681])).
fof(f15681,plain,(
  s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) != s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) | ($spl0 | $spl5)),
  inference(backward_demodulation,[],[f53,f58])).
fof(f15680,plain,(
  $false | ($spl1 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f15679,f55])).
fof(f15679,plain,(
  s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) != s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f15677,f41])).
fof(f41,plain,(
  ( ! [X0] : (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.pair.FST_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))) )),
  inference(cnf_transformation,[],[f15])).
fof(f15,plain,(
  ! [X0] : s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.pair.FST_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.realax.real','const.pair.FST_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6)))),
  file('thm.complex.COMPLEX_RE_IM_EQ','thm.complex.RE')).
fof(f15677,plain,(
  s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) != s('type.realax.real','const.pair.FST_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) | ($spl1 | $spl4)),
  inference(unit_resulting_resolution,[],[f54,f57,f1740])).
fof(f1740,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.pair.FST_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),X1)) )),
  inference(superposition,[],[f294,f42])).
fof(f42,plain,(
  ( ! [X0] : (s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))) )),
  inference(cnf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0] : s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6)))),
  file('thm.complex.COMPLEX_RE_IM_EQ','thm.complex.IM')).
fof(f294,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) != s('type.realax.real','const.pair.FST_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),X1)) )),
  inference(forward_demodulation,[],[f281,f41])).
fof(f281,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),X1) | s('type.realax.real','const.pair.FST_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.pair.FST_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) )),
  inference(superposition,[],[f51,f42])).
fof(f51,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) != s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))) | s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),X3) | s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) != s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3)))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2,X3] : ((s('type.pair.prod'(X0,X1),X2) != s('type.pair.prod'(X0,X1),X3) | (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))) & s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))))) & (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) != s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))) | s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) != s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))) | s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),X3)))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2,X3] : ((s('type.pair.prod'(X0,X1),X2) != s('type.pair.prod'(X0,X1),X3) | (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))) & s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))))) & ((s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) != s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))) | s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) != s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))) | s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),X3)))),
  inference(nnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2,X3] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),X3) <=> (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))) & s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ! [X0,X1] : ! [X2,X3] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),X3) <=> (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))) & s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8,X11] : ! [X12,X13] : (s('type.pair.prod'(X8,X11),X12) = s('type.pair.prod'(X8,X11),X13) <=> (s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X11),X12))) = s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X11),X13))) & s(X11,'const.pair.SND_1'(s('type.pair.prod'(X8,X11),X12))) = s(X11,'const.pair.SND_1'(s('type.pair.prod'(X8,X11),X13)))))),
  file('thm.complex.COMPLEX_RE_IM_EQ','thm.pair.PAIR_FST_SND_EQ')).
fof(f15147,plain,(
  ( ! [X167,X168] : (s(fun(bool,bool),X167) = s(fun(bool,bool),X168) | s(bool,'F_0') = s(bool,sK2(X167,X168,bool,bool))) ) | $spl142),
  inference(cnf_transformation,[],[f15147_D])).
fof(f15147_D,plain,(
  ( ! [X167,X168] : (s(fun(bool,bool),X167) = s(fun(bool,bool),X168) | s(bool,'F_0') = s(bool,sK2(X167,X168,bool,bool))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f15148,plain,(
  ( ! [X169] : (s(bool,'F_0') != s(bool,f(s(bool,X169))) | s(bool,'F_0') = s(bool,X169)) ) | $spl144),
  inference(cnf_transformation,[],[f15148_D])).
fof(f15148_D,plain,(
  ( ! [X169] : (s(bool,'F_0') != s(bool,f(s(bool,X169))) | s(bool,'F_0') = s(bool,X169)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f14737,plain,(
  ( ! [X167,X169,X168] : (s(bool,'F_0') != s(bool,f(s(bool,X169))) | s(fun(bool,bool),X167) = s(fun(bool,bool),X168) | s(bool,'F_0') = s(bool,sK2(X167,X168,bool,bool)) | s(bool,'F_0') = s(bool,X169)) ) | ($spl12 | $spl56)),
  inference(superposition,[],[f1548,f14083])).
fof(f14083,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl56),
  inference(superposition,[],[f5197,f5197])).
fof(f14663,plain,(
  ( ! [X167,X169,X168] : (s(bool,'F_0') != s(bool,f(s(bool,X169))) | s(fun(bool,bool),X167) = s(fun(bool,bool),X168) | s(bool,'F_0') = s(bool,X169) | s(bool,'F_0') = s(bool,sK2(X167,X168,bool,bool))) ) | ($spl12 | $spl56)),
  inference(superposition,[],[f1548,f14083])).
fof(f2363,plain,(
  ( ! [X64] : (s(bool,'T_0') = s(bool,X64)) ) | $spl32),
  inference(cnf_transformation,[],[f2363_D])).
fof(f2363_D,plain,(
  ( ! [X64] : (s(bool,'T_0') = s(bool,X64)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f14206,plain,(
  ( ! [X46] : (p(s(bool,X46)) | s(bool,'F_0') = s(bool,X46)) ) | $spl140),
  inference(cnf_transformation,[],[f14206_D])).
fof(f14206_D,plain,(
  ( ! [X46] : (p(s(bool,X46)) | s(bool,'F_0') = s(bool,X46)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f14205,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) ) | ($spl56 | $spl138)),
  inference(subsumption_resolution,[],[f14161,f11770])).
fof(f14161,plain,(
  ( ! [X47,X46] : (p(s(bool,X46)) | s(bool,X47) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) ) | $spl56),
  inference(superposition,[],[f129,f5197])).
fof(f129,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f38,f59])).
fof(f59,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f43,f43])).
fof(f43,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.complex.COMPLEX_RE_IM_EQ',aHL_BOOL_CASES)).
fof(f38,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.complex.COMPLEX_RE_IM_EQ',aHL_TRUTH)).
fof(f65,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl6),
  inference(cnf_transformation,[],[f65_D])).
fof(f65_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f14204,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl9 | $spl34 | $spl51 | $spl56 | $spl81 | $spl88 | $spl113 | $spl138)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f294,f282,f284,f1360,f1218,f377,f482,f484,f518,f520,f844,f1419,f536,f538,f541,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1740,f2928,f2967,f3033,f3034,f3036,f3055,f3056,f3088,f3097,f3099,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5406,f5408,f5410,f5412,f5414,f5416,f5418,f5420,f5422,f5424,f5426,f5428,f5431,f5435,f3570,f3567,f5437,f5440,f3527,f3528,f3529,f3530,f3531,f3534,f3536,f3537,f3539,f2956,f2955,f2954,f2953,f2952,f2951,f2950,f2949,f2948,f2947,f2946,f2945,f3538,f3535,f3533,f3532,f3192,f3189,f5499,f2978,f5763,f5818,f6014,f6016,f6047,f6053,f6166,f6168,f6251,f6323,f2364,f6369,f6454,f6455,f6457,f6487,f6485,f8403,f6514,f6512,f6442,f9907,f9909,f5329,f11832,f741,f5739,f11868,f11893,f12199,f12272,f12446,f12448,f557,f535,f517,f481,f260,f258,f5767,f5765,f5398,f2406,f2405,f2404,f2402,f2399,f2398,f2397,f2393,f2392,f2391,f2389,f2386,f2385,f2384,f1540,f1512,f1220,f286,f272,f262,f257,f40,f11770,f5197,f13506,f13519,f13468,f13520,f14069,f14183,f14083,f14144])).
fof(f14144,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | ($spl56 | $spl138)),
  inference(superposition,[],[f11770,f5197])).
fof(f14183,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X28) = s(bool,g(s(bool,X27))) | s(bool,'F_0') = s(bool,X28) | s(bool,X27) = s(bool,X29) | s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | ($spl56 | $spl81 | $spl88)),
  inference(subsumption_resolution,[],[f14079,f11832])).
fof(f14079,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X28) = s(bool,g(s(bool,X27))) | s(bool,'F_0') = s(bool,X28) | s(bool,X27) = s(bool,X29) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | $spl56),
  inference(superposition,[],[f5197,f97])).
fof(f14069,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,g(s(bool,X0))) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl56),
  inference(superposition,[],[f5197,f43])).
fof(f13520,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X28) != s(bool,g(s(bool,X27))) | s(bool,'T_0') = s(bool,X28) | s(bool,X27) = s(bool,X29) | s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | ($spl81 | $spl88 | $spl138)),
  inference(subsumption_resolution,[],[f13478,f11832])).
fof(f13478,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,X28) != s(bool,g(s(bool,X27))) | s(bool,'T_0') = s(bool,X28) | s(bool,X27) = s(bool,X29) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X29) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) ) | $spl138),
  inference(superposition,[],[f11770,f97])).
fof(f13468,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | $spl138),
  inference(superposition,[],[f11770,f43])).
fof(f13519,plain,(
  ( ! [X45] : (s(bool,X45) != s(bool,g(s(bool,X45))) | s(bool,'T_0') = s(bool,X45)) ) | ($spl81 | $spl88 | $spl138)),
  inference(subsumption_resolution,[],[f13462,f11832])).
fof(f13462,plain,(
  ( ! [X45] : (s(bool,X45) != s(bool,g(s(bool,X45))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X45)) ) | $spl138),
  inference(superposition,[],[f11770,f43])).
fof(f13506,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,g(s(bool,'F_0'))) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) ) | $spl138),
  inference(condensation,[],[f13505])).
fof(f13505,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,X26) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,X27) | s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl138),
  inference(duplicate_literal_removal,[],[f13458])).
fof(f13458,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,X26) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X25) | s(bool,X27) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) ) | $spl138),
  inference(superposition,[],[f11770,f97])).
fof(f40,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f14])).
fof(f14,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.complex.COMPLEX_RE_IM_EQ',aHL_FALSITY)).
fof(f257,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,sK2(X16,X17,X18,bool)) | s(fun(bool,X18),X16) = s(fun(bool,X18),X17) | s(X18,f(s(bool,'F_0'))) != s(X18,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f48,f43])).
fof(f262,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,g(s(X16,sK2(X17,X18,bool,X16)))) | s(fun(X16,bool),X17) = s(fun(X16,bool),X18) | s(bool,'F_0') != s(bool,f(s(X16,sK2(X17,X18,bool,X16))))) )),
  inference(superposition,[],[f48,f43])).
fof(f272,plain,(
  ( ! [X26,X24,X25] : (s(X26,f(s(bool,'F_0'))) != s(X26,g(s(bool,'F_0'))) | s(fun(bool,X26),X24) = s(fun(bool,X26),X25) | s(X26,f(s(bool,'T_0'))) != s(X26,g(s(bool,'T_0')))) )),
  inference(duplicate_literal_removal,[],[f271])).
fof(f271,plain,(
  ( ! [X26,X24,X25] : (s(X26,f(s(bool,'T_0'))) != s(X26,g(s(bool,'T_0'))) | s(fun(bool,X26),X24) = s(fun(bool,X26),X25) | s(fun(bool,X26),X24) = s(fun(bool,X26),X25) | s(X26,f(s(bool,'F_0'))) != s(X26,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f48,f257])).
fof(f286,plain,(
  ( ! [X19,X20,X18] : (s(X18,'const.pair.FST_1'(s('type.pair.prod'(X18,bool),X19))) != s(X18,'const.pair.FST_1'(s('type.pair.prod'(X18,bool),X20))) | s('type.pair.prod'(X18,bool),X19) = s('type.pair.prod'(X18,bool),X20) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X18,bool),X20))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X18,bool),X19)))) )),
  inference(superposition,[],[f51,f43])).
fof(f1220,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,sK2(X2,X3,X4,bool)) | s(bool,X0) = s(bool,X1) | s(fun(bool,X4),X2) = s(fun(bool,X4),X3) | s(X4,f(s(bool,'F_0'))) != s(X4,g(s(bool,'F_0')))) )),
  inference(condensation,[],[f1219])).
fof(f1219,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,sK2(X108,X109,X110,bool)) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113) | s(fun(bool,X110),X108) = s(fun(bool,X110),X109) | s(X110,f(s(bool,'F_0'))) != s(X110,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f357])).
fof(f357,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,sK2(X108,X109,X110,bool)) | s(bool,'T_0') = s(bool,X111) | s(bool,X112) = s(bool,X113) | s(bool,'T_0') = s(bool,X113) | s(fun(bool,X110),X108) = s(fun(bool,X110),X109) | s(X110,f(s(bool,'F_0'))) != s(X110,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f97,f257])).
fof(f1512,plain,(
  ( ! [X72,X71,X73] : (s(bool,'F_0') != s(bool,f(s(X71,sK2(X72,X73,bool,X71)))) | s(fun(X71,bool),X72) = s(fun(X71,bool),X73) | s(bool,'T_0') != s(bool,f(s(X71,sK2(X72,X73,bool,X71))))) )),
  inference(duplicate_literal_removal,[],[f1511])).
fof(f1511,plain,(
  ( ! [X72,X71,X73] : (s(bool,'T_0') != s(bool,f(s(X71,sK2(X72,X73,bool,X71)))) | s(fun(X71,bool),X72) = s(fun(X71,bool),X73) | s(fun(X71,bool),X72) = s(fun(X71,bool),X73) | s(bool,'F_0') != s(bool,f(s(X71,sK2(X72,X73,bool,X71))))) )),
  inference(superposition,[],[f48,f262])).
fof(f1540,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(fun(X1,bool),X2) = s(fun(X1,bool),X3) | s(bool,'F_0') != s(bool,f(s(X1,sK2(X2,X3,bool,X1)))) | s(bool,X4) = s(bool,g(s(X1,sK2(X2,X3,bool,X1)))) | s(bool,X0) = s(bool,X4)) )),
  inference(condensation,[],[f1539])).
fof(f1539,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s(bool,'T_0') = s(bool,X38) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') != s(bool,f(s(X35,sK2(X36,X37,bool,X35)))) | s(bool,X39) = s(bool,g(s(X35,sK2(X36,X37,bool,X35)))) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(duplicate_literal_removal,[],[f1495])).
fof(f1495,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s(bool,'T_0') = s(bool,X38) | s(fun(X35,bool),X36) = s(fun(X35,bool),X37) | s(bool,'F_0') != s(bool,f(s(X35,sK2(X36,X37,bool,X35)))) | s(bool,X39) = s(bool,g(s(X35,sK2(X36,X37,bool,X35)))) | s(bool,'T_0') = s(bool,X38) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f262,f97])).
fof(f2384,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X1))) | s('type.pair.prod'(bool,bool),X0) = s('type.pair.prod'(bool,bool),X1) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X1))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X0)))) )),
  inference(superposition,[],[f286,f43])).
fof(f2385,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X3))) | s('type.pair.prod'(bool,bool),X2) = s('type.pair.prod'(bool,bool),X3) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X3))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X2))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f286,f59])).
fof(f2386,plain,(
  ( ! [X6,X7,X5] : (s(bool,X6) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X7))) | s('type.pair.prod'(bool,bool),X5) = s('type.pair.prod'(bool,bool),X7) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X7))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X5))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X5)))) )),
  inference(superposition,[],[f286,f59])).
fof(f2389,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,X15) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X16))) | s('type.pair.prod'(bool,bool),X14) = s('type.pair.prod'(bool,bool),X16) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X16))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X14))) | s(bool,X15) = s(bool,X17) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X14))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f286,f97])).
fof(f2391,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X25))) | s('type.pair.prod'(bool,bool),X24) = s('type.pair.prod'(bool,bool),X25) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X25))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X24))) | s(bool,X26) = s(bool,X27) | s(bool,X27) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X24))) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f286,f97])).
fof(f2392,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X31))) | s('type.pair.prod'(bool,bool),X29) = s('type.pair.prod'(bool,bool),X31) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X31))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X29))) | s(bool,X32) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f286,f97])).
fof(f2393,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X36))) | s('type.pair.prod'(bool,bool),X34) = s('type.pair.prod'(bool,bool),X36) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X36))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X34))) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X34))) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f286,f97])).
fof(f2397,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X1))) | s('type.pair.prod'(bool,bool),X0) = s('type.pair.prod'(bool,bool),X1) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X0))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X1))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X0)))) )),
  inference(superposition,[],[f286,f43])).
fof(f2398,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X3))) | s('type.pair.prod'(bool,bool),X2) = s('type.pair.prod'(bool,bool),X3) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X2))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X3))) | s(bool,X4) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f286,f59])).
fof(f2399,plain,(
  ( ! [X6,X7,X5] : (s(bool,X6) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X7))) | s('type.pair.prod'(bool,bool),X5) = s('type.pair.prod'(bool,bool),X7) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X5))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X7))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X5)))) )),
  inference(superposition,[],[f286,f59])).
fof(f2402,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,X15) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X16))) | s('type.pair.prod'(bool,bool),X14) = s('type.pair.prod'(bool,bool),X16) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X14))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X16))) | s(bool,X15) = s(bool,X17) | s(bool,'T_0') = s(bool,X17) | s(bool,X18) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X14))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f286,f97])).
fof(f2404,plain,(
  ( ! [X28,X26,X24,X27,X25] : (s(bool,'T_0') != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X25))) | s('type.pair.prod'(bool,bool),X24) = s('type.pair.prod'(bool,bool),X25) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X24))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X25))) | s(bool,X26) = s(bool,X27) | s(bool,X27) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X24))) | s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f286,f97])).
fof(f2405,plain,(
  ( ! [X30,X33,X31,X29,X32] : (s(bool,X30) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X31))) | s('type.pair.prod'(bool,bool),X29) = s('type.pair.prod'(bool,bool),X31) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X29))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X31))) | s(bool,X32) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f286,f97])).
fof(f2406,plain,(
  ( ! [X37,X35,X38,X36,X34] : (s(bool,X35) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X36))) | s('type.pair.prod'(bool,bool),X34) = s('type.pair.prod'(bool,bool),X36) | s(bool,'F_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X34))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(bool,bool),X36))) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,bool),X34))) | s(bool,X37) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f286,f97])).
fof(f5765,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,f(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) = s(bool,X1)) ) | $spl51),
  inference(superposition,[],[f2978,f97])).
fof(f5767,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X8))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl51),
  inference(superposition,[],[f2978,f59])).
fof(f258,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,f(s(X0,sK2(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,X3) = s(bool,g(s(X0,sK2(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f48,f59])).
fof(f260,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X11) != s(bool,f(s(X8,sK2(X9,X10,bool,X8)))) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK2(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f48,f59])).
fof(f481,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,X70) != s(bool,f(s(X67,sK2(X68,X69,bool,X67)))) | s(fun(X67,bool),X68) = s(fun(X67,bool),X69) | s(bool,X70) = s(bool,X71) | s(bool,'T_0') = s(bool,X71) | s(bool,X72) = s(bool,g(s(X67,sK2(X68,X69,bool,X67)))) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f48,f97])).
fof(f517,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,X70) != s(bool,f(s(X67,sK2(X68,X69,bool,X67)))) | s(fun(X67,bool),X68) = s(fun(X67,bool),X69) | s(bool,X71) = s(bool,g(s(X67,sK2(X68,X69,bool,X67)))) | s(bool,'T_0') = s(bool,X70) | s(bool,X71) = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f48,f97])).
fof(f535,plain,(
  ( ! [X68,X66,X64,X69,X67,X65] : (s(bool,X67) != s(bool,f(s(X64,sK2(X65,X66,bool,X64)))) | s(fun(X64,bool),X65) = s(fun(X64,bool),X66) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,g(s(X64,sK2(X65,X66,bool,X64)))) | s(bool,X68) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f48,f97])).
fof(f557,plain,(
  ( ! [X59,X57,X54,X58,X56,X55] : (s(bool,'T_0') != s(bool,f(s(X54,sK2(X55,X56,bool,X54)))) | s(fun(X54,bool),X55) = s(fun(X54,bool),X56) | s(bool,X57) = s(bool,X58) | s(bool,X58) = s(bool,g(s(X54,sK2(X55,X56,bool,X54)))) | s(bool,X57) = s(bool,X59) | s(bool,'T_0') = s(bool,X59)) )),
  inference(superposition,[],[f48,f97])).
fof(f12448,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,f(s(bool,X8))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl51),
  inference(superposition,[],[f5739,f59])).
fof(f12446,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) = s(bool,X1)) ) | $spl51),
  inference(superposition,[],[f5739,f97])).
fof(f12272,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,f(s(bool,X17))) | s(bool,X17) = s(bool,X18) | s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | ($spl81 | $spl88)),
  inference(subsumption_resolution,[],[f12209,f11832])).
fof(f12209,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,f(s(bool,X17))) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | $spl88),
  inference(superposition,[],[f5398,f97])).
fof(f12199,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl88),
  inference(superposition,[],[f5398,f43])).
fof(f11893,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (s(X66,f(s(bool,X65))) != s(X66,g(s(bool,X65))) | s(fun(bool,X66),X67) = s(fun(bool,X66),X68) | s(X66,f(s(bool,'T_0'))) != s(X66,g(s(bool,'T_0'))) | s(bool,X65) = s(bool,X69) | s(bool,X69) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) ) | ($spl81 | $spl88)),
  inference(subsumption_resolution,[],[f10554,f11832])).
fof(f10554,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (s(X66,f(s(bool,X65))) != s(X66,g(s(bool,X65))) | s(fun(bool,X66),X67) = s(fun(bool,X66),X68) | s(X66,f(s(bool,'T_0'))) != s(X66,g(s(bool,'T_0'))) | s(bool,X65) = s(bool,X69) | s(bool,X69) = s(bool,X70) | s(bool,'T_0') = s(bool,X70) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f272,f97])).
fof(f11868,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X18) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,X18) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X22)) ) | $spl51),
  inference(forward_demodulation,[],[f10562,f5739])).
fof(f10562,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,X18) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f272,f97])).
fof(f5739,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl51),
  inference(unit_resulting_resolution,[],[f40,f2978,f129])).
fof(f741,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X0) = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f622])).
fof(f622,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X0) = s(bool,X1) | s(bool,X0) = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(equality_factoring,[],[f97])).
fof(f11832,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl81 | $spl88)),
  inference(backward_demodulation,[],[f5398,f5329])).
fof(f9909,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) = s(bool,g(s(bool,X10))) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | ($spl9 | $spl51 | $spl113)),
  inference(superposition,[],[f6487,f59])).
fof(f9907,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X0) = s(bool,X2)) ) | ($spl9 | $spl51 | $spl113)),
  inference(superposition,[],[f6487,f97])).
fof(f6442,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl9 | $spl51 | $spl113)),
  inference(unit_resulting_resolution,[],[f67,f6369,f6369,f6369,f97])).
fof(f6512,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,g(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) = s(bool,X1)) ) | ($spl9 | $spl51 | $spl113)),
  inference(superposition,[],[f6442,f97])).
fof(f6514,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,g(s(bool,X8))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl9 | $spl51 | $spl113)),
  inference(superposition,[],[f6442,f59])).
fof(f8403,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool)))) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl51 | $spl113)),
  inference(subsumption_resolution,[],[f8279,f6369])).
fof(f8279,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool)))) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f8205])).
fof(f8205,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f262,f257])).
fof(f6485,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) = s(bool,X11) | s(bool,X11) = s(bool,g(s(bool,'T_0'))) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | ($spl9 | $spl51 | $spl113)),
  inference(trivial_inequality_removal,[],[f6478])).
fof(f6478,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X10) = s(bool,X11) | s(bool,X11) = s(bool,g(s(bool,'T_0'))) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | ($spl9 | $spl51 | $spl113)),
  inference(superposition,[],[f6369,f97])).
fof(f6487,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl51 | $spl113)),
  inference(trivial_inequality_removal,[],[f6472])).
fof(f6472,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl51 | $spl113)),
  inference(superposition,[],[f6369,f59])).
fof(f6457,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,g(s(bool,X8))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl9 | $spl51 | $spl113)),
  inference(superposition,[],[f6369,f59])).
fof(f6455,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,g(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) = s(bool,X1)) ) | ($spl9 | $spl51 | $spl113)),
  inference(superposition,[],[f6369,f97])).
fof(f6454,plain,(
  ~p(s(bool,g(s(bool,'T_0')))) | ($spl9 | $spl51 | $spl113)),
  inference(unit_resulting_resolution,[],[f6369,f6369,f6369,f844])).
fof(f6369,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | ($spl9 | $spl51 | $spl113)),
  inference(forward_demodulation,[],[f5499,f5763])).
fof(f6323,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,f(s(bool,X17))) | s(bool,X17) = s(bool,X18) | s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | ($spl9 | $spl81)),
  inference(subsumption_resolution,[],[f6261,f67])).
fof(f6261,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,f(s(bool,X17))) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | ($spl9 | $spl81)),
  inference(superposition,[],[f6014,f97])).
fof(f6251,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl81)),
  inference(superposition,[],[f6014,f43])).
fof(f6168,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,f(s(bool,X8))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl9 | $spl51)),
  inference(superposition,[],[f5763,f59])).
fof(f6166,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) = s(bool,X1)) ) | ($spl9 | $spl51)),
  inference(superposition,[],[f5763,f97])).
fof(f6053,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,'F_0'))) | s(bool,X0) = s(bool,X1)) ) | $spl81),
  inference(condensation,[],[f6052])).
fof(f6052,plain,(
  ( ! [X14,X15,X13] : (s(bool,X14) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl81),
  inference(subsumption_resolution,[],[f6038,f43])).
fof(f6038,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X14) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl81),
  inference(superposition,[],[f5329,f97])).
fof(f6047,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,f(s(bool,X17))) | s(bool,X17) = s(bool,X18) | s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | ($spl9 | $spl81)),
  inference(subsumption_resolution,[],[f6026,f67])).
fof(f6026,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,f(s(bool,X17))) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) ) | $spl81),
  inference(superposition,[],[f5329,f97])).
fof(f6016,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl81),
  inference(superposition,[],[f5329,f43])).
fof(f6014,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | ($spl9 | $spl81)),
  inference(unit_resulting_resolution,[],[f38,f67,f5329,f844])).
fof(f5818,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,X1)) ) | $spl51),
  inference(condensation,[],[f5817])).
fof(f5817,plain,(
  ( ! [X14,X15,X13] : (s(bool,X14) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl51),
  inference(subsumption_resolution,[],[f5789,f43])).
fof(f5789,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X14) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl51),
  inference(superposition,[],[f2978,f97])).
fof(f5763,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl9 | $spl51)),
  inference(unit_resulting_resolution,[],[f38,f67,f2978,f844])).
fof(f3539,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(bool,X71) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X72) = s(fun(bool,bool),X73) | s(bool,X74) = s(bool,sK2(X72,X73,bool,bool)) | s(bool,'T_0') = s(bool,X74) | s(bool,X71) = s(bool,X75) | s(bool,X75) = s(bool,X76) | s(bool,'T_0') = s(bool,X76) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) )),
  inference(superposition,[],[f253,f97])).
fof(f3536,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(bool,X53) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X54) = s(fun(bool,bool),X55) | s(bool,X56) = s(bool,sK2(X54,X55,bool,bool)) | s(bool,'T_0') = s(bool,X56) | s(bool,X53) = s(bool,X57) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X57) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f253,f97])).
fof(f3530,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,X21) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X22) = s(fun(bool,bool),X23) | s(bool,X24) = s(bool,sK2(X22,X23,bool,bool)) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f253,f59])).
fof(f3529,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,X17) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19) | s(bool,X20) = s(bool,sK2(X18,X19,bool,bool)) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) )),
  inference(superposition,[],[f253,f59])).
fof(f3527,plain,(
  ( ! [X12,X10,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,X12) = s(bool,sK2(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,g(s(bool,'T_0')))) )),
  inference(superposition,[],[f253,f43])).
fof(f5428,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X53) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X54) = s(fun(bool,bool),X55) | s(bool,X53) = s(bool,X56) | s(bool,X56) = s(bool,X57) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl51)),
  inference(forward_demodulation,[],[f2956,f3033])).
fof(f5426,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X48) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X49) = s(fun(bool,bool),X50) | s(bool,X51) = s(bool,g(s(bool,'F_0'))) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52) | s(bool,'T_0') = s(bool,X48)) ) | ($spl9 | $spl51)),
  inference(forward_demodulation,[],[f2955,f3033])).
fof(f5424,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X43) = s(fun(bool,bool),X44) | s(bool,X45) = s(bool,X46) | s(bool,X46) = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,X45) = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl51)),
  inference(forward_demodulation,[],[f2954,f3033])).
fof(f5422,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X38) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X39) = s(fun(bool,bool),X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) ) | ($spl9 | $spl51)),
  inference(forward_demodulation,[],[f2953,f3033])).
fof(f5420,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X33) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X34) = s(fun(bool,bool),X35) | s(bool,X36) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X33) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | ($spl9 | $spl51)),
  inference(forward_demodulation,[],[f2952,f3033])).
fof(f5418,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X28) = s(fun(bool,bool),X29) | s(bool,X30) = s(bool,X31) | s(bool,X31) = s(bool,g(s(bool,'F_0'))) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) ) | ($spl9 | $spl51)),
  inference(forward_demodulation,[],[f2951,f3033])).
fof(f5416,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X23) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X24) = s(fun(bool,bool),X25) | s(bool,X26) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X26) | s(bool,X23) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) ) | ($spl9 | $spl51)),
  inference(forward_demodulation,[],[f2950,f3033])).
fof(f5414,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X18) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,X18) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X22)) ) | ($spl9 | $spl51)),
  inference(forward_demodulation,[],[f2949,f3033])).
fof(f5412,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X15) = s(fun(bool,bool),X16) | s(bool,X17) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl9 | $spl51)),
  inference(forward_demodulation,[],[f2948,f3033])).
fof(f5410,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X12) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X13) = s(fun(bool,bool),X14) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X12)) ) | ($spl9 | $spl51)),
  inference(forward_demodulation,[],[f2947,f3033])).
fof(f5408,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl51)),
  inference(forward_demodulation,[],[f2946,f3033])).
fof(f5406,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,X8) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) ) | ($spl9 | $spl51)),
  inference(forward_demodulation,[],[f2945,f3033])).
fof(f3508,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X1,f(s(bool,X0))) != s(X1,g(s(bool,X0))) | s(fun(bool,X1),X2) = s(fun(bool,X1),X3) | s(bool,X4) = s(bool,sK2(X2,X3,X1,bool)) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X0) = s(bool,X5)) )),
  inference(superposition,[],[f253,f97])).
fof(f3099,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,f(s(bool,X8))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl9 | $spl51)),
  inference(superposition,[],[f3033,f59])).
fof(f3097,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) = s(bool,X1)) ) | ($spl9 | $spl51)),
  inference(superposition,[],[f3033,f97])).
fof(f3088,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,X1)) ) | $spl51),
  inference(condensation,[],[f3087])).
fof(f3087,plain,(
  ( ! [X14,X15,X13] : (s(bool,X14) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl51),
  inference(subsumption_resolution,[],[f3058,f43])).
fof(f3058,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,X14) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl51),
  inference(superposition,[],[f2978,f97])).
fof(f3036,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X8))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl51),
  inference(superposition,[],[f2978,f59])).
fof(f3034,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,f(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) = s(bool,X1)) ) | $spl51),
  inference(superposition,[],[f2978,f97])).
fof(f3033,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | ($spl9 | $spl51)),
  inference(unit_resulting_resolution,[],[f38,f67,f2978,f844])).
fof(f2967,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(X48,f(s(bool,X47))) != s(X48,g(s(bool,X47))) | s(fun(bool,X48),X49) = s(fun(bool,X48),X50) | s(X48,f(s(bool,'T_0'))) != s(X48,g(s(bool,'T_0'))) | s(bool,X47) = s(bool,X51) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) ) | $spl9),
  inference(subsumption_resolution,[],[f2938,f67])).
fof(f2938,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(X48,f(s(bool,X47))) != s(X48,g(s(bool,X47))) | s(fun(bool,X48),X49) = s(fun(bool,X48),X50) | s(X48,f(s(bool,'T_0'))) != s(X48,g(s(bool,'T_0'))) | s(bool,X47) = s(bool,X51) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f272,f97])).
fof(f2928,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(bool,X0))) != s(X1,g(s(bool,X0))) | s(fun(bool,X1),X2) = s(fun(bool,X1),X3) | s(X1,f(s(bool,'T_0'))) != s(X1,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f272,f43])).
fof(f657,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2)) )),
  inference(equality_factoring,[],[f97])).
fof(f1422,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f650,f59])).
fof(f650,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X2)) )),
  inference(equality_factoring,[],[f97])).
fof(f740,plain,(
  ( ! [X2,X3,X1] : (s(bool,X1) = s(bool,X2) | s(bool,X2) != s(bool,X3) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X3)) )),
  inference(condensation,[],[f623])).
fof(f623,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X1) != s(bool,X2) | s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(equality_factoring,[],[f97])).
fof(f560,plain,(
  ( ! [X76,X74,X72,X77,X75,X73] : (s(X74,f(s(bool,'T_0'))) != s(X74,g(s(bool,'T_0'))) | s(fun(bool,X74),X72) = s(fun(bool,X74),X73) | s(bool,X75) = s(bool,X76) | s(bool,X76) = s(bool,sK2(X72,X73,X74,bool)) | s(bool,X75) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f48,f97])).
fof(f558,plain,(
  ( ! [X61,X64,X62,X60,X65,X63] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X60,bool),X62))) | s('type.pair.prod'(X60,bool),X61) = s('type.pair.prod'(X60,bool),X62) | s(X60,'const.pair.FST_1'(s('type.pair.prod'(X60,bool),X61))) != s(X60,'const.pair.FST_1'(s('type.pair.prod'(X60,bool),X62))) | s(bool,X63) = s(bool,X64) | s(bool,X64) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X60,bool),X61))) | s(bool,X63) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f51,f97])).
fof(f541,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X13)) | s(bool,X14) = s(bool,X15) | s(bool,X13) = s(bool,X15) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f38,f97])).
fof(f538,plain,(
  ( ! [X87,X85,X83,X86,X84,X82] : (s(X84,f(s(bool,X85))) != s(X84,g(s(bool,X85))) | s(fun(bool,X84),X82) = s(fun(bool,X84),X83) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,sK2(X82,X83,X84,bool)) | s(bool,X86) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f48,f97])).
fof(f536,plain,(
  ( ! [X70,X74,X72,X71,X75,X73] : (s(bool,X72) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X70,bool),X73))) | s('type.pair.prod'(X70,bool),X71) = s('type.pair.prod'(X70,bool),X73) | s(X70,'const.pair.FST_1'(s('type.pair.prod'(X70,bool),X71))) != s(X70,'const.pair.FST_1'(s('type.pair.prod'(X70,bool),X73))) | s(bool,X72) = s(bool,X74) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X70,bool),X71))) | s(bool,X74) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f51,f97])).
fof(f1419,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') != s(bool,X44) | s(bool,X44) = s(bool,X45) | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) ) | $spl9),
  inference(subsumption_resolution,[],[f531,f67])).
fof(f531,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') != s(bool,X44) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) ) | $spl9),
  inference(superposition,[],[f67,f97])).
fof(f844,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f843])).
fof(f843,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X25)) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,X26) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(duplicate_literal_removal,[],[f526])).
fof(f526,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,X25)) | s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,X26) | s(bool,'T_0') = s(bool,X24) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f61,f97])).
fof(f520,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (s(X87,f(s(bool,X88))) != s(X87,g(s(bool,X88))) | s(fun(bool,X87),X85) = s(fun(bool,X87),X86) | s(bool,X89) = s(bool,sK2(X85,X86,X87,bool)) | s(bool,'T_0') = s(bool,X88) | s(bool,X89) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f48,f97])).
fof(f518,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (s(bool,X75) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),X76))) | s('type.pair.prod'(X73,bool),X74) = s('type.pair.prod'(X73,bool),X76) | s(X73,'const.pair.FST_1'(s('type.pair.prod'(X73,bool),X74))) != s(X73,'const.pair.FST_1'(s('type.pair.prod'(X73,bool),X76))) | s(bool,X77) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),X74))) | s(bool,'T_0') = s(bool,X75) | s(bool,X77) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f51,f97])).
fof(f484,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (s(X87,f(s(bool,X88))) != s(X87,g(s(bool,X88))) | s(fun(bool,X87),X85) = s(fun(bool,X87),X86) | s(bool,X88) = s(bool,X89) | s(bool,'T_0') = s(bool,X89) | s(bool,X90) = s(bool,sK2(X85,X86,X87,bool)) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f48,f97])).
fof(f482,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (s(bool,X75) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),X76))) | s('type.pair.prod'(X73,bool),X74) = s('type.pair.prod'(X73,bool),X76) | s(X73,'const.pair.FST_1'(s('type.pair.prod'(X73,bool),X74))) != s(X73,'const.pair.FST_1'(s('type.pair.prod'(X73,bool),X76))) | s(bool,X75) = s(bool,X77) | s(bool,'T_0') = s(bool,X77) | s(bool,X78) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X73,bool),X74))) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f51,f97])).
fof(f377,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X1) = s(bool,X3) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X0) = s(bool,X5)) )),
  inference(superposition,[],[f97,f97])).
fof(f1218,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2)) )),
  inference(condensation,[],[f1217])).
fof(f1217,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(duplicate_literal_removal,[],[f358])).
fof(f358,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f97,f43])).
fof(f1360,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X1) = s(bool,X3)) )),
  inference(condensation,[],[f1359])).
fof(f1359,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s(bool,X80) = s(bool,X81) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X83) | s(bool,X84) = s(bool,X85) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,X80) = s(bool,X84)) )),
  inference(duplicate_literal_removal,[],[f314])).
fof(f314,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s(bool,X80) = s(bool,X81) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X83) | s(bool,'T_0') = s(bool,X83) | s(bool,X84) = s(bool,X85) | s(bool,X85) = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,X80) = s(bool,X84)) )),
  inference(superposition,[],[f97,f97])).
fof(f284,plain,(
  ( ! [X12,X10,X13,X11] : (s(bool,X12) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X10,bool),X13))) | s('type.pair.prod'(X10,bool),X11) = s('type.pair.prod'(X10,bool),X13) | s(X10,'const.pair.FST_1'(s('type.pair.prod'(X10,bool),X11))) != s(X10,'const.pair.FST_1'(s('type.pair.prod'(X10,bool),X13))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X10,bool),X11))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f51,f59])).
fof(f282,plain,(
  ( ! [X4,X2,X5,X3] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),X4))) | s('type.pair.prod'(X2,bool),X3) = s('type.pair.prod'(X2,bool),X4) | s(X2,'const.pair.FST_1'(s('type.pair.prod'(X2,bool),X3))) != s(X2,'const.pair.FST_1'(s('type.pair.prod'(X2,bool),X4))) | s(bool,X5) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X2,bool),X3))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f51,f59])).
fof(f255,plain,(
  ( ! [X10,X8,X11,X9] : (s(X10,f(s(bool,X11))) != s(X10,g(s(bool,X11))) | s(fun(bool,X10),X8) = s(fun(bool,X10),X9) | s(bool,'T_0') = s(bool,sK2(X8,X9,X10,bool)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f48,f59])).
fof(f253,plain,(
  ( ! [X2,X0,X3,X1] : (s(X2,f(s(bool,'T_0'))) != s(X2,g(s(bool,'T_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,X3) = s(bool,sK2(X0,X1,X2,bool)) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f48,f59])).
fof(f97,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f59,f59])).
fof(f61,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f40,f43])).
fof(f49,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'(X0,X1),X2) != s('type.pair.prod'(X0,X1),X3) | s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3)))) )),
  inference(cnf_transformation,[],[f34])).
fof(f50,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'(X0,X1),X2) != s('type.pair.prod'(X0,X1),X3) | s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))) )),
  inference(cnf_transformation,[],[f34])).
fof(f48,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK2(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK2(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f21])).
fof(f21,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.complex.COMPLEX_RE_IM_EQ',aHL_EXT)).
fof(f12983,plain,(
  $false | ($spl88 | $spl126)),
  inference(subsumption_resolution,[],[f12933,f38])).
fof(f12933,plain,(
  ~p(s(bool,'T_0')) | ($spl88 | $spl126)),
  inference(backward_demodulation,[],[f12894,f40])).
fof(f12894,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl88 | $spl126)),
  inference(duplicate_literal_removal,[],[f12756])).
fof(f12756,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl88 | $spl126)),
  inference(superposition,[],[f6639,f5398])).
fof(f12925,plain,(
  $false | ($spl15 | $spl51 | $spl81 | $spl88 | $spl126)),
  inference(subsumption_resolution,[],[f12924,f11832])).
fof(f12924,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl15 | $spl51 | $spl81 | $spl88 | $spl126)),
  inference(forward_demodulation,[],[f12923,f5398])).
fof(f12923,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl15 | $spl51 | $spl81 | $spl88 | $spl126)),
  inference(subsumption_resolution,[],[f12737,f11869])).
fof(f11869,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | ($spl15 | $spl51)),
  inference(global_subsumption,[],[f11868,f37,f36,f35,f38,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f97,f129,f253,f255,f294,f282,f284,f1360,f1218,f377,f482,f484,f518,f520,f844,f536,f538,f541,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1740,f2495,f2496,f2498,f2529,f2527,f2928,f3241,f3243,f3508,f4348,f4305,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f3527,f5489,f3528,f3529,f5492,f3530,f5494,f3531,f3534,f3536,f5495,f3537,f3539,f5497,f2956,f2955,f2954,f2953,f2952,f2951,f2950,f2949,f2948,f2947,f2946,f2945,f3538,f3535,f3533,f3532,f3192,f3189,f1549,f5614,f5615,f5617,f5647,f5645,f557,f535,f517,f481,f260,f258,f2406,f2405,f2404,f2402,f2399,f2398,f2397,f2393,f2392,f2391,f2389,f2386,f2385,f2384,f1540,f1512,f1220,f286,f272,f262,f257,f40,f741])).
fof(f5645,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) = s(bool,X11) | s(bool,X11) = s(bool,g(s(bool,'T_0'))) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl15),
  inference(trivial_inequality_removal,[],[f5638])).
fof(f5638,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X10) = s(bool,X11) | s(bool,X11) = s(bool,g(s(bool,'T_0'))) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl15),
  inference(superposition,[],[f1549,f97])).
fof(f5647,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl15),
  inference(trivial_inequality_removal,[],[f5632])).
fof(f5632,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl15),
  inference(superposition,[],[f1549,f59])).
fof(f5617,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,g(s(bool,X8))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl15),
  inference(superposition,[],[f1549,f59])).
fof(f5615,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,g(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) = s(bool,X1)) ) | $spl15),
  inference(superposition,[],[f1549,f97])).
fof(f5614,plain,(
  ~p(s(bool,g(s(bool,'T_0')))) | $spl15),
  inference(unit_resulting_resolution,[],[f1549,f1549,f1549,f844])).
fof(f1549,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | $spl15),
  inference(cnf_transformation,[],[f1549_D])).
fof(f1549_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f4305,plain,(
  ( ! [X64,X63] : (s(bool,X63) != s(bool,g(s(bool,'T_0'))) | s(bool,X63) = s(bool,X64) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,g(s(bool,X63)))) ) | $spl15),
  inference(superposition,[],[f2498,f2529])).
fof(f4348,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,X33) = s(bool,X34) | s(bool,'T_0') = s(bool,X34) | s(bool,X35) = s(bool,X36) | s(bool,X36) = s(bool,g(s(bool,X33))) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl15),
  inference(subsumption_resolution,[],[f4299,f59])).
fof(f4299,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,'T_0') != s(bool,X33) | s(bool,X33) = s(bool,X34) | s(bool,'T_0') = s(bool,X34) | s(bool,X35) = s(bool,X36) | s(bool,X36) = s(bool,g(s(bool,X33))) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | $spl15),
  inference(superposition,[],[f2498,f97])).
fof(f3243,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) = s(bool,g(s(bool,X10))) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl15),
  inference(superposition,[],[f2529,f59])).
fof(f3241,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,X1) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X0) = s(bool,X2)) ) | $spl15),
  inference(superposition,[],[f2529,f97])).
fof(f2527,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) = s(bool,X11) | s(bool,X11) = s(bool,g(s(bool,'T_0'))) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl15),
  inference(trivial_inequality_removal,[],[f2520])).
fof(f2520,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X10) = s(bool,X11) | s(bool,X11) = s(bool,g(s(bool,'T_0'))) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl15),
  inference(superposition,[],[f1549,f97])).
fof(f2529,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl15),
  inference(trivial_inequality_removal,[],[f2514])).
fof(f2514,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl15),
  inference(superposition,[],[f1549,f59])).
fof(f2498,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,g(s(bool,X8))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl15),
  inference(superposition,[],[f1549,f59])).
fof(f2496,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,g(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) = s(bool,X1)) ) | $spl15),
  inference(superposition,[],[f1549,f97])).
fof(f2495,plain,(
  ~p(s(bool,g(s(bool,'T_0')))) | $spl15),
  inference(unit_resulting_resolution,[],[f1549,f1549,f1549,f844])).
fof(f12737,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl15 | $spl51 | $spl81 | $spl88 | $spl126)),
  inference(superposition,[],[f6639,f12350])).
fof(f12350,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl15 | $spl51 | $spl81 | $spl88)),
  inference(unit_resulting_resolution,[],[f11832,f11869,f11869,f11869,f97])).
fof(f12922,plain,(
  $false | ($spl81 | $spl86 | $spl88 | $spl126)),
  inference(subsumption_resolution,[],[f12921,f11832])).
fof(f12921,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl81 | $spl86 | $spl88 | $spl126)),
  inference(condensation,[],[f12920])).
fof(f12920,plain,(
  ( ! [X46] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X46)) ) | ($spl81 | $spl86 | $spl88 | $spl126)),
  inference(subsumption_resolution,[],[f12734,f12713])).
fof(f12713,plain,(
  ( ! [X45] : (s(bool,X45) != s(bool,f(s(bool,X45))) | s(bool,'T_0') = s(bool,X45)) ) | ($spl81 | $spl86 | $spl88)),
  inference(subsumption_resolution,[],[f12667,f11832])).
fof(f12667,plain,(
  ( ! [X45] : (s(bool,X45) != s(bool,f(s(bool,X45))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X45)) ) | $spl86),
  inference(superposition,[],[f5338,f43])).
fof(f12734,plain,(
  ( ! [X46] : (s(bool,X46) = s(bool,f(s(bool,X46))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X46)) ) | $spl126),
  inference(superposition,[],[f6639,f43])).
fof(f12919,plain,(
  $false | ($spl81 | $spl88 | $spl126)),
  inference(subsumption_resolution,[],[f12918,f11832])).
fof(f12918,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl88 | $spl126)),
  inference(condensation,[],[f12917])).
fof(f12917,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl88 | $spl126)),
  inference(forward_demodulation,[],[f12911,f5398])).
fof(f12911,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl126),
  inference(duplicate_literal_removal,[],[f12721])).
fof(f12721,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl126),
  inference(superposition,[],[f6639,f43])).
fof(f12916,plain,(
  $false | ($spl81 | $spl88 | $spl126)),
  inference(subsumption_resolution,[],[f12915,f11832])).
fof(f12915,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl81 | $spl88 | $spl126)),
  inference(forward_demodulation,[],[f12719,f5398])).
fof(f12719,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl81 | $spl88 | $spl126)),
  inference(unit_resulting_resolution,[],[f11832,f6639])).
fof(f12914,plain,(
  $false | ($spl15 | $spl51 | $spl81 | $spl88 | $spl126)),
  inference(subsumption_resolution,[],[f12913,f11832])).
fof(f12913,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl15 | $spl51 | $spl81 | $spl88 | $spl126)),
  inference(forward_demodulation,[],[f12912,f5398])).
fof(f12912,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl15 | $spl51 | $spl81 | $spl88 | $spl126)),
  inference(forward_demodulation,[],[f12720,f12350])).
fof(f12720,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,g(s(bool,'T_0'))) | ($spl15 | $spl51 | $spl126)),
  inference(unit_resulting_resolution,[],[f11869,f6639])).
fof(f11862,plain,(
  $false | ($spl6 | $spl51)),
  inference(subsumption_resolution,[],[f11861,f38])).
fof(f11861,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl51)),
  inference(forward_demodulation,[],[f40,f11853])).
fof(f11853,plain,(
  ( ! [X12] : (s(bool,'T_0') = s(bool,X12)) ) | ($spl6 | $spl51)),
  inference(subsumption_resolution,[],[f5806,f65])).
fof(f5806,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X12)) ) | $spl51),
  inference(duplicate_literal_removal,[],[f5769])).
fof(f5769,plain,(
  ( ! [X12] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl51),
  inference(superposition,[],[f2978,f59])).
fof(f11860,plain,(
  $false | ($spl6 | $spl15 | $spl51)),
  inference(subsumption_resolution,[],[f1549,f11853])).
fof(f11855,plain,(
  $false | ($spl6 | $spl113)),
  inference(subsumption_resolution,[],[f11854,f65])).
fof(f11854,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | ($spl6 | $spl113)),
  inference(forward_demodulation,[],[f5499,f65])).
fof(f11852,plain,(
  $false | ($spl6 | $spl51)),
  inference(subsumption_resolution,[],[f2978,f65])).
fof(f11850,plain,(
  $false | ($spl88 | $spl92 | $spl113 | $spl126)),
  inference(subsumption_resolution,[],[f11849,f11834])).
fof(f11834,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1)) ) | ($spl88 | $spl92 | $spl126)),
  inference(global_subsumption,[],[f6639,f37,f36,f35,f38,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f97,f129,f253,f255,f294,f282,f284,f1360,f1218,f377,f482,f484,f518,f520,f844,f536,f538,f541,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1740,f2928,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f3527,f3528,f3529,f3530,f3531,f3534,f3536,f3537,f3539,f2956,f2955,f2954,f2953,f2952,f2951,f2950,f2949,f2948,f2947,f2946,f2945,f3538,f3535,f3533,f3532,f3192,f3189,f557,f535,f517,f481,f260,f258,f5401,f2406,f2405,f2404,f2402,f2399,f2398,f2397,f2393,f2392,f2391,f2389,f2386,f2385,f2384,f1540,f1512,f1220,f286,f272,f262,f257,f40,f10697,f10728,f10789,f10695,f10694,f10736,f5398])).
fof(f10736,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,f(s(bool,'F_0')))) ) | $spl92),
  inference(condensation,[],[f10735])).
fof(f10735,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl92),
  inference(duplicate_literal_removal,[],[f10734])).
fof(f10734,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl92),
  inference(condensation,[],[f10733])).
fof(f10733,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X1)) ) | $spl92),
  inference(condensation,[],[f10681])).
fof(f10681,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X40) | s(bool,X39) = s(bool,X40) | s(bool,X41) = s(bool,X42) | s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43) | s(bool,X39) = s(bool,X41)) ) | $spl92),
  inference(superposition,[],[f5401,f97])).
fof(f10694,plain,(
  ( ! [X78,X76,X77,X75] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X78) | s(bool,X78) = s(bool,g(s(X75,sK2(X76,X77,bool,X75)))) | s(fun(X75,bool),X76) = s(fun(X75,bool),X77) | s(bool,'F_0') != s(bool,f(s(X75,sK2(X76,X77,bool,X75))))) ) | $spl92),
  inference(superposition,[],[f5401,f262])).
fof(f10695,plain,(
  ( ! [X80,X81,X79,X82] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X82) | s(bool,X82) = s(bool,sK2(X79,X80,X81,bool)) | s(fun(bool,X81),X79) = s(fun(bool,X81),X80) | s(X81,f(s(bool,'F_0'))) != s(X81,g(s(bool,'F_0')))) ) | $spl92),
  inference(superposition,[],[f5401,f257])).
fof(f10789,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X0) = s(bool,X2) | s(bool,X1) = s(bool,f(s(bool,'F_0')))) ) | $spl92),
  inference(condensation,[],[f10788])).
fof(f10788,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X1) = s(bool,f(s(bool,'F_0')))) ) | $spl92),
  inference(condensation,[],[f10787])).
fof(f10787,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,X47) | s(bool,X46) = s(bool,X47) | s(bool,X48) = s(bool,X49) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,X48) = s(bool,f(s(bool,'F_0')))) ) | $spl92),
  inference(subsumption_resolution,[],[f10724,f59])).
fof(f10724,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'T_0') != s(bool,X46) | s(bool,'T_0') = s(bool,X47) | s(bool,X46) = s(bool,X47) | s(bool,X48) = s(bool,X49) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,X48) = s(bool,f(s(bool,'F_0')))) ) | $spl92),
  inference(superposition,[],[f5401,f97])).
fof(f10728,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,f(s(bool,'F_0')))) ) | $spl92),
  inference(equality_resolution,[],[f5401])).
fof(f10697,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X2) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) ) | $spl92),
  inference(superposition,[],[f5401,f43])).
fof(f11849,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl88 | $spl92 | $spl113 | $spl126)),
  inference(forward_demodulation,[],[f5499,f11834])).
fof(f11848,plain,(
  $false | ($spl15 | $spl88 | $spl92 | $spl126)),
  inference(subsumption_resolution,[],[f1549,f11834])).
fof(f11847,plain,(
  $false | ($spl88 | $spl92 | $spl126)),
  inference(subsumption_resolution,[],[f11846,f38])).
fof(f11846,plain,(
  ~p(s(bool,'T_0')) | ($spl88 | $spl92 | $spl126)),
  inference(forward_demodulation,[],[f40,f11834])).
fof(f11829,plain,(
  $false | ($spl9 | $spl51 | $spl81 | $spl92 | $spl113)),
  inference(subsumption_resolution,[],[f10665,f11777])).
fof(f11777,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl81 | $spl92)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f294,f282,f284,f1360,f1218,f377,f482,f484,f518,f520,f844,f1419,f536,f538,f541,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1740,f2928,f2967,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5329,f3527,f3528,f3529,f3530,f3531,f3534,f3536,f3537,f3539,f2956,f2955,f2954,f2953,f2952,f2951,f2950,f2949,f2948,f2947,f2946,f2945,f3538,f3535,f3533,f3532,f3192,f3189,f6014,f6016,f6047,f6053,f6251,f6323,f557,f535,f517,f481,f260,f258,f5401,f2406,f2405,f2404,f2402,f2399,f2398,f2397,f2393,f2392,f2391,f2389,f2386,f2385,f2384,f1540,f1512,f1220,f286,f272,f262,f257,f40,f10697,f10764,f10728,f10789,f10695,f11773,f10694,f11774,f10736])).
fof(f11774,plain,(
  ( ! [X78,X76,X77,X75] : (s(fun(X75,bool),X76) = s(fun(X75,bool),X77) | s(bool,'F_0') != s(bool,f(s(X75,sK2(X76,X77,bool,X75)))) | s(bool,X78) = s(bool,g(s(X75,sK2(X76,X77,bool,X75)))) | s(bool,'T_0') = s(bool,X78)) ) | ($spl9 | $spl81 | $spl92)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f294,f282,f284,f1360,f1218,f377,f482,f484,f518,f520,f844,f1419,f536,f538,f541,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1740,f2928,f2967,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5329,f3527,f3528,f3529,f3530,f3531,f3534,f3536,f3537,f3539,f2956,f2955,f2954,f2953,f2952,f2951,f2950,f2949,f2948,f2947,f2946,f2945,f3538,f3535,f3533,f3532,f3192,f3189,f6014,f6016,f6047,f6053,f6251,f6323,f557,f535,f517,f481,f260,f258,f5401,f2406,f2405,f2404,f2402,f2399,f2398,f2397,f2393,f2392,f2391,f2389,f2386,f2385,f2384,f1540,f1512,f1220,f286,f272,f262,f257,f40,f10697,f10764,f10728,f10789,f10695,f11773,f10694])).
fof(f11773,plain,(
  ( ! [X80,X81,X79,X82] : (s(fun(bool,X81),X79) = s(fun(bool,X81),X80) | s(X81,f(s(bool,'F_0'))) != s(X81,g(s(bool,'F_0'))) | s(bool,X82) = s(bool,sK2(X79,X80,X81,bool)) | s(bool,'T_0') = s(bool,X82)) ) | ($spl9 | $spl81 | $spl92)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f294,f282,f284,f1360,f1218,f377,f482,f484,f518,f520,f844,f1419,f536,f538,f541,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1740,f2928,f2967,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5329,f3527,f3528,f3529,f3530,f3531,f3534,f3536,f3537,f3539,f2956,f2955,f2954,f2953,f2952,f2951,f2950,f2949,f2948,f2947,f2946,f2945,f3538,f3535,f3533,f3532,f3192,f3189,f6014,f6016,f6047,f6053,f6251,f6323,f557,f535,f517,f481,f260,f258,f5401,f2406,f2405,f2404,f2402,f2399,f2398,f2397,f2393,f2392,f2391,f2389,f2386,f2385,f2384,f1540,f1512,f1220,f286,f272,f262,f257,f40,f10697,f10764,f10728,f10789,f10695])).
fof(f10764,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(bool,X38) != s(bool,f(s(bool,X37))) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X39) | s(bool,X37) = s(bool,X40) | s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) ) | ($spl9 | $spl92)),
  inference(subsumption_resolution,[],[f10707,f67])).
fof(f10707,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s(bool,X38) != s(bool,f(s(bool,X37))) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X39) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X40) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) ) | $spl92),
  inference(superposition,[],[f5401,f97])).
fof(f10665,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | ($spl9 | $spl51 | $spl92 | $spl113)),
  inference(unit_resulting_resolution,[],[f6369,f6369,f5401])).
fof(f11828,plain,(
  $false | ($spl9 | $spl81 | $spl92)),
  inference(subsumption_resolution,[],[f10664,f11777])).
fof(f10664,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | ($spl9 | $spl92)),
  inference(unit_resulting_resolution,[],[f67,f67,f5401])).
fof(f11826,plain,(
  $false | ($spl9 | $spl81 | $spl92)),
  inference(subsumption_resolution,[],[f11778,f38])).
fof(f11778,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl81 | $spl92)),
  inference(backward_demodulation,[],[f11777,f40])).
fof(f11769,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'F_0')))) ) | $spl136),
  inference(cnf_transformation,[],[f11769_D])).
fof(f11769_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'F_0')))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f11770,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl138),
  inference(cnf_transformation,[],[f11770_D])).
fof(f11770_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f11312,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,g(s(bool,'F_0'))) | s(bool,X0) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl96),
  inference(equality_factoring,[],[f5403])).
fof(f5270,plain,(
  ( ! [X80] : (s(bool,X80) = s(bool,f(s(bool,X80))) | s(bool,X80) = s(bool,g(s(bool,'F_0')))) ) | $spl60),
  inference(cnf_transformation,[],[f5270_D])).
fof(f5270_D,plain,(
  ( ! [X80] : (s(bool,X80) = s(bool,f(s(bool,X80))) | s(bool,X80) = s(bool,g(s(bool,'F_0')))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f11343,plain,(
  ( ! [X83,X82] : (s(bool,X83) = s(bool,f(s(bool,X83))) | s(bool,'T_0') = s(bool,X82) | s(bool,X83) = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl51 | $spl89 | $spl96 | $spl113)),
  inference(duplicate_literal_removal,[],[f11263])).
fof(f11263,plain,(
  ( ! [X83,X82] : (s(bool,X83) = s(bool,f(s(bool,X83))) | s(bool,'T_0') = s(bool,X82) | s(bool,X83) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X82)) ) | ($spl9 | $spl51 | $spl89 | $spl96 | $spl113)),
  inference(superposition,[],[f7464,f5403])).
fof(f7464,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl51 | $spl89 | $spl113)),
  inference(superposition,[],[f6823,f43])).
fof(f6823,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl9 | $spl51 | $spl89 | $spl113)),
  inference(forward_demodulation,[],[f6773,f6442])).
fof(f6773,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,g(s(bool,'T_0'))) | ($spl9 | $spl51 | $spl89 | $spl113)),
  inference(unit_resulting_resolution,[],[f6369,f5399,f5399,f5399,f97])).
fof(f5268,plain,(
  ( ! [X76] : (s(bool,'T_0') != s(bool,f(s(bool,X76))) | s(bool,X76) = s(bool,g(s(bool,'F_0')))) ) | $spl58),
  inference(cnf_transformation,[],[f5268_D])).
fof(f5268_D,plain,(
  ( ! [X76] : (s(bool,'T_0') != s(bool,f(s(bool,X76))) | s(bool,X76) = s(bool,g(s(bool,'F_0')))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f11347,plain,(
  ( ! [X76,X77] : (s(bool,'T_0') != s(bool,f(s(bool,X77))) | s(bool,'T_0') = s(bool,X76) | s(bool,X77) = s(bool,g(s(bool,'F_0')))) ) | ($spl89 | $spl96)),
  inference(duplicate_literal_removal,[],[f11260])).
fof(f11260,plain,(
  ( ! [X76,X77] : (s(bool,'T_0') != s(bool,f(s(bool,X77))) | s(bool,'T_0') = s(bool,X76) | s(bool,X77) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X76)) ) | ($spl89 | $spl96)),
  inference(superposition,[],[f6791,f5403])).
fof(f6791,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl89),
  inference(superposition,[],[f5399,f43])).
fof(f5196,plain,(
  ( ! [X51,X49] : (s(bool,X51) = s(bool,g(s(bool,'F_0'))) | s(bool,X49) = s(bool,g(s(bool,'F_0'))) | s(bool,X49) = s(bool,X51)) ) | $spl54),
  inference(cnf_transformation,[],[f5196_D])).
fof(f5196_D,plain,(
  ( ! [X51,X49] : (s(bool,X51) = s(bool,g(s(bool,'F_0'))) | s(bool,X49) = s(bool,g(s(bool,'F_0'))) | s(bool,X49) = s(bool,X51)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f11507,plain,(
  ( ! [X61,X62,X60] : (s(bool,X61) = s(bool,X62) | s(bool,X62) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X60) | s(bool,X61) = s(bool,g(s(bool,'F_0')))) ) | $spl96),
  inference(duplicate_literal_removal,[],[f10990])).
fof(f10990,plain,(
  ( ! [X61,X62,X60] : (s(bool,X61) = s(bool,X62) | s(bool,X62) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X60) | s(bool,X61) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X60)) ) | $spl96),
  inference(superposition,[],[f5403,f5403])).
fof(f11548,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') = s(bool,X64) | s(bool,X63) = s(bool,g(s(bool,'F_0'))) | s(bool,X65) = s(bool,g(s(bool,'F_0'))) | s(bool,X63) = s(bool,X65)) ) | $spl96),
  inference(duplicate_literal_removal,[],[f10954])).
fof(f10954,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') = s(bool,X64) | s(bool,X63) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X64) | s(bool,X65) = s(bool,g(s(bool,'F_0'))) | s(bool,X63) = s(bool,X65)) ) | $spl96),
  inference(superposition,[],[f5403,f5403])).
fof(f5197,plain,(
  ( ! [X62] : (s(bool,X62) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X62)) ) | $spl56),
  inference(cnf_transformation,[],[f5197_D])).
fof(f5197_D,plain,(
  ( ! [X62] : (s(bool,X62) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X62)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f11665,plain,(
  ( ! [X64,X63] : (s(bool,'F_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X63) | s(bool,X64) = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl51 | $spl89 | $spl96 | $spl113)),
  inference(forward_demodulation,[],[f11569,f6823])).
fof(f11569,plain,(
  ( ! [X64,X63] : (s(bool,'T_0') = s(bool,X63) | s(bool,X64) = s(bool,g(s(bool,'F_0'))) | s(bool,X64) = s(bool,f(s(bool,'F_0')))) ) | ($spl89 | $spl96)),
  inference(duplicate_literal_removal,[],[f10932])).
fof(f10932,plain,(
  ( ! [X64,X63] : (s(bool,'T_0') = s(bool,X63) | s(bool,X64) = s(bool,g(s(bool,'F_0'))) | s(bool,X64) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X63)) ) | ($spl89 | $spl96)),
  inference(superposition,[],[f5403,f6821])).
fof(f6821,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl89),
  inference(trivial_inequality_removal,[],[f6806])).
fof(f6806,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl89),
  inference(superposition,[],[f5399,f59])).
fof(f11570,plain,(
  ( ! [X61,X62] : (s(bool,'T_0') = s(bool,X61) | s(bool,X62) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X62)) ) | $spl96),
  inference(duplicate_literal_removal,[],[f10931])).
fof(f10931,plain,(
  ( ! [X61,X62] : (s(bool,'T_0') = s(bool,X61) | s(bool,X62) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X61)) ) | $spl96),
  inference(superposition,[],[f5403,f43])).
fof(f11571,plain,(
  ( ! [X47,X45,X46] : (s(bool,'T_0') = s(bool,X46) | s(bool,X47) = s(bool,g(s(bool,'F_0'))) | s(bool,X45) = s(bool,X47) | s(bool,X45) = s(bool,g(s(bool,'F_0')))) ) | $spl96),
  inference(duplicate_literal_removal,[],[f10927])).
fof(f10927,plain,(
  ( ! [X47,X45,X46] : (s(bool,'T_0') = s(bool,X46) | s(bool,X47) = s(bool,g(s(bool,'F_0'))) | s(bool,X45) = s(bool,X47) | s(bool,X45) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X46)) ) | $spl96),
  inference(superposition,[],[f5403,f5403])).
fof(f11658,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X1)) ) | ($spl9 | $spl48 | $spl51 | $spl96 | $spl113)),
  inference(condensation,[],[f11657])).
fof(f11657,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') = s(bool,X64) | s(bool,X63) = s(bool,X64) | s(bool,X65) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X65)) ) | ($spl9 | $spl48 | $spl51 | $spl96 | $spl113)),
  inference(subsumption_resolution,[],[f11656,f2976])).
fof(f11656,plain,(
  ( ! [X64,X65,X63] : (s(bool,'F_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X64) | s(bool,X63) = s(bool,X64) | s(bool,X65) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X65)) ) | ($spl9 | $spl51 | $spl96 | $spl113)),
  inference(forward_demodulation,[],[f10874,f6442])).
fof(f10874,plain,(
  ( ! [X64,X65,X63] : (s(bool,X63) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X64) | s(bool,X63) = s(bool,X64) | s(bool,X65) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X65)) ) | $spl96),
  inference(superposition,[],[f5403,f5403])).
fof(f9280,plain,(
  $false | ($spl9 | $spl40 | $spl48 | $spl51 | $spl89 | $spl113)),
  inference(subsumption_resolution,[],[f9257,f9229])).
fof(f9229,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl40 | $spl48 | $spl51 | $spl113)),
  inference(condensation,[],[f9228])).
fof(f9228,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | ($spl9 | $spl40 | $spl48 | $spl51 | $spl113)),
  inference(subsumption_resolution,[],[f9227,f2976])).
fof(f9227,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X1)) ) | ($spl9 | $spl40 | $spl51 | $spl113)),
  inference(condensation,[],[f9226])).
fof(f9226,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl9 | $spl40 | $spl51 | $spl113)),
  inference(condensation,[],[f9225])).
fof(f9225,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'F_0') = s(bool,X9) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X9) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | ($spl9 | $spl40 | $spl51 | $spl113)),
  inference(subsumption_resolution,[],[f9142,f6457])).
fof(f9142,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,X10) = s(bool,g(s(bool,X10))) | s(bool,'F_0') = s(bool,X9) | s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X9) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl40),
  inference(superposition,[],[f2371,f97])).
fof(f9257,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl9 | $spl40 | $spl48 | $spl51 | $spl89 | $spl113)),
  inference(backward_demodulation,[],[f9229,f5399])).
fof(f9278,plain,(
  $false | ($spl9 | $spl40 | $spl48 | $spl51 | $spl113)),
  inference(subsumption_resolution,[],[f9230,f38])).
fof(f9230,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl40 | $spl48 | $spl51 | $spl113)),
  inference(backward_demodulation,[],[f9229,f40])).
fof(f9224,plain,(
  $false | ($spl9 | $spl40 | $spl48 | $spl51 | $spl113)),
  inference(subsumption_resolution,[],[f9223,f6369])).
fof(f9223,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl40 | $spl48)),
  inference(duplicate_literal_removal,[],[f9222])).
fof(f9222,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl40 | $spl48)),
  inference(condensation,[],[f9221])).
fof(f9221,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | ($spl40 | $spl48)),
  inference(subsumption_resolution,[],[f9141,f2976])).
fof(f9141,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X7) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl40),
  inference(superposition,[],[f2371,f59])).
fof(f9220,plain,(
  $false | ($spl9 | $spl40 | $spl48 | $spl51 | $spl113)),
  inference(subsumption_resolution,[],[f9219,f6369])).
fof(f9219,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl40 | $spl48)),
  inference(duplicate_literal_removal,[],[f9218])).
fof(f9218,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl40 | $spl48)),
  inference(condensation,[],[f9217])).
fof(f9217,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl40 | $spl48)),
  inference(subsumption_resolution,[],[f9138,f2976])).
fof(f9138,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl40),
  inference(superposition,[],[f2371,f59])).
fof(f9216,plain,(
  $false | ($spl9 | $spl40 | $spl51 | $spl113)),
  inference(subsumption_resolution,[],[f9134,f6369])).
fof(f9134,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl9 | $spl40)),
  inference(unit_resulting_resolution,[],[f67,f2371])).
fof(f9215,plain,(
  $false | ($spl9 | $spl40 | $spl51 | $spl113)),
  inference(subsumption_resolution,[],[f9135,f67])).
fof(f9135,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl9 | $spl40 | $spl51 | $spl113)),
  inference(unit_resulting_resolution,[],[f6369,f2371])).
fof(f9136,plain,(
  $false | ($spl9 | $spl40 | $spl51 | $spl113)),
  inference(unit_resulting_resolution,[],[f67,f6369,f2371])).
fof(f8402,plain,(
  $false | ($spl9 | $spl51 | $spl86 | $spl89 | $spl113)),
  inference(subsumption_resolution,[],[f8379,f8349])).
fof(f8349,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl51 | $spl86 | $spl89 | $spl113)),
  inference(subsumption_resolution,[],[f8342,f6823])).
fof(f8342,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl86),
  inference(duplicate_literal_removal,[],[f8285])).
fof(f8285,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl86),
  inference(superposition,[],[f5338,f43])).
fof(f8379,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl9 | $spl51 | $spl86 | $spl89 | $spl113)),
  inference(backward_demodulation,[],[f8349,f5399])).
fof(f8400,plain,(
  $false | ($spl9 | $spl51 | $spl86 | $spl89 | $spl113)),
  inference(subsumption_resolution,[],[f8350,f38])).
fof(f8350,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl51 | $spl86 | $spl89 | $spl113)),
  inference(backward_demodulation,[],[f8349,f40])).
fof(f8348,plain,(
  $false | ($spl9 | $spl51 | $spl86 | $spl89 | $spl113)),
  inference(subsumption_resolution,[],[f8280,f6823])).
fof(f8280,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | ($spl9 | $spl86)),
  inference(unit_resulting_resolution,[],[f67,f5338])).
fof(f8347,plain,(
  $false | ($spl9 | $spl51 | $spl86 | $spl89 | $spl113)),
  inference(subsumption_resolution,[],[f8346,f6823])).
fof(f8346,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | ($spl9 | $spl51 | $spl86 | $spl113)),
  inference(forward_demodulation,[],[f8281,f6442])).
fof(f8281,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,g(s(bool,'T_0'))))) | ($spl9 | $spl51 | $spl86 | $spl113)),
  inference(unit_resulting_resolution,[],[f6369,f5338])).
fof(f8345,plain,(
  $false | ($spl9 | $spl51 | $spl86 | $spl89 | $spl113)),
  inference(subsumption_resolution,[],[f8344,f6823])).
fof(f8344,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | ($spl9 | $spl51 | $spl86 | $spl89 | $spl113)),
  inference(forward_demodulation,[],[f8282,f6823])).
fof(f8282,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,f(s(bool,'F_0'))))) | ($spl86 | $spl89)),
  inference(unit_resulting_resolution,[],[f5399,f5338])).
fof(f8343,plain,(
  $false | ($spl9 | $spl51 | $spl86 | $spl89 | $spl113)),
  inference(subsumption_resolution,[],[f8283,f67])).
fof(f8283,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl9 | $spl51 | $spl86 | $spl89 | $spl113)),
  inference(unit_resulting_resolution,[],[f6823,f5338])).
fof(f8284,plain,(
  $false | ($spl9 | $spl51 | $spl86 | $spl89 | $spl113)),
  inference(unit_resulting_resolution,[],[f67,f6823,f5338])).
fof(f7680,plain,(
  $false | ($spl9 | $spl42 | $spl51 | $spl89 | $spl113)),
  inference(subsumption_resolution,[],[f7656,f7626])).
fof(f7626,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl42 | $spl51 | $spl113)),
  inference(subsumption_resolution,[],[f7625,f38])).
fof(f7625,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'T_0'))) ) | ($spl9 | $spl42 | $spl51 | $spl113)),
  inference(condensation,[],[f7624])).
fof(f7624,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl9 | $spl42 | $spl51 | $spl113)),
  inference(subsumption_resolution,[],[f7599,f6457])).
fof(f7599,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl42),
  inference(superposition,[],[f2378,f59])).
fof(f7656,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl9 | $spl42 | $spl51 | $spl89 | $spl113)),
  inference(backward_demodulation,[],[f7626,f5399])).
fof(f7678,plain,(
  $false | ($spl9 | $spl42 | $spl51 | $spl113)),
  inference(subsumption_resolution,[],[f7627,f38])).
fof(f7627,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl42 | $spl51 | $spl113)),
  inference(backward_demodulation,[],[f7626,f40])).
fof(f7623,plain,(
  $false | ($spl9 | $spl42 | $spl51 | $spl113)),
  inference(subsumption_resolution,[],[f7595,f6369])).
fof(f7595,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl42),
  inference(unit_resulting_resolution,[],[f38,f2378])).
fof(f7622,plain,(
  $false | ($spl9 | $spl42 | $spl51 | $spl113)),
  inference(subsumption_resolution,[],[f7596,f38])).
fof(f7596,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl42 | $spl51 | $spl113)),
  inference(unit_resulting_resolution,[],[f6369,f2378])).
fof(f7597,plain,(
  $false | ($spl9 | $spl42 | $spl51 | $spl113)),
  inference(unit_resulting_resolution,[],[f38,f6369,f2378])).
fof(f7590,plain,(
  ( ! [X37,X38] : (s(fun(fun(bool,bool),bool),X37) = s(fun(fun(bool,bool),bool),X38) | s(bool,'F_0') != s(bool,f(s(fun(bool,bool),sK2(X37,X38,bool,fun(bool,bool)))))) ) | $spl132),
  inference(cnf_transformation,[],[f7590_D])).
fof(f7590_D,plain,(
  ( ! [X37,X38] : (s(fun(fun(bool,bool),bool),X37) = s(fun(fun(bool,bool),bool),X38) | s(bool,'F_0') != s(bool,f(s(fun(bool,bool),sK2(X37,X38,bool,fun(bool,bool)))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f7591,plain,(
  ( ! [X39] : (s(bool,'T_0') = s(bool,g(s(fun(bool,bool),X39)))) ) | $spl134),
  inference(cnf_transformation,[],[f7591_D])).
fof(f7591_D,plain,(
  ( ! [X39] : (s(bool,'T_0') = s(bool,g(s(fun(bool,bool),X39)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f7587,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,g(s(fun(bool,bool),X39))) | s(fun(fun(bool,bool),bool),X37) = s(fun(fun(bool,bool),bool),X38) | s(bool,'F_0') != s(bool,f(s(fun(bool,bool),sK2(X37,X38,bool,fun(bool,bool)))))) ) | $spl46),
  inference(superposition,[],[f262,f2975])).
fof(f7578,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,g(s(fun(bool,bool),X39))) | s(fun(fun(bool,bool),bool),X37) = s(fun(fun(bool,bool),bool),X38) | s(bool,'F_0') != s(bool,f(s(fun(bool,bool),sK2(X37,X38,bool,fun(bool,bool)))))) ) | $spl46),
  inference(superposition,[],[f262,f2975])).
fof(f6938,plain,(
  $false | ($spl89 | $spl98)),
  inference(subsumption_resolution,[],[f6937,f38])).
fof(f6937,plain,(
  ~p(s(bool,'T_0')) | ($spl89 | $spl98)),
  inference(forward_demodulation,[],[f40,f6920])).
fof(f6920,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1)) ) | ($spl89 | $spl98)),
  inference(duplicate_literal_removal,[],[f6919])).
fof(f6919,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl89 | $spl98)),
  inference(condensation,[],[f6918])).
fof(f6918,plain,(
  ( ! [X68,X69,X67] : (s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X69) | s(bool,X68) = s(bool,X69)) ) | ($spl89 | $spl98)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f97,f129,f253,f255,f294,f282,f284,f1360,f1218,f377,f482,f484,f518,f520,f844,f536,f538,f541,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1740,f2928,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f517,f481,f260,f258,f2406,f2405,f2404,f2402,f2399,f2398,f2397,f2393,f2392,f2391,f2389,f2386,f2385,f2384,f1540,f1512,f1220,f286,f272,f262,f257,f40,f3527,f3528,f3529,f3530,f3531,f3534,f3536,f3537,f3539,f2956,f2955,f2954,f2953,f2952,f2951,f2950,f2949,f2948,f2947,f2946,f2945,f3538,f3535,f3533,f3532,f3192,f3189,f5404,f5399,f6790,f6791,f6818,f6914,f557,f535])).
fof(f6914,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl89),
  inference(duplicate_literal_removal,[],[f6913])).
fof(f6913,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl89),
  inference(condensation,[],[f6912])).
fof(f6912,plain,(
  ( ! [X8,X7,X9] : (s(bool,X8) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl89),
  inference(subsumption_resolution,[],[f6811,f59])).
fof(f6811,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') != s(bool,X7) | s(bool,X8) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl89),
  inference(superposition,[],[f5399,f97])).
fof(f6818,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) = s(bool,X20) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,X19) = s(bool,f(s(bool,'F_0')))) ) | $spl89),
  inference(trivial_inequality_removal,[],[f6815])).
fof(f6815,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X19) = s(bool,X20) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,X19) = s(bool,f(s(bool,'F_0')))) ) | $spl89),
  inference(superposition,[],[f5399,f97])).
fof(f6790,plain,(
  ~p(s(bool,f(s(bool,'F_0')))) | $spl89),
  inference(unit_resulting_resolution,[],[f5399,f5399,f5399,f844])).
fof(f6929,plain,(
  $false | ($spl89 | $spl98)),
  inference(subsumption_resolution,[],[f5399,f6920])).
fof(f6911,plain,(
  $false | ($spl89 | $spl90)),
  inference(subsumption_resolution,[],[f6884,f6854])).
fof(f6854,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl89 | $spl90)),
  inference(subsumption_resolution,[],[f6821,f5400])).
fof(f6884,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl89 | $spl90)),
  inference(backward_demodulation,[],[f6854,f5399])).
fof(f6909,plain,(
  $false | ($spl89 | $spl90)),
  inference(subsumption_resolution,[],[f6855,f38])).
fof(f6855,plain,(
  ~p(s(bool,'T_0')) | ($spl89 | $spl90)),
  inference(backward_demodulation,[],[f6854,f40])).
fof(f2974,plain,(
  ( ! [X22] : (s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X22)) ) | $spl44),
  inference(cnf_transformation,[],[f2974_D])).
fof(f2974_D,plain,(
  ( ! [X22] : (s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X22)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f5329,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl81),
  inference(cnf_transformation,[],[f5329_D])).
fof(f5329_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f6629,plain,(
  ( ! [X34] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,g(s(bool,'F_0')))) ) | ($spl10 | $spl48)),
  inference(superposition,[],[f2976,f1546])).
fof(f1546,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,g(s(bool,'F_0'))) | $spl10),
  inference(cnf_transformation,[],[f1546_D])).
fof(f1546_D,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,g(s(bool,'F_0'))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1553,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl18),
  inference(cnf_transformation,[],[f1553_D])).
fof(f1553_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f6641,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X18) | s(bool,X17) = s(bool,X18) | s(bool,X16) = s(bool,X17)) ) | $spl130),
  inference(cnf_transformation,[],[f6641_D])).
fof(f6641_D,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X18) | s(bool,X17) = s(bool,X18) | s(bool,X16) = s(bool,X17)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f6625,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,f(s(bool,'F_0'))) | s(bool,X25) = s(bool,X26) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl10),
  inference(superposition,[],[f97,f1546])).
fof(f5403,plain,(
  ( ! [X14,X15] : (s(bool,X14) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X15) | s(bool,X14) = s(bool,X15)) ) | $spl96),
  inference(cnf_transformation,[],[f5403_D])).
fof(f5403_D,plain,(
  ( ! [X14,X15] : (s(bool,X14) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X15) | s(bool,X14) = s(bool,X15)) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f6639,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl126),
  inference(cnf_transformation,[],[f6639_D])).
fof(f6639_D,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f6624,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) = s(bool,f(s(bool,'F_0'))) | s(bool,X23) = s(bool,g(s(bool,'F_0'))) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X22)) ) | $spl10),
  inference(superposition,[],[f97,f1546])).
fof(f5402,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,X10) = s(bool,X12) | s(bool,X11) = s(bool,g(s(bool,'F_0')))) ) | $spl94),
  inference(cnf_transformation,[],[f5402_D])).
fof(f5402_D,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) = s(bool,X11) | s(bool,'T_0') = s(bool,X12) | s(bool,X10) = s(bool,X12) | s(bool,X11) = s(bool,g(s(bool,'F_0')))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f5398,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | $spl88),
  inference(cnf_transformation,[],[f5398_D])).
fof(f5398_D,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f6623,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X19) = s(bool,X20) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,X19) = s(bool,g(s(bool,'F_0')))) ) | $spl10),
  inference(superposition,[],[f97,f1546])).
fof(f6622,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) = s(bool,f(s(bool,'F_0'))) | s(bool,X16) = s(bool,X17) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) ) | $spl10),
  inference(superposition,[],[f97,f1546])).
fof(f6621,plain,(
  ( ! [X14,X15,X13] : (s(bool,X13) = s(bool,f(s(bool,'F_0'))) | s(bool,X14) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl10),
  inference(superposition,[],[f97,f1546])).
fof(f6620,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X10) = s(bool,X11) | s(bool,X11) = s(bool,g(s(bool,'F_0'))) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl10),
  inference(superposition,[],[f97,f1546])).
fof(f6640,plain,(
  ( ! [X4,X5] : (s(bool,X4) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5) | s(bool,X4) = s(bool,X5)) ) | $spl128),
  inference(cnf_transformation,[],[f6640_D])).
fof(f6640_D,plain,(
  ( ! [X4,X5] : (s(bool,X4) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5) | s(bool,X4) = s(bool,X5)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f6619,plain,(
  ( ! [X8,X7,X9] : (s(bool,X7) = s(bool,f(s(bool,'F_0'))) | s(bool,X8) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(superposition,[],[f97,f1546])).
fof(f6618,plain,(
  ( ! [X6,X4,X5] : (s(bool,X4) = s(bool,f(s(bool,'F_0'))) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(superposition,[],[f97,f1546])).
fof(f6616,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X3) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f59,f1546])).
fof(f6615,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f59,f1546])).
fof(f6614,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl10),
  inference(superposition,[],[f59,f1546])).
fof(f6613,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X0) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f59,f1546])).
fof(f5328,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | $spl80),
  inference(cnf_transformation,[],[f5328_D])).
fof(f5328_D,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f6612,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl10),
  inference(superposition,[],[f43,f1546])).
fof(f6611,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) = s(bool,f(s(bool,'F_0'))) | s(bool,X25) = s(bool,X26) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl10),
  inference(superposition,[],[f1546,f97])).
fof(f6610,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) = s(bool,f(s(bool,'F_0'))) | s(bool,X23) = s(bool,g(s(bool,'F_0'))) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X22)) ) | $spl10),
  inference(superposition,[],[f1546,f97])).
fof(f6609,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X19) = s(bool,X20) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,X19) = s(bool,g(s(bool,'F_0')))) ) | $spl10),
  inference(superposition,[],[f1546,f97])).
fof(f6608,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) = s(bool,f(s(bool,'F_0'))) | s(bool,X16) = s(bool,X17) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) ) | $spl10),
  inference(superposition,[],[f1546,f97])).
fof(f6607,plain,(
  ( ! [X14,X15,X13] : (s(bool,X13) = s(bool,f(s(bool,'F_0'))) | s(bool,X14) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl10),
  inference(superposition,[],[f1546,f97])).
fof(f6606,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X10) = s(bool,X11) | s(bool,X11) = s(bool,g(s(bool,'F_0'))) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl10),
  inference(superposition,[],[f1546,f97])).
fof(f6605,plain,(
  ( ! [X8,X7,X9] : (s(bool,X7) = s(bool,f(s(bool,'F_0'))) | s(bool,X8) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(superposition,[],[f1546,f97])).
fof(f6604,plain,(
  ( ! [X6,X4,X5] : (s(bool,X4) = s(bool,f(s(bool,'F_0'))) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl10),
  inference(superposition,[],[f1546,f97])).
fof(f6603,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X3) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f1546,f59])).
fof(f6602,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl10),
  inference(superposition,[],[f1546,f59])).
fof(f6601,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl10),
  inference(superposition,[],[f1546,f59])).
fof(f6600,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,f(s(bool,'F_0'))) | s(bool,X0) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(superposition,[],[f1546,f59])).
fof(f6599,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl10),
  inference(superposition,[],[f1546,f43])).
fof(f6368,plain,(
  $false | ($spl9 | $spl14 | $spl51 | $spl113)),
  inference(subsumption_resolution,[],[f6367,f5763])).
fof(f6367,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl14 | $spl113)),
  inference(forward_demodulation,[],[f5499,f1550])).
fof(f6366,plain,(
  $false | ($spl9 | $spl10 | $spl14 | $spl51 | $spl81 | $spl113)),
  inference(subsumption_resolution,[],[f6365,f5763])).
fof(f6365,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl9 | $spl10 | $spl14 | $spl81 | $spl113)),
  inference(forward_demodulation,[],[f5499,f6364])).
fof(f6364,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl9 | $spl10 | $spl14 | $spl81)),
  inference(global_subsumption,[],[f6363,f37,f36,f35,f38,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f97,f129,f253,f255,f294,f282,f284,f1360,f1218,f377,f482,f484,f518,f520,f844,f536,f538,f541,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1550,f1560,f1562,f1624,f1740,f2337,f1841,f2327,f2323,f2311,f2293,f2275,f2268,f2267,f2250,f1913,f1914,f2219,f2218,f1966,f2163,f2158,f1999,f2000,f2002,f2157,f2152,f2139,f2036,f2037,f2039,f2041,f2046,f2082,f2083,f2085,f2089,f2103,f2928,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f557,f535,f517,f481,f260,f258,f2406,f2405,f2404,f2402,f2399,f2398,f2397,f2393,f2392,f2391,f2389,f2386,f2385,f2384,f1540,f1512,f1220,f286,f272,f262,f257,f40,f3527,f3528,f3529,f3530,f3531,f3534,f3536,f3537,f3539,f2956,f2955,f2954,f2953,f2952,f2951,f2950,f2949,f2948,f2947,f2946,f2945,f3538,f3535,f3533,f3532,f3192,f3189])).
fof(f2103,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f2095])).
fof(f2095,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(equality_factoring,[],[f1562])).
fof(f2089,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl14),
  inference(equality_factoring,[],[f1562])).
fof(f2085,plain,(
  ( ! [X101,X99,X100,X98] : (s(X100,f(s(bool,'T_0'))) != s(X100,g(s(bool,'T_0'))) | s(fun(bool,X100),X98) = s(fun(bool,X100),X99) | s(bool,X101) = s(bool,g(s(bool,X101))) | s(bool,X101) = s(bool,sK2(X98,X99,X100,bool))) ) | $spl14),
  inference(superposition,[],[f48,f1562])).
fof(f2083,plain,(
  ( ! [X92,X90,X93,X91] : (s(bool,'T_0') != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X90,bool),X92))) | s('type.pair.prod'(X90,bool),X91) = s('type.pair.prod'(X90,bool),X92) | s(X90,'const.pair.FST_1'(s('type.pair.prod'(X90,bool),X91))) != s(X90,'const.pair.FST_1'(s('type.pair.prod'(X90,bool),X92))) | s(bool,X93) = s(bool,g(s(bool,X93))) | s(bool,X93) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X90,bool),X91)))) ) | $spl14),
  inference(superposition,[],[f51,f1562])).
fof(f2082,plain,(
  ( ! [X88,X87,X89,X86] : (s(bool,'T_0') != s(bool,f(s(X86,sK2(X87,X88,bool,X86)))) | s(fun(X86,bool),X87) = s(fun(X86,bool),X88) | s(bool,X89) = s(bool,g(s(bool,X89))) | s(bool,X89) = s(bool,g(s(X86,sK2(X87,X88,bool,X86))))) ) | $spl14),
  inference(superposition,[],[f48,f1562])).
fof(f2046,plain,(
  ( ! [X24,X23] : (s(bool,X23) = s(bool,g(s(bool,X23))) | s(bool,X24) = s(bool,g(s(bool,X24))) | s(bool,X23) = s(bool,X24)) ) | $spl14),
  inference(superposition,[],[f1550,f1562])).
fof(f2041,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,X0) = s(bool,X1)) ) | $spl14),
  inference(superposition,[],[f38,f1562])).
fof(f2039,plain,(
  ( ! [X118,X120,X119,X117] : (s(X119,f(s(bool,X120))) != s(X119,g(s(bool,X120))) | s(fun(bool,X119),X117) = s(fun(bool,X119),X118) | s(bool,X120) = s(bool,g(s(bool,X120))) | s(bool,'T_0') = s(bool,sK2(X117,X118,X119,bool))) ) | $spl14),
  inference(superposition,[],[f48,f1562])).
fof(f2037,plain,(
  ( ! [X111,X109,X112,X110] : (s(bool,X111) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X109,bool),X112))) | s('type.pair.prod'(X109,bool),X110) = s('type.pair.prod'(X109,bool),X112) | s(X109,'const.pair.FST_1'(s('type.pair.prod'(X109,bool),X110))) != s(X109,'const.pair.FST_1'(s('type.pair.prod'(X109,bool),X112))) | s(bool,X111) = s(bool,g(s(bool,X111))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X109,bool),X110)))) ) | $spl14),
  inference(superposition,[],[f51,f1562])).
fof(f2036,plain,(
  ( ! [X107,X105,X108,X106] : (s(bool,X108) != s(bool,f(s(X105,sK2(X106,X107,bool,X105)))) | s(fun(X105,bool),X106) = s(fun(X105,bool),X107) | s(bool,X108) = s(bool,g(s(bool,X108))) | s(bool,'T_0') = s(bool,g(s(X105,sK2(X106,X107,bool,X105))))) ) | $spl14),
  inference(superposition,[],[f48,f1562])).
fof(f2002,plain,(
  ( ! [X118,X120,X119,X117] : (s(X119,f(s(bool,X120))) != s(X119,g(s(bool,X120))) | s(fun(bool,X119),X117) = s(fun(bool,X119),X118) | s(bool,sK2(X117,X118,X119,bool)) = s(bool,g(s(bool,sK2(X117,X118,X119,bool)))) | s(bool,'T_0') = s(bool,X120)) ) | $spl14),
  inference(superposition,[],[f48,f1562])).
fof(f2000,plain,(
  ( ! [X111,X109,X112,X110] : (s(bool,X111) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X109,bool),X112))) | s('type.pair.prod'(X109,bool),X110) = s('type.pair.prod'(X109,bool),X112) | s(X109,'const.pair.FST_1'(s('type.pair.prod'(X109,bool),X110))) != s(X109,'const.pair.FST_1'(s('type.pair.prod'(X109,bool),X112))) | s(bool,'const.pair.SND_1'(s('type.pair.prod'(X109,bool),X110))) = s(bool,g(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X109,bool),X110))))) | s(bool,'T_0') = s(bool,X111)) ) | $spl14),
  inference(superposition,[],[f51,f1562])).
fof(f1999,plain,(
  ( ! [X107,X105,X108,X106] : (s(bool,X108) != s(bool,f(s(X105,sK2(X106,X107,bool,X105)))) | s(fun(X105,bool),X106) = s(fun(X105,bool),X107) | s(bool,g(s(X105,sK2(X106,X107,bool,X105)))) = s(bool,g(s(bool,g(s(X105,sK2(X106,X107,bool,X105)))))) | s(bool,'T_0') = s(bool,X108)) ) | $spl14),
  inference(superposition,[],[f48,f1562])).
fof(f1966,plain,(
  ( ! [X72,X71,X73] : (s(bool,sK2(X71,X72,bool,bool)) != s(bool,f(s(bool,sK2(X71,X72,bool,bool)))) | s(fun(bool,bool),X71) = s(fun(bool,bool),X72) | s(bool,X73) = s(bool,sK2(X71,X72,bool,bool)) | s(bool,'T_0') = s(bool,X73)) ) | $spl14),
  inference(superposition,[],[f48,f1562])).
fof(f1914,plain,(
  ( ! [X6,X4,X8,X7,X5,X9] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,g(s(bool,X6))) | s(bool,X5) = s(bool,X6) | s(bool,X7) = s(bool,X8) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9) | s(bool,X4) = s(bool,X7)) ) | $spl14),
  inference(superposition,[],[f1562,f97])).
fof(f1913,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,g(s(bool,X2))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,g(s(bool,X3))) | s(bool,X0) = s(bool,X3)) ) | $spl14),
  inference(superposition,[],[f1562,f1562])).
fof(f2311,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,g(s(bool,X0))) = s(bool,g(s(bool,g(s(bool,X0)))))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f2310])).
fof(f2310,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,g(s(bool,X0))) = s(bool,g(s(bool,g(s(bool,X0))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl14),
  inference(condensation,[],[f1856])).
fof(f1856,plain,(
  ( ! [X68,X69,X67] : (s(bool,X67) = s(bool,X68) | s(bool,X67) = s(bool,X69) | s(bool,'T_0') = s(bool,X69) | s(bool,g(s(bool,X67))) = s(bool,g(s(bool,g(s(bool,X67))))) | s(bool,'T_0') = s(bool,X68)) ) | $spl14),
  inference(superposition,[],[f1562,f1562])).
fof(f2323,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,g(s(bool,X0))) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f2322])).
fof(f2322,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X1) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl14),
  inference(condensation,[],[f1849])).
fof(f1849,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,X36) = s(bool,X37) | s(bool,X36) = s(bool,X38) | s(bool,'T_0') = s(bool,X38) | s(bool,X39) = s(bool,g(s(bool,X36))) | s(bool,'T_0') = s(bool,X37) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | $spl14),
  inference(superposition,[],[f1562,f97])).
fof(f2327,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,X0)))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f2326])).
fof(f2326,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,X0)))) ) | $spl14),
  inference(condensation,[],[f1843])).
fof(f1843,plain,(
  ( ! [X14,X12,X13] : (s(bool,X12) = s(bool,X13) | s(bool,X12) = s(bool,X14) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,g(s(bool,X12)))) ) | $spl14),
  inference(superposition,[],[f1562,f59])).
fof(f1841,plain,(
  ( ! [X8,X7] : (s(bool,'F_0') = s(bool,X7) | s(bool,X7) = s(bool,X8) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,g(s(bool,X7)))) ) | $spl14),
  inference(superposition,[],[f1562,f43])).
fof(f1624,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,g(s(bool,'T_0'))) | s(bool,X0) = s(bool,X1)) ) | $spl14),
  inference(condensation,[],[f1623])).
fof(f1623,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,X8) | s(bool,X9) = s(bool,g(s(bool,'T_0'))) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f1581])).
fof(f1581,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,X8) | s(bool,X9) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X8) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | $spl14),
  inference(superposition,[],[f1550,f97])).
fof(f1562,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,g(s(bool,X8))) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl14),
  inference(superposition,[],[f1550,f59])).
fof(f1560,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,g(s(bool,X0))) | s(bool,X1) = s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X0) = s(bool,X1)) ) | $spl14),
  inference(superposition,[],[f1550,f97])).
fof(f6363,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool))))) ) | ($spl9 | $spl10 | $spl81)),
  inference(subsumption_resolution,[],[f6362,f6014])).
fof(f6362,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool))))) ) | ($spl9 | $spl10 | $spl81)),
  inference(forward_demodulation,[],[f5977,f6045])).
fof(f6045,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | ($spl9 | $spl10 | $spl81)),
  inference(backward_demodulation,[],[f6014,f1546])).
fof(f5977,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool)))) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f5903])).
fof(f5903,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f262,f257])).
fof(f6358,plain,(
  ( ! [X35,X31,X34,X32] : (s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,X35) | s(bool,X34) = s(bool,sK2(X31,X32,bool,bool))) ) | $spl124),
  inference(cnf_transformation,[],[f6358_D])).
fof(f6358_D,plain,(
  ( ! [X35,X31,X34,X32] : (s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,X35) | s(bool,X34) = s(bool,sK2(X31,X32,bool,bool))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f5338,plain,(
  ( ! [X48] : (s(bool,'F_0') != s(bool,f(s(bool,X48))) | s(bool,'T_0') = s(bool,X48)) ) | $spl86),
  inference(cnf_transformation,[],[f5338_D])).
fof(f5338_D,plain,(
  ( ! [X48] : (s(bool,'F_0') != s(bool,f(s(bool,X48))) | s(bool,'T_0') = s(bool,X48)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f6339,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'F_0') != s(bool,f(s(bool,X48))) | s(fun(bool,bool),X46) = s(fun(bool,bool),X47) | s(bool,X49) = s(bool,sK2(X46,X47,bool,bool)) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X48)) ) | $spl12),
  inference(superposition,[],[f1548,f97])).
fof(f6336,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s(bool,'F_0') != s(bool,f(s(bool,X33))) | s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(bool,X34) = s(bool,sK2(X31,X32,bool,bool)) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl12),
  inference(superposition,[],[f1548,f97])).
fof(f5510,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool))))) ) | $spl122),
  inference(cnf_transformation,[],[f5510_D])).
fof(f5510_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f2978,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl51),
  inference(cnf_transformation,[],[f2978_D])).
fof(f2978_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f5509,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool))))) ) | $spl10),
  inference(subsumption_resolution,[],[f3208,f1546])).
fof(f3208,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool)))) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f3181])).
fof(f3181,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1512,f257])).
fof(f5433,plain,(
  ( ! [X61,X59,X60] : (s(fun(bool,bool),X59) = s(fun(bool,bool),X60) | s(bool,'T_0') = s(bool,X61) | s(bool,X61) = s(bool,sK2(X59,X60,bool,bool))) ) | $spl102),
  inference(cnf_transformation,[],[f5433_D])).
fof(f5433_D,plain,(
  ( ! [X61,X59,X60] : (s(fun(bool,bool),X59) = s(fun(bool,bool),X60) | s(bool,'T_0') = s(bool,X61) | s(bool,X61) = s(bool,sK2(X59,X60,bool,bool))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f5508,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,X6) = s(bool,sK2(X4,X5,bool,bool)) | s(bool,'T_0') = s(bool,X6)) ) | $spl20),
  inference(subsumption_resolution,[],[f3183,f4214])).
fof(f4214,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,X5) | s(bool,X5) = s(bool,sK2(X3,X4,bool,bool)) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,f(s(bool,sK2(X3,X4,bool,bool))))) ) | $spl20),
  inference(trivial_inequality_removal,[],[f4199])).
fof(f4199,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X5) | s(bool,X5) = s(bool,sK2(X3,X4,bool,bool)) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,f(s(bool,sK2(X3,X4,bool,bool))))) ) | $spl20),
  inference(superposition,[],[f1554,f43])).
fof(f3183,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X4,X5,bool,bool)))) | s(bool,X6) = s(bool,sK2(X4,X5,bool,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1512,f59])).
fof(f5507,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X13) = s(fun(bool,bool),X14) | s(bool,X15) = s(bool,sK2(X13,X14,bool,bool)) | s(bool,'T_0') = s(bool,X15)) ) | $spl20),
  inference(subsumption_resolution,[],[f3186,f4214])).
fof(f3186,plain,(
  ( ! [X14,X15,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X13) = s(fun(bool,bool),X14) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X13,X14,bool,bool)))) | s(bool,X15) = s(bool,sK2(X13,X14,bool,bool)) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f1512,f59])).
fof(f5500,plain,(
  ( ! [X4,X1] : (s(bool,'F_0') != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X4) | s(bool,X1) = s(bool,X4)) ) | $spl114),
  inference(cnf_transformation,[],[f5500_D])).
fof(f5500_D,plain,(
  ( ! [X4,X1] : (s(bool,'F_0') != s(bool,f(s(bool,X1))) | s(bool,'T_0') = s(bool,X4) | s(bool,X1) = s(bool,X4)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f5506,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'F_0') != s(bool,f(s(bool,X18))) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17) | s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,sK2(X16,X17,bool,bool)) | s(bool,'T_0') = s(bool,X20)) ) | $spl20),
  inference(subsumption_resolution,[],[f3187,f4214])).
fof(f3187,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'F_0') != s(bool,f(s(bool,X18))) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X16,X17,bool,bool)))) | s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,sK2(X16,X17,bool,bool)) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f1512,f97])).
fof(f5505,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,'F_0') != s(bool,f(s(bool,X23))) | s(fun(bool,bool),X21) = s(fun(bool,bool),X22) | s(bool,X24) = s(bool,sK2(X21,X22,bool,bool)) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) ) | $spl20),
  inference(subsumption_resolution,[],[f3188,f4214])).
fof(f3188,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,'F_0') != s(bool,f(s(bool,X23))) | s(fun(bool,bool),X21) = s(fun(bool,bool),X22) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X21,X22,bool,bool)))) | s(bool,X24) = s(bool,sK2(X21,X22,bool,bool)) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f1512,f97])).
fof(f5504,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(fun(bool,bool),X41) = s(fun(bool,bool),X42) | s(bool,X43) = s(bool,sK2(X41,X42,bool,bool)) | s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X45) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X41,X42,bool,bool))))) ) | $spl120),
  inference(cnf_transformation,[],[f5504_D])).
fof(f5504_D,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(fun(bool,bool),X41) = s(fun(bool,bool),X42) | s(bool,X43) = s(bool,sK2(X41,X42,bool,bool)) | s(bool,X43) = s(bool,X44) | s(bool,'T_0') = s(bool,X45) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X41,X42,bool,bool))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f3189,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X26,X27,bool,bool)))) | s(bool,X28) = s(bool,X29) | s(bool,X29) = s(bool,sK2(X26,X27,bool,bool)) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f1512,f97])).
fof(f3192,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X41) = s(fun(bool,bool),X42) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X41,X42,bool,bool)))) | s(bool,X43) = s(bool,X44) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,X43) = s(bool,sK2(X41,X42,bool,bool))) )),
  inference(superposition,[],[f1512,f97])).
fof(f5436,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X7)) ) | $spl104),
  inference(cnf_transformation,[],[f5436_D])).
fof(f5436_D,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X7)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f5503,plain,(
  ( ! [X35,X40] : (s(bool,X35) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X40) | s(bool,X35) = s(bool,X40)) ) | $spl118),
  inference(cnf_transformation,[],[f5503_D])).
fof(f5503_D,plain,(
  ( ! [X35,X40] : (s(bool,X35) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X40) | s(bool,X35) = s(bool,X40)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f3532,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s(bool,X29) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X30) = s(fun(bool,bool),X31) | s(bool,X32) = s(bool,sK2(X30,X31,bool,bool)) | s(bool,'T_0') = s(bool,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f253,f97])).
fof(f3533,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s(bool,X35) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X36) = s(fun(bool,bool),X37) | s(bool,X38) = s(bool,sK2(X36,X37,bool,bool)) | s(bool,'T_0') = s(bool,X38) | s(bool,X39) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f253,f97])).
fof(f5502,plain,(
  ( ! [X70,X69] : (s(bool,X69) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X70) | s(bool,X69) = s(bool,X70)) ) | $spl116),
  inference(cnf_transformation,[],[f5502_D])).
fof(f5502_D,plain,(
  ( ! [X70,X69] : (s(bool,X69) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X70) | s(bool,X69) = s(bool,X70)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f5493,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,X17) | s(bool,X17) != s(bool,f(s(bool,'T_0')))) ) | $spl108),
  inference(cnf_transformation,[],[f5493_D])).
fof(f5493_D,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,X17) | s(bool,X17) != s(bool,f(s(bool,'T_0')))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f3535,plain,(
  ( ! [X47,X52,X50,X48,X51,X49] : (s(bool,X47) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X48) = s(fun(bool,bool),X49) | s(bool,X50) = s(bool,sK2(X48,X49,bool,bool)) | s(bool,'T_0') = s(bool,X50) | s(bool,X51) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X47) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f253,f97])).
fof(f3538,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (s(bool,X65) != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X66) = s(fun(bool,bool),X67) | s(bool,X68) = s(bool,sK2(X66,X67,bool,bool)) | s(bool,'T_0') = s(bool,X68) | s(bool,X69) = s(bool,g(s(bool,'T_0'))) | s(bool,X69) = s(bool,X70) | s(bool,'T_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f253,f97])).
fof(f5501,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X5) | s(bool,X5) = s(bool,sK2(X3,X4,bool,bool)) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4)) ) | ($spl10 | $spl20)),
  inference(subsumption_resolution,[],[f4229,f1546])).
fof(f4229,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X5) | s(bool,X5) = s(bool,sK2(X3,X4,bool,bool)) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) ) | $spl20),
  inference(duplicate_literal_removal,[],[f4182])).
fof(f4182,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X5) | s(bool,X5) = s(bool,sK2(X3,X4,bool,bool)) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) ) | $spl20),
  inference(superposition,[],[f1554,f257])).
fof(f4228,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,sK2(X1,X2,bool,bool)) | s(fun(bool,bool),X1) = s(fun(bool,bool),X2)) ) | $spl20),
  inference(duplicate_literal_removal,[],[f4227])).
fof(f4227,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,sK2(X1,X2,bool,bool)) | s(fun(bool,bool),X1) = s(fun(bool,bool),X2) | s(bool,X0) = s(bool,sK2(X1,X2,bool,bool))) ) | $spl20),
  inference(condensation,[],[f4184])).
fof(f4184,plain,(
  ( ! [X12,X10,X11,X9] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X11) | s(bool,X11) = s(bool,sK2(X9,X10,bool,bool)) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,X12) = s(bool,sK2(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,X12)) ) | $spl20),
  inference(superposition,[],[f1554,f59])).
fof(f4226,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,sK2(X1,X2,bool,bool)) | s(fun(bool,bool),X1) = s(fun(bool,bool),X2)) ) | $spl20),
  inference(duplicate_literal_removal,[],[f4225])).
fof(f4225,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,sK2(X1,X2,bool,bool)) | s(fun(bool,bool),X1) = s(fun(bool,bool),X2) | s(bool,X0) = s(bool,sK2(X1,X2,bool,bool))) ) | $spl20),
  inference(condensation,[],[f4187])).
fof(f4187,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X23) | s(bool,X23) = s(bool,sK2(X21,X22,bool,bool)) | s(fun(bool,bool),X21) = s(fun(bool,bool),X22) | s(bool,X24) = s(bool,sK2(X21,X22,bool,bool)) | s(bool,'T_0') = s(bool,X24)) ) | $spl20),
  inference(superposition,[],[f1554,f59])).
fof(f4224,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,sK2(X2,X3,bool,bool)) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,X1) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl20),
  inference(duplicate_literal_removal,[],[f4223])).
fof(f4223,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,sK2(X2,X3,bool,bool)) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,X1) = s(bool,X4) | s(bool,'T_0') = s(bool,X4) | s(bool,X0) = s(bool,sK2(X2,X3,bool,bool))) ) | $spl20),
  inference(condensation,[],[f4188])).
fof(f4188,plain,(
  ( ! [X30,X28,X26,X29,X27,X25] : (s(bool,'F_0') != s(bool,f(s(bool,X27))) | s(bool,'T_0') = s(bool,X28) | s(bool,X28) = s(bool,sK2(X25,X26,bool,bool)) | s(fun(bool,bool),X25) = s(fun(bool,bool),X26) | s(bool,X27) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,X30) = s(bool,sK2(X25,X26,bool,bool)) | s(bool,'T_0') = s(bool,X30)) ) | $spl20),
  inference(superposition,[],[f1554,f97])).
fof(f4222,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,sK2(X2,X3,bool,bool)) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,X1) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl20),
  inference(duplicate_literal_removal,[],[f4221])).
fof(f4221,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,X1))) | s(bool,X0) = s(bool,sK2(X2,X3,bool,bool)) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,X0) = s(bool,sK2(X2,X3,bool,bool)) | s(bool,X1) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl20),
  inference(condensation,[],[f4189])).
fof(f4189,plain,(
  ( ! [X35,X33,X31,X36,X34,X32] : (s(bool,'F_0') != s(bool,f(s(bool,X33))) | s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,sK2(X31,X32,bool,bool)) | s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(bool,X35) = s(bool,sK2(X31,X32,bool,bool)) | s(bool,'T_0') = s(bool,X35) | s(bool,X33) = s(bool,X36) | s(bool,'T_0') = s(bool,X36)) ) | $spl20),
  inference(superposition,[],[f1554,f97])).
fof(f4220,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,sK2(X1,X2,bool,bool)) | s(fun(bool,bool),X1) = s(fun(bool,bool),X2)) ) | $spl20),
  inference(condensation,[],[f4219])).
fof(f4219,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,sK2(X1,X2,bool,bool)) | s(fun(bool,bool),X1) = s(fun(bool,bool),X2) | s(bool,'T_0') = s(bool,X3)) ) | $spl20),
  inference(duplicate_literal_removal,[],[f4218])).
fof(f4218,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,sK2(X1,X2,bool,bool)) | s(fun(bool,bool),X1) = s(fun(bool,bool),X2) | s(bool,X0) = s(bool,sK2(X1,X2,bool,bool)) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl20),
  inference(condensation,[],[f4190])).
fof(f4190,plain,(
  ( ! [X39,X37,X41,X38,X42,X40] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X39) | s(bool,X39) = s(bool,sK2(X37,X38,bool,bool)) | s(fun(bool,bool),X37) = s(fun(bool,bool),X38) | s(bool,X40) = s(bool,X41) | s(bool,X41) = s(bool,sK2(X37,X38,bool,bool)) | s(bool,X40) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) ) | $spl20),
  inference(superposition,[],[f1554,f97])).
fof(f4217,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X0) = s(bool,sK2(X1,X2,bool,bool)) | s(fun(bool,bool),X1) = s(fun(bool,bool),X2)) ) | $spl20),
  inference(condensation,[],[f4216])).
fof(f4216,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X4) = s(bool,sK2(X5,X6,bool,bool)) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,X7)) ) | $spl20),
  inference(duplicate_literal_removal,[],[f4215])).
fof(f4215,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'T_0') = s(bool,X4) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,X4) = s(bool,sK2(X5,X6,bool,bool)) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X7) | s(bool,X4) = s(bool,sK2(X5,X6,bool,bool))) ) | $spl20),
  inference(condensation,[],[f4193])).
fof(f4193,plain,(
  ( ! [X59,X57,X60,X58,X56,X55] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X57) | s(bool,X57) = s(bool,sK2(X55,X56,bool,bool)) | s(fun(bool,bool),X55) = s(fun(bool,bool),X56) | s(bool,X58) = s(bool,X59) | s(bool,X59) = s(bool,X60) | s(bool,'T_0') = s(bool,X60) | s(bool,X58) = s(bool,sK2(X55,X56,bool,bool))) ) | $spl20),
  inference(superposition,[],[f1554,f97])).
fof(f5499,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | $spl113),
  inference(cnf_transformation,[],[f5499_D])).
fof(f5499_D,plain,(
  s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f2975,plain,(
  ( ! [X19,X20] : (s(fun(bool,bool),X19) = s(fun(bool,bool),X20)) ) | $spl46),
  inference(cnf_transformation,[],[f2975_D])).
fof(f2975_D,plain,(
  ( ! [X19,X20] : (s(fun(bool,bool),X19) = s(fun(bool,bool),X20)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f5399,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl89),
  inference(cnf_transformation,[],[f5399_D])).
fof(f5399_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f2945,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,X8) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f272,f59])).
fof(f5400,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1)) ) | $spl90),
  inference(cnf_transformation,[],[f5400_D])).
fof(f5400_D,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1)) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f2946,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f272,f59])).
fof(f2947,plain,(
  ( ! [X14,X12,X13] : (s(bool,X12) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X13) = s(fun(bool,bool),X14) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f272,f59])).
fof(f2948,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X15) = s(fun(bool,bool),X16) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,X17) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f272,f59])).
fof(f5401,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5) | s(bool,X4) = s(bool,X5)) ) | $spl92),
  inference(cnf_transformation,[],[f5401_D])).
fof(f5401_D,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5) | s(bool,X4) = s(bool,X5)) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f2949,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,X18) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,X18) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f272,f97])).
fof(f2950,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(bool,X23) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X24) = s(fun(bool,bool),X25) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,X26) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X26) | s(bool,X23) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f272,f97])).
fof(f2951,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X28) = s(fun(bool,bool),X29) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,X30) = s(bool,X31) | s(bool,X31) = s(bool,g(s(bool,'F_0'))) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f272,f97])).
fof(f2952,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,X33) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X34) = s(fun(bool,bool),X35) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,X36) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X33) | s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f272,f97])).
fof(f5404,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X18) | s(bool,X17) = s(bool,X18) | s(bool,X16) = s(bool,X17)) ) | $spl98),
  inference(cnf_transformation,[],[f5404_D])).
fof(f5404_D,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X18) | s(bool,X17) = s(bool,X18) | s(bool,X16) = s(bool,X17)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f2953,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(bool,X38) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X39) = s(fun(bool,bool),X40) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f272,f97])).
fof(f2954,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X43) = s(fun(bool,bool),X44) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,X45) = s(bool,X46) | s(bool,X46) = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,X45) = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f272,f97])).
fof(f2955,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,X48) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X49) = s(fun(bool,bool),X50) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,X51) = s(bool,g(s(bool,'F_0'))) | s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X52) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f272,f97])).
fof(f2956,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X53) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X54) = s(fun(bool,bool),X55) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,X53) = s(bool,X56) | s(bool,X56) = s(bool,X57) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f272,f97])).
fof(f5496,plain,(
  ( ! [X57,X58,X53] : (s(bool,X53) != s(bool,f(s(bool,'T_0'))) | s(bool,X57) = s(bool,X58) | s(bool,'T_0') = s(bool,X58) | s(bool,X53) = s(bool,X57)) ) | $spl110),
  inference(cnf_transformation,[],[f5496_D])).
fof(f5496_D,plain,(
  ( ! [X57,X58,X53] : (s(bool,X53) != s(bool,f(s(bool,'T_0'))) | s(bool,X57) = s(bool,X58) | s(bool,'T_0') = s(bool,X58) | s(bool,X53) = s(bool,X57)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f5497,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (s(bool,X75) = s(bool,X76) | s(bool,'T_0') = s(bool,X76) | s(bool,'T_0') = s(bool,X74) | s(fun(bool,bool),X72) = s(fun(bool,bool),X73) | s(bool,X71) = s(bool,X75) | s(bool,X74) = s(bool,sK2(X72,X73,bool,bool)) | s(bool,X71) != s(bool,f(s(bool,'T_0')))) ) | $spl15),
  inference(global_subsumption,[],[f37,f36,f35,f38,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f97,f129,f253,f255,f294,f282,f284,f1360,f1218,f377,f482,f484,f518,f520,f844,f536,f538,f541,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1740,f1549,f2495,f2496,f2498,f2529,f2527,f2928,f3241,f3243,f3508,f4348,f4305,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f557,f535,f517,f481,f260,f258,f2406,f2405,f2404,f2402,f2399,f2398,f2397,f2393,f2392,f2391,f2389,f2386,f2385,f2384,f1540,f1512,f1220,f286,f272,f262,f257,f40,f3527,f5489,f3528,f3529,f5492,f3530,f5494,f3531,f3534,f3536,f5495,f3537,f3539])).
fof(f5432,plain,(
  ( ! [X64,X62,X63] : (s(bool,X62) = s(bool,X63) | s(bool,X62) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X64) | s(bool,X63) = s(bool,X64)) ) | $spl100),
  inference(cnf_transformation,[],[f5432_D])).
fof(f5432_D,plain,(
  ( ! [X64,X62,X63] : (s(bool,X62) = s(bool,X63) | s(bool,X62) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X64) | s(bool,X63) = s(bool,X64)) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f5491,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | $spl107),
  inference(cnf_transformation,[],[f5491_D])).
fof(f5491_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f3537,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X59) = s(fun(bool,bool),X60) | s(bool,X61) = s(bool,sK2(X59,X60,bool,bool)) | s(bool,'T_0') = s(bool,X61) | s(bool,X62) = s(bool,X63) | s(bool,X63) = s(bool,X64) | s(bool,'T_0') = s(bool,X64) | s(bool,X62) = s(bool,g(s(bool,'T_0')))) )),
  inference(superposition,[],[f253,f97])).
fof(f5495,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(bool,X56) = s(bool,sK2(X54,X55,bool,bool)) | s(bool,X53) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X56) | s(bool,X53) = s(bool,X57) | s(bool,'T_0') = s(bool,X58) | s(fun(bool,bool),X54) = s(fun(bool,bool),X55) | s(bool,X57) = s(bool,X58)) ) | $spl15),
  inference(global_subsumption,[],[f37,f36,f35,f38,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f97,f129,f253,f255,f294,f282,f284,f1360,f1218,f377,f482,f484,f518,f520,f844,f536,f538,f541,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1740,f1549,f2495,f2496,f2498,f2529,f2527,f2928,f3241,f3243,f3508,f4348,f4305,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f557,f535,f517,f481,f260,f258,f2406,f2405,f2404,f2402,f2399,f2398,f2397,f2393,f2392,f2391,f2389,f2386,f2385,f2384,f1540,f1512,f1220,f286,f272,f262,f257,f40,f3527,f5489,f3528,f3529,f5492,f3530,f5494,f3531,f3534,f3536])).
fof(f3534,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X41) = s(fun(bool,bool),X42) | s(bool,X43) = s(bool,sK2(X41,X42,bool,bool)) | s(bool,'T_0') = s(bool,X43) | s(bool,X44) = s(bool,X45) | s(bool,X45) = s(bool,g(s(bool,'T_0'))) | s(bool,X44) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f253,f97])).
fof(f3531,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X25) = s(fun(bool,bool),X26) | s(bool,X27) = s(bool,sK2(X25,X26,bool,bool)) | s(bool,'T_0') = s(bool,X27) | s(bool,X28) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f253,f59])).
fof(f5494,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,'T_0') = s(bool,X21) | s(bool,X21) != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X24) | s(fun(bool,bool),X22) = s(fun(bool,bool),X23) | s(bool,X24) = s(bool,sK2(X22,X23,bool,bool))) ) | $spl15),
  inference(global_subsumption,[],[f37,f36,f35,f38,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f97,f129,f253,f255,f294,f282,f284,f1360,f1218,f377,f482,f484,f518,f520,f844,f536,f538,f541,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1740,f1549,f2495,f2496,f2498,f2529,f2527,f2928,f3241,f3243,f3508,f4348,f4305,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f557,f535,f517,f481,f260,f258,f2406,f2405,f2404,f2402,f2399,f2398,f2397,f2393,f2392,f2391,f2389,f2386,f2385,f2384,f1540,f1512,f1220,f286,f272,f262,f257,f40,f3527,f5489,f3528,f3529,f5492,f3530])).
fof(f5492,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,X17) | s(bool,X17) != s(bool,f(s(bool,'T_0'))) | s(bool,X20) = s(bool,sK2(X18,X19,bool,bool)) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19)) ) | $spl15),
  inference(global_subsumption,[],[f37,f36,f35,f38,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f97,f129,f253,f255,f294,f282,f284,f1360,f1218,f377,f482,f484,f518,f520,f844,f536,f538,f541,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1740,f1549,f2495,f2496,f2498,f2529,f2527,f2928,f3241,f3243,f3508,f4348,f4305,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f557,f535,f517,f481,f260,f258,f2406,f2405,f2404,f2402,f2399,f2398,f2397,f2393,f2392,f2391,f2389,f2386,f2385,f2384,f1540,f1512,f1220,f286,f272,f262,f257,f40,f3527,f5489,f3528,f3529])).
fof(f3528,plain,(
  ( ! [X14,X15,X13,X16] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X13) = s(fun(bool,bool),X14) | s(bool,X15) = s(bool,sK2(X13,X14,bool,bool)) | s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f253,f59])).
fof(f5489,plain,(
  ( ! [X12,X10,X11] : (s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X12) | s(bool,X12) = s(bool,sK2(X10,X11,bool,bool))) ) | $spl15),
  inference(global_subsumption,[],[f37,f36,f35,f38,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f97,f129,f253,f255,f294,f282,f284,f1360,f1218,f377,f482,f484,f518,f520,f844,f536,f538,f541,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1740,f1549,f2495,f2496,f2498,f2529,f2527,f2928,f3241,f3243,f3508,f4348,f4305,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f557,f535,f517,f481,f260,f258,f2406,f2405,f2404,f2402,f2399,f2398,f2397,f2393,f2392,f2391,f2389,f2386,f2385,f2384,f1540,f1512,f1220,f286,f272,f262,f257,f40,f3527])).
fof(f5487,plain,(
  $false | ($spl32 | $spl81)),
  inference(subsumption_resolution,[],[f5470,f2363])).
fof(f5470,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl32 | $spl81)),
  inference(backward_demodulation,[],[f2363,f5329])).
fof(f5485,plain,(
  $false | $spl32),
  inference(subsumption_resolution,[],[f5442,f38])).
fof(f5442,plain,(
  ~p(s(bool,'T_0')) | $spl32),
  inference(backward_demodulation,[],[f2363,f40])).
fof(f5440,plain,(
  ( ! [X14,X15,X13,X16] : (s(fun(bool,bool),X13) = s(fun(bool,bool),X14) | s(bool,X15) = s(bool,sK2(X13,X14,bool,bool)) | s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl9 | $spl51)),
  inference(subsumption_resolution,[],[f3528,f3033])).
fof(f5437,plain,(
  ( ! [X28,X26,X27,X25] : (s(fun(bool,bool),X25) = s(fun(bool,bool),X26) | s(bool,X27) = s(bool,sK2(X25,X26,bool,bool)) | s(bool,'T_0') = s(bool,X27) | s(bool,X28) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X28)) ) | ($spl9 | $spl51)),
  inference(subsumption_resolution,[],[f3531,f3033])).
fof(f3567,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X3) = s(bool,sK2(X4,X5,bool,bool)) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl51),
  inference(duplicate_literal_removal,[],[f3566])).
fof(f3566,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,X3) = s(bool,sK2(X4,X5,bool,bool)) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl51),
  inference(condensation,[],[f3565])).
fof(f3565,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s(fun(bool,bool),X30) = s(fun(bool,bool),X31) | s(bool,X32) = s(bool,sK2(X30,X31,bool,bool)) | s(bool,'T_0') = s(bool,X32) | s(bool,X29) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X34)) ) | $spl51),
  inference(subsumption_resolution,[],[f3532,f3088])).
fof(f3570,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,sK2(X5,X6,bool,bool)) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,X4) | s(bool,X7) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X7)) ) | $spl51),
  inference(duplicate_literal_removal,[],[f3569])).
fof(f3569,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,sK2(X5,X6,bool,bool)) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,X4) | s(bool,X7) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,X4)) ) | $spl51),
  inference(condensation,[],[f3568])).
fof(f3568,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s(fun(bool,bool),X36) = s(fun(bool,bool),X37) | s(bool,X38) = s(bool,sK2(X36,X37,bool,bool)) | s(bool,'T_0') = s(bool,X38) | s(bool,X39) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | $spl51),
  inference(subsumption_resolution,[],[f3533,f3088])).
fof(f5435,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s(fun(bool,bool),X41) = s(fun(bool,bool),X42) | s(bool,X43) = s(bool,sK2(X41,X42,bool,bool)) | s(bool,'T_0') = s(bool,X43) | s(bool,X44) = s(bool,X45) | s(bool,X45) = s(bool,g(s(bool,'T_0'))) | s(bool,X44) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) ) | ($spl9 | $spl51)),
  inference(subsumption_resolution,[],[f3534,f3033])).
fof(f5431,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (s(fun(bool,bool),X59) = s(fun(bool,bool),X60) | s(bool,X61) = s(bool,sK2(X59,X60,bool,bool)) | s(bool,'T_0') = s(bool,X61) | s(bool,X62) = s(bool,X63) | s(bool,X63) = s(bool,X64) | s(bool,'T_0') = s(bool,X64) | s(bool,X62) = s(bool,g(s(bool,'T_0')))) ) | ($spl9 | $spl51)),
  inference(subsumption_resolution,[],[f3537,f3033])).
fof(f5429,plain,(
  ( ! [X57,X54,X56,X55,X53] : (s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(fun(bool,bool),X54) = s(fun(bool,bool),X55) | s(bool,X53) != s(bool,f(s(bool,'F_0'))) | s(bool,X56) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) ) | ($spl9 | $spl14 | $spl51)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f40,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f257,f258,f260,f262,f272,f294,f282,f284,f286,f1360,f1220,f1218,f377,f481,f482,f484,f517,f518,f520,f844,f1419,f535,f536,f538,f541,f557,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1540,f1512,f1550,f1560,f1562,f1624,f1740,f2337,f1841,f2327,f2323,f2311,f2293,f2275,f2268,f2267,f2250,f1913,f1914,f2219,f2218,f1966,f2163,f2158,f1999,f2000,f2002,f2157,f2152,f2139,f2379,f2380,f2036,f2037,f2039,f2041,f2046,f2111,f2082,f2083,f2085,f2089,f2103,f2384,f2385,f2386,f2389,f2391,f2392,f2393,f2397,f2398,f2399,f2402,f2404,f2405,f2406,f2928,f2967,f2978,f3033,f3034,f3036,f3055,f3056,f3088,f3097,f3099,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418,f5419,f5420,f5421,f5422,f5423,f5424,f5425,f5426,f5427,f5428])).
fof(f2111,plain,(
  ( ! [X72] : (s(bool,X72) = s(bool,g(s(bool,X72))) | s(bool,'F_0') = s(bool,X72)) ) | ($spl9 | $spl14)),
  inference(trivial_inequality_removal,[],[f2079])).
fof(f2079,plain,(
  ( ! [X72] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X72) = s(bool,g(s(bool,X72))) | s(bool,'F_0') = s(bool,X72)) ) | ($spl9 | $spl14)),
  inference(superposition,[],[f67,f1562])).
fof(f2380,plain,(
  ( ! [X88] : (~p(s(bool,X88)) | s(bool,X88) = s(bool,g(s(bool,X88)))) ) | ($spl9 | $spl14)),
  inference(subsumption_resolution,[],[f2030,f67])).
fof(f2030,plain,(
  ( ! [X88] : (~p(s(bool,X88)) | s(bool,X88) = s(bool,g(s(bool,X88))) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl14),
  inference(superposition,[],[f40,f1562])).
fof(f2379,plain,(
  ( ! [X85] : (s(bool,'T_0') != s(bool,X85) | s(bool,X85) = s(bool,g(s(bool,X85)))) ) | ($spl9 | $spl14)),
  inference(subsumption_resolution,[],[f2028,f67])).
fof(f2028,plain,(
  ( ! [X85] : (s(bool,'T_0') != s(bool,X85) | s(bool,X85) = s(bool,g(s(bool,X85))) | s(bool,'T_0') = s(bool,'F_0')) ) | ($spl9 | $spl14)),
  inference(superposition,[],[f67,f1562])).
fof(f5427,plain,(
  ( ! [X52,X50,X48,X51,X49] : (s(bool,X51) = s(bool,X52) | s(bool,'T_0') = s(bool,X48) | s(fun(bool,bool),X49) = s(fun(bool,bool),X50) | s(bool,X48) != s(bool,f(s(bool,'F_0'))) | s(bool,X51) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X52)) ) | ($spl9 | $spl14 | $spl51)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f40,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f257,f258,f260,f262,f272,f294,f282,f284,f286,f1360,f1220,f1218,f377,f481,f482,f484,f517,f518,f520,f844,f1419,f535,f536,f538,f541,f557,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1540,f1512,f1550,f1560,f1562,f1624,f1740,f2337,f1841,f2327,f2323,f2311,f2293,f2275,f2268,f2267,f2250,f1913,f1914,f2219,f2218,f1966,f2163,f2158,f1999,f2000,f2002,f2157,f2152,f2139,f2379,f2380,f2036,f2037,f2039,f2041,f2046,f2111,f2082,f2083,f2085,f2089,f2103,f2384,f2385,f2386,f2389,f2391,f2392,f2393,f2397,f2398,f2399,f2402,f2404,f2405,f2406,f2928,f2967,f2978,f3033,f3034,f3036,f3055,f3056,f3088,f3097,f3099,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418,f5419,f5420,f5421,f5422,f5423,f5424,f5425,f5426])).
fof(f5425,plain,(
  ( ! [X47,X45,X43,X46,X44] : (s(bool,X45) = s(bool,X46) | s(bool,'T_0') = s(bool,X47) | s(bool,X45) = s(bool,g(s(bool,'F_0'))) | s(fun(bool,bool),X43) = s(fun(bool,bool),X44) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,X46) = s(bool,X47)) ) | ($spl9 | $spl14 | $spl51)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f40,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f257,f258,f260,f262,f272,f294,f282,f284,f286,f1360,f1220,f1218,f377,f481,f482,f484,f517,f518,f520,f844,f1419,f535,f536,f538,f541,f557,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1540,f1512,f1550,f1560,f1562,f1624,f1740,f2337,f1841,f2327,f2323,f2311,f2293,f2275,f2268,f2267,f2250,f1913,f1914,f2219,f2218,f1966,f2163,f2158,f1999,f2000,f2002,f2157,f2152,f2139,f2379,f2380,f2036,f2037,f2039,f2041,f2046,f2111,f2082,f2083,f2085,f2089,f2103,f2384,f2385,f2386,f2389,f2391,f2392,f2393,f2397,f2398,f2399,f2402,f2404,f2405,f2406,f2928,f2967,f2978,f3033,f3034,f3036,f3055,f3056,f3088,f3097,f3099,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418,f5419,f5420,f5421,f5422,f5423,f5424])).
fof(f5423,plain,(
  ( ! [X39,X41,X38,X42,X40] : (s(fun(bool,bool),X39) = s(fun(bool,bool),X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X42) | s(bool,X38) != s(bool,f(s(bool,'F_0'))) | s(bool,X41) = s(bool,X42)) ) | ($spl9 | $spl14 | $spl51)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f40,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f257,f258,f260,f262,f272,f294,f282,f284,f286,f1360,f1220,f1218,f377,f481,f482,f484,f517,f518,f520,f844,f1419,f535,f536,f538,f541,f557,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1540,f1512,f1550,f1560,f1562,f1624,f1740,f2337,f1841,f2327,f2323,f2311,f2293,f2275,f2268,f2267,f2250,f1913,f1914,f2219,f2218,f1966,f2163,f2158,f1999,f2000,f2002,f2157,f2152,f2139,f2379,f2380,f2036,f2037,f2039,f2041,f2046,f2111,f2082,f2083,f2085,f2089,f2103,f2384,f2385,f2386,f2389,f2391,f2392,f2393,f2397,f2398,f2399,f2402,f2404,f2405,f2406,f2928,f2967,f2978,f3033,f3034,f3036,f3055,f3056,f3088,f3097,f3099,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418,f5419,f5420,f5421,f5422])).
fof(f5421,plain,(
  ( ! [X37,X35,X33,X36,X34] : (s(bool,X36) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,X36) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X33) | s(fun(bool,bool),X34) = s(fun(bool,bool),X35) | s(bool,X33) != s(bool,f(s(bool,'F_0')))) ) | ($spl9 | $spl14 | $spl51)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f40,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f257,f258,f260,f262,f272,f294,f282,f284,f286,f1360,f1220,f1218,f377,f481,f482,f484,f517,f518,f520,f844,f1419,f535,f536,f538,f541,f557,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1540,f1512,f1550,f1560,f1562,f1624,f1740,f2337,f1841,f2327,f2323,f2311,f2293,f2275,f2268,f2267,f2250,f1913,f1914,f2219,f2218,f1966,f2163,f2158,f1999,f2000,f2002,f2157,f2152,f2139,f2379,f2380,f2036,f2037,f2039,f2041,f2046,f2111,f2082,f2083,f2085,f2089,f2103,f2384,f2385,f2386,f2389,f2391,f2392,f2393,f2397,f2398,f2399,f2402,f2404,f2405,f2406,f2928,f2967,f2978,f3033,f3034,f3036,f3055,f3056,f3088,f3097,f3099,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418,f5419,f5420])).
fof(f5419,plain,(
  ( ! [X30,X28,X31,X29,X32] : (s(fun(bool,bool),X28) = s(fun(bool,bool),X29) | s(bool,X30) = s(bool,X32) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X32) | s(bool,X30) = s(bool,X31) | s(bool,X31) = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl14 | $spl51)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f40,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f257,f258,f260,f262,f272,f294,f282,f284,f286,f1360,f1220,f1218,f377,f481,f482,f484,f517,f518,f520,f844,f1419,f535,f536,f538,f541,f557,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1540,f1512,f1550,f1560,f1562,f1624,f1740,f2337,f1841,f2327,f2323,f2311,f2293,f2275,f2268,f2267,f2250,f1913,f1914,f2219,f2218,f1966,f2163,f2158,f1999,f2000,f2002,f2157,f2152,f2139,f2379,f2380,f2036,f2037,f2039,f2041,f2046,f2111,f2082,f2083,f2085,f2089,f2103,f2384,f2385,f2386,f2389,f2391,f2392,f2393,f2397,f2398,f2399,f2402,f2404,f2405,f2406,f2928,f2967,f2978,f3033,f3034,f3036,f3055,f3056,f3088,f3097,f3099,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416,f5417,f5418])).
fof(f5417,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(bool,X23) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,X26) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X26) | s(fun(bool,bool),X24) = s(fun(bool,bool),X25) | s(bool,X23) != s(bool,f(s(bool,'F_0')))) ) | ($spl9 | $spl14 | $spl51)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f40,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f257,f258,f260,f262,f272,f294,f282,f284,f286,f1360,f1220,f1218,f377,f481,f482,f484,f517,f518,f520,f844,f1419,f535,f536,f538,f541,f557,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1540,f1512,f1550,f1560,f1562,f1624,f1740,f2337,f1841,f2327,f2323,f2311,f2293,f2275,f2268,f2267,f2250,f1913,f1914,f2219,f2218,f1966,f2163,f2158,f1999,f2000,f2002,f2157,f2152,f2139,f2379,f2380,f2036,f2037,f2039,f2041,f2046,f2111,f2082,f2083,f2085,f2089,f2103,f2384,f2385,f2386,f2389,f2391,f2392,f2393,f2397,f2398,f2399,f2402,f2404,f2405,f2406,f2928,f2967,f2978,f3033,f3034,f3036,f3055,f3056,f3088,f3097,f3099,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414,f5415,f5416])).
fof(f5415,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X22) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,X18) != s(bool,f(s(bool,'F_0'))) | s(bool,X18) = s(bool,X21)) ) | ($spl9 | $spl14 | $spl51)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f40,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f257,f258,f260,f262,f272,f294,f282,f284,f286,f1360,f1220,f1218,f377,f481,f482,f484,f517,f518,f520,f844,f1419,f535,f536,f538,f541,f557,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1540,f1512,f1550,f1560,f1562,f1624,f1740,f2337,f1841,f2327,f2323,f2311,f2293,f2275,f2268,f2267,f2250,f1913,f1914,f2219,f2218,f1966,f2163,f2158,f1999,f2000,f2002,f2157,f2152,f2139,f2379,f2380,f2036,f2037,f2039,f2041,f2046,f2111,f2082,f2083,f2085,f2089,f2103,f2384,f2385,f2386,f2389,f2391,f2392,f2393,f2397,f2398,f2399,f2402,f2404,f2405,f2406,f2928,f2967,f2978,f3033,f3034,f3036,f3055,f3056,f3088,f3097,f3099,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5406,f5407,f5408,f5409,f5410,f5411,f5412,f5413,f5414])).
fof(f5413,plain,(
  ( ! [X17,X15,X16] : (s(fun(bool,bool),X15) = s(fun(bool,bool),X16) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X17) | s(bool,X17) = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl14 | $spl51)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f40,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f257,f258,f260,f262,f272,f294,f282,f284,f286,f1360,f1220,f1218,f377,f481,f482,f484,f517,f518,f520,f844,f1419,f535,f536,f538,f541,f557,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1540,f1512,f1550,f1560,f1562,f1624,f1740,f2337,f1841,f2327,f2323,f2311,f2293,f2275,f2268,f2267,f2250,f1913,f1914,f2219,f2218,f1966,f2163,f2158,f1999,f2000,f2002,f2157,f2152,f2139,f2379,f2380,f2036,f2037,f2039,f2041,f2046,f2111,f2082,f2083,f2085,f2089,f2103,f2384,f2385,f2386,f2389,f2391,f2392,f2393,f2397,f2398,f2399,f2402,f2404,f2405,f2406,f2928,f2967,f2978,f3033,f3034,f3036,f3055,f3056,f3088,f3097,f3099,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5406,f5407,f5408,f5409,f5410,f5411,f5412])).
fof(f5411,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X12) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X12) | s(fun(bool,bool),X13) = s(fun(bool,bool),X14)) ) | ($spl9 | $spl14 | $spl51)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f40,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f257,f258,f260,f262,f272,f294,f282,f284,f286,f1360,f1220,f1218,f377,f481,f482,f484,f517,f518,f520,f844,f1419,f535,f536,f538,f541,f557,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1540,f1512,f1550,f1560,f1562,f1624,f1740,f2337,f1841,f2327,f2323,f2311,f2293,f2275,f2268,f2267,f2250,f1913,f1914,f2219,f2218,f1966,f2163,f2158,f1999,f2000,f2002,f2157,f2152,f2139,f2379,f2380,f2036,f2037,f2039,f2041,f2046,f2111,f2082,f2083,f2085,f2089,f2103,f2384,f2385,f2386,f2389,f2391,f2392,f2393,f2397,f2398,f2399,f2402,f2404,f2405,f2406,f2928,f2967,f2978,f3033,f3034,f3036,f3055,f3056,f3088,f3097,f3099,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5406,f5407,f5408,f5409,f5410])).
fof(f5409,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11)) ) | ($spl9 | $spl14 | $spl51)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f40,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f257,f258,f260,f262,f272,f294,f282,f284,f286,f1360,f1220,f1218,f377,f481,f482,f484,f517,f518,f520,f844,f1419,f535,f536,f538,f541,f557,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1540,f1512,f1550,f1560,f1562,f1624,f1740,f2337,f1841,f2327,f2323,f2311,f2293,f2275,f2268,f2267,f2250,f1913,f1914,f2219,f2218,f1966,f2163,f2158,f1999,f2000,f2002,f2157,f2152,f2139,f2379,f2380,f2036,f2037,f2039,f2041,f2046,f2111,f2082,f2083,f2085,f2089,f2103,f2384,f2385,f2386,f2389,f2391,f2392,f2393,f2397,f2398,f2399,f2402,f2404,f2405,f2406,f2928,f2967,f2978,f3033,f3034,f3036,f3055,f3056,f3088,f3097,f3099,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5406,f5407,f5408])).
fof(f5407,plain,(
  ( ! [X6,X8,X7] : (s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,X8) = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl14 | $spl51)),
  inference(global_subsumption,[],[f37,f36,f35,f38,f40,f41,f42,f43,f48,f51,f50,f49,f59,f61,f64,f67,f97,f129,f253,f255,f257,f258,f260,f262,f272,f294,f282,f284,f286,f1360,f1220,f1218,f377,f481,f482,f484,f517,f518,f520,f844,f1419,f535,f536,f538,f541,f557,f558,f560,f740,f1422,f657,f1545,f1477,f1478,f1479,f1480,f1481,f1482,f1483,f1484,f1485,f1486,f1487,f1488,f1489,f1540,f1512,f1550,f1560,f1562,f1624,f1740,f2337,f1841,f2327,f2323,f2311,f2293,f2275,f2268,f2267,f2250,f1913,f1914,f2219,f2218,f1966,f2163,f2158,f1999,f2000,f2002,f2157,f2152,f2139,f2379,f2380,f2036,f2037,f2039,f2041,f2046,f2111,f2082,f2083,f2085,f2089,f2103,f2384,f2385,f2386,f2389,f2391,f2392,f2393,f2397,f2398,f2399,f2402,f2404,f2405,f2406,f2928,f2967,f2978,f3033,f3034,f3036,f3055,f3056,f3088,f3097,f3099,f3508,f3965,f3964,f3963,f3962,f3961,f3960,f3959,f3958,f3957,f3956,f3955,f3954,f3953,f3193,f3190,f5406])).
fof(f5395,plain,(
  ( ! [X26,X27,X25] : (s(bool,X25) != s(bool,f(s(bool,'F_0'))) | s(bool,X25) = s(bool,X26) | s(bool,X26) = s(bool,X27) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl11),
  inference(superposition,[],[f1547,f97])).
fof(f5394,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) != s(bool,f(s(bool,'F_0'))) | s(bool,X23) = s(bool,g(s(bool,'F_0'))) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X22)) ) | $spl11),
  inference(superposition,[],[f1547,f97])).
fof(f5393,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,X19) = s(bool,X20) | s(bool,X20) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,X19) = s(bool,g(s(bool,'F_0')))) ) | $spl11),
  inference(superposition,[],[f1547,f97])).
fof(f5392,plain,(
  ( ! [X17,X18,X16] : (s(bool,X16) != s(bool,f(s(bool,'F_0'))) | s(bool,X16) = s(bool,X17) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X17) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) ) | $spl11),
  inference(superposition,[],[f1547,f97])).
fof(f5391,plain,(
  ( ! [X14,X15,X13] : (s(bool,X13) != s(bool,f(s(bool,'F_0'))) | s(bool,X14) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl11),
  inference(superposition,[],[f1547,f97])).
fof(f5390,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,X10) = s(bool,X11) | s(bool,X11) = s(bool,g(s(bool,'F_0'))) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) ) | $spl11),
  inference(superposition,[],[f1547,f97])).
fof(f5389,plain,(
  ( ! [X8,X7,X9] : (s(bool,X7) != s(bool,f(s(bool,'F_0'))) | s(bool,X8) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl11),
  inference(superposition,[],[f1547,f97])).
fof(f5388,plain,(
  ( ! [X6,X4,X5] : (s(bool,X4) != s(bool,f(s(bool,'F_0'))) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl11),
  inference(superposition,[],[f1547,f97])).
fof(f5387,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,X3) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X3)) ) | $spl11),
  inference(superposition,[],[f1547,f59])).
fof(f5386,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl11),
  inference(superposition,[],[f1547,f59])).
fof(f5385,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) ) | $spl11),
  inference(superposition,[],[f1547,f59])).
fof(f5384,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(bool,X0) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl11),
  inference(superposition,[],[f1547,f59])).
fof(f5318,plain,(
  ( ! [X83,X82] : (s(bool,'T_0') = s(bool,sK2(X82,X83,bool,bool)) | s(fun(bool,bool),X82) = s(fun(bool,bool),X83)) ) | $spl62),
  inference(cnf_transformation,[],[f5318_D])).
fof(f5318_D,plain,(
  ( ! [X83,X82] : (s(bool,'T_0') = s(bool,sK2(X82,X83,bool,bool)) | s(fun(bool,bool),X82) = s(fun(bool,bool),X83)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f2801,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK2(X2,X3,bool,bool)) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2764])).
fof(f2764,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK2(X2,X3,bool,bool)) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK2(X2,X3,bool,bool))) ) | $spl16),
  inference(superposition,[],[f1551,f43])).
fof(f2800,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,sK2(X7,X8,bool,bool)) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,X9)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2766])).
fof(f2766,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X9))) | s(bool,'T_0') = s(bool,sK2(X7,X8,bool,bool)) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK2(X7,X8,bool,bool))) ) | $spl16),
  inference(superposition,[],[f1551,f59])).
fof(f2799,plain,(
  ( ! [X12,X10,X11] : (s(bool,'F_0') != s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,sK2(X10,X11,bool,bool)) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,X12)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2767])).
fof(f2767,plain,(
  ( ! [X12,X10,X11] : (s(bool,'F_0') != s(bool,f(s(bool,X12))) | s(bool,'T_0') = s(bool,sK2(X10,X11,bool,bool)) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK2(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X12)) ) | $spl16),
  inference(superposition,[],[f1551,f59])).
fof(f2798,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,X2))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(condensation,[],[f2797])).
fof(f2797,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,X2))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X3)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2796])).
fof(f2796,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,X2))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) ) | $spl16),
  inference(condensation,[],[f2772])).
fof(f2772,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s(bool,'F_0') != s(bool,f(s(bool,X33))) | s(bool,'T_0') = s(bool,sK2(X31,X32,bool,bool)) | s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(bool,X34) = s(bool,sK2(X31,X32,bool,bool)) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | $spl16),
  inference(superposition,[],[f1551,f97])).
fof(f5336,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') != s(bool,f(s(bool,X53))) | s(bool,'T_0') = s(bool,X55) | s(bool,X54) = s(bool,X55) | s(bool,X53) = s(bool,X54)) ) | $spl82),
  inference(cnf_transformation,[],[f5336_D])).
fof(f5336_D,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') != s(bool,f(s(bool,X53))) | s(bool,'T_0') = s(bool,X55) | s(bool,X54) = s(bool,X55) | s(bool,X53) = s(bool,X54)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f2795,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'F_0') != s(bool,f(s(bool,X38))) | s(bool,'T_0') = s(bool,sK2(X36,X37,bool,bool)) | s(fun(bool,bool),X36) = s(fun(bool,bool),X37) | s(bool,X38) = s(bool,X39) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2773])).
fof(f2773,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'F_0') != s(bool,f(s(bool,X38))) | s(bool,'T_0') = s(bool,sK2(X36,X37,bool,bool)) | s(fun(bool,bool),X36) = s(fun(bool,bool),X37) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,sK2(X36,X37,bool,bool)) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | $spl16),
  inference(superposition,[],[f1551,f97])).
fof(f2794,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,X2))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(condensation,[],[f2793])).
fof(f2793,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,X2))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2792])).
fof(f2792,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,sK2(X0,X1,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,X2))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(condensation,[],[f2775])).
fof(f2775,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'F_0') != s(bool,f(s(bool,X48))) | s(bool,'T_0') = s(bool,sK2(X46,X47,bool,bool)) | s(fun(bool,bool),X46) = s(fun(bool,bool),X47) | s(bool,X49) = s(bool,sK2(X46,X47,bool,bool)) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X48)) ) | $spl16),
  inference(superposition,[],[f1551,f97])).
fof(f2791,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,'F_0') != s(bool,f(s(bool,X53))) | s(bool,'T_0') = s(bool,sK2(X51,X52,bool,bool)) | s(fun(bool,bool),X51) = s(fun(bool,bool),X52) | s(bool,X53) = s(bool,X54) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2776])).
fof(f2776,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,'F_0') != s(bool,f(s(bool,X53))) | s(bool,'T_0') = s(bool,sK2(X51,X52,bool,bool)) | s(fun(bool,bool),X51) = s(fun(bool,bool),X52) | s(bool,X53) = s(bool,X54) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,sK2(X51,X52,bool,bool))) ) | $spl16),
  inference(superposition,[],[f1551,f97])).
fof(f5342,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK2(X2,X3,bool,bool))) ) | ($spl9 | $spl15 | $spl16)),
  inference(subsumption_resolution,[],[f3182,f5334])).
fof(f5334,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') = s(bool,sK2(X48,X49,bool,bool)) | s(fun(bool,bool),X48) = s(fun(bool,bool),X49) | s(bool,'T_0') = s(bool,f(s(bool,sK2(X48,X49,bool,bool))))) ) | ($spl9 | $spl15 | $spl16)),
  inference(subsumption_resolution,[],[f3291,f2483])).
fof(f2483,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl9 | $spl15)),
  inference(unit_resulting_resolution,[],[f67,f1549,f1549,f1549,f97])).
fof(f3291,plain,(
  ( ! [X48,X49] : (s(bool,'F_0') != s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,sK2(X48,X49,bool,bool)) | s(fun(bool,bool),X48) = s(fun(bool,bool),X49) | s(bool,'T_0') = s(bool,f(s(bool,sK2(X48,X49,bool,bool))))) ) | ($spl15 | $spl16)),
  inference(superposition,[],[f1551,f2529])).
fof(f3182,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X2,X3,bool,bool)))) | s(bool,'T_0') = s(bool,sK2(X2,X3,bool,bool))) )),
  inference(superposition,[],[f1512,f43])).
fof(f5341,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X9))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK2(X7,X8,bool,bool))) ) | ($spl9 | $spl15 | $spl16)),
  inference(subsumption_resolution,[],[f3184,f5334])).
fof(f3184,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X9))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X7,X8,bool,bool)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK2(X7,X8,bool,bool))) )),
  inference(superposition,[],[f1512,f59])).
fof(f5340,plain,(
  ( ! [X12,X10,X11] : (s(bool,'F_0') != s(bool,f(s(bool,X12))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK2(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X12)) ) | ($spl9 | $spl15 | $spl16)),
  inference(subsumption_resolution,[],[f3185,f5334])).
fof(f3185,plain,(
  ( ! [X12,X10,X11] : (s(bool,'F_0') != s(bool,f(s(bool,X12))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X10,X11,bool,bool)))) | s(bool,'T_0') = s(bool,sK2(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f1512,f59])).
fof(f5337,plain,(
  ( ! [X47,X50,X46,X49] : (s(fun(bool,bool),X46) = s(fun(bool,bool),X47) | s(bool,'T_0') = s(bool,X50) | s(bool,X49) = s(bool,X50) | s(bool,X49) = s(bool,sK2(X46,X47,bool,bool)) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X46,X47,bool,bool))))) ) | $spl84),
  inference(cnf_transformation,[],[f5337_D])).
fof(f5337_D,plain,(
  ( ! [X47,X50,X46,X49] : (s(fun(bool,bool),X46) = s(fun(bool,bool),X47) | s(bool,'T_0') = s(bool,X50) | s(bool,X49) = s(bool,X50) | s(bool,X49) = s(bool,sK2(X46,X47,bool,bool)) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X46,X47,bool,bool))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f3190,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s(bool,'F_0') != s(bool,f(s(bool,X33))) | s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X31,X32,bool,bool)))) | s(bool,X34) = s(bool,sK2(X31,X32,bool,bool)) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f1512,f97])).
fof(f5339,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'F_0') != s(bool,f(s(bool,X38))) | s(fun(bool,bool),X36) = s(fun(bool,bool),X37) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,sK2(X36,X37,bool,bool)) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) ) | ($spl9 | $spl15 | $spl16)),
  inference(subsumption_resolution,[],[f3191,f5334])).
fof(f3191,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'F_0') != s(bool,f(s(bool,X38))) | s(fun(bool,bool),X36) = s(fun(bool,bool),X37) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X36,X37,bool,bool)))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,sK2(X36,X37,bool,bool)) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1512,f97])).
fof(f3193,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'F_0') != s(bool,f(s(bool,X48))) | s(fun(bool,bool),X46) = s(fun(bool,bool),X47) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X46,X47,bool,bool)))) | s(bool,X49) = s(bool,sK2(X46,X47,bool,bool)) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f1512,f97])).
fof(f5335,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,'F_0') != s(bool,f(s(bool,X53))) | s(fun(bool,bool),X51) = s(fun(bool,bool),X52) | s(bool,X53) = s(bool,X54) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,sK2(X51,X52,bool,bool))) ) | ($spl9 | $spl15 | $spl16)),
  inference(subsumption_resolution,[],[f3194,f5334])).
fof(f3194,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,'F_0') != s(bool,f(s(bool,X53))) | s(fun(bool,bool),X51) = s(fun(bool,bool),X52) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X51,X52,bool,bool)))) | s(bool,X53) = s(bool,X54) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,sK2(X51,X52,bool,bool))) )),
  inference(superposition,[],[f1512,f97])).
fof(f5333,plain,(
  ( ! [X61,X60] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,sK2(X60,X61,bool,bool)) | s(fun(bool,bool),X60) = s(fun(bool,bool),X61)) ) | ($spl9 | $spl15 | $spl16)),
  inference(forward_demodulation,[],[f3347,f2483])).
fof(f3347,plain,(
  ( ! [X61,X60] : (s(bool,'F_0') != s(bool,f(s(bool,g(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,sK2(X60,X61,bool,bool)) | s(fun(bool,bool),X60) = s(fun(bool,bool),X61)) ) | ($spl15 | $spl16)),
  inference(duplicate_literal_removal,[],[f3299])).
fof(f3299,plain,(
  ( ! [X61,X60] : (s(bool,'F_0') != s(bool,f(s(bool,g(s(bool,'T_0'))))) | s(bool,'T_0') = s(bool,sK2(X60,X61,bool,bool)) | s(fun(bool,bool),X60) = s(fun(bool,bool),X61) | s(bool,'T_0') = s(bool,sK2(X60,X61,bool,bool))) ) | ($spl15 | $spl16)),
  inference(superposition,[],[f1551,f2529])).
fof(f5327,plain,(
  ( ! [X83,X84] : (s(fun(bool,bool),X83) = s(fun(bool,bool),X84) | s(bool,'T_0') = s(bool,sK2(X83,X84,bool,bool)) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X83,X84,bool,bool))))) ) | $spl78),
  inference(cnf_transformation,[],[f5327_D])).
fof(f5327_D,plain,(
  ( ! [X83,X84] : (s(fun(bool,bool),X83) = s(fun(bool,bool),X84) | s(bool,'T_0') = s(bool,sK2(X83,X84,bool,bool)) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X83,X84,bool,bool))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f5326,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X83) = s(fun(bool,bool),X84) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X83,X84,bool,bool)))) | s(bool,'T_0') = s(bool,sK2(X83,X84,bool,bool))) ) | ($spl9 | $spl15)),
  inference(forward_demodulation,[],[f3307,f2483])).
fof(f3307,plain,(
  ( ! [X83,X84] : (s(bool,'F_0') != s(bool,f(s(bool,g(s(bool,'T_0'))))) | s(fun(bool,bool),X83) = s(fun(bool,bool),X84) | s(bool,'T_0') != s(bool,f(s(bool,sK2(X83,X84,bool,bool)))) | s(bool,'T_0') = s(bool,sK2(X83,X84,bool,bool))) ) | $spl15),
  inference(superposition,[],[f1512,f2529])).
fof(f5319,plain,(
  ( ! [X81] : (s(bool,'F_0') != s(bool,f(s(bool,X81))) | s(bool,'T_0') = s(bool,g(s(bool,X81))) | s(bool,'T_0') = s(bool,X81)) ) | $spl64),
  inference(cnf_transformation,[],[f5319_D])).
fof(f5319_D,plain,(
  ( ! [X81] : (s(bool,'F_0') != s(bool,f(s(bool,X81))) | s(bool,'T_0') = s(bool,g(s(bool,X81))) | s(bool,'T_0') = s(bool,X81)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f3953,plain,(
  ( ! [X14,X15,X16] : (s(bool,'F_0') != s(bool,f(s(bool,X14))) | s(fun(bool,bool),X15) = s(fun(bool,bool),X16) | s(bool,'T_0') = s(bool,sK2(X15,X16,bool,bool)) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,g(s(bool,X14)))) )),
  inference(superposition,[],[f255,f43])).
fof(f5324,plain,(
  ( ! [X29,X32] : (s(bool,'T_0') != s(bool,f(s(bool,X29))) | s(bool,'T_0') = s(bool,X32) | s(bool,X32) = s(bool,g(s(bool,X29))) | s(bool,'T_0') = s(bool,X29)) ) | $spl74),
  inference(cnf_transformation,[],[f5324_D])).
fof(f5324_D,plain,(
  ( ! [X29,X32] : (s(bool,'T_0') != s(bool,f(s(bool,X29))) | s(bool,'T_0') = s(bool,X32) | s(bool,X32) = s(bool,g(s(bool,X29))) | s(bool,'T_0') = s(bool,X29)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3954,plain,(
  ( ! [X19,X17,X20,X18] : (s(bool,'T_0') != s(bool,f(s(bool,X17))) | s(fun(bool,bool),X18) = s(fun(bool,bool),X19) | s(bool,'T_0') = s(bool,sK2(X18,X19,bool,bool)) | s(bool,'T_0') = s(bool,X17) | s(bool,X20) = s(bool,g(s(bool,X17))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f255,f59])).
fof(f5325,plain,(
  ( ! [X26,X25] : (s(bool,X26) != s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,g(s(bool,X25))) | s(bool,'T_0') = s(bool,X25)) ) | $spl76),
  inference(cnf_transformation,[],[f5325_D])).
fof(f5325_D,plain,(
  ( ! [X26,X25] : (s(bool,X26) != s(bool,f(s(bool,X25))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,g(s(bool,X25))) | s(bool,'T_0') = s(bool,X25)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f3955,plain,(
  ( ! [X24,X23,X21,X22] : (s(bool,X22) != s(bool,f(s(bool,X21))) | s(fun(bool,bool),X23) = s(fun(bool,bool),X24) | s(bool,'T_0') = s(bool,sK2(X23,X24,bool,bool)) | s(bool,'T_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X22) | s(bool,'T_0') = s(bool,g(s(bool,X21)))) )),
  inference(superposition,[],[f255,f59])).
fof(f3956,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,X26) != s(bool,f(s(bool,X25))) | s(fun(bool,bool),X27) = s(fun(bool,bool),X28) | s(bool,'T_0') = s(bool,sK2(X27,X28,bool,bool)) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,g(s(bool,X25))) | s(bool,'T_0') = s(bool,X26)) )),
  inference(superposition,[],[f255,f59])).
fof(f3957,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'T_0') != s(bool,f(s(bool,X29))) | s(fun(bool,bool),X30) = s(fun(bool,bool),X31) | s(bool,'T_0') = s(bool,sK2(X30,X31,bool,bool)) | s(bool,'T_0') = s(bool,X29) | s(bool,X32) = s(bool,g(s(bool,X29))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f255,f59])).
fof(f5323,plain,(
  ( ! [X39,X43,X44,X40] : (s(bool,X40) != s(bool,f(s(bool,X39))) | s(bool,'T_0') = s(bool,X44) | s(bool,X40) = s(bool,X44) | s(bool,'T_0') = s(bool,X43) | s(bool,X43) = s(bool,g(s(bool,X39))) | s(bool,'T_0') = s(bool,X39)) ) | $spl72),
  inference(cnf_transformation,[],[f5323_D])).
fof(f5323_D,plain,(
  ( ! [X39,X43,X44,X40] : (s(bool,X40) != s(bool,f(s(bool,X39))) | s(bool,'T_0') = s(bool,X44) | s(bool,X40) = s(bool,X44) | s(bool,'T_0') = s(bool,X43) | s(bool,X43) = s(bool,g(s(bool,X39))) | s(bool,'T_0') = s(bool,X39)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3958,plain,(
  ( ! [X37,X35,X33,X38,X36,X34] : (s(bool,X34) != s(bool,f(s(bool,X33))) | s(fun(bool,bool),X35) = s(fun(bool,bool),X36) | s(bool,'T_0') = s(bool,sK2(X35,X36,bool,bool)) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37) | s(bool,X38) = s(bool,g(s(bool,X33))) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f255,f97])).
fof(f3959,plain,(
  ( ! [X39,X43,X41,X44,X42,X40] : (s(bool,X40) != s(bool,f(s(bool,X39))) | s(fun(bool,bool),X41) = s(fun(bool,bool),X42) | s(bool,'T_0') = s(bool,sK2(X41,X42,bool,bool)) | s(bool,'T_0') = s(bool,X39) | s(bool,X43) = s(bool,g(s(bool,X39))) | s(bool,'T_0') = s(bool,X43) | s(bool,X40) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f255,f97])).
fof(f5322,plain,(
  ( ! [X68,X66,X67,X63] : (s(bool,'T_0') != s(bool,f(s(bool,X63))) | s(bool,X66) = s(bool,X67) | s(bool,X66) = s(bool,g(s(bool,X63))) | s(bool,'T_0') = s(bool,X68) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X63)) ) | $spl70),
  inference(cnf_transformation,[],[f5322_D])).
fof(f5322_D,plain,(
  ( ! [X68,X66,X67,X63] : (s(bool,'T_0') != s(bool,f(s(bool,X63))) | s(bool,X66) = s(bool,X67) | s(bool,X66) = s(bool,g(s(bool,X63))) | s(bool,'T_0') = s(bool,X68) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X63)) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f3960,plain,(
  ( ! [X47,X45,X50,X48,X46,X49] : (s(bool,'T_0') != s(bool,f(s(bool,X45))) | s(fun(bool,bool),X46) = s(fun(bool,bool),X47) | s(bool,'T_0') = s(bool,sK2(X46,X47,bool,bool)) | s(bool,'T_0') = s(bool,X45) | s(bool,X48) = s(bool,X49) | s(bool,X49) = s(bool,g(s(bool,X45))) | s(bool,X48) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f255,f97])).
fof(f5321,plain,(
  ( ! [X70,X74,X69,X73] : (s(bool,X70) != s(bool,f(s(bool,X69))) | s(bool,'T_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X74) | s(bool,X73) = s(bool,X74) | s(bool,X73) = s(bool,g(s(bool,X69))) | s(bool,'T_0') = s(bool,X69)) ) | $spl68),
  inference(cnf_transformation,[],[f5321_D])).
fof(f5321_D,plain,(
  ( ! [X70,X74,X69,X73] : (s(bool,X70) != s(bool,f(s(bool,X69))) | s(bool,'T_0') = s(bool,X70) | s(bool,'T_0') = s(bool,X74) | s(bool,X73) = s(bool,X74) | s(bool,X73) = s(bool,g(s(bool,X69))) | s(bool,'T_0') = s(bool,X69)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f3961,plain,(
  ( ! [X54,X52,X56,X55,X53,X51] : (s(bool,X52) != s(bool,f(s(bool,X51))) | s(fun(bool,bool),X53) = s(fun(bool,bool),X54) | s(bool,'T_0') = s(bool,sK2(X53,X54,bool,bool)) | s(bool,'T_0') = s(bool,X51) | s(bool,X55) = s(bool,g(s(bool,X51))) | s(bool,'T_0') = s(bool,X52) | s(bool,X55) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f255,f97])).
fof(f5320,plain,(
  ( ! [X80,X76,X79,X75] : (s(bool,X76) != s(bool,f(s(bool,X75))) | s(bool,'T_0') = s(bool,g(s(bool,X75))) | s(bool,'T_0') = s(bool,X80) | s(bool,X79) = s(bool,X80) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X75)) ) | $spl66),
  inference(cnf_transformation,[],[f5320_D])).
fof(f5320_D,plain,(
  ( ! [X80,X76,X79,X75] : (s(bool,X76) != s(bool,f(s(bool,X75))) | s(bool,'T_0') = s(bool,g(s(bool,X75))) | s(bool,'T_0') = s(bool,X80) | s(bool,X79) = s(bool,X80) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X75)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f3962,plain,(
  ( ! [X61,X59,X57,X62,X60,X58] : (s(bool,X58) != s(bool,f(s(bool,X57))) | s(fun(bool,bool),X59) = s(fun(bool,bool),X60) | s(bool,'T_0') = s(bool,sK2(X59,X60,bool,bool)) | s(bool,'T_0') = s(bool,X57) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,g(s(bool,X57))) | s(bool,X61) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f255,f97])).
fof(f3963,plain,(
  ( ! [X68,X66,X64,X67,X65,X63] : (s(bool,'T_0') != s(bool,f(s(bool,X63))) | s(fun(bool,bool),X64) = s(fun(bool,bool),X65) | s(bool,'T_0') = s(bool,sK2(X64,X65,bool,bool)) | s(bool,'T_0') = s(bool,X63) | s(bool,X66) = s(bool,X67) | s(bool,X67) = s(bool,X68) | s(bool,'T_0') = s(bool,X68) | s(bool,X66) = s(bool,g(s(bool,X63)))) )),
  inference(superposition,[],[f255,f97])).
fof(f3964,plain,(
  ( ! [X70,X74,X72,X71,X69,X73] : (s(bool,X70) != s(bool,f(s(bool,X69))) | s(fun(bool,bool),X71) = s(fun(bool,bool),X72) | s(bool,'T_0') = s(bool,sK2(X71,X72,bool,bool)) | s(bool,'T_0') = s(bool,X69) | s(bool,X73) = s(bool,g(s(bool,X69))) | s(bool,X73) = s(bool,X74) | s(bool,'T_0') = s(bool,X74) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f255,f97])).
fof(f3965,plain,(
  ( ! [X80,X78,X76,X79,X77,X75] : (s(bool,X76) != s(bool,f(s(bool,X75))) | s(fun(bool,bool),X77) = s(fun(bool,bool),X78) | s(bool,'T_0') = s(bool,sK2(X77,X78,bool,bool)) | s(bool,'T_0') = s(bool,X75) | s(bool,X76) = s(bool,X79) | s(bool,X79) = s(bool,X80) | s(bool,'T_0') = s(bool,X80) | s(bool,'T_0') = s(bool,g(s(bool,X75)))) )),
  inference(superposition,[],[f255,f97])).
fof(f5317,plain,(
  ( ! [X83,X81,X82] : (s(bool,'F_0') != s(bool,f(s(bool,X81))) | s(fun(bool,bool),X82) = s(fun(bool,bool),X83) | s(bool,'T_0') = s(bool,sK2(X82,X83,bool,bool)) | s(bool,'T_0') = s(bool,X81) | s(bool,'T_0') = s(bool,g(s(bool,X81)))) ) | ($spl9 | $spl15)),
  inference(forward_demodulation,[],[f3966,f2483])).
fof(f3966,plain,(
  ( ! [X83,X81,X82] : (s(bool,g(s(bool,'T_0'))) != s(bool,f(s(bool,X81))) | s(fun(bool,bool),X82) = s(fun(bool,bool),X83) | s(bool,'T_0') = s(bool,sK2(X82,X83,bool,bool)) | s(bool,'T_0') = s(bool,X81) | s(bool,'T_0') = s(bool,g(s(bool,X81)))) ) | $spl15),
  inference(superposition,[],[f255,f2529])).
fof(f4922,plain,(
  ( ! [X80,X79] : (s(bool,X80) = s(bool,f(s(bool,X80))) | s(bool,'T_0') = s(bool,X79) | s(bool,X80) = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl11 | $spl15 | $spl22)),
  inference(duplicate_literal_removal,[],[f4721])).
fof(f4721,plain,(
  ( ! [X80,X79] : (s(bool,X80) = s(bool,f(s(bool,X80))) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,X79) | s(bool,X80) = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl11 | $spl15 | $spl22)),
  inference(superposition,[],[f2884,f1804])).
fof(f1804,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,g(s(bool,'F_0'))) | s(bool,X0) = s(bool,X1)) ) | ($spl9 | $spl22)),
  inference(condensation,[],[f1803])).
fof(f1803,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,X8) | s(bool,X9) = s(bool,g(s(bool,'F_0'))) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | ($spl9 | $spl22)),
  inference(duplicate_literal_removal,[],[f1760])).
fof(f1760,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,X8) | s(bool,X9) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) ) | ($spl9 | $spl22)),
  inference(superposition,[],[f1643,f97])).
fof(f1643,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | ($spl9 | $spl22)),
  inference(unit_resulting_resolution,[],[f67,f1555])).
fof(f2884,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl11 | $spl15 | $spl22)),
  inference(superposition,[],[f2727,f43])).
fof(f2727,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl9 | $spl11 | $spl15 | $spl22)),
  inference(forward_demodulation,[],[f2677,f2483])).
fof(f2677,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,g(s(bool,'T_0'))) | ($spl9 | $spl11 | $spl15 | $spl22)),
  inference(unit_resulting_resolution,[],[f1549,f2383,f2383,f2383,f97])).
fof(f2383,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | ($spl9 | $spl11 | $spl22)),
  inference(forward_demodulation,[],[f1547,f1643])).
fof(f4924,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') != s(bool,f(s(bool,X76))) | s(bool,'T_0') = s(bool,X75) | s(bool,X76) = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl11 | $spl22)),
  inference(duplicate_literal_removal,[],[f4719])).
fof(f4719,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') != s(bool,f(s(bool,X76))) | s(bool,'T_0') = s(bool,X75) | s(bool,'T_0') = s(bool,X75) | s(bool,X76) = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl11 | $spl22)),
  inference(superposition,[],[f2695,f1804])).
fof(f2695,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl11 | $spl22)),
  inference(superposition,[],[f2383,f43])).
fof(f5070,plain,(
  ( ! [X59,X57,X58] : (s(bool,'T_0') = s(bool,X58) | s(bool,X57) = s(bool,g(s(bool,'F_0'))) | s(bool,X59) = s(bool,g(s(bool,'F_0'))) | s(bool,X57) = s(bool,X59)) ) | ($spl9 | $spl22)),
  inference(duplicate_literal_removal,[],[f4509])).
fof(f4509,plain,(
  ( ! [X59,X57,X58] : (s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,X58) | s(bool,X57) = s(bool,g(s(bool,'F_0'))) | s(bool,X59) = s(bool,g(s(bool,'F_0'))) | s(bool,X57) = s(bool,X59)) ) | ($spl9 | $spl22)),
  inference(superposition,[],[f1804,f1804])).
fof(f5089,plain,(
  ( ! [X61,X62,X60] : (s(bool,X61) = s(bool,X62) | s(bool,'T_0') = s(bool,X60) | s(bool,X62) = s(bool,g(s(bool,'F_0'))) | s(bool,X61) = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl22)),
  inference(duplicate_literal_removal,[],[f4473])).
fof(f4473,plain,(
  ( ! [X61,X62,X60] : (s(bool,X61) = s(bool,X62) | s(bool,'T_0') = s(bool,X60) | s(bool,X62) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X60) | s(bool,X61) = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl22)),
  inference(superposition,[],[f1804,f1804])).
fof(f5198,plain,(
  ( ! [X64,X63] : (s(bool,'F_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X63) | s(bool,X64) = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl11 | $spl15 | $spl22)),
  inference(forward_demodulation,[],[f5163,f2727])).
fof(f5163,plain,(
  ( ! [X64,X63] : (s(bool,'T_0') = s(bool,X63) | s(bool,X64) = s(bool,g(s(bool,'F_0'))) | s(bool,X64) = s(bool,f(s(bool,'F_0')))) ) | ($spl9 | $spl11 | $spl22)),
  inference(duplicate_literal_removal,[],[f4389])).
fof(f4389,plain,(
  ( ! [X64,X63] : (s(bool,'T_0') = s(bool,X63) | s(bool,X64) = s(bool,g(s(bool,'F_0'))) | s(bool,X64) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X63)) ) | ($spl9 | $spl11 | $spl22)),
  inference(superposition,[],[f1804,f2725])).
fof(f2725,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl11 | $spl22)),
  inference(trivial_inequality_removal,[],[f2710])).
fof(f2710,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl9 | $spl11 | $spl22)),
  inference(superposition,[],[f2383,f59])).
fof(f5164,plain,(
  ( ! [X61,X62] : (s(bool,'T_0') = s(bool,X61) | s(bool,X62) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X62)) ) | ($spl9 | $spl22)),
  inference(duplicate_literal_removal,[],[f4388])).
fof(f4388,plain,(
  ( ! [X61,X62] : (s(bool,'T_0') = s(bool,X61) | s(bool,X62) = s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,X62) | s(bool,'T_0') = s(bool,X61)) ) | ($spl9 | $spl22)),
  inference(superposition,[],[f1804,f43])).
fof(f5165,plain,(
  ( ! [X50,X51,X49] : (s(bool,'T_0') = s(bool,X50) | s(bool,X51) = s(bool,g(s(bool,'F_0'))) | s(bool,X49) = s(bool,X51) | s(bool,X49) = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl22)),
  inference(duplicate_literal_removal,[],[f4385])).
fof(f4385,plain,(
  ( ! [X50,X51,X49] : (s(bool,'T_0') = s(bool,X50) | s(bool,X51) = s(bool,g(s(bool,'F_0'))) | s(bool,X49) = s(bool,X51) | s(bool,'T_0') = s(bool,X50) | s(bool,X49) = s(bool,g(s(bool,'F_0')))) ) | ($spl9 | $spl22)),
  inference(superposition,[],[f1804,f1804])).
fof(f3086,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl52),
  inference(cnf_transformation,[],[f3086_D])).
fof(f3086_D,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2976,plain,(
  ( ! [X21,X18] : (s(bool,'F_0') != s(bool,X18) | s(bool,'T_0') = s(bool,X21) | s(bool,X18) = s(bool,X21)) ) | $spl48),
  inference(cnf_transformation,[],[f2976_D])).
fof(f2976_D,plain,(
  ( ! [X21,X18] : (s(bool,'F_0') != s(bool,X18) | s(bool,'T_0') = s(bool,X21) | s(bool,X18) = s(bool,X21)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f3056,plain,(
  ( ! [X8,X7,X9] : (s(bool,'F_0') != s(bool,X7) | s(bool,X8) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X8) | s(bool,X7) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl51),
  inference(superposition,[],[f2978,f97])).
fof(f3055,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,X4) = s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,f(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl51),
  inference(superposition,[],[f2978,f97])).
fof(f2980,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') != s(bool,X23) | s(fun(bool,bool),X24) = s(fun(bool,bool),X25) | s(bool,X26) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X26) | s(bool,X23) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) ) | ($spl9 | $spl11 | $spl15 | $spl22)),
  inference(forward_demodulation,[],[f2979,f2483])).
fof(f2979,plain,(
  ( ! [X26,X24,X23,X27,X25] : (s(bool,'F_0') != s(bool,X23) | s(fun(bool,bool),X24) = s(fun(bool,bool),X25) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,X26) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X26) | s(bool,X23) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) ) | ($spl9 | $spl11 | $spl15 | $spl22)),
  inference(forward_demodulation,[],[f2950,f2727])).
fof(f2973,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(bool,'F_0') != s(bool,X18) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,X18) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X22)) ) | ($spl9 | $spl11 | $spl15 | $spl22)),
  inference(forward_demodulation,[],[f2972,f2483])).
fof(f2972,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'F_0') != s(bool,X18) | s(fun(bool,bool),X19) = s(fun(bool,bool),X20) | s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(bool,X18) = s(bool,X21) | s(bool,'T_0') = s(bool,X21) | s(bool,X22) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X22)) ) | ($spl9 | $spl11 | $spl15 | $spl22)),
  inference(forward_demodulation,[],[f2949,f2727])).
fof(f2378,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) = s(bool,g(s(bool,X15)))) ) | $spl42),
  inference(cnf_transformation,[],[f2378_D])).
fof(f2378_D,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) = s(bool,g(s(bool,X15)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2139,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | ~p(s(bool,X3)) | s(bool,X3) = s(bool,g(s(bool,X3)))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f2138])).
fof(f2138,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | ~p(s(bool,X3)) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,g(s(bool,X3))) | s(bool,'T_0') = s(bool,X2)) ) | $spl14),
  inference(condensation,[],[f2019])).
fof(f2019,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | s(bool,'T_0') = s(bool,X61) | s(bool,X59) = s(bool,X62) | s(bool,X61) = s(bool,X62) | s(bool,X60) = s(bool,g(s(bool,X60))) | s(bool,'T_0') = s(bool,X59)) ) | $spl14),
  inference(superposition,[],[f844,f1562])).
fof(f2152,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,g(s(bool,X15)))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f2009])).
fof(f2009,plain,(
  ( ! [X14,X15] : (~p(s(bool,X15)) | s(bool,'T_0') = s(bool,X14) | s(bool,X15) = s(bool,g(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) ) | $spl14),
  inference(superposition,[],[f61,f1562])).
fof(f2371,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) ) | $spl40),
  inference(cnf_transformation,[],[f2371_D])).
fof(f2371_D,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2157,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,g(s(bool,X1)))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f2004])).
fof(f2004,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f43,f1562])).
fof(f2370,plain,(
  ( ! [X97,X95,X96] : (s(bool,sK2(X95,X96,X97,bool)) = s(bool,g(s(bool,sK2(X95,X96,X97,bool)))) | s(X97,f(s(bool,'F_0'))) != s(X97,g(s(bool,'F_0'))) | s(fun(bool,X97),X95) = s(fun(bool,X97),X96)) ) | $spl38),
  inference(cnf_transformation,[],[f2370_D])).
fof(f2370_D,plain,(
  ( ! [X97,X95,X96] : (s(bool,sK2(X95,X96,X97,bool)) = s(bool,g(s(bool,sK2(X95,X96,X97,bool)))) | s(X97,f(s(bool,'F_0'))) != s(X97,g(s(bool,'F_0'))) | s(fun(bool,X97),X95) = s(fun(bool,X97),X96)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2158,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,'T_0') = s(bool,X100) | s(fun(bool,X99),X97) = s(fun(bool,X99),X98) | s(X99,f(s(bool,'F_0'))) != s(X99,g(s(bool,'F_0'))) | s(bool,sK2(X97,X98,X99,bool)) = s(bool,g(s(bool,sK2(X97,X98,X99,bool))))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f1998])).
fof(f1998,plain,(
  ( ! [X99,X97,X100,X98] : (s(bool,'T_0') = s(bool,X100) | s(fun(bool,X99),X97) = s(fun(bool,X99),X98) | s(X99,f(s(bool,'F_0'))) != s(X99,g(s(bool,'F_0'))) | s(bool,sK2(X97,X98,X99,bool)) = s(bool,g(s(bool,sK2(X97,X98,X99,bool)))) | s(bool,'T_0') = s(bool,X100)) ) | $spl14),
  inference(superposition,[],[f257,f1562])).
fof(f2365,plain,(
  ( ! [X85,X86,X84] : (s(bool,g(s(X84,sK2(X85,X86,bool,X84)))) = s(bool,g(s(bool,g(s(X84,sK2(X85,X86,bool,X84)))))) | s(bool,'F_0') != s(bool,f(s(X84,sK2(X85,X86,bool,X84)))) | s(fun(X84,bool),X85) = s(fun(X84,bool),X86)) ) | $spl36),
  inference(cnf_transformation,[],[f2365_D])).
fof(f2365_D,plain,(
  ( ! [X85,X86,X84] : (s(bool,g(s(X84,sK2(X85,X86,bool,X84)))) = s(bool,g(s(bool,g(s(X84,sK2(X85,X86,bool,X84)))))) | s(bool,'F_0') != s(bool,f(s(X84,sK2(X85,X86,bool,X84)))) | s(fun(X84,bool),X85) = s(fun(X84,bool),X86)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2163,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,X92) | s(fun(X89,bool),X90) = s(fun(X89,bool),X91) | s(bool,'F_0') != s(bool,f(s(X89,sK2(X90,X91,bool,X89)))) | s(bool,g(s(X89,sK2(X90,X91,bool,X89)))) = s(bool,g(s(bool,g(s(X89,sK2(X90,X91,bool,X89))))))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f1994])).
fof(f1994,plain,(
  ( ! [X92,X90,X91,X89] : (s(bool,'T_0') = s(bool,X92) | s(fun(X89,bool),X90) = s(fun(X89,bool),X91) | s(bool,'F_0') != s(bool,f(s(X89,sK2(X90,X91,bool,X89)))) | s(bool,g(s(X89,sK2(X90,X91,bool,X89)))) = s(bool,g(s(bool,g(s(X89,sK2(X90,X91,bool,X89)))))) | s(bool,'T_0') = s(bool,X92)) ) | $spl14),
  inference(superposition,[],[f262,f1562])).
fof(f2218,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') = s(bool,X54) | s(bool,X55) = s(bool,g(s(bool,X55))) | s(bool,'F_0') = s(bool,X55)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f1944])).
fof(f1944,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') = s(bool,X54) | s(bool,X55) = s(bool,g(s(bool,X55))) | s(bool,'F_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X54)) ) | $spl14),
  inference(superposition,[],[f1562,f43])).
fof(f2364,plain,(
  ( ! [X65,X63] : (s(bool,X63) = s(bool,g(s(bool,X63))) | s(bool,X63) = s(bool,X65) | s(bool,X65) = s(bool,g(s(bool,X65)))) ) | $spl34),
  inference(cnf_transformation,[],[f2364_D])).
fof(f2364_D,plain,(
  ( ! [X65,X63] : (s(bool,X63) = s(bool,g(s(bool,X63))) | s(bool,X63) = s(bool,X65) | s(bool,X65) = s(bool,g(s(bool,X65)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2219,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') = s(bool,X47) | s(bool,X48) = s(bool,g(s(bool,X48))) | s(bool,X46) = s(bool,X48) | s(bool,X46) = s(bool,g(s(bool,X46)))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f1943])).
fof(f1943,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') = s(bool,X47) | s(bool,X48) = s(bool,g(s(bool,X48))) | s(bool,X46) = s(bool,X48) | s(bool,X46) = s(bool,g(s(bool,X46))) | s(bool,'T_0') = s(bool,X47)) ) | $spl14),
  inference(superposition,[],[f1562,f1562])).
fof(f2250,plain,(
  ( ! [X59,X57,X58] : (s(bool,X58) = s(bool,X59) | s(bool,X59) = s(bool,g(s(bool,X59))) | s(bool,'T_0') = s(bool,X57) | s(bool,X58) = s(bool,g(s(bool,X58)))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f1899])).
fof(f1899,plain,(
  ( ! [X59,X57,X58] : (s(bool,X58) = s(bool,X59) | s(bool,X59) = s(bool,g(s(bool,X59))) | s(bool,'T_0') = s(bool,X57) | s(bool,X58) = s(bool,g(s(bool,X58))) | s(bool,'T_0') = s(bool,X57)) ) | $spl14),
  inference(superposition,[],[f1562,f1562])).
fof(f2267,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f1885])).
fof(f1885,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f1562,f43])).
fof(f2268,plain,(
  ( ! [X97,X95,X98,X96] : (s(bool,'T_0') = s(bool,X98) | s(bool,sK2(X95,X96,X97,bool)) = s(bool,g(s(bool,sK2(X95,X96,X97,bool)))) | s(fun(bool,X97),X95) = s(fun(bool,X97),X96) | s(X97,f(s(bool,'F_0'))) != s(X97,g(s(bool,'F_0')))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f1884])).
fof(f1884,plain,(
  ( ! [X97,X95,X98,X96] : (s(bool,'T_0') = s(bool,X98) | s(bool,sK2(X95,X96,X97,bool)) = s(bool,g(s(bool,sK2(X95,X96,X97,bool)))) | s(bool,'T_0') = s(bool,X98) | s(fun(bool,X97),X95) = s(fun(bool,X97),X96) | s(X97,f(s(bool,'F_0'))) != s(X97,g(s(bool,'F_0')))) ) | $spl14),
  inference(superposition,[],[f1562,f257])).
fof(f2275,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') = s(bool,X87) | s(bool,g(s(X84,sK2(X85,X86,bool,X84)))) = s(bool,g(s(bool,g(s(X84,sK2(X85,X86,bool,X84)))))) | s(fun(X84,bool),X85) = s(fun(X84,bool),X86) | s(bool,'F_0') != s(bool,f(s(X84,sK2(X85,X86,bool,X84))))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f1879])).
fof(f1879,plain,(
  ( ! [X87,X85,X86,X84] : (s(bool,'T_0') = s(bool,X87) | s(bool,g(s(X84,sK2(X85,X86,bool,X84)))) = s(bool,g(s(bool,g(s(X84,sK2(X85,X86,bool,X84)))))) | s(bool,'T_0') = s(bool,X87) | s(fun(X84,bool),X85) = s(fun(X84,bool),X86) | s(bool,'F_0') != s(bool,f(s(X84,sK2(X85,X86,bool,X84))))) ) | $spl14),
  inference(superposition,[],[f1562,f262])).
fof(f2293,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') = s(bool,X55) | s(bool,X54) = s(bool,g(s(bool,X54))) | s(bool,X56) = s(bool,g(s(bool,X56))) | s(bool,X54) = s(bool,X56)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f1870])).
fof(f1870,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') = s(bool,X55) | s(bool,X54) = s(bool,g(s(bool,X54))) | s(bool,'T_0') = s(bool,X55) | s(bool,X56) = s(bool,g(s(bool,X56))) | s(bool,X54) = s(bool,X56)) ) | $spl14),
  inference(superposition,[],[f1562,f1562])).
fof(f2337,plain,(
  ( ! [X64,X65,X63] : (s(bool,X63) = s(bool,g(s(bool,X63))) | s(bool,'T_0') = s(bool,X64) | s(bool,X65) = s(bool,g(s(bool,X65))) | s(bool,X63) = s(bool,X65)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f1826])).
fof(f1826,plain,(
  ( ! [X64,X65,X63] : (s(bool,X63) = s(bool,g(s(bool,X63))) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,X64) | s(bool,X65) = s(bool,g(s(bool,X65))) | s(bool,X63) = s(bool,X65)) ) | $spl14),
  inference(superposition,[],[f1562,f1562])).
fof(f1551,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,sK2(X2,X3,bool,bool)))) | s(bool,'T_0') = s(bool,sK2(X2,X3,bool,bool)) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3)) ) | $spl16),
  inference(cnf_transformation,[],[f1551_D])).
fof(f1551_D,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,sK2(X2,X3,bool,bool)))) | s(bool,'T_0') = s(bool,sK2(X2,X3,bool,bool)) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1559,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') = s(bool,g(s(bool,X38))) | s(bool,'T_0') = s(bool,X40) | s(bool,X39) = s(bool,X40) | s(bool,X38) = s(bool,X39)) ) | $spl30),
  inference(cnf_transformation,[],[f1559_D])).
fof(f1559_D,plain,(
  ( ! [X39,X38,X40] : (s(bool,'T_0') = s(bool,g(s(bool,X38))) | s(bool,'T_0') = s(bool,X40) | s(bool,X39) = s(bool,X40) | s(bool,X38) = s(bool,X39)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1489,plain,(
  ( ! [X54,X52,X55,X53,X51] : (s(bool,'T_0') = s(bool,g(s(bool,X53))) | s(fun(bool,bool),X51) = s(fun(bool,bool),X52) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X51,X52,bool,bool)))) | s(bool,X53) = s(bool,X54) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,sK2(X51,X52,bool,bool))) )),
  inference(superposition,[],[f262,f97])).
fof(f1558,plain,(
  ( ! [X35,X31,X34,X32] : (s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,X35) | s(bool,X34) = s(bool,sK2(X31,X32,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X31,X32,bool,bool))))) ) | $spl28),
  inference(cnf_transformation,[],[f1558_D])).
fof(f1558_D,plain,(
  ( ! [X35,X31,X34,X32] : (s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(bool,'T_0') = s(bool,X35) | s(bool,X34) = s(bool,X35) | s(bool,X34) = s(bool,sK2(X31,X32,bool,bool)) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X31,X32,bool,bool))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1555,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,g(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) | $spl22),
  inference(cnf_transformation,[],[f1555_D])).
fof(f1555_D,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,g(s(bool,X9))) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1488,plain,(
  ( ! [X47,X50,X48,X46,X49] : (s(bool,'T_0') = s(bool,g(s(bool,X48))) | s(fun(bool,bool),X46) = s(fun(bool,bool),X47) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X46,X47,bool,bool)))) | s(bool,X49) = s(bool,sK2(X46,X47,bool,bool)) | s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f262,f97])).
fof(f1557,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,'T_0') = s(bool,X30) | s(bool,X28) = s(bool,X30) | s(bool,X29) = s(bool,sK2(X26,X27,bool,bool)) | s(bool,X28) = s(bool,X29) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X26,X27,bool,bool))))) ) | $spl26),
  inference(cnf_transformation,[],[f1557_D])).
fof(f1557_D,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,'T_0') = s(bool,X30) | s(bool,X28) = s(bool,X30) | s(bool,X29) = s(bool,sK2(X26,X27,bool,bool)) | s(bool,X28) = s(bool,X29) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X26,X27,bool,bool))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1550,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | $spl14),
  inference(cnf_transformation,[],[f1550_D])).
fof(f1550_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1487,plain,(
  ( ! [X45,X43,X41,X44,X42] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X41) = s(fun(bool,bool),X42) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X41,X42,bool,bool)))) | s(bool,X43) = s(bool,X44) | s(bool,X44) = s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,X43) = s(bool,sK2(X41,X42,bool,bool))) )),
  inference(superposition,[],[f262,f97])).
fof(f1486,plain,(
  ( ! [X39,X37,X38,X36,X40] : (s(bool,'T_0') = s(bool,g(s(bool,X38))) | s(fun(bool,bool),X36) = s(fun(bool,bool),X37) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X36,X37,bool,bool)))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,sK2(X36,X37,bool,bool)) | s(bool,X39) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f262,f97])).
fof(f1485,plain,(
  ( ! [X35,X33,X31,X34,X32] : (s(bool,'T_0') = s(bool,g(s(bool,X33))) | s(fun(bool,bool),X31) = s(fun(bool,bool),X32) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X31,X32,bool,bool)))) | s(bool,X34) = s(bool,sK2(X31,X32,bool,bool)) | s(bool,'T_0') = s(bool,X33) | s(bool,X34) = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) )),
  inference(superposition,[],[f262,f97])).
fof(f1484,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X26) = s(fun(bool,bool),X27) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X26,X27,bool,bool)))) | s(bool,X28) = s(bool,X29) | s(bool,X29) = s(bool,sK2(X26,X27,bool,bool)) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f262,f97])).
fof(f1554,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') != s(bool,f(s(bool,sK2(X4,X5,bool,bool)))) | s(bool,'T_0') = s(bool,X6) | s(bool,X6) = s(bool,sK2(X4,X5,bool,bool)) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5)) ) | $spl20),
  inference(cnf_transformation,[],[f1554_D])).
fof(f1554_D,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') != s(bool,f(s(bool,sK2(X4,X5,bool,bool)))) | s(bool,'T_0') = s(bool,X6) | s(bool,X6) = s(bool,sK2(X4,X5,bool,bool)) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1556,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') = s(bool,g(s(bool,X18))) | s(bool,'T_0') = s(bool,X19) | s(bool,X18) = s(bool,X19)) ) | $spl24),
  inference(cnf_transformation,[],[f1556_D])).
fof(f1556_D,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') = s(bool,g(s(bool,X18))) | s(bool,'T_0') = s(bool,X19) | s(bool,X18) = s(bool,X19)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1483,plain,(
  ( ! [X24,X23,X21,X25,X22] : (s(bool,'T_0') = s(bool,g(s(bool,X23))) | s(fun(bool,bool),X21) = s(fun(bool,bool),X22) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X21,X22,bool,bool)))) | s(bool,X24) = s(bool,sK2(X21,X22,bool,bool)) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f262,f97])).
fof(f1482,plain,(
  ( ! [X19,X17,X20,X18,X16] : (s(bool,'T_0') = s(bool,g(s(bool,X18))) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X16,X17,bool,bool)))) | s(bool,X18) = s(bool,X19) | s(bool,'T_0') = s(bool,X19) | s(bool,X20) = s(bool,sK2(X16,X17,bool,bool)) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f262,f97])).
fof(f1481,plain,(
  ( ! [X14,X15,X13] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X13) = s(fun(bool,bool),X14) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X13,X14,bool,bool)))) | s(bool,X15) = s(bool,sK2(X13,X14,bool,bool)) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f262,f59])).
fof(f1480,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,g(s(bool,X12))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X10,X11,bool,bool)))) | s(bool,'T_0') = s(bool,sK2(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f262,f59])).
fof(f1479,plain,(
  ( ! [X8,X7,X9] : (s(bool,'T_0') = s(bool,g(s(bool,X9))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X7,X8,bool,bool)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,sK2(X7,X8,bool,bool))) )),
  inference(superposition,[],[f262,f59])).
fof(f1478,plain,(
  ( ! [X6,X4,X5] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X4) = s(fun(bool,bool),X5) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X4,X5,bool,bool)))) | s(bool,X6) = s(bool,sK2(X4,X5,bool,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f262,f59])).
fof(f1477,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X2,X3,bool,bool)))) | s(bool,'T_0') = s(bool,sK2(X2,X3,bool,bool))) )),
  inference(superposition,[],[f262,f43])).
fof(f1547,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0'))) | $spl11),
  inference(cnf_transformation,[],[f1547_D])).
fof(f1547_D,plain,(
  s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0'))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f1548,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl12),
  inference(cnf_transformation,[],[f1548_D])).
fof(f1548_D,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1545,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool)))) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(duplicate_literal_removal,[],[f1476])).
fof(f1476,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'F_0') != s(bool,f(s(bool,sK2(X0,X1,bool,bool)))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,f(s(bool,'F_0'))) != s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f262,f257])).
fof(f67,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl9),
  inference(cnf_transformation,[],[f67_D])).
fof(f67_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f64,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f43])).
fof(f53,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK0) = s('type.pair.prod'('type.realax.real','type.realax.real'),sK1) | $spl0),
  inference(cnf_transformation,[],[f53_D])).
fof(f53_D,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK0) = s('type.pair.prod'('type.realax.real','type.realax.real'),sK1) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f55,plain,(
  s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) | $spl2),
  inference(cnf_transformation,[],[f55_D])).
fof(f55_D,plain,(
  s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f35,plain,(
  s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),sK0) = s('type.pair.prod'('type.realax.real','type.realax.real'),sK1)),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  (s('type.pair.prod'('type.realax.real','type.realax.real'),sK0) = s('type.pair.prod'('type.realax.real','type.realax.real'),sK1) | (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) & s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))))) & (s('type.pair.prod'('type.realax.real','type.realax.real'),sK0) != s('type.pair.prod'('type.realax.real','type.realax.real'),sK1) | s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) != s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) | s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) != s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f27])).
fof(f27,plain,(
  ? [X0,X1] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),X1) | (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) & s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))) & (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),X1) | s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ? [X0,X1] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),X1) | (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) & s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))) & (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),X1) | (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))))),
  inference(nnf_transformation,[],[f24])).
fof(f24,plain,(
  ? [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),X1) <~> (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) & s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))),
  inference(ennf_transformation,[],[f13])).
fof(f13,plain,(
  ~! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),X1) <=> (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) & s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),X7) <=> (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) & s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),X7) <=> (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) & s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7)))))),
  file('thm.complex.COMPLEX_RE_IM_EQ','thm.complex.COMPLEX_RE_IM_EQ')).
fof(f57,plain,(
  s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) | $spl4),
  inference(cnf_transformation,[],[f57_D])).
fof(f57_D,plain,(
  s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f36,plain,(
  s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),sK0) = s('type.pair.prod'('type.realax.real','type.realax.real'),sK1)),
  inference(cnf_transformation,[],[f28])).
fof(f54,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK0) != s('type.pair.prod'('type.realax.real','type.realax.real'),sK1) | $spl1),
  inference(cnf_transformation,[],[f54_D])).
fof(f54_D,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK0) != s('type.pair.prod'('type.realax.real','type.realax.real'),sK1) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f56,plain,(
  s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) != s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) | $spl3),
  inference(cnf_transformation,[],[f56_D])).
fof(f56_D,plain,(
  s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) != s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f58,plain,(
  s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) != s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) | $spl5),
  inference(cnf_transformation,[],[f58_D])).
fof(f58_D,plain,(
  s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) != s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f37,plain,(
  s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) != s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) | s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK0))) != s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),sK0) != s('type.pair.prod'('type.realax.real','type.realax.real'),sK1)),
  inference(cnf_transformation,[],[f28])).
% SZS output end Proof for thm.complex
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 143
% Passive clauses: 708
% Generated clauses: 16305
% Final active clauses: 44
% Final passive clauses: 112
% Input formulas: 11
% Initial clauses: 17
% 
% Duplicate literals: 2099
% Trivial inequalities: 154
% Fw subsumption resolutions: 561
% Fw demodulations: 709
% Bw demodulations: 341
% Condensations: 2880
% Global subsumptions: 28
% 
% Simple tautologies: 230
% Equational tautologies: 864
% Forward subsumptions: 7944
% Fw demodulations to eq. taut.: 1
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 1573
% Forward superposition: 3856
% Backward superposition: 2635
% Self superposition: 229
% Equality factoring: 200
% Equality resolution: 13
% 
% Split clauses: 185
% Split components: 420
% SAT solver clauses: 1154
% SAT solver unit clauses: 12
% SAT solver binary clauses: 96
% 
% Sat splits: 185
% Sat splitting refutations: 46
% 
% Memory used [KB]: 5245
% Time elapsed: 1.113 s
% ------------------------------
% ------------------------------
% Success in time 2.421 s
