Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.list
% (62643)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 104
% Passive clauses: 3748
% Generated clauses: 13218
% Final active clauses: 96
% Final passive clauses: 3467
% Input formulas: 37
% Initial clauses: 95
% 
% Duplicate literals: 582
% Trivial inequalities: 253
% Fw subsumption resolutions: 2074
% Fw demodulations: 468
% Bw demodulations: 4
% 
% Simple tautologies: 361
% Equational tautologies: 81
% Forward subsumptions: 4446
% Bw demodulations to eq. taut.: 4
% 
% Binary resolution: 560
% Factoring: 3
% Forward superposition: 3705
% Backward superposition: 4898
% Self superposition: 232
% Equality factoring: 197
% Equality resolution: 23
% 
% Split clauses: 476
% Split components: 965
% SAT solver clauses: 1183
% SAT solver unit clauses: 38
% SAT solver binary clauses: 392
% 
% Sat splits: 476
% 
% TWLsolver clauses: 1183
% TWLsolver calls for satisfiability: 205
% 
% Memory used [KB]: 7675
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (62661)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 149
% Passive clauses: 923
% Generated clauses: 17567
% Final active clauses: 134
% Final passive clauses: 684
% Input formulas: 37
% Initial clauses: 95
% 
% Duplicate literals: 446
% Trivial inequalities: 390
% Fw subsumption resolutions: 2073
% Fw demodulations: 3017
% Bw demodulations: 70
% Condensations: 751
% Global subsumptions: 13
% 
% Simple tautologies: 264
% Equational tautologies: 31
% Forward subsumptions: 8126
% Bw demodulations to eq. taut.: 35
% 
% Binary resolution: 425
% Factoring: 5
% Forward superposition: 4829
% Backward superposition: 4317
% Self superposition: 121
% Equality factoring: 67
% Equality resolution: 16
% 
% Split clauses: 67
% Split components: 135
% SAT solver clauses: 1649
% SAT solver unit clauses: 93
% SAT solver binary clauses: 67
% SAT solver learnt clauses: 26
% 
% Sat splits: 67
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1597
% TWLsolver calls for satisfiability: 420
% 
% Memory used [KB]: 14583
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% (62698)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 736
% Passive clauses: 6448
% Generated clauses: 253158
% Final active clauses: 52
% Final passive clauses: 873
% Input formulas: 37
% Initial clauses: 95
% 
% Duplicate literals: 7021
% Trivial inequalities: 2755
% Fw subsumption resolutions: 18772
% Fw demodulations: 52303
% Bw demodulations: 2838
% Condensations: 14271
% Global subsumptions: 1367
% 
% Simple tautologies: 2149
% Equational tautologies: 3842
% Forward subsumptions: 141017
% Fw demodulations to eq. taut.: 6
% Bw demodulations to eq. taut.: 130
% 
% Unit resulting resolution: 40412
% Factoring: 2
% Forward superposition: 46527
% Backward superposition: 43114
% Self superposition: 684
% Equality factoring: 347
% Equality resolution: 136
% 
% Split clauses: 651
% Split components: 1306
% SAT solver clauses: 11036
% SAT solver unit clauses: 47
% SAT solver binary clauses: 480
% 
% Sat splits: 651
% Sat splitting refutations: 1916
% 
% Memory used [KB]: 128953
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f1047,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f144,f240,f143,f242,f142,f244,f238,f430,f1046])).
fof(f1046,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f832,f588])).
fof(f588,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK8(sK2,sK4,sK5,sK3,sK1,sK0)))),s(sK1,sK4)))) | $spl0),
  inference(resolution,[],[f238,f147])).
fof(f147,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,sK8(X5,X4,X3,X2,X1,X0)))),s(X1,X4)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X3))))))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X3)))))) | (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK8(X5,X4,X3,X2,X1,X0)),s('type.list.list'(X0),sK9(X5,X4,X3,X2,X1,X0)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,sK8(X5,X4,X3,X2,X1,X0)))),s(X1,X4)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),sK9(X5,X4,X3,X2,X1,X0)),s('type.list.list'(X1),X3)))))) & (! [X6,X7] : (s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X6))),s(X1,X4)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X7),s('type.list.list'(X1),X3))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X3))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9])],[f104])).
fof(f104,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X3)))))) | ? [X8,X9] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X8),s('type.list.list'(X0),X9))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X8))),s(X1,X4)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X9),s('type.list.list'(X1),X3)))))) & (! [X6,X7] : (s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X6))),s(X1,X4)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X7),s('type.list.list'(X1),X3))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X3))))))))),
  inference(rectify,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X3)))))) | ? [X6,X7] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X6))),s(X1,X4)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X7),s('type.list.list'(X1),X3)))))) & (! [X6,X7] : (s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X6))),s(X1,X4)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X7),s('type.list.list'(X1),X3))))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X3))))))))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X3)))))) <=> ? [X6,X7] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X6))),s(X1,X4)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X7),s('type.list.list'(X1),X3))))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X3)))))) <=> ? [X6,X7] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X6))),s(X1,X4)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X7),s('type.list.list'(X1),X3))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X14,X15] : ! [X13,X16,X17,X18] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X14,fun(X15,bool)),X18),s('type.list.list'(X14),X13),s('type.list.list'(X15),'const.list.CONS_2'(s(X15,X17),s('type.list.list'(X15),X16)))))) <=> ? [X19,X20] : (s('type.list.list'(X14),X13) = s('type.list.list'(X14),'const.list.CONS_2'(s(X14,X19),s('type.list.list'(X14),X20))) & p(s(bool,happ(s(fun(X15,bool),happ(s(fun(X14,fun(X15,bool)),X18),s(X14,X19))),s(X15,X17)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X14,fun(X15,bool)),X18),s('type.list.list'(X14),X20),s('type.list.list'(X15),X16))))))),
  file('thm.list.EVERY2_THM|split|3','thm.list.LIST_REL_CONS2')).
fof(f832,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK8(sK2,sK4,sK5,sK3,sK1,sK0)))),s(sK1,sK4)))) | $spl0),
  inference(subsumption_resolution,[],[f831,f589])).
fof(f589,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK9(sK2,sK4,sK5,sK3,sK1,sK0)),s('type.list.list'(sK1),sK5)))) | $spl0),
  inference(resolution,[],[f238,f148])).
fof(f148,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),sK9(X5,X4,X3,X2,X1,X0)),s('type.list.list'(X1),X3)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X3))))))) )),
  inference(cnf_transformation,[],[f105])).
fof(f831,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK8(sK2,sK4,sK5,sK3,sK1,sK0)))),s(sK1,sK4)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK9(sK2,sK4,sK5,sK3,sK1,sK0)),s('type.list.list'(sK1),sK5)))) | $spl0),
  inference(trivial_inequality_removal,[],[f795])).
fof(f795,plain,(
  s('type.list.list'(sK0),sK3) != s('type.list.list'(sK0),sK3) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK8(sK2,sK4,sK5,sK3,sK1,sK0)))),s(sK1,sK4)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK9(sK2,sK4,sK5,sK3,sK1,sK0)),s('type.list.list'(sK1),sK5)))) | $spl0),
  inference(superposition,[],[f236,f587])).
fof(f587,plain,(
  s('type.list.list'(sK0),sK3) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK8(sK2,sK4,sK5,sK3,sK1,sK0)),s('type.list.list'(sK0),sK9(sK2,sK4,sK5,sK3,sK1,sK0)))) | $spl0),
  inference(resolution,[],[f238,f146])).
fof(f146,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK8(X5,X4,X3,X2,X1,X0)),s('type.list.list'(X0),sK9(X5,X4,X3,X2,X1,X0)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X3))))))) )),
  inference(cnf_transformation,[],[f105])).
fof(f236,plain,(
  ( ! [X6,X7] : (s('type.list.list'(sK0),sK3) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,X6),s('type.list.list'(sK0),X7))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,X6))),s(sK1,sK4)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),X7),s('type.list.list'(sK1),sK5))))) )),
  inference(subsumption_resolution,[],[f145,f149])).
fof(f149,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X2),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X3)))))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X5),s('type.list.list'(X0),X7),s('type.list.list'(X1),X3)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X5),s(X0,X6))),s(X1,X4)))) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7)))) )),
  inference(cnf_transformation,[],[f105])).
fof(f145,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),X7),s('type.list.list'(sK1),sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,X6))),s(sK1,sK4)))) | s('type.list.list'(sK0),sK3) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,X6),s('type.list.list'(sK0),X7))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK4),s('type.list.list'(sK1),sK5))))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  (p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK4),s('type.list.list'(sK1),sK5)))))) | (s('type.list.list'(sK0),sK3) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7))) & p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK6))),s(sK1,sK4)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK7),s('type.list.list'(sK1),sK5)))))) & (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK4),s('type.list.list'(sK1),sK5)))))) | ! [X6,X7] : (s('type.list.list'(sK0),sK3) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,X6),s('type.list.list'(sK0),X7))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,X6))),s(sK1,sK4)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),X7),s('type.list.list'(sK1),sK5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7])],[f101])).
fof(f101,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X5)))))) | ? [X8,X9] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X8),s('type.list.list'(X0),X9))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X8))),s(X1,X4)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X9),s('type.list.list'(X1),X5)))))) & (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X5)))))) | ! [X6,X7] : (s('type.list.list'(X0),X3) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X4)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X5)))))))),
  inference(rectify,[],[f100])).
fof(f100,plain,(
  ? [X0,X1,X2,X3,X4,X5] : ((p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X5)))))) | ? [X6,X7] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X4)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X5)))))) & (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X5)))))) | ! [X6,X7] : (s('type.list.list'(X0),X3) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X4)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X5)))))))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X5)))))) <~> ? [X6,X7] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X4)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X5))))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X5)))))) <=> ? [X6,X7] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X4)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X5))))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X4),s('type.list.list'(X1),X5)))))) <=> ? [X6,X7] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X6),s('type.list.list'(X0),X7))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X6))),s(X1,X4)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X5))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10,X11] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X8),s('type.list.list'(X6),X9),s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X10),s('type.list.list'(X7),X11)))))) <=> ? [X12,X13] : (s('type.list.list'(X6),X9) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X12),s('type.list.list'(X6),X13))) & p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X8),s(X6,X12))),s(X7,X10)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X8),s('type.list.list'(X6),X13),s('type.list.list'(X7),X11))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10,X11] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X8),s('type.list.list'(X6),X9),s('type.list.list'(X7),'const.list.CONS_2'(s(X7,X10),s('type.list.list'(X7),X11)))))) <=> ? [X12,X13] : (s('type.list.list'(X6),X9) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X12),s('type.list.list'(X6),X13))) & p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X6,fun(X7,bool)),X8),s(X6,X12))),s(X7,X10)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X6,fun(X7,bool)),X8),s('type.list.list'(X6),X13),s('type.list.list'(X7),X11))))))),
  file('thm.list.EVERY2_THM|split|3','thm.list.EVERY2_THM|split|3')).
fof(f430,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f429,f240])).
fof(f429,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK7),s('type.list.list'(sK1),sK5)))) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f427,f242])).
fof(f427,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK6))),s(sK1,sK4)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK7),s('type.list.list'(sK1),sK5)))) | $spl6),
  inference(trivial_inequality_removal,[],[f425])).
fof(f425,plain,(
  s('type.list.list'(sK0),sK3) != s('type.list.list'(sK0),sK3) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK6))),s(sK1,sK4)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK7),s('type.list.list'(sK1),sK5)))) | $spl6),
  inference(superposition,[],[f236,f244])).
fof(f238,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK4),s('type.list.list'(sK1),sK5)))))) | $spl0),
  inference(cnf_transformation,[],[f238_D])).
fof(f238_D,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK4),s('type.list.list'(sK1),sK5)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f244,plain,(
  s('type.list.list'(sK0),sK3) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7))) | $spl6),
  inference(cnf_transformation,[],[f244_D])).
fof(f244_D,plain,(
  s('type.list.list'(sK0),sK3) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f142,plain,(
  s('type.list.list'(sK0),sK3) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK6),s('type.list.list'(sK0),sK7))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK4),s('type.list.list'(sK1),sK5))))))),
  inference(cnf_transformation,[],[f102])).
fof(f242,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK6))),s(sK1,sK4)))) | $spl4),
  inference(cnf_transformation,[],[f242_D])).
fof(f242_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK6))),s(sK1,sK4)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f143,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK0,fun(sK1,bool)),sK2),s(sK0,sK6))),s(sK1,sK4)))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK4),s('type.list.list'(sK1),sK5))))))),
  inference(cnf_transformation,[],[f102])).
fof(f240,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK7),s('type.list.list'(sK1),sK5)))) | $spl2),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK7),s('type.list.list'(sK1),sK5)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f144,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK7),s('type.list.list'(sK1),sK5)))) | p(s(bool,'const.list.LIST_REL_3'(s(fun(sK0,fun(sK1,bool)),sK2),s('type.list.list'(sK0),sK3),s('type.list.list'(sK1),'const.list.CONS_2'(s(sK1,sK4),s('type.list.list'(sK1),sK5))))))),
  inference(cnf_transformation,[],[f102])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 63
% Passive clauses: 154
% Generated clauses: 812
% Final active clauses: 55
% Final passive clauses: 105
% Input formulas: 37
% Initial clauses: 92
% 
% Selected by SInE selection: 35
% SInE iterations: 1
% 
% Duplicate literals: 21
% Trivial inequalities: 27
% Fw subsumption resolutions: 78
% Fw demodulations: 155
% Bw demodulations: 5
% 
% Simple tautologies: 75
% Equational tautologies: 3
% Forward subsumptions: 360
% 
% Binary resolution: 91
% Forward superposition: 154
% Backward superposition: 271
% Equality resolution: 1
% 
% Split clauses: 3
% Split components: 6
% SAT solver clauses: 6
% SAT solver unit clauses: 1
% SAT solver binary clauses: 3
% 
% Sat splits: 3
% Sat splitting refutations: 2
% 
% Memory used [KB]: 1535
% Time elapsed: 0.018 s
% ------------------------------
% ------------------------------
% Success in time 11.79 s
