Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.numRing
% (82097)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 159
% Passive clauses: 3966
% Generated clauses: 20320
% Final active clauses: 151
% Final passive clauses: 3673
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 553
% Trivial inequalities: 349
% Fw subsumption resolutions: 959
% Fw demodulations: 1198
% Bw demodulations: 4
% 
% Simple tautologies: 328
% Equational tautologies: 22
% Forward subsumptions: 12569
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 31
% Forward superposition: 9525
% Backward superposition: 7216
% Self superposition: 211
% Equality factoring: 54
% Equality resolution: 29
% 
% Split clauses: 234
% Split components: 468
% SAT solver clauses: 463
% SAT solver unit clauses: 5
% SAT solver binary clauses: 100
% 
% Sat splits: 234
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 463
% TWLsolver calls for satisfiability: 58
% 
% Memory used [KB]: 20724
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.numRing
% (82127)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 133
% Passive clauses: 1381
% Generated clauses: 16202
% Final active clauses: 121
% Final passive clauses: 1104
% Input formulas: 118
% Initial clauses: 123
% 
% Duplicate literals: 589
% Trivial inequalities: 198
% Fw subsumption resolutions: 559
% Fw demodulations: 1024
% Bw demodulations: 121
% Condensations: 896
% 
% Simple tautologies: 175
% Equational tautologies: 25
% Forward subsumptions: 9202
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 26
% 
% Binary resolution: 43
% Forward superposition: 6259
% Backward superposition: 4681
% Self superposition: 215
% Equality factoring: 74
% Equality resolution: 19
% 
% Split clauses: 106
% Split components: 212
% SAT solver clauses: 2108
% SAT solver unit clauses: 235
% SAT solver binary clauses: 50
% 
% Sat splits: 106
% Sat splitting refutations: 30
% 
% TWLsolver clauses: 2108
% TWLsolver calls for satisfiability: 267
% 
% Memory used [KB]: 22899
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.numRing
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numRing
% SZS output start Proof for thm.numRing
fof(f118799,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f118717])).
fof(f118717,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.|48|_0')),
  inference(superposition,[],[f307,f118535])).
fof(f118535,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numRing.num_interp_cs_2'(s('type.quote.varmap'('type.num.num'),X2),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))) )),
  inference(forward_demodulation,[],[f118444,f371])).
fof(f371,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X1) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))) )),
  inference(cnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,X1) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(flattening,[],[f203])).
fof(f203,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,X1) = s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),'const.semi_ring.recordtype|46|semi_ring_4'(s(X0,X1),s(X0,X2),s(fun(X0,fun(X0,X0)),X3),s(fun(X0,fun(X0,X0)),X4)))))),
  inference(rectify,[],[f79])).
fof(f79,axiom,(
  ! [X7] : ! [X33,X42,X43,X44] : s(X7,X33) = s(X7,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X7),'const.semi_ring.recordtype|46|semi_ring_4'(s(X7,X33),s(X7,X42),s(fun(X7,fun(X7,X7)),X43),s(fun(X7,fun(X7,X7)),X44)))))),
  file('thm.numRing.num_ring_thms|split|47','thm.semi_ring.semi_ring_SR|48|')).
fof(f118444,plain,(
  ( ! [X2] : (s('type.num.num','const.numRing.num_interp_cs_2'(s('type.quote.varmap'('type.num.num'),X2),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0'))) = s('type.num.num','const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))) )),
  inference(superposition,[],[f23985,f339])).
fof(f339,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(flattening,[],[f141])).
fof(f141,plain,(
  ! [X0] : ! [X1,X2] : s(X0,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X0),X1))) = s(X0,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X7] : ! [X12,X6] : s(X7,'const.semi_ring.semi_ring_SR|48|_1'(s('type.semi_ring.semi_ring'(X7),X12))) = s(X7,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X7),X12),s('type.quote.varmap'(X7),X6),s('type.canonical.canonical_sum'(X7),'const.canonical.Nil_monom_0')))),
  file('thm.numRing.num_ring_thms|split|47','thm.canonical.interp_cs_def|split|0')).
fof(f23985,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numRing.num_interp_cs_2'(s('type.quote.varmap'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.num.num','const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.quote.varmap'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1)))) )),
  inference(forward_demodulation,[],[f23951,f416])).
fof(f416,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numRing.num_interp_cs_2'(s('type.quote.varmap'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.num.num',happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.num.num'),happ(s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.num.num')),'const.numRing.num_interp_cs_0'),s('type.quote.varmap'('type.num.num'),X0))),s('type.canonical.canonical_sum'('type.num.num'),X1)))) )),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1] : s('type.num.num','const.numRing.num_interp_cs_2'(s('type.quote.varmap'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.num.num',happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.num.num'),happ(s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.num.num')),'const.numRing.num_interp_cs_0'),s('type.quote.varmap'('type.num.num'),X0))),s('type.canonical.canonical_sum'('type.num.num'),X1)))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X45,X46] : s('type.num.num','const.numRing.num_interp_cs_2'(s('type.quote.varmap'('type.num.num'),X45),s('type.canonical.canonical_sum'('type.num.num'),X46))) = s('type.num.num',happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.num.num'),happ(s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.num.num')),'const.numRing.num_interp_cs_0'),s('type.quote.varmap'('type.num.num'),X45))),s('type.canonical.canonical_sum'('type.num.num'),X46)))),
  file('thm.numRing.num_ring_thms|split|47','ARITY_0_2_const.numRing.num_interp_cs')).
fof(f23951,plain,(
  ( ! [X0,X1] : (s('type.num.num',happ(s(fun('type.canonical.canonical_sum'('type.num.num'),'type.num.num'),happ(s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.num.num')),'const.numRing.num_interp_cs_0'),s('type.quote.varmap'('type.num.num'),X0))),s('type.canonical.canonical_sum'('type.num.num'),X1))) = s('type.num.num','const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'))),s('type.quote.varmap'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),X1)))) )),
  inference(superposition,[],[f362,f313])).
fof(f313,plain,(
  s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.num.num')),'const.numRing.num_interp_cs_0') = s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.num.num')),'const.canonical.interp_cs_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  inference(cnf_transformation,[],[f62])).
fof(f62,axiom,(
  s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.num.num')),'const.numRing.num_interp_cs_0') = s(fun('type.quote.varmap'('type.num.num'),fun('type.canonical.canonical_sum'('type.num.num'),'type.num.num')),'const.canonical.interp_cs_1'(s('type.semi_ring.semi_ring'('type.num.num'),'const.semi_ring.recordtype|46|semi_ring_4'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.+_0'),s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')))))),
  file('thm.numRing.num_ring_thms|split|47','thm.numRing.num_interp_cs_def')).
fof(f362,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.canonical.interp_cs_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.canonical.interp_cs_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0] : ! [X1,X2,X3] : s(X0,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X0),X1),s('type.quote.varmap'(X0),X2),s('type.canonical.canonical_sum'(X0),X3))) = s(X0,happ(s(fun('type.canonical.canonical_sum'(X0),X0),happ(s(fun('type.quote.varmap'(X0),fun('type.canonical.canonical_sum'(X0),X0)),'const.canonical.interp_cs_1'(s('type.semi_ring.semi_ring'(X0),X1))),s('type.quote.varmap'(X0),X2))),s('type.canonical.canonical_sum'(X0),X3)))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X7] : ! [X45,X46,X47] : s(X7,'const.canonical.interp_cs_3'(s('type.semi_ring.semi_ring'(X7),X45),s('type.quote.varmap'(X7),X46),s('type.canonical.canonical_sum'(X7),X47))) = s(X7,happ(s(fun('type.canonical.canonical_sum'(X7),X7),happ(s(fun('type.quote.varmap'(X7),fun('type.canonical.canonical_sum'(X7),X7)),'const.canonical.interp_cs_1'(s('type.semi_ring.semi_ring'(X7),X45))),s('type.quote.varmap'(X7),X46))),s('type.canonical.canonical_sum'(X7),X47)))),
  file('thm.numRing.num_ring_thms|split|47','ARITY_1_3_const.canonical.interp_cs')).
fof(f307,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numRing.num_interp_cs_2'(s('type.quote.varmap'('type.num.num'),sK0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numRing.num_interp_cs_2'(s('type.quote.varmap'('type.num.num'),sK0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f298])).
fof(f298,plain,(
  ? [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.numRing.num_interp_cs_2'(s('type.quote.varmap'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(ennf_transformation,[],[f120])).
fof(f120,plain,(
  ~! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numRing.num_interp_cs_2'(s('type.quote.varmap'('type.num.num'),X0),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numRing.num_interp_cs_2'(s('type.quote.varmap'('type.num.num'),X6),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.numRing.num_interp_cs_2'(s('type.quote.varmap'('type.num.num'),X6),s('type.canonical.canonical_sum'('type.num.num'),'const.canonical.Nil_monom_0')))),
  file('thm.numRing.num_ring_thms|split|47','thm.numRing.num_ring_thms|split|47')).
% SZS output end Proof for thm.numRing
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 532
% Passive clauses: 9463
% Generated clauses: 121830
% Final active clauses: 215
% Final passive clauses: 1877
% Input formulas: 118
% Initial clauses: 123
% Discarded non-redundant clauses: 3264
% 
% Duplicate literals: 3443
% Trivial inequalities: 1175
% Fw subsumption resolutions: 3130
% Fw demodulations: 13768
% Bw demodulations: 1815
% Condensations: 4727
% Global subsumptions: 207
% 
% Simple tautologies: 597
% Equational tautologies: 261
% Forward subsumptions: 84011
% Fw demodulations to eq. taut.: 23
% Bw demodulations to eq. taut.: 116
% 
% Unit resulting resolution: 2720
% Forward superposition: 45240
% Backward superposition: 40148
% Self superposition: 712
% Equality factoring: 214
% Equality resolution: 48
% 
% Split clauses: 338
% Split components: 681
% SAT solver clauses: 10582
% SAT solver unit clauses: 1186
% SAT solver binary clauses: 2724
% 
% Sat splits: 338
% Sat splitting refutations: 234
% 
% Memory used [KB]: 156074
% Time elapsed: 6.966 s
% ------------------------------
% ------------------------------
% Success in time 8.299 s
